---
id: cf1111c1c11feddfaeb3bdef
title: 加法运算
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number`是 JavaScript 中的一种数据类型，表示数值。

现在让我们来尝试在 JavaScript 中做加法运算。

JavaScript 中使用`+`号进行加法运算。

**示例**

```js
myVar = 5 + 10; // assigned 15
```

# --instructions--

改变数字`0`让变量 sum 的值为`20`。

# --hints--

`sum`应该等于`20`。

```js
assert(sum === 20);
```

要使用`+`运算符。

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
var sum = 10 + 0;
```

# --solutions--

```js
var sum = 10 + 10;
```
