---
id: 56533eb9ac21ba0edf2244a9
title: 使用赋值运算符初始化变量
challengeType: 1
videoUrl: 'https://scrimba.com/c/cWJ4Bfb'
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

通常在声明变量的时候会给变量<dfn>初始化</dfn>一个初始值。

```js
var myVar = 0;
```

创建一个名为 `myVar` 的变量，并指定其初始值为 `0`。

# --instructions--

通过关键字 `var` 定义一个变量 `a`，并且给它一个初始值 `9`。

# --hints--

应该初始化 `a` 的值为 `9`。

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(code));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
