---
id: cf1231c1c11feddfaeb5bdef
title: 乘法运算
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
---

# --description--

我们也可在 JavaScript 中使用乘法运算。

JavaScript 使用`*`符号表示两数相乘。

**示例**

```js
myVar = 13 * 13; // assigned 169
```

# --instructions--

改变数值`0`来让变量 product 的值等于`80`。

# --hints--

要使`product`的值等于 80。

```js
assert(product === 80);
```

使用`*`运算符。

```js
assert(/\*/.test(code));
```

# --solutions--

