---
id: cf1111c1c11feddfaeb7bdef
title: 将一个数组嵌套在另一个数组中
challengeType: 1
videoUrl: 'https://scrimba.com/c/crZQZf8'
forumTopicId: 18247
---

# --description--

你也可以在数组中包含其他数组，例如：`[["Bulls", 23], ["White Sox", 45]]`。这被称为一个<dfn>多维数组<dfn>。</dfn></dfn>

# --instructions--

创建一个名为`myArray`的多维数组。

# --hints--

应该包含至少一个嵌入的数组。

```js
assert(Array.isArray(myArray) && myArray.some(Array.isArray));
```

# --solutions--

