---
id: 56533eb9ac21ba0edf2244bc
title: 购物清单
challengeType: 1
videoUrl: 'https://scrimba.com/c/c9MEKHZ'
forumTopicId: 18280
---

# --description--

创建一个名叫`myList`的购物清单，清单的数据格式就是多维数组。

每个子数组中的第一个元素应该是购买的物品名称，第二个元素应该是物品的数量，类似于：

`["Chocolate Bar", 15]`

任务：你的购物清单至少应该有 5 个子数组。

# --hints--

`myList`应该一个数组。

```js
assert(isArray);
```

你的每个子数组的第一个元素的类型都应该是字符串。

```js
assert(hasString);
```

你的每个子数组的第二个元素的类型都应该是数字。

```js
assert(hasNumber);
```

你的列表中至少要包含 5 个元素。

```js
assert(count > 4);
```

# --solutions--

