---
id: bd7123c9c441eddfaeb5bdef
title: 理解布尔值
challengeType: 1
videoUrl: 'https://scrimba.com/c/c9Me8t4'
forumTopicId: 301176
---

# --description--

另一种数据类型是<dfn>布尔</dfn>（Boolean）。`布尔`值要么是`true`要么是`false`。它非常像电路开关，`true`是 “开”，`false`是 “关”。这两种状态是互斥的。

**注意**  
`布尔值`是不带引号的，`"true"`和`"false"`是`字符串`而不是`布尔值`，在 JavaScript 中也没有特殊含义。

# --instructions--

修改`welcomeToBooleans`函数，让它返回`true`而不是`false`。

# --hints--

`welcomeToBooleans()`函数应该返回一个布尔值 (true/false)。

```js
assert(typeof welcomeToBooleans() === 'boolean');
```

`welcomeToBooleans()`应该返回 true。

```js
assert(welcomeToBooleans() === true);
```

# --solutions--

