---
id: bd7123c9c452eddfaeb5bdef
title: 使用方括号查找字符串中的第N个字符到最后一个字符
challengeType: 1
videoUrl: 'https://scrimba.com/c/cw4vkh9'
forumTopicId: 18344
---

# --description--

我们既可以获取字符串的最后一个字符，也可以用获取字符串的倒数第N个字符。

例如，你可以这样`firstName[firstName.length - 3]`操作来获得`var firstName = "Charles"`字符串中的倒数第三个字符。

# --instructions--

使用<dfn>方括号</dfn>来获得`lastName`字符串中的倒数第二个字符。

**提示**  
如果你遇到困难了，不妨看看`thirdToLastLetterOfFirstName`变量是如何做到的。

# --hints--

`secondToLastLetterOfLastName`应该是"c"。

```js
assert(secondToLastLetterOfLastName === 'c');
```

你需要使用`.length`获取倒数第二个字符。

```js
assert(code.match(/\.length/g).length === 2);
```

# --solutions--

