---
id: 61fd70336ebb3e4f62ee81ba
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Таблиці HTML використовують елемент `caption`, щоб описати зміст таблиці. Елемент `caption` завжди повинен бути першим дочірнім елементом `table`, але його можна розмістити з допомогою CSS-властивості `caption-side`.

Додайте елемент `caption` до свого першого `table` та надайте йому текст `Assets`.

# --hints--

Ваш перший елемент `table` повинен мати елемент `caption`.

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

Ваш елемент `caption` повинен мати текст `Assets`.

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
