---
id: 61fd5a93fd62bb35968adeab
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Налаштуйте свій HTML з допомогою елементів `DOCTYPE`, `html` вказуючи, що цей документ англійською мовою, `head` та `body`.

Надайте елементу `head` відповідні елементи `meta` для `charset` та `viewport`, елемент `title` з відповідною назвою та елемент `link` для своєї таблиці стилів.

# --hints--

Ваш код повинен мати декларацію `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ваш код повинен мати елемент `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Ваш елемент `html` повинен мати початковий теґ з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Елемент `html` повинен мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Ваш код повинен мати елемент `head` в межах елемента `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Ваш код повинен мати елемент `body` в межах елемента `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Ваш елемент `head` повинен йти перед елементом `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Ви повинні мати два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Один елемент `meta` повинен мати `name` зі значенням `viewport` та `content` зі значенням `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Інший елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш `title` повинен мати якийсь текст.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Ваш код повинен мати елемент `link`.

```js
assert(/<link/.test(code))
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
