---
id: 61fd6ab779390f49148773bb
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Під елементом `h1` створіть елемент `div`. Надайте йому атрибут `id` зі значенням `years`. Потрібно приховати цей конкретний елемент від читачів екрана, тому надайте йому атрибут `aria-hidden` зі значенням `true`.

# --hints--

Ви повинні створити новий елемент `div` після елемента `h1`.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

Ваш новий елемент `div` повинен мати атрибут `id` зі значенням `years`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

Ваш новий елемент `div` повинен мати атрибут `aria-hidden` зі значенням `true`.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
