---
id: 61fd75ea7f663457612dba02
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

В кожному з ваших нових `th` елементів вкладіть елемент `span`, що має `class` зі значенням `sr-only year`. Надайте їм такий текст (за порядком): `2019`, `2020` та `2021`.

Надайте третьому елементу `th` атрибут `class` зі значенням `current`.

Залиште елемент `td` порожнім. Цей елемент існує лише для того, щоб ваша таблиця мала макет із чотирма стовпцями та пов’язувала заголовки з правильними стовпцями.

# --hints--

Кожен з ваших елементів `th` повинен мати елемент `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

Кожен з ваших нових елементів `span` повинен мати атрибут `class` зі значенням `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

Ваш перший елемент `span` повинен мати текст `2019`.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

Ваш другий елемент `span` повинен мати текст `2020`.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

Ваш третій елемент `span` повинен мати текст `2021`.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

Ваш третій елемент `th` повинен мати `class` зі значенням `current`.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

Ваш елемент `td` повинен бути порожнім.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
