---
id: 61fd6b7c83dbf54a08cf0498
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

В елемент `div` додайте три елементи `span`. Надайте кожному з них атрибут `class` зі значенням `year` та додайте такий текст (за порядком): `2019`, `2020` та `2021`.

# --hints--

Ваш елемент `div` повинен мати три елементи `span`.

```js
assert(document.querySelector('div')?.children?.length === 3);
```

Кожен елемент `span` повинен мати атрибут `class` зі значенням `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

Ваш перший `span` повинен мати текст `2019`.

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

Ваш другий `span` повинен мати текст `2020`.

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

Ваш третій `span` повинен мати текст `2021`.

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
