---
id: 5d822fd413a79914d39e98d3
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

أدخل أربع عناصر `div` في حاوية `.bb1`. أعطيهم الفئات `bb1a`, و`bb1b`, و`bb1c`, و`bb1d` بهذا الترتيب. سيتكون هذا المبنى من أربع أقسام.

# --hints--

يجب عليك إضافة أربع عناصر `div` جديدة.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

يجب عليك إعطاء واحدة من عناصر `div` الجديدة فئة `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

يجب عليك إعطاء واحدة من عناصر `div` الجديدة فئة `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

يجب عليك إعطاء واحدة من عناصر `div` الجديدة فئة `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

يجب عليك إعطاء واحدة من عناصر `div` الجديدة فئة `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

يجب عليك وضع عناصر `div` الجديدة في الترتيب الصحيح.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

يجب عليك وضع عناصر `div` الجديدة ضمن عنصر `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

