---
id: 5d822fd413a79914d39e98cf
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

يمكنك رؤية `body` (هو الصندوق الأكثر داخلاً في صفحتك)؛ والصندوق المحيط به هو عنصر `html`. اجعل `body` يملأ كامل نافذة العرض بواسطة إعطائه `height` بقيمة `100vh`. إزالة `margin` الافتراضي من `body` عن طريق تعيين `margin` إلى `0`. وأخيرا، عيّن خاصية `overflow` إلى `hidden` لإخفاء أي شريط تمرير يظهر عندما يمتد شيء ما إلى ما وراء نافذة العرض.

# --hints--

يجب عليك استخدام منتقي `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` الخاص بك يجب أن يكون له `height` بقيمة `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` الخاص بك يجب أن يكون له `margin` بقيمة `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

يجب أن يكون عنصر `body` الخاص بك يحمل خاصية `overflow` بقيمة `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

