---
id: 5d822fd413a79914d39e98c9
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

مرحبا بكم في متغيرات CSS بمشروع Skyline! ابدأ بإضافة تعريف `!DOCTYPE html` في الجزء العلوي من المستند حتى يعرف المتصفح نوع المستند الذي يقرأه.

# --hints--

يجب أن يحتوي كودك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` مع `>` بعد تحديد النوع.

```js
assert(code.match(/html\s*>/gi));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
