---
id: 5d822fd413a79914d39e98ca
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

اضف علامات فتح وإغلاق `html` أسفل `DOCTYPE` حتى يكون لديك مكان لبدء وضع بعض الكود. تحقق من تعيين اللغة الإنجليزية.

# --hints--

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة افتتاحية مع خاصية `lang` من `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة الأغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن تكون العلامات الخاصة بعنصر `html` في الترتيب الصحيح.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

يجب أن يكون لديك عنصر `html` واحد فقط.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
