---
id: 5d822fd413a79914d39e98d4
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

أمنح أجزاء من المبنى الخاص بك خاصتين `width` و `height` مع هذا القيم: `70%` و `10%` إلى فئة `.bb1a` وهذا القيم: `80%` و`10%` إلى فئة `.bb1b`, و هذا القيم: `90%`>و `10%` إلى فئة `.bb1c`, وهذا القيم: `100%` و `70%` إلى فئة `.bb1d`. تذكر أن النسب المئوية تتنسب إلى الوالد و لاحظ أن الارتفاع سوف يصل إلى 100٪ - ملئ الحاوية عمودياً.

# --hints--

يجب عليك استخدام منتقي الفئة لتصميم `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

يجب عليك إعطاء `.bb1a` الـ `width` بقيمة `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

يجب عليك إعطاء `.bb1a` الـ `height` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

يجب عليك استخدام منتقي الفئة لتصميم `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

يجب عليك إعطاء `.bb1b` الـ `width` بقيمة `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

يجب عليك إعطاء `.bb1b` الـ `height` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

يجب عليك استخدام منتقي الفئة لتصميم `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

يجب عليك إعطاء `.bb1c` الـ `width` بقيمة `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

يجب عليك إعطاء `.bb1c` الـ `height` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

يجب عليك استخدام منتقي الفئة لتصميم `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

يجب عليك إعطاء `.bb1d`. الـ `width` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

يجب عليك إعطاء `.bb1d` الـ `height` بقيمة `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

