---
id: 5d822fd413a79914d39e98ce
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

إضافة `box-sizing` بقيمة `border-box` إلى كل شيء. هذا سيجعل الحدود لا يتم إضافتها ألى أحجام عناصرك.

# --hints--

يجب عليك استخدام خاصية `box-sizing`.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

يجب عليك استخدام منتقي `*` الموجود.

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

يجب أن تحتوي جميع العناصر على `box-sizing` بقيمة `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

