---
id: 5d822fd413a79914d39e98d8
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

استخدم نفس المٌتغير مثل `background-color` في فئات `.bb1b`, و`.bb1c`، و `.bb1d` ملأ بقية المبنى.

# --hints--

يجب تعيين `background-color` لعنصر `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b')?.backgroundColor)
```

يجب عليك استخدام `var(--building-color1)` لتعيين `background-color` عنصر `.bb1b`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.getPropVal('background-color', true), 'var(--building-color1)');
```

يجب تعيين `background-color` لعنصر `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c')?.backgroundColor)
```

يجب عليك استخدام `var(--building-color1)` لتعيين `background-color` عنصر `.bb1c`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.getPropVal('background-color', true), 'var(--building-color1)');
```

يجب تعيين `background-color` لعنصر `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d')?.backgroundColor)
```

يجب عليك استخدام `var(--building-color1)` لتعيين `background-color` عنصر `.bb1d`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
--fcc-editable-region--

```
