---
id: 5d822fd413a79914d39e98cb
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بعد ذلك، ضف وسمات فتح وإغلاق `head` و `body` داخل عنصر `html`.

# --hints--

يجب أن يكون لديك وسم فتح للعنصر `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم أغلاق للعنصر `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم فتح للعنصر `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم أغلاق للعنصر `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عناصر `head` و `body` متلازمين (siblings).

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">

</html>
--fcc-editable-region--

```

