---
id: 5d822fd413a79914d39e98d2
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

ادخل `div` مع فئة `bb1` في حاوية المباني الخلفية. افتح مِلَفّك `styles.css`، وأعطي `.bb1` خاصية `width` بقيمة `10%`, وخاصية `height` بقيمة `70%`. "bb" هو اختصار إلى "background building"، سيكون هذا أول مبنيك.

# --hints--

يجب عليك إضافة عنصر `div` جديد.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

يجب أن تعطي `div` جديد فئة تسمى `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

يجب عليك استخدام منتقي الفئة `.bb1` لتصميم العنصر.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

يجب عليك إعطاء عنصر `.bb1` خاصية `width` بقيمة `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

يجب عليك إعطاء عنصر `.bb1` خاصية `height` بقيمة `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

