---
id: 5d822fd413a79914d39e98e1
title: الخطوة 25
challengeType: 0
dashedName: step-25
---

# --description--

لم يفعل ذلك، المتغيرات التي أعلنتها في `.bb1` لا تتسلسل إلى عنصران متوازيان `.bb2` و `.bb3`. هذه هي طريقة عمل CSS. بسبب هذا، غالباً ما يتم الإعلان عن المتغيرات في منتقي `:root`. هذا هو أعلى مستوى منتقى في CSS؛ وضع المتغيرات هناك سيجعلها قابلة للاستخدام في كل مكان. أضف منتقي `:root` إلى الجزء العلوي من stylesheet، ونقل جميع إعلانات متغيرك إلى هناك.

# --hints--

يجب أن تعلن منتقي `:root` في الجزء العلوي من stylesheet.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(':root'));
```

يجب عليك تعريف `--building-color1` بقيمة `#aa80ff` في منتقي `:root`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color1')?.trim(), '#aa80ff');
```

يجب عليك تعريف `--building-color2` بقيمة `#66cc99` في منتقي `:root`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color2')?.trim(), '#66cc99');
```

يجب عليك تعريف `--building-color3` بقيمة `#cc6699` في منتقي `:root`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color3')?.trim(), '#cc6699');
```

يجب عليك إزالة متغيرات الخاصية المخصصة من `.bb1`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color2'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2, green);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3, pink);
}

.bb4 {
  width: 11%;
  height: 58%;
}

```

