---
id: 5d822fd413a79914d39e98d0
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

أنشئ عنصر `div` في `body` مع فئة `background-buildings`. وستكون هذه الحاوية لمجموعة من المباني.

# --hints--

يجب عليك إنشاء عنصر `div`.

```js
assert.exists(document.querySelector('div'));
```

يجب أن يكون عنصرك `div` داخل `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

يجب أن تحتوي عنصر `div` على فئة تسمي `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

