---
id: 5d822fd413a79914d39e98cc
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

داخل `head`، أدخل عنصر `meta` مع خاصية `charset` بقيمة `UTF-8`، وعنصر `title` بعنوان `City Skyline`، و `link` الذي يربط مِلَفّ `styles.css`.

# --hints--

يجب عليك إنشاء عنصر `meta` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

يجب عليك إعطاء عنصر `meta` خاصية `charset` بقيمة `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

يجب أن يحتوي كودك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون عنصر `title` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > title'));
```

يجب أن يكون لمشروعك العنوان `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

تذكر أن الحالة الحرف سواء كبيرا أو صغيرا, والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

يجب أن يحتوي كودك على عنصر `link`.

```js
assert.match(code, /<link/)
```

لا يجب عليك تغيير عنصر `head` الحالي. تحقق أنك لم تقم بحذف وسم الإغلاق.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

يجب أن يحتوي كودك على عنصر `link` مغلق ذاتيا.

```js
assert(document.querySelectorAll('link').length === 1);
```

يجب أن يكون عنصرك `link` داخل عنصرك `head`.

```js
assert.exists(document.querySelector('head > link'));
```

يجب أن يحتوي عنصرك `link` على سمة `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

يجب أن يحتوي عنصرك `link` على خاصية `href` مع القيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
