---
id: 5d822fd413a79914d39e98d1
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

اعطي عنصر `.background-buildings` خصائص `width` و `height` من `100%` لجعله العرض والارتفاع مطبقان لحاوية `body`.

# --hints--

يجب عليك استخدام فئة `background-buildings` لتحديد العنصر الصحيح.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

يجب أن يكون لعنصر `.background-buildings` الخاص بك `width` من `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

يجب أن يكون لعنصر `.background-buildings` الخاص بك `height` من `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

