---
id: 69446f18fdcb11aa516f7c5c
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will be creating a `countup` function, that creates an array with numbers from 1 to a given number using recursion.

Start by creating an empty function named `countup` that takes a single parameter.

# --hints--

You should have a function named `countup`.

```js
assert.isFunction(countup);
```

Your `countup` function should accept a single parameter.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
