---
id: 6948705160b66cf93a916e29
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Declare an empty array named `countArray` inside the `countup` function just above your `if` statement.

This array will store the value of `number` at each function call.

# --hints--

You should declare an empty array named `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);

```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
