---
id: 642ddfdea4200e313f80a4b6
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Ricorda che l'oggetto `document` ha un metodo `.createElement()` che consente di creare dinamicamente nuovi elementi HTML.

Nella funzione `createLabel`, dichiara una variabile `label` e assegnale un nuovo elemento `div`.

# --hints--

Dovresti dichiarare una variabile `label` nella funzione `createLabel`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

La variabile `label` dovrebbe essere dichiarata con `const`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

Dovresti usare il metodo `.createElement()` dell'oggetto `document`.

```js
assert.match(code, /document\.createElement\(/);
```

Dovresti passare la stringa `div` al metodo `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

Dovresti assegnare il nuovo elemento `div` a `label`.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
