---
id: 642df9df4b5216350de7b0d2
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Avrai bisogno di una funzione per generare un intervallo di numeri.

Dichiara una funzione `range` vuota che richiede dei parametri `start` e `end`. Usa il costruttore `Array()` e restituisci implicitamente un array vuoto.

# --hints--

Dovresti dichiarare una variabile `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Dovresti usare `const` per dichiarare la variabile `range`.

```js
assert.match(code, /const\s+range/);
```

La variabile `range` dovrebbe essere una funzione.

```js
assert.isFunction(range);
```

La funzione `range` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

La funzione `range` dovrebbe prendere `start` come primo parametro.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

La funzione `range` dovrebbe prendere `end` come secondo parametro.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

La funzione `range` dovrebbe utilizzare un return implicito. Ricorda che ciò significa che non devi usare parentesi graffe.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

La funzione `range` dovrebbe utilizzare il costruttore `Array()`. I costruttori primitivi non hanno bisogno della parola chiave `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

Non dovresti passare nulla al costruttore `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
