---
id: 642e02be7845f13b014cd2b0
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Ora che hai una funzione `range`, puoi usarla anche per creare un intervallo di lettere.

Dichiara una funzione `charRange` utilizzando `const` e la sintassi freccia. Dovrebbe prendere dei parametri `start` e `end`. La funzione dovrebbe restituire implicitamente il risultato della chiamata di `range()` con `start` e `end` come argomenti.

# --hints--

Dovresti dichiarare una variabile `charRange`.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

La variabile `charRange` dovrebbe essere dichiarata con `const`.

```js
assert.match(code, /const\s+charRange/);
```

La variabile `charRange` dovrebbe essere una funzione.

```js
assert.isFunction(charRange);
```

La funzione `charRange` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

La funzione `charRange` dovrebbe prendere `start` come primo parametro.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

La funzione `charRange` dovrebbe prendere `end` come secondo parametro.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

La funzione `charRange` dovrebbe utilizzare un return implicito.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

La funzione `charRange` dovrebbe chiamare la funzione `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

Dovresti passare `start` e `end` come argomenti alla chiamata di `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
