---
id: 643498755d54c6279ba09078
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

La maggior parte dei programmi di fogli di calcolo includono funzioni integrate per il calcolo.

Dichiara una funzione `sum` che prende un parametro `nums`, che sarà un array di numeri. Dovrebbe restituire il risultato della chiamata di `reduce` sull'array per sommare tutti i numeri.

# --hints--

Dovresti dichiarare una variabile `sum`.

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

Dovresti usare `const` per dichiarare la variabile `sum`.

```js
assert.match(code, /const\s+sum/);
```

La variabile `sum` dovrebbe essere una funzione.

```js
assert.isFunction(sum);
```

La funzione `sum` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

La funzione `sum` dovrebbe avere un parametro `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

La funzione `sum` dovrebbe utilizzare un return implicito.

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

La funzione `sum` dovrebbe restituire il risultato della chiamata `.reduce()` su `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

La funzione `sum` dovrebbe restituire la somma di tutti i numeri in `nums`.

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
