---
id: 643475e13dc727231acd0f72
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Ora nella chiamata `.forEach()` annidata, dichiara una variabile `input`. Usa il metodo `.createElement()` dell'oggetto `document` per creare un elemento `input`. Imposta l'attributo `type` su `text` e l'attributo `id` su `letter + number`.

# --hints--

Dovresti dichiarare una variabile `input`.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

Dovresti usare `const` per dichiarare la variabile `input`.

```js
assert.match(code, /const\s+input/)
```

Dovresti chiamare il metodo `.createElement()` dell'oggetto `document`.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

Dovresti passare la stringa `input` al metodo `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Dovresti assegnare il nuovo elemento `input` alla variabile `input`.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Dovresti accedere alla proprietà `type` dell'elemento `input`.

```js
assert.match(code, /input\.type/);
```

Dovresti impostare l'attributo `type` dell'elemento `input` su `text`.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

Dovresti accedere alla proprietà `id` dell'elemento `input`.

```js
assert.match(code, /input\.id/);
```

Dovresti impostare l'attributo `id` dell'elemento `input` su `letter + number`.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
