---
id: 6434552bcc0a951a0a99df3b
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

La funzione `range` si aspetta dei numeri, ma i valori `start` e `end` saranno stringhe (in particolare, saranno singoli caratteri come `A`).

Converti i valori `start` e `end` nella chiamata `range()` in numeri utilizzando su di essi il metodo `.charCodeAt()`, passando al metodo il numero `0` come argomento.

# --hints--

Dovresti usare il metodo `.charCodeAt()`.

```js
assert.match(code, /\.charCodeAt\(/);
```

Dovresti chiamare il metodo `.charCodeAt()` su `start`.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Dovresti passare `0` al metodo `.charCodeAt()` di `start`.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Dovresti chiamare il metodo `.charCodeAt()` su `end`.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Dovresti passare `0` al metodo `.charCodeAt()` di `end`.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Dovresti usare i metodi `.charCodeAt()` direttamente nella chiamata di `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
