---
id: 643715013330824ecaa70442
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

La prossima funzione calcolerà il valore mediano di un array di numeri. Inizia dichiarando una funzione freccia `median` che prende un parametro `nums`.

Nella funzione, dichiara una variabile `sorted` e assegnale il valore dell'ordinamento di una copia dell'array `nums`.

Dovresti usare il metodo `slice()` per creare una copia superficiale dell'array.

# --hints--

Dovresti dichiarare una variabile `median`.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

Dovresti usare `const` per dichiarare la variabile `median`.

```js
assert.match(code, /const\s+median/);
```

La variabile `median` dovrebbe essere una funzione.

```js
assert.isFunction(median);
```

La funzione `median` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

La funzione `median` dovrebbe avere un parametro `nums`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

La funzione `median` non dovrebbe utilizzare un return implicito.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

La funzione `median` dovrebbe avere una variabile `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

Dovresti usare `const` per dichiarare la variabile `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

Dovresti usare `.slice()` per assegnare una copia dell'array `nums` a `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

Dovresti concatenare il metodo `.sort()` al metodo `.slice()`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

Dovresti passare una funzione callback al metodo `sort` per ordinare accuratamente i numeri in ordine crescente. Usa un return implicito per chiarezza.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
