---
id: 6437124c4c03dd4c8fb35d56
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Dichiara una funzione `isEven`, che prende un parametro `num` e restituisce `true` se il numero è pari, e `false` in caso contrario. Usa l'operatore di modulo `%` per determinare se un numero è pari o dispari.

# --hints--

Dovresti dichiarare una variabile `isEven`.

```js
assert.match(code, /(?:let|const|var)\s+isEven/);
```

Dovresti usare `const` per dichiarare la variabile `isEven`.

```js
assert.match(code, /const\s+isEven/);
```

La variabile `isEven` dovrebbe essere una funzione.

```js
assert.isFunction(isEven);
```

La funzione `isEven` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s+isEven\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Your `isEven` function should have a `num` parameter.

```js
assert.match(code, /const\s+isEven\s*=\s*(\(\s*num\s*\)|num)\s*=>/);
```

La funzione `isEven` dovrebbe utilizzare l'operatore di modulo `%`.

```js
assert.match(isEven.toString(),  /%/);
```

La funzione `isEven` dovrebbe restituire `true` per numeri pari.

```js
assert.isTrue(isEven(2));
assert.isTrue(isEven(1000));
assert.isTrue(isEven(42));
```

La funzione `isEven` dovrebbe restituire `false` per numeri dispari.

```js
assert.isFalse(isEven(1));
assert.isFalse(isEven(333));
assert.isFalse(isEven(777777777));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
const sum = nums => nums.reduce((acc, el) => acc + el, 0);

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
