---
id: 6437133052eaf04d7300e622
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Dichiara una funzione `average` che prende un array di numeri come parametro `nums`. Dovrebbe restituire la media di tutti i numeri nell'array.

La media può essere calcolata dividendo la somma di tutti i numeri nell'array per la lunghezza dell'array. Ricorda che hai una funzione `sum` che puoi usare.

# --hints--

Dovresti dichiarare una variabile `average`.

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

Dovresti usare `const` per dichiarare la variabile `average`.

```js
assert.match(code, /const\s+average/);
```

La variabile `average` dovrebbe essere una funzione.

```js
assert.isFunction(average);
```

La funzione `average` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

La funzione `average` dovrebbe avere un parametro `nums`.

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

La funzione `average` dovrebbe utilizzare un return implicito.

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

La funzione `average` dovrebbe restituire il valore medio dell'array `nums`.

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
