---
id: 64496df724dd3716a71fe971
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Per tenere traccia di tutte le funzioni del foglio di calcolo, dichiara un oggetto `spreadsheetFunctions`. Usando la sintassi di destrutturazione, imposta `sum`, `average` e `median` come proprietà e valori sull'oggetto `spreadsheetFunctions`.

# --hints--

Dovresti dichiarare una variabile `spreadsheetFunctions`.

```js
assert.match(code, /(?:const|let|var)\s+spreadsheetFunctions/);
```

Dovresti usare `const` per dichiarare la variabile `spreadsheetFunctions`.

```js
assert.match(code, /const\s+spreadsheetFunctions/);
```

La variabile `spreadsheetFunctions` dovrebbe essere un oggetto.

```js
assert.isObject(spreadsheetFunctions);
```

L'oggetto `spreadsheetFunctions` dovrebbe avere una proprietà `sum`.

```js
assert.property(spreadsheetFunctions, "sum");
```

La proprietà `sum` dovrebbe essere la tua funzione `sum`.

```js
assert.equal(spreadsheetFunctions?.sum, sum);
```

L'oggetto `spreadsheetFunctions` dovrebbe avere una proprietà `average`.

```js
assert.property(spreadsheetFunctions, "average");
```

La proprietà `average` dovrebbe essere la tua funzione `average`.

```js
assert.equal(spreadsheetFunctions?.average, average);
```

L'oggetto `spreadsheetFunctions` dovrebbe avere una proprietà `median`.

```js
assert.property(spreadsheetFunctions, "median");
```

La proprietà `median` dovrebbe essere la tua funzione `median`.

```js
assert.equal(spreadsheetFunctions?.median, median);
```

Dovresti usare la sintassi di destrutturazione per assegnare le tue proprietà.

```js
const objectText = code.replace(/.*const\s*spreadsheetFunctions\s*=\s*\{([^}]*)}.*/s, "$1");
assert.include(objectText, "sum");
assert.include(objectText, "average");
assert.include(objectText, "median");
assert.notInclude(objectText, ":");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
