---
id: 64497e764135bd24b7960dd3
title: Step 32
challengeType: 0
dashedName: step-32
---

# --description--

Per eseguire le funzioni del foglio di calcolo, è necessario essere in grado di elaborare e valutare la stringa di input. Questo è un ottimo momento per usare un'altra funzione.

Dichiara una funzione freccia `evalFormula` che prende i parametri `x` e `cells`.

# --hints--

Dovresti dichiarare una variabile `evalFormula`.

```js
assert.match(code, /(?:let|const|var)\s*evalFormula/);
```

Dovresti usare `const` per dichiarare la variabile `evalFormula`.

```js
assert.match(code, /const\s*evalFormula/);
```

La variabile `evalFormula` dovrebbe essere una funzione.

```js
assert.isFunction(evalFormula);
```

La funzione `evalFormula` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /const\s*evalFormula\s*=\s*\(/);
```

La funzione `evalFormula` dovrebbe avere `x` come primo parametro.

```js
assert.match(code, /const\s*evalFormula\s*=\s*\(\s*x/);
```

La funzione `evalFormula` dovrebbe avere `cells` come secondo parametro.

```js
assert.match(code, /const\s*evalFormula\s*=\s*\(\s*x\s*,\s*cells\s*\)\s*=>/);
```

La funzione `evalFormula` dovrebbe essere vuota.

```js
assert.match(code, /const\s*evalFormula\s*=\s*\(\s*x\s*,\s*cells\s*\)\s*=>\s*{\s*}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

const spreadsheetFunctions = {
  sum,
  average,
  median
}

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      input.onchange = update;
      container.appendChild(input);
    })
  })
}

const update = event => {
  const element = event.target;
  const value = element.value.replace(/\s/g, "");
  if (!value.includes(element.id) && value.startsWith('=')) {

  }
}
```
