---
id: 646d3d80c3b4aebc4103618e
title: Step 66
challengeType: 0
dashedName: step-66
---

# --description--

Seguendo lo stesso modello, aggiungi una proprietà per la moltiplicazione `*` e la divisione `/` con le funzioni appropriate.

# --hints--

L'oggetto `infixToFunction` dovrebbe avere una proprietà `*`.

```js
assert.property(infixToFunction, '*');
```

La proprietà `*` dovrebbe essere una funzione.

```js
assert.isFunction(infixToFunction['*']);
```

La funzione `*` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /('|"|`)\*\1\s*:\s*\(/);
```

La funzione `*` dovrebbe avere `x` come primo parametro.

```js
assert.match(code, /('|"|`)\*\1\s*:\s*\(\s*x/);
```

La funzione `*` dovrebbe avere `y` come secondo parametro.

```js
assert.match(code, /('|"|`)\*\1\s*:\s*\(\s*x\s*,\s*y\s*\)/);
```

La funzione `*` dovrebbe utilizzare un return implicito.

```js
assert.notMatch(code, /('|"|`)\*\1\s*:\s*\(\s*x\s*,\s*y\s*\)\s*\{/);
```

La funzione `*` dovrebbe restituire il risultato della moltiplicazione di `x` per `y`.

```js
assert.equal(infixToFunction['*'](2, 5), 10);
```

L'oggetto `infixToFunction` dovrebbe avere una proprietà `/`.

```js
assert.property(infixToFunction, '/');
```

La proprietà `/` dovrebbe essere una funzione.

```js
assert.isFunction(infixToFunction['/']);
```

La funzione `/` dovrebbe utilizzare la sintassi freccia.

```js
assert.match(code, /('|"|`)\/\1\s*:\s*\(/);
```

La funzione `/` dovrebbe avere `x` come primo parametro.

```js
assert.match(code, /('|"|`)\/\1\s*:\s*\(\s*x/);
```

La funzione `/` dovrebbe avere `y` come secondo parametro.

```js
assert.match(code, /('|"|`)\/\1\s*:\s*\(\s*x\s*,\s*y\s*\)/);
```

La funzione `/` dovrebbe utilizzare un return implicito.

```js
assert.notMatch(code, /('|"|`)\/\1\s*:\s*\(\s*x\s*,\s*y\s*\)\s*\{/);
```

La funzione `/` dovrebbe restituire il risultato della divisione di `x` per `y`.

```js
assert.equal(infixToFunction['/'](10, 2), 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
const infixToFunction = {
  "+": (x, y) => x + y,
  "-": (x, y) => x - y,

}
--fcc-editable-region--

const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

const median = nums => {
  const sorted = nums.slice().sort((a, b) => a - b);
  const length = sorted.length;
  const middle = length / 2 - 1;
  return isEven(length)
    ? average([sorted[middle], sorted[middle + 1]])
    : sorted[Math.ceil(middle)];
}

const spreadsheetFunctions = {
  sum,
  average,
  median
}

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

const evalFormula = (x, cells) => {
  const idToText = id => cells.find(cell => cell.id === id).value;
  const rangeRegex = /([A-J])([1-9][0-9]?):([A-J])([1-9][0-9]?)/gi;
  const rangeFromString = (num1, num2) => range(parseInt(num1), parseInt(num2));
  const elemValue = num => character => idToText(character + num);
  const addCharacters = character1 => character2 => num => charRange(character1, character2).map(elemValue(num));
  const rangeExpanded = x.replace(rangeRegex, (_match, char1, num1, char2, num2) => rangeFromString(num1, num2).map(addCharacters(char1)(char2)));
  const cellRegex = /[A-J][1-9][0-9]?/gi;
  const cellExpanded = rangeExpanded.replace(cellRegex, match => idToText(match.toUpperCase()));
}

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      input.onchange = update;
      container.appendChild(input);
    })
  })
}

const update = event => {
  const element = event.target;
  const value = element.value.replace(/\s/g, "");
  if (!value.includes(element.id) && value.startsWith('=')) {

  }
}
```
