---
id: 61537485c4f2a624f18d7794
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Comience con tu plantilla estándar de HTML. Añade una declaración `DOCTYPE`, un elemento `html`, un elemento `head` y un elemento `body`.

Añade el atributo `lang` a la etiqueta `<html>` de apertura con el valor `en`.

# --hints--

Tu código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debes definir `html` como tipo de documento.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu `html` debe tener una etiqueta de apertura.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu etiqueta de apertura `<html>` debe tener el atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Tu declaración `DOCTYPE` debe estar al inicio de tu código HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debes tener una etiqueta `head` de apertura.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta de cierre del `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta `body` de apertura.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta `body` de cierre.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

El elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
