---
id: 61537bb9b1a29430ac15ad38
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro de tu elemento `head`, agrega un elemento `title` con el texto establecido en `Photo Gallery`, y un elemento `link` para agregar tu archivo `styles.css` a la página.

# --hints--

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu proyecto debe tener el título `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
