---
id: 61537d86bdc3dd343688fceb
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

A continuación, dale a cada `img` un atributo `src` de acuerdo con su orden en el documento. El primer `img` debe tener un atributo `src` con el valor `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. El resto debe ser igual, excepto que hay que sustituir el `1` por el número en el que se encuentra el `img` en el documento.

# --hints--

Los nueve elementos `img` deben tener un atributo `src`.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

Tu primer elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Tu segundo elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Tu tercer elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Tu cuarto elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Tu quinto elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Tu sexto elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Tu séptimo elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Tu octavo elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Tu noveno elemento `img` debe tener `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` establecido como valor del atributo `src`.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
