---
id: 61537c5f81f0cf325b4a854c
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Añade un elemento `header` dentro del elemento `body` y asígnale la clase `header`.

Dentro del `header`, crea un `h1` que tenga el texto `css flexbox photo gallery`.

# --hints--

Debes tener un elemento `header` dentro de tu elemento `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Tu elemento `header` debe tener un `class` con el valor `header`.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Tu elemento `header` debe tener un elemento `h1` dentro de él.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Tu elemento `h1` debe tener el texto `css flexbox photo gallery` dentro de él.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
