---
id: 61537c9eecea6a335db6da79
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Debajo del elemento `.header`, crea un nuevo elemento `div` y asigna el valor `gallery` a su atributo `class`. Este `div` actuará como contenedor de las imágenes de la galería.

Dentro de ese elemento `.gallery`, crea nueve elementos `img`.

# --hints--

Debes crear un elemento `div` en tu elemento `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Tu nuevo elemento `div` debe tener un atributo `class` con el valor `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Tu nuevo elemento `div` debe estar después de tu elemento `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Tu elemento `.gallery` debe tener nueve elementos `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
