---
id: 61537a8054753e2f1f2a1574
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu elemento `head`, agrega una etiqueta `meta` con el `name` establecido a `viewport` y el `content` establecido a `width=device-width, initial-scale=1`.

También añade una etiqueta `meta` y a `charset` dale el valor `UTF-8`.

# --hints--

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Un elemento `meta` debe tener un `name` establecido en `viewport`, y el `content` establecido a `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

Tu otro elemento `meta` debe tener un atributo `charset` establecido a `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
