---
id: 64dd076a1b33c86d84b4232f
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Para visualizar mejor el tamaño de los elementos, puede ser útil añadir un borde.

Dale al elemento `.gallery` un `width` de `50%` y un `border` de `5px solid red`.

Luego da a tus elementos `img` un `width` de `100%`, un `padding` de `5px` y un `border`de `5px solid blue`.

# --hints--

Debes tener un selector `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Tu selector `.gallery` debe tener una propiedad `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Tu selector `.gallery` debe tener una propiedad `width` establecida en `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Tu selector `.gallery` debe tener una propiedad `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Tu selector `.gallery` debe tener una propiedad `border` establecida en `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Debes tener un selector `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Tu selector `img` debe tener una propiedad `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Tu selector `img` debe tener una propiedad `width` establecida en `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Tu selector `img` debe tener una propiedad `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Tu selector `img` debe tener una propiedad `padding` establecida en `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Tu selector `img` debe tener una propiedad `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Tu selector `img` debe tener una propiedad `border` establecida en `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
