---
id: 6930c6ffe3de7934afe5e782
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now that you understand the data structures, let's start building the helper functions that will support the simulation.

First, you'll create a measureThroughput function that analyzes how many attendees each gate successfully processed. This function will take an array of gates and return an array of objects containing throughput statistics.

Each returned object should have:

id: the gate's identifier
totalServed: the total number of attendees processed (from gate.throughput)
perTick: a copy of the per-tick throughput array (from gate.throughputPerTick)
Create a function called measureThroughput that takes a parameter called gatesList. 


```js
function measureThroughput(gatesList) {

}
```

Start by creating a function called `measureThroughput`.

# --hints--

Use .slice() to create a copy of the throughputPerTick array.

```js
assert.isFunction(measureThroughput);
```

Your `measureThroughput` function should use `.slice()` to create a copy of the `throughputPerTick` array.

```js
assert.match(code, /\.slice\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
