---
id: 69babf3f446433af8d065bb5
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Create an empty `while` loop that continues as long as there are attendees waiting (`currentTickQueue > 0`) and the gate has remaining capacity (`processed < gate.capacity`).

# --hints--

Your `while` loop's condition should be `currentTickQueue > 0` and `processed < gate.capacity`.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const viableSolutions = [
  "while(currentTickQueue>0&&processed<gate.capacity){}}",
  "while((currentTickQueue>0)&&(processed<gate.capacity)){}}",
  "while(processed<gate.capacity&&currentTickQueue>0){}}",
  "while((processed<gate.capacity)&&(currentTickQueue>0)){}}",
  "while(currentTickQueue>0&&processed<gate.capacity){};}",
  "while((currentTickQueue>0)&&(processed<gate.capacity)){};}",
  "while(processed<gate.capacity&&currentTickQueue>0){};}",
  "while((processed<gate.capacity)&&(currentTickQueue>0)){};}"
];
const endsCorrectly = viableSolutions.some(sol => cleaned.endsWith(sol));
assert(
  endsCorrectly,
  "Your while loop must have a condition with both `currentTickQueue > 0` and `processed < gate.capacity`, in any order."
);
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
--fcc-editable-region--

--fcc-editable-region--
}
```
