---
id: 69bac03ae5faf2979a3339ff
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Now that you know which gate to reroute to, you can move the overflow attendees.

Take the `overflowAmount` and add it to the next gate’s queue at the same `tickIndex`.

Specifically, you must:

- Access the next gate using `nextGateIndex`.
- Access its `queue`.
- Add the `overflowAmount` to the value at position `tickIndex`.

# --hints--

Your `rerouteOverflow` function should add `overflowAmount` to the value at `tickIndex` in the `queue` of the next gate.

```js
const testGates1 = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const northGate = testGates1[0];
const originalEastQueue = testGates1[1].queue[0];
rerouteOverflow(testGates1, northGate, 0, 5);
assert.strictEqual(
  testGates1[1].queue[0],
  originalEastQueue + 5,
  "Rerouting from North gate should add overflow to East gate's queue at tickIndex 0"
);

const testGates2 = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const westGate = testGates2[3];
const originalNorthQueue = testGates2[0].queue[2];
rerouteOverflow(testGates2, westGate, 2, 3);
assert.strictEqual(
  testGates2[0].queue[2],
  originalNorthQueue + 3,
  "Rerouting from West gate should wrap around and add overflow to North gate's queue at tickIndex 2"
);

const testGates3 = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const southGate = testGates3[2];
const originalQueues = testGates3.map(gate => [...gate.queue]);
rerouteOverflow(testGates3, southGate, 1, 2);
assert.deepStrictEqual(testGates3[0].queue, originalQueues[0], "North gate queue should not be modified");
assert.deepStrictEqual(testGates3[1].queue, originalQueues[1], "East gate queue should not be modified");
assert.deepStrictEqual(testGates3[2].queue, originalQueues[2], "South gate queue should not be modified");
assert.strictEqual(testGates3[3].queue[1], originalQueues[3][1] + 2, "Only West gate at tickIndex 1 should be modified");
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
--fcc-editable-region--

--fcc-editable-region--
}
```
