---
id: 69bac073b1de8b484bbccc94
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

First, log the following message to the console: `"\nProcessing " + gate.id + "..."`.

# --hints--

Your `handleGateAtTick` function should first log the following message to the console, `"\nProcessing " + gate.id + "..."`.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };

const spy = __helpers.spyOn(console, "log");
try {
  const northGate = testGates[0];
  handleGateAtTick(testGates, northGate, 0, throughputSummary);
  const expectedCall = ["\nProcessing North..."];
  assert.deepEqual(
    spy.calls[0],
    expectedCall,
    "handleGateAtTick should first log the processing message with the gate id"
  );
} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Your `handleGateAtTick` function should log the correct gate ID in the message.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };

const spy = __helpers.spyOn(console, "log");
try {
  const southGate = testGates[2];
  handleGateAtTick(testGates, southGate, 1, throughputSummary);
  const expectedCall = ["\nProcessing South..."];
  assert.deepEqual(
    spy.calls[0],
    expectedCall,
    "handleGateAtTick should log the processing message with the correct gate id"
  );
} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    "  " + overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
--fcc-editable-region--

--fcc-editable-region--
}
```
