---
id: 69bac12483f656898154f969
title: Step 23
challengeType: 1
dashedName: step-23
---

# --description--

Use a `for...in` loop to iterate through every gate in the input `summary` object and log the total attendees processed.  

For each gate, log a message in this format: `gate + ": " + summary[gate] + " attendees processed"`.

As an example, for an input `summary` object defined as `{ North: 5, East: 3, South: 4, West: 2 }`, you should log: 

```js
North: 5 attendees processed
East: 3 attendees processed
South: 4 attendees processed
West: 2 attendees processed
```

# --hints--

You should loop through every element in the `summary` object and log the following message once per element: `element + ": " + summary[element] + " attendees processed`.

```js
const testSummary = { North: 12, East: 8, South: 5, West: 4 };

const spy = __helpers.spyOn(console, "log");

try {
  printSummary(testSummary);
  
  assert.lengthOf(
    spy.calls,
    5,
    "Should log 1 header message + 4 gate messages"
  );
  
  assert.deepEqual(
    spy.calls[0],
    ["\nThroughput Summary"],
    "First log should be the header"
  );
  
  assert.match(
    spy.calls[1][0],
    /North: 12 attendees processed/,
    "Should log throughput for North gate"
  );
  assert.match(
    spy.calls[2][0],
    /East: 8 attendees processed/,
    "Should log throughput for East gate"
  );
  assert.match(
    spy.calls[3][0],
    /South: 5 attendees processed/,
    "Should log throughput for South gate"
  );
  assert.match(
    spy.calls[4][0],
    /West: 4 attendees processed/,
    "Should log throughput for West gate"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
--fcc-editable-region--

--fcc-editable-region--
}
```
