---
id: 69bac17a4c53a3812b595eb9
title: Step 26
challengeType: 1
dashedName: step-26
---

# --description--

Before starting the simulation, you need an object to track how many attendees each gate processes over all ticks. Call `initializeThroughput()` with `gates` and assign the result to a variable named `throughputSummary`.

# --hints--

You should call `initializeThroughput()` with `gates` and assign the result to a variable named `throughputSummary`.

```js
const rawFuncStr = __helpers.removeJSComments(simulateFestival.toString());
assert.match(
  rawFuncStr,
  /(var|const|let)\s+throughputSummary/,
  "You must have at least one space between the declaration keyword (let/const) and 'throughputSummary'."
);

const cleaned = __helpers.removeWhiteSpace(rawFuncStr);
const solution = /(var|const|let)throughputSummary=initializeThroughput\(gates\);?}$/;

assert(
  solution.test(cleaned),
  "Inside your `simulateFestival` function, call `initializeThroughput()` with `gates` and assign the result to a variable named `throughputSummary`."
);
```

Your `simulateFestival` function should call the `initializeThroughput()` function with `gates`.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];

const originalInitializeThroughput = initializeThroughput;
let called = false;
let callArgs = null;

try {
  initializeThroughput = function (...args) {
    called = true;
    callArgs = args;
    return originalInitializeThroughput(...args);
  };

  const consoleSpy = __helpers.spyOn(console, "log");
  
  try {
    simulateFestival(testGates, "Morning");
    
    assert.isTrue(
      called,
      "simulateFestival should call initializeThroughput()"
    );
    assert.deepEqual(
      callArgs,
      [testGates],
      "initializeThroughput should be called with gates parameter"
    );
    
  } finally {
    consoleSpy.restore();
  }

} catch (err) {
  assert.fail(err);
} finally {
  initializeThroughput = originalInitializeThroughput;
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
  console.log(
    gate.queue[tickIndex] + " attendees arriving."
  );
  const result = processGateFlow(gate, tickIndex);
  throughputSummary[gate.id] += result.processed;
  if (result.overflow > 0) {
    console.log(
      "Overflow of " + result.overflow +
      " attendees. Rerouting..."
    );
    rerouteOverflow(gates, gate, tickIndex, result.overflow);
  }
}

function printSummary(summary) {
  console.log("\nThroughput Summary");
  for (const gateId in summary) {
    console.log(
      gateId + ": " + summary[gateId] +
      " attendees processed"
    );
  }
}

function simulateFestival(gates, timeBlock) {
  console.log("\n" + timeBlock + " Simulation");
--fcc-editable-region--

--fcc-editable-region--
}
```
