---
id: 69bac08105bec16db267dcbf
title: Step 17
challengeType: 1
dashedName: step-17
---

# --description--

Next, log the following message to the console: `gate.queue[tickIndex] + " attendees arriving."`.

# --hints--

Your `handleGateAtTick` function should log the following message to the console: `gate.queue[tickIndex] + " attendees arriving."`.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };

const spy = __helpers.spyOn(console, "log");
try {
  const northGate = testGates[0];
  handleGateAtTick(testGates, northGate, 0, throughputSummary);
  const expectedFirstCall = ["\nProcessing North..."];
  const expectedSecondCall = ["3 attendees arriving."];
  assert.deepEqual(
    spy.calls[0],
    expectedFirstCall,
    "First log should be the processing message"
  );
  assert.deepEqual(
    spy.calls[1],
    expectedSecondCall,
    "Second log should show the number of attendees arriving"
  );

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Your `handleGateAtTick` function should log the correct number of attendees based on the gate's queue at the given tickIndex.

```js
const testGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] }
];
const throughputSummary = { North: 0, East: 0, South: 0, West: 0 };

const spy = __helpers.spyOn(console, "log");
try {
  const eastGate = testGates[1];
  handleGateAtTick(testGates, eastGate, 3, throughputSummary);
  const expectedFirstCall = ["\nProcessing East..."];
  const expectedSecondCall = ["5 attendees arriving."];
  assert.deepEqual(
    spy.calls[0],
    expectedFirstCall,
    "First log should be the processing message with the correct gate"
  );
  assert.deepEqual(
    spy.calls[1],
    expectedSecondCall,
    "Second log should show the correct number of attendees from queue at tickIndex 3"
  );
} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js
const morningGates = [
  { id: "North", capacity: 5, queue: [3, 6, 2, 4] },
  { id: "East", capacity: 3, queue: [2, 4, 3, 5] },
  { id: "South", capacity: 4, queue: [1, 2, 3, 1] },
  { id: "West", capacity: 2, queue: [4, 1, 2, 3] },
];

const nightGates = [
  { id: "North", capacity: 4, queue: [6, 2, 5, 1] },
  { id: "East", capacity: 2, queue: [3, 3, 4, 2] },
  { id: "South", capacity: 5, queue: [2, 1, 2, 3] },
  { id: "West", capacity: 3, queue: [5, 2, 1, 4] },
];

function initializeThroughput(gates) {
  const summary = {};
  for (const gate of gates) {
    summary[gate.id] = 0;
  };
  return summary;
}

function processGateFlow(gate, tickIndex) {
  let currentTickQueue = gate.queue[tickIndex];
  let processed = 0;
  while (currentTickQueue > 0 && processed < gate.capacity) {
    currentTickQueue--;
    processed++;
  }
  return {
    processed: processed,
    overflow: currentTickQueue
  };
}

function rerouteOverflow(gates, currentGate, tickIndex, overflowAmount) {
  const currentIndex = gates.indexOf(currentGate);
  const nextGateIndex = (currentIndex + 1) % gates.length;
  gates[nextGateIndex].queue[tickIndex] += overflowAmount;
  console.log(
    overflowAmount + " attendees rerouted to " +
    gates[nextGateIndex].id
  );
}

function handleGateAtTick(gates, gate, tickIndex, throughputSummary) {
  console.log("\nProcessing " + gate.id + "...");
--fcc-editable-region--

--fcc-editable-region--
}
```
