---
id: 6140d0069049f5139d78da40
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

يتم استخدام خاصية `transform-origin` لتعيين النقطة التي يتم حولها تطبيق تحويل (transformation) CSS. على سبيل المثال، عند تنفيذ `rotate` (الذي ستفعله لاحقا في هذا المشروع)، يحدد `transform-origin` النقطة التي يتم تدوير العنصر حولها.

قم بإعطاء منتقي `.line` خاصية `transform-origin` بقيمة `0% 0%`. سيؤدي هذا إلى موازنة تعويض (offset) النقطة الأصل من اليسار بقيمة `0%` ومن الأعلى `0%`، مع تعيينها في الزاوية اليسرى العلوية للعنصر.

# --hints--

يجب أن يكون لمنتقي `.line` الخاص بك الخاصية `transform-origin` بقيمة `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
