---
id: 6140c7e645d8e905819f1dd4
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

ابدأ بالنموذج المعتاد. أضف تعريف `DOCTYPE`، وعنصر `html` مع تعين لغة الصفحة إلى ألانجليزية (en) وعنصر `head`, وعنصر `body`.

أضف عنصر `meta` الخاص بك إلى `charset` الصحيحة، عنصر `title`، وعنصر `link` لمِلَفّ `./styles.css`.

عيّن `title` إلى `Ferris Wheel`.

# --hints--

يجب أن يحتوي الكود الخاص بك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` مع `>` بعد إضافة type.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة افتتاحية مع سمة `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على closing tag.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك العلامة مفتوحة (opening tag) الآتية للعنصر `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتية للعنصر `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك العلامة مفتوحة (opening tag) الآتية للعنصر `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتية للعنصر `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

عناصر `head` و `body` يجب أن يكونا أشقاء (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

عنصر `head` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

عنصر `body` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

الكود الخاص بك يجب أن يحتوي على عنصر `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

يجب أن يحتوي عنصر `meta` الخاص بك على خاصية `charset` مع القيمة `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لمشروعك عنوان `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

تذكر أن casing، أي حالة الحرف سواء كبير أو صغير (capital or small) والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
assert.match(code, /<link/)
```

الكود الخاص بك يجب أن يحتوي على عنصر `link` مغلق ذاتيا.

```js
assert(document.querySelectorAll('link').length === 1);
```

عنصر `link` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert.exists(document.querySelector('head > link'));
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
