---
id: 6140cdebd39d6a101e747529
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

قم بإنشاء منتقي لعنصر `.line` الخاص بك. ابدأ من خلال تعيين `background-color` إلى `black` و `width` إلى `50%` و `height` إلى `2px`.

# --hints--

يجب أن يكون لديك منتقي `.line`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

يجب أن يحتوي منتقي `.line` على خاصية `background-color` تم تعيينها إلى `black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

يجب أن يكون لمنتقي `.line` الخاص بك الخاصية `width` بقيمة `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

يجب أن يكون لمنتقي `.line` الخاص بك الخاصية `height` بقيمة `2px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
