---
id: 6140cd32d018ed0f600eefce
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

عيين قيمة خاصية `position` في منتقي `.wheel` إلى `absolute`. قم بتعيين `height` و `width` إلى `55vw`.

# --hints--

يجب أن يكون لمنتقي `.wheel` الخاص بك الخاصية `position` بقيمة `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

يجب أن يكون لمنتقي `.wheel` الخاص بك الخاصية `height` بقيمة `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

يجب أن يكون لمنتقي `.wheel` الخاص بك الخاصية `width` بقيمة `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
