---
id: 6140cbeec34e970dfe75e710
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

قم بإنشاء منتقي لعنصر `.wheel` الخاص بك. ابدأ بتعيين `border` إلى `2px solid black`، و `border-radius`الي `50%`، و `margin-left` إلى `50px`.

# --hints--

يجب أن يكون لديك منتقي `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

يجب أن يكون لمنتقي `.wheel` الخاص بك خاصية `border` بقيمة `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

يجب أن يكون لمنتقي `.wheel` الخاص بك الخاصية `border-radius` بقيمة `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

يجب أن يكون لمنتقي `.wheel` الخاص بك الخاصية `margin-left` بقيمة `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
