---
id: 6140c7e645d8e905819f1dd4
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

ابدأ بالنموذج المعتاد. أضف تعريف `DOCTYPE`، وعنصر `html` مع تعين لغة الصفحة إلى ألانجليزية (en) وعنصر `head`, وعنصر `body`.

أضف عنصر `meta` الخاص بك إلى `charset` الصحيحة، عنصر `title`، وعنصر `link` لمِلَفّ `./styles.css`.

عيّن `title` إلى `Ferris Wheel`.

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

يجب عليك إغلاق تعريف `DOCTYPE` مع `>` بعد إضافة type.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة افتتاحية مع سمة `lang` بقيمة `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

يجب أن يحتوي عنصر `html` الخاص بك على closing tag.

```js
assert.match(code, /<\/html\s*>/);
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

يجب أن يكون لديك العلامة مفتوحة (opening tag) الآتية للعنصر `head`.

```js
assert.match(code, /<head\s*>/i);
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتية للعنصر `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

يجب أن يكون لديك العلامة مفتوحة (opening tag) الآتية للعنصر `body`.

```js
assert.match(code, /<body\s*>/i);
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتية للعنصر `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

الكود الخاص بك يجب أن يحتوي على عنصر `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

يجب أن يحتوي عنصر `meta` الخاص بك على خاصية `charset` مع القيمة `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لمشروعك عنوان `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

تذكر أن casing، أي حالة الحرف سواء كبير أو صغير (capital or small) والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

عنصر `link` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert.exists(document.querySelector('head > link'));
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `rel` بقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
