---
id: 6140d94b5fab7f1d73c9bedb
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

تستخدم القاعدة `@keyframes` لتحديد تدفق الرسوم المتحركة (animation) لـ CSS. ضمن قاعدة `@keyframes`، يمكنك إنشاء selectors لنقاط محددة في سلسلة الرسوم المتحركة (animation sequence)، مثل `0%` أو `25%`، أو استخدام `from` و `to` لتحديد بداية ونهاية التسلسل.

تتطلب قواعد `@keyframes` تعيين اسم لها، والذي تستخدمه في القواعد الأخرى للاشارة إليه. على سبيل المثال، قاعدة `@keyframes freeCodeCamp { }` سيتم تسميتها `freeCodeCamp`.

حان الوقت لبدء الرسوم المتحركة (animating). قم بإنشاء قاعدة `@keyframes` اسمها `wheel`.

# --hints--

يجب أن يكون لديك قاعدة `@keyframes`.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.length === 1);
```

قاعدة `@keyframes` الجديدة الخاصة بك يجب أن تسمى `wheel`.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name === 'wheel');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

--fcc-editable-region--

--fcc-editable-region--
```
