---
id: 6140c9d35015ae0ba0c250e8
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف `div` داخل عنصر `body` وأعطيه `class` يسمي `wheel`.

داخل `div` الجديد، أضف ستة عناصر `span` مع `class` يسمي `line`، وستة عناصر `div` مع `class` يسمي `cabin`.

# --hints--

يجب عليك إنشاء عنصر `div` جديد في عنصر `body` الخاص بك.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

`div` داخل عنصر `body` يجب ان يحتوي علي `class` يسمي `wheel`.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

عنصر `.wheel` الخاص بك يجب أن يحتوي على ستة عناصر `span`.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

يجب أن تحتوي عناصر `span` الستة الخاصة بك داخل عنصر `.wheel` على `class` يسمي `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

عنصر `.wheel` الخاص بك يجب أن يحتوي على ستة عناصر `div`.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

يجب أن تحتوي عناصر `div` الستة الخاصة بك داخل عنصر `.wheel` على `class` يسمي `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
