---
id: 6143869bb45bd85e3b1926aa
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

В елементі `.heading` створіть елемент `header`, що має `class` зі значенням `hero`.

В цьому елементі створіть елемент `img`, що має `src` зі значенням `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, `alt` зі значенням `freecodecamp logo` та `class` зі значенням `hero-img`.

Атрибут `loading` в елементі `img` можна встановити на `lazy`, щоб сказати браузеру не завантажувати ресурс зображення, допоки не знадобиться (як-от, коли користувач прокручує зображення в поле зору). Як додаткова перевага, ліниво завантажені елементи не завантажуватимуться, доки не завантажаться не ліниві елементи – це означає, що користувачі з повільним підключенням до інтернету можуть переглядати вміст вашої сторінки, не чекаючи, поки завантажаться зображення.

Надайте елементу `img` атрибут `loading` зі значенням `lazy`.

# --hints--

Ви повинні створити елемент `header`.

```js
assert.exists(document.querySelector('header'));
```

Ваш елемент `header` повинен бути в межах елемента `.heading`.

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

Ваш елемент `header` повинен мати `class` зі значенням `hero`.

```js
assert(document.querySelector('header')?.className === 'hero');
```

Ви повинні створити елемент `img`.

```js
assert.exists(document.querySelector('img'));
```

Ваш елемент `img` повинен бути в межах елемента `header`.

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

Ваш елемент `img` повинен мати `src` зі значенням `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Ваш елемент `img` повинен мати `alt` зі значенням `freecodecamp logo`.

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

Ваш елемент `img` повинен мати атрибут `loading` зі значенням `lazy`.

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

Ваш елемент `img` повинен мати `class` зі значенням `hero-img`.

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
