---
id: 614389f601bb4f611db98563
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Під елементом `.author` створіть новий елемент `div`, що має клас `social-icons`.

Додайте п’ять елементів `a` в цей новий `div` та надайте їм наступні атрибути `href`.

- Перший елемент `a` повинен мати `href` зі значенням `https://www.facebook.com/freecodecamp`.
- Другий елемент `a` повинен мати `href` зі значенням `https://twitter.com/freecodecamp`.
- Третій елемент `a` повинен мати `href` зі значенням `https://instagram.com/freecodecamp`.
- Четвертий елемент `a` повинен мати `href` зі значенням `https://www.linkedin.com/school/free-code-camp`.
- П'ятий елемент `a` повинен мати `href` зі значенням `https://www.youtube.com/freecodecamp`.

# --hints--

Ви повинні створити новий елемент `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Ваш новий елемент `div` повинен йти після елемента `.author`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

Ваш новий елемент `div` повинен мати клас `social-icons`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

Ваш елемент `.social-icons` повинен мати п'ять елементів `a`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

Ваш перший елемент `a` повинен мати `href` зі значенням `https://www.facebook.com/freecodecamp`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

Ваш другий елемент `a` повинен мати `href` зі значенням `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

Ваш третій елемент `a` повинен мати `href` зі значенням `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

Ваш четвертий елемент `a` повинен мати `href` зі значенням `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

Ваш п'ятий елемент `a` повинен мати `href` зі значенням `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
