---
id: 614385513d91ae5c251c2052
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте елемент `title` з текстом `Magazine`, елемент `link` для таблиці стилів шрифтів `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, `link` для таблиці стилів FontAwesome `https://use.fontawesome.com/releases/v5.8.2/css/all.css` та `link` для своєї таблиці стилів `./styles.css`.

Ваша таблиця стилів шрифтів завантажить три окремі шрифти: `Anton`, `Baskervville` та `Raleway`.

# --hints--

Ваш код повинен мати три самозакривні елементи `link`.

```js
assert(document.querySelectorAll('link').length === 3);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert(document.querySelectorAll('head > link').length === 3);
```

Ваші три елементи `link` повинні мати атрибут `rel` зі значенням `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Один із ваших елементів зв'язку повинен мати `href` зі значенням `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Один із ваших елементів зв'язку повинен мати `href` зі значенням `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Один з ваших елементів `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/)
```

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш проєкт повинен мати титул `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

Пам’ятайте, що вживання великої літери та написання мають значення.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
