---
id: 6165d3b702a5d92ad970b30c
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Після свого елемента `img` додайте елемент `h1` з `class` зі значенням `hero-title` та текстом `OUR NEW CURRICULUM`, а потім елемент `p` з `class` зі значенням `hero-subtitle` та текстом `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

Ви повинні створити елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` повинен йти після елемента `img`.

```js
assert(document.querySelector('h1')?.previousElementSibling?.localName === 'img');
```

Ваш елемент `h1` повинен мати `class` зі значенням `hero-title`.

```js
assert(document.querySelector('h1')?.className === 'hero-title');
```

Ваш елемент `h1` повинен мати текст `OUR NEW CURRICULUM`.

```js
assert(document.querySelector('h1')?.textContent === 'OUR NEW CURRICULUM');
```

Ви повинні створити новий елемент `p`.

```js
assert.exists(document.querySelector('p'));
```

Ваш елемент `p` повинен йти після елемента `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Ваш елемент `p` повинен мати `class` зі значенням `hero-subtitle`.

```js
assert(document.querySelector('p')?.className === 'hero-subtitle');
```

Ваш елемент `p` повинен мати текст `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
