---
id: 6169cd8a558aa8434e0ad7f6
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

HTTP заголовок `Referer` містить інформацію про адресу або URL-адресу сторінки, з якої може перейти користувач. Цю інформацію можна використати в аналітиці, щоб відстежувати, наприклад, скільки користувачів вашої сторінки відвідують freecodecamp.org. Встановлення атрибута `rel` на `noreferrer` пропускає цю інформацію з запиту HTTP. Надайте своєму елементу `a` атрибут `rel` зі значенням `noreferrer`.

# --hints--

Ваш елемент `a` повинен мати `rel` зі значенням `noreferrer`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel') === 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
