---
id: 64afc37bf3b37856e035b85e
title: Bevorstehendes Python-Projekt
challengeType: 20
dashedName: upcoming-python-project
---

# --description--

Anweisungen für Python-Projekt

# --hints--

Test 1

```js
({ input: ["Beau", "Carnes"], test: () => {
  assert.equal( "Beau", __userGlobals.get("name"));
  assert.equal( "Carnes", __userGlobals.get("last_name"));
  } })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
name = input('What is your name?')
print('Hi ' + name)
--fcc-editable-region--
```

# --solutions--

```py
name = input('What is your name?')
last_name = input('What is your last name?')
print('Hi ' + name + ' ' + last_name)
```
