---
id: 681dc623b2b18887266777b1
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

In prior lessons, you learned about the `map()` method which creates a new array by applying a given function to each element of the original array.

Here is an example:

```js
const developers = [
  { name: "Alice", city: "New York", title: "Frontend Developer" },
  { name: "Bob", city: "San Francisco", title: "Backend Developer" }
];

console.log(developers.map(dev => dev.name));
// ["Alice", "Bob"]
```

`dev` in this example represents each object in the `developers` array. Then, dot notation is used to get the `name` from the object. Lastly, the result will be a new array of names. 

Inside the `getBookInformation` function, use the `map()` method on the `catalog` parameter to return a new array of just book titles. Refer to the example if you need help.

# --hints--

Your `getBookInformation` function should return an array.

```js
assert.isArray(getBookInformation(library));
```

Your `getBookInformation` function should return an array of strings which represents the book titles.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = ["Title A", "Title B", "Title C"];

assert.deepEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  
}
--fcc-editable-region--
```
