---
id: 69b2762b17d23fbee40a6c00
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

The merge sort algorithm performs three actions:

- Divide an unsorted sequence of items into sub-parts
- Sort the items in the sub-parts
- Merge the sorted sub-parts

The above happens recursively until the sub-parts are merged into the complete sorted sequence. Let's start by dividing the sequence.

First, declare a variable `middlePoint` using `const` and assign it `Math.floor(array.length / 2)` so you can split the sequence in half before recursing.

# --hints--

You should declare a `middlePoint` variable and assign `Math.floor(array.length / 2)` to it.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
