---
id: 69b279bae0f20eda63163cff
title: Step 21
challengeType: 1
dashedName: step-21
---

# --description--

Outside the function definition, create an array `numbers` with the following values: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

You should create an array called `numbers` with these values: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
const explorer = await __helpers.Explorer(code);
const { numbers } = explorer.variables;
assert.exists(
  numbers,
  'You should create a numbers array outside of mergeSort'
);
assert.isTrue(
  numbers.value.matches('[4, 10, 6, 14, 2, 1, 8, 5]'),
  'numbers should equal [4, 10, 6, 14, 2, 1, 8, 5]'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  if (array.length <= 1) {
    return;
  }
  
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }
  
  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
