---
id: 69b2762b17d23fbee40a6c00
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

The merge sort algorithm performs three actions:

- Divide an unsorted sequence of items into sub-parts
- Sort the items in the sub-parts
- Merge the sorted sub-parts

The above happens recursively until the sub-parts are merged into the complete sorted sequence. Start by dividing the sequence.

First, inside the `mergeSort` function, declare a variable `middlePoint` using `const` and assign it `Math.floor(array.length / 2)` so you can split the sequence in half before recursing.

# --hints--

You should declare a `middlePoint` variable and assign `Math.floor(array.length / 2)` to it.

```js
const explorer = await __helpers.Explorer(
  __helpers.removeJSComments(mergeSort.toString())
);

const { middlePoint } = explorer.functions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
