---
id: 69b279bae0f20eda63163cf0
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

At the bottom of your function body, call the `mergeSort()` function again. This time, pass in `rightPart` as the argument to the function call.

# --hints--

You should call `mergeSort()` with the argument `rightPart` at the bottom of your function body.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
