---
id: 69b279bae0f20eda63163cf1
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now it's time to sort and merge the lists (`leftPart` and `rightPart`) into the original `array`. 

You can do this by comparing elements on both lists, and merging the smaller element to the main list. You are going to do this comparison for all the indexes in `leftPart` and `rightPart`.

Create three variables: `leftArrayIndex`, `rightArrayIndex`, and `sortedIndex`, and set their values to `0`. These variables will help you keep track of each index during the sorting process.

# --hints--

You should declare a `leftArrayIndex` variable initialized to `0`.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

You should declare a `rightArrayIndex` variable initialized to `0`.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

You should declare a `sortedIndex` variable initialized to `0`.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
