---
id: 69b278d82f0e514f32acb724
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

After slicing the left half, you still need the right side to complete the merge step.

Use the slice syntax to extract the right half of `array` and assign it to a variable named `rightPart`.

# --hints--

You should have a `rightPart` variable inside the `mergeSort` function.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
```

You should use the slice syntax to extract the right half of the array and assign it to the `rightPart` variable.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
assert.isTrue(
  rightPart.value.matches('array.slice(middlePoint)'),
  'You should assign array.slice(middlePoint) to rightPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
