---
id: 69b279bae0f20eda63163cf9
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

The `while` loop you created compares one element from `leftPart` with another in `rightPart`, then adds the smaller element to the main `array` list.

It will continue this operation until there are no elements left to be compared. But `leftPart` may still have elements left while `rightPart` has none, and vice versa.

Create another `while` loop that runs when `leftArrayIndex` is less than `len(leftPart)`. In the next steps, you'll use it to copy the remaining elements in `leftPart` into the `array` list.

# --hints--

Your `while` loop condition should compare `leftArrayIndex` with `leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
