---
id: 69b279bae0f20eda63163cf4
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

When the `if` condition evaluates to `True`, it means that the element in the `leftPart` list is smaller than the element it is being compared to in the `rightPart` list.

In that case, you can put the element found at `leftArrayIndex` in `leftPart` within the sorted array.

Inside the `if` block assign `left_part[left_array_index]` to `array[sorted_index]`.

# --hints--

You should assign `leftPart[leftArrayIndex]` to `array[sortedIndex]` inside the `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
