---
id: 69b279bae0f20eda63163cf6
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

In a previous step, you assigned the element in the `leftPart` to the `array` list because it was smaller. But this will not always be the case. In some comparison cases, the element on the right could be smaller.

Create an `else` clause to execute when the element in `leftPart` is not less than the element in `rightPart`.

Inside the `else` block, assign `rightPart[rightArrayIndex]` to `array[sortedIndex]`.

# --hints--

You should create an `else` block that contains `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
