---
id: 69b278d82f0e514f32acb724
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

After slicing the left half, you still need the right side to complete the merge step.

Use the slice syntax to extract the right half of `array` and assign it to a variable named `right_part`.

# --hints--

You should have a `rightPart` variable inside `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

You should use the slice syntax to extract the right half of the array and assign it to the `rightPart` variable.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
