---
id: 69b279bae0f20eda63163cfe
title: Step 20
challengeType: 1
dashedName: step-20
---

# --description--

Before testing the `mergeSort()` function, you need to create a base case that stops the function execution when the length of `array` is less than or equal to `1`.

This base case will stop the recursion call. Without it, the merge sort operation would continue to run even when the list has been sorted or has no elements in it.

Right after the function declaration, create an `if` statement that checks if `array.length` is less than or equal to 1. Within the body of the `if` statement, add a `return` statement to stop the execution of the function.

# --hints--

You should create an `if` statement with the condition `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

You should have a `return` statement in the body of the `if` statement.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
