---
id: 5dc174fcf86c76b9248c6eb2
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Елементи HTML мають початкові теґи, як-от `<h1>`, та кінцеві теґи, як-от `</h1>`.

Текст для елемента йде між його початковим та кінцевим теґами.

Знайдіть елемент `h1` та змініть його текст на:

`CatPhotoApp`

# --hints--

Текст `CatPhotoApp` повинен бути в коді. Можливо, доведеться перевірити написання.

```js
assert(code.match(/catphotoapp/i));
```

Ваш елемент `h1` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Ваш елемент `h1` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/h1\>/));
```

Ви маєте більше ніж один елемент `h1`. Видаліть зайвий елемент `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Текст вашого елемента `h1` повинен бути `CatPhotoApp`. Ви пропустили текст, маєте друкарську помилку, або його немає між початковим та кінцевим теґами елемента `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Здається, ви використовуєте розширення браузера, яке змінює сторінку. Переконайтеся, що вимкнули усі розширення браузера.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

