---
id: 5dc17d3bf86c76b9248c6eb4
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Елемент `p` використаний для створення абзаців тексту на сайтах. Створіть елемент `p` під своїм елементом `h2` та надайте йому такий текст:

`See more cat photos in our gallery.`

# --hints--

Ваш елемент `p` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Ваш елемент `p` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/p\>/));
```

Текстом вашого елемента `p` повинен бути `See more cat photos in our gallery.` Ви або не написали текст, або маєте друкарську помилку.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

Ваш елемент `p` повинен бути під елементом `h2`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

