---
id: 5dc17dc8f86c76b9248c6eb5
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Коментування дозволяє вам залишати повідомлення без впливу на показ браузера. Це також дозволяє зробити код неактивним. Коментар в HTML починається із `<!--`, містить будь-яку кількість рядків тексту та закінчується на `-->`. Наприклад, коментар `<!-- TODO: Remove h1 -->` містить текст `TODO: Remove h1`.

Додайте коментар над елементом `p` з таким текстом:

`TODO: Add link to cat photos`

# --hints--

Ваш коментар повинен починатися з `<!--`. Не вистачає одного чи більше символів, що визначають початок коментаря.

```js
assert(code.match(/<!--/));
```

Ваш коментар повинен закінчуватися на `-->`.  Не вистачає одного чи більше символів, що визначають кінець коментаря.

```js
assert(code.match(/-->/));
```

Ваш код не повинен мати зайвих початкових/кінцевих знаків коментаря. Ви маєте зайвий `<!--` або `-->`, що відображається в браузері.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Ваш коментар повинен містити текст `TODO: Add link to cat photos`.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

Ваш коментар повинен бути над елементом `p`. Вони знаходяться в неправильному порядку.

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

