---
id: 5dc1798ff86c76b9248c6eb3
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Заголовкові елементи від `h1` до `h6` використовують для позначення важливості вмісту під ними. Чим менше число, тим вища важливість; отже елементи `h2` мають меншу важливість, ніж елементи `h1`. Використовуйте лише один елемент `h1` на сторінці та розмістіть заголовки меншої важливості під заголовками вищої важливості.

Під елементом `h1` додайте елемент `h2` з таким текстом:

`Cat Photos`

# --hints--

Ваш елемент `h1` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Ваш елемент `h1` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/h1\>/));
```

Ви повинні мати тільки один елемент `h1`. Видаліть зайві.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Текст вашого елемента `h1` повинен бути 'CatPhotoApp'. Ви або не написали текст, або маєте друкарську помилку.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Ваш елемент `h2` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Елемент `h2` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/h2\>/));
```

Текст вашого елемента `h2` повинен бути 'Cat Photos'. Розмістіть текст `Cat Photos` між початковими та кінцевими теґами `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Ваш елемент `h2` повинен бути під елементом `h1`. Елемент `h1` має більшу важливість та повинен бути над елементом `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

