---
id: 5dc23991f86c76b9248c6eb8
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

В попередньому кроці ви вклали елемент `h2`, коментар та елемент `p` в елемент `main`. Абзацний відступ вкладених елементів, що більший на два пробіли за їхній батьківський елемент, покращує читабельність:

```html
<ul> 
  <li> This `li` element indented </li>
  <li> This `li` element is also indented </li>
</ul>
```

Додайте ще два пробіли перед елементом `h2`, коментарем та елементами `p`, щоб ваш HTML був читабельнішим.

# --hints--

Ваш код повинен мати елемент `h2` з текстом `Cat Photos`. Можливо, ви випадково видалили його, в ньому відсутні обидва початковий та кінцевий теґи, або змінився текст.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Ви не повинні додавати елементи `ul` або `li` з прикладу.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

Ваш елемент `h2` повинен бути під початковим теґом елемента `main`, а його початковий теґ повинен починатися після 6 пробілів від початку рядка.

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s{6}<h2>/));
```

Ваш код повинен мати коментар. Ви видалили коментар з попереднього кроку.

```js
assert(code.match(/<!--.*-->/));
```

Текст коментаря повинен бути `TODO: Add link to cat photos`. Не змінюйте текст та пробіли коментаря.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Ваш коментар повинен бути під елементом `h2` та починатись після 6 пробілів з початку рядка.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Ваш код повинен мати елемент `p`. Ви видалили елемент `p` з попереднього кроку.

```js
assert(document.querySelector('p'));
```

Текст елемента `p` повинен бути `Click here to view more cat photos.` Не змінюйте текст, пробіли чи знаки пунктуації елемента `p`.

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

Ваш елемент `p` повинен бути під коментарем, а його початковий теґ повинен починатися через 6 пробілів від початку рядка.

```js
assert(code.toLowerCase().match(/-->\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

