---
id: 5dc2385ff86c76b9248c6eb7
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 має деякі елементи, що ідентифікують різні області вмісту. Ці елементи полегшують прочитання вашого HTML, а також допомагають з оптимізацією пошукової системи (SEO) та доступністю.

Ідентифікуйте головну секцію цієї сторінки, додавши початковий теґ `<main>` після елемента `h1` та кінцевий теґ `</main>` після елемента `p`.

# --hints--

Ваш елемент `main` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('main'));
```

Ваш елемент `main` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/main\>/));
```

Початковий теґ вашого елемента `main` повинен бути під елементом `h1`. Вони знаходяться в неправильному порядку.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling?.tagName, "H1");
```

Початковий теґ вашого елемента `main` повинен бути над елементом `h2`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('main,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H2'));
```

Кінцевий теґ вашого елемента `main` повинен бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/click here to view more cat photos/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
--fcc-editable-region--
  </body>
</html>
```

