---
id: 5dc23f9bf86c76b9248c6eba
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ви можете додати зображення до свого сайту з допомогою елемента `img`. Елементи `img` мають початковий теґ без кінцевого теґа. Теґ для елемента без кінцевого теґа відомий як <dfn>самозакривний теґ</dfn>.

Додайте елемент `img` під елементом `p`. На цю мить жодне зображення не буде відображатися в браузері.

# --hints--

Ваш елемент `img` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Ваш елемент `img` не повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(!code.match(/<\/img\>/));
```

Ви повинні мати тільки один елемент `img`. Видаліть зайві.

```js
assert(document.querySelectorAll('img').length === 1);
```

Ваш елемент `img` повинен бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

