---
id: 5dc24614f86c76b9248c6ebd
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Ви можете посилати на іншу сторінку з елементом прив'язки (`a`). Наприклад, `<a href='https://freecodecamp.org'></a>` посилає на `freecodecamp.org`.

Додайте елемент прив'язки після абзацу, який посилає на `https://freecatphotoapp.com`. На цей момент посилання не з’явиться в попередньому перегляді.

# --hints--

Ваш елемент прив'язки `a` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Ваш елемент прив'язки (`a`) повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/a\>/));
```

Ваш елемент прив'язки (`a`) повинен бути під елементом `p`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Ваш елемент прив'язки (`a`) не має атрибута `href`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Ваш елемент прив'язки (`a`) повинен посилати на `https://freecatphotoapp.com`. Ви або не написали URL-адресу, або маєте друкарську помилку.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Хоча ви встановили атрибут `href` елемента прив'язки ('a') як правильне посилання, краще завжди писати значення атрибута в лапках.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

