---
id: 5dc24073f86c76b9248c6ebb
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>атрибути</dfn> – це спеціальні слова, які використовують всередині початкового теґа елемента для керування поведінкою елемента. Атрибут `src` в елементі `img` зазначає URL-адресу зображення (де локалізовано зображення). Приклад використання елементом `img` атрибута `src`: `<img src="https://www.example.com/the-image.jpg">`.

Всередині наявного елемента `img` додайте атрибут `src` з такою URL-адресою:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Ваш код повинен мати елемент `img`. Можливо, ви видалили елемент `img` або ви не оточили значення атрибута `src` в лапки.

```js
assert(document.querySelector('img'));
```

Елемент `img` повинен мати атрибут `src`. Ви або не додали атрибут, або маєте друкарську помилку. Переконайтеся, що між назвою елемента та назвою атрибуту є пробіл.

```js
assert(document.querySelector('img').src);
```

Атрибут `src` вашого елемента `img` повинен мати значення '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Ви або не написали URL-адресу, або маєте друкарську помилку. Важливо, з якої літери написана URL-адреса (з великої чи малої).

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Хоча ви встановили атрибут `src` елемента `img` як правильну URL-адресу, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

