---
id: 5dc24165f86c76b9248c6ebc
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Всі елементи `img` повинні мати атрибут `alt`. Текст атрибута `alt` використовують для читачів екрана, щоб покращити доступність, та він відображається, якщо зображення не завантажується. Наприклад, `<img src="cat.jpg" alt="A cat">` має атрибут `alt` з текстом `A cat`.

Всередині наявного елемента `img` додайте атрибут `alt` з таким текстом:

`A cute orange cat lying on its back`

# --hints--

Ваш код повинен мати елемент `img`. Ви видалили елемент `img` з попереднього кроку.

```js
assert(document.querySelector('img'));
```

Ваш елемент `img` не має атрибута `alt`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

Значення атрибута `alt` вашого елемента `img` встановлене на щось інше, а не на 'A cute orange cat lying on its back'. Переконайтеся, що значення атрибута `alt` написане в лапках.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

