---
id: 5ddbd81294d8ddc1510a8e56
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Текст посилання повинен бути розміщений між початковим та кінцевим теґами елемента anchor (`a`). Наприклад, `<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` є посиланням з текстом `click here to go to freeCodeCamp.org`.

Додайте текст anchor `cat photos` до елемента anchor. Це буде текстом посилання.

# --hints--

Ваш елемент anchor (`a`) повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Ваш елемент anchor (`a`) повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/a\>/));
```

Текст вашого елемент anchor (`a`) повинен бути `cat photos`. Переконайтеся, що розмістили текст посилання між початковим та кінцевим теґами елемента anchor (`a`).

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Click here to view more cat photos.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

