---
id: 5dfa22d1b521be39a3de7be0
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

У попередньому кроці ви перетворили слова `link to cat pictures` на посилання, розмістивши їх між початковим та кінцевим теґами прив’язки (`a`). Ви можете зробити те саме зі словами всередині елемента, наприклад елемента `p`.

У тексті свого елемента `p` перетворіть слова `cat photos` на посилання до `https://freecatphotoapp.com`, додавши початковий та кінцевий теґи прив'язки (`a`) навколо цих слів.


# --hints--

Ви повинні вкласти новий елемент прив'язки (`a`) в межах елемента `p`.

```js
assert($('p > a').length);
```

Значенням посилання `href` повинне бути `https://freecatphotoapp.com`. Ви або не написали значення `href`, або маєте друкарську помилку.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Текст посилання повинен бути `cat photos`. Ви або не написали текст, або маєте друкарську помилку.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Після вкладення елемента прив'язки (`a`) єдиним видимим вмістом елемента `p` в браузері повинен бути `See more cat photos in our gallery.` Перевірте текст, пробіли та пунктуацію обох елементів (`p` та вкладеного елемента прив'язки).

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

