---
id: 5dfa22d1b521be39a3de7be0
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Тепер в попередньому перегляді ви можете побачити підкреслені слова `cat photos` поруч із зображенням. Це ваше посилання; сміливо клацайте. В тексті вашого елемента `p` перетворіть слова `cat photos` на посилання `https://freecatphotoapp.com`. Коли ви закінчите, видаліть старий anchor тег та текст під абзацом.

# --hints--

Ваш код повинен мати лише один елемент anchor (`a`). Видаліть зайві елементи anchor.

```js
assert(document.querySelectorAll('a').length === 1);
```

Ваш елемент anchor (`a`) повинен бути вкладеним в межах елемента `p`.

```js
assert($('p > a').length);
```

Значення href посилання повинне бути `https://freecatphotoapp.com`. Ви або не написали значення href, або маєте друкарську помилку.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Текст посилання повинен бути `cat photos`. Ви або не написали текст, або маєте друкарську помилку.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Після вкладення елемента anchor (`a`) єдиним видимим вмістом елемента `p` в браузері повинен бути `Click here to view more cat photos.` Двічі перевірте текст, пробіли або пунктуацію обох елементів `p` та вкладеного елемента anchor.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

