---
id: 5dfa2407b521be39a3de7be1
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Додайте атрибут `target` зі значенням `_blank` до початкового теґа елемента прив'язки (`a`) так, щоб посилання відкривалося в новій вкладці.

# --hints--

Ваш елемент `p` повинен мати вкладений елемент прив'язки (`a`) з текстом `cat photos`. Можливо, ви видалили його або допустили помилку.

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Ваш елемент прив'язки (`a`) не має атрибута `target`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
assert(document.querySelector('a').hasAttribute('target'));
```

Значення атрибута `target` повинне бути `_blank`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

