---
id: 5dfa22d1b521be39a3de7be0
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Будь-який текст можна перетворити на посилання. Наприклад, текст всередині елемента `p`.

``` html
<p>I think <a href="https://www.freecodecamp.org">freeCodeCamp</a> is great.</p>
```

Перетворіть слова `cat photos` в тексті елемента `p` на посилання, додавши початковий та кінцевий теги посилання (`a`) навколо цих слів. Потім встановіть атрибут `href` на `https://freecatphotoapp.com`


# --hints--

Ви повинні вкласти новий елемент прив'язки (`a`) в межах елемента `p`.

```js
const nestedAnchor = document.querySelector(`p > a`);
assert.isNotNull(nestedAnchor)
```

Значенням посилання `href` повинне бути `https://freecatphotoapp.com`. Ви або не написали значення `href`, або маєте друкарську помилку.

```js
assert(
  document.querySelector('p > a').href === 'https://freecatphotoapp.com/'
);
```

Текст посилання повинен бути `cat photos`. Ви або не написали текст, або маєте друкарську помилку.

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isTrue(innerContent.trim() === 'cat photos');
```

The text inside your anchor element has extra leading or trailing whitespace. The only space in the anchor text should be between the word `cat` and the word `photos`.

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

After nesting the anchor (`a`) element, the only `p` element content visible in the browser should be `See more cat photos in our gallery.` Double check the text, spacing, or punctuation of both the `p` and nested anchor element.

```js
assert.match(code, /<p>see more <a[^>]*>cat photos<\/a> in our gallery\.?<\/p>/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

