---
id: 5dfb6a35eacea3f48c6300b4
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Елемент підпису до рисунка (`figcaption`) використовують для додавання підпису, щоб описати зображення, яке міститься в елементі `figure`. Наприклад, `<figcaption>A cute cat</figcaption>` додає підпис `A cute cat`.

Після зображення, вкладеного в елемент `figure`, додайте елемент `figcaption` з таким текстом:

`Cats love lasagna.`

# --hints--

Ваш елемент `figcaption` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

Ваш елемент `figcaption` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

Елемент `figcaption` повинен бути вкладеним в елементі `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

Елемент `img` «лазанья» повинен бути вкладеним в елементі `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Елемент `figcaption`, вкладений в елемент `figure`, повинен бути під елементом `img`. Вони знаходяться в неправильному порядку.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

Текст елемента `figcaption` повинен бути `Cats love lasagna.` Ви або не написали текст, або маєте друкарську помилку.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

