---
id: 5dfb6250eacea3f48c6300b2
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Після невпорядкованого списку додайте нове зображення зі значенням атрибута `src`:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Та зі значенням атрибута `alt`:

`A slice of lasagna on a plate.`

# --hints--

Елемент `img` повинен бути відразу після кінцевого теґа `</ul>`.

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

Нове зображення не має атрибута `alt`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

Нове зображення повинне мати `alt` зі значенням `A slice of lasagna on a plate.` Переконайтеся, що значення атрибута `alt` написане в лапках.

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

Нове зображення не має атрибута `src`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

Нове зображення повинне мати `src` зі значенням `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Переконайтеся, що значення атрибута `src` написане в лапках.

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Хоча ви встановили атрибут `src` нового зображення як правильну URL-адресу, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

