---
id: 5ef9b03c81a63668521804d2
title: Крок 27
challengeType: 0
dashedName: step-27
---

# --description--

Код для впорядкованого списку (`ol`) подібний до невпорядкованого списку, але елементи списку в упорядкованому списку нумеруються під час зображення.

Після останнього елемента `h3` другого елементу `section` додайте впорядкований список із цими трьома частинами списку:

`flea treatment` `thunder` `other cats`

# --hints--

Ваш елемент `ol` повинен містити початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

Ваш елемент `ol` повинен містити кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/ol>/));
```

Ваш елемент `ol` повинен бути над кінцевим теґом другого елемента `section`. Вони знаходяться в неправильному порядку.

```js
assert($('main > section')[1].lastElementChild.nodeName === 'OL');
```

Три елементи `li` повинні бути вкладені в елемент `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Ви повинні мати три елементи `li` з текстом `flea treatment`, `thunder` та `other cats` в будь-якому порядку.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Ви повинні мати лише один елемент `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

