---
id: 5ef9b03c81a63668521804e8
title: Крок 63
challengeType: 0
dashedName: step-63
---

# --description--

Вкладіть елемент `p` із текстом `No Copyright - freeCodeCamp.org` в елемент `footer`.

# --hints--

Ви або видалили елемент `footer`, або в ньому відсутній початковий чи кінцевий теґ.

```js
assert(document.querySelector('footer') && code.match(/<\/footer>/));
```

Ваш елемент `footer` повинен мати елемент `p`. Переконайтеся, що додали початковий та кінцевий теґи для елемента `p`.

```js
assert(document.querySelector('footer > p'));
```

Ваш елемент `footer` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert(pElemClosingTags && pElemClosingTags.length === 2);
```

Текст вашого елемента `p` повинен бути `No Copyright - freeCodeCamp.org`. Ви пропустили текст, маєте друкарську помилку, або його немає між початковим та кінцевим теґами елемента `p`.

```js
const extraSpacesRemoved = $('footer > p')[0].innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/No Copyright - freeCodeCamp\.org$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```

