---
id: 5ef9b03c81a63668521804d7
title: Крок 36
challengeType: 0
dashedName: step-36
---

# --description--

Атрибут `action` вказує, куди потрібно надсилати дані форми. Наприклад, `<form action="/submit-url"></form>` повідомляє браузеру, що дані форми треба надіслати на шлях `/submit-url`.

Додайте атрибут `action` зі значенням `https://freecatphotoapp.com/submit-cat-photo` до елемента `form`.

# --hints--

Ваш елемент `form` повинен мати початковий та кінцевий теґи в правильному порядку. Можливо, вам не вистачає одного/обох необхідних теґів або вони розміщені в неправильному порядку.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Ваш елемент `form`, вкладений в останній елемент `section`, повинен бути під елементом `h2`. Елементи `h2` та `form` розташовані в неправильному порядку.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Ваш елемент `form` не повинен мати вмісту. Видаліть будь-які елементи HTML або текст між теґами елемента `form`.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

Ваш елемент `form` не має атрибута `action`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Ваш елемент `form` повинен мати атрибут `action` зі значенням `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Хоча ви встановили атрибут `action` як правильну URL-адресу, краще завжди писати значення атрибута в лапках.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

