---
id: 5ef9b03c81a63668521804e1
title: Крок 50
challengeType: 0
dashedName: step-50
---

# --description--

Елемент `fieldset` використовують для групування пов’язаних ведених даних та міток у вебформі. Елементи `fieldset` є <dfn>елементами на рівні блоків</dfn>, тобто вони з’являються в новому рядку.

Вкладіть радіокнопки `Indoor` та `Outdoor` в елемент `fieldset`, та не забудьте зробити відступи для радіокнопок.

# --hints--

Обидві радіокнопки досі повинні бути між початковим та кінцевим теґами елемента `label`.

```js
const labelChildNodes = [...$('label')].map((node) => [...node.childNodes]);
assert(
  labelChildNodes.filter(
    childNodes =>
      childNodes.filter(node => node.nodeName === 'INPUT').length === 1
  ).length === 2
);
```

Ваш елемент `fieldset` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('fieldset'));
```

Ваш елемент `fieldset` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/fieldset\>/));
```

Обидві радіокнопки та пов’язані мітки повинні бути всередині початкового та кінцевого теґів елемента `fieldset`.

```js
const radioButtons = [...$('input[type="radio"]')];
assert(
  radioButtons.every((btn) => btn.parentNode.parentNode.nodeName === 'FIELDSET')
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

