---
id: 5ef9b03c81a63668521804dc
title: Крок 44
challengeType: 0
dashedName: step-44
---

# --description--

Ви можете використовувати радіокнопки для запитань, де хочете тільки одну відповідь з декількох варіантів.

Ось приклад радіокнопки з опцією `cat`: `<input type="radio"> cat`. Пам'ятайте, що елементи `input` є <dfn>самозакривними</dfn>.

Перед полем тексту додайте радіокнопку з опцією, встановленою як:

`Indoor`

# --hints--

Ви повинні створити елемент `input` для своєї радіокнопки. Перевірте синтаксис.

```js
assert($('form > input').length >= 2);
```

Ваш елемент `input` повинен мати початковий теґ, але не кінцевий теґ.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

Ви повинні додати лише один елемент `input` для своєї радіокнопки. Видаліть зайві.

```js
assert($('form > input').length === 2);
```

Ваш новий елемент `input` повинен бути над наявним `input` з атрибутом `type`, встановленим на `text`. Вони знаходяться в неправильному порядку.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Ваш новий елемент `input` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового теґа.

```js
assert($('input')[0].hasAttribute('type'));
```

Ваш новий елемент `input` повинен мати лише один атрибут `type`. Видаліть зайві.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Ваш новий елемент `input` повинен мати атрибут `type` зі значенням `radio`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Хоча ви встановили атрибут `type` нового елемента `input` як `radio`, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Текст `Indoor` кнопки `radio` повинен бути розташованим після неї, а не перед нею.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

Текст `Indoor` повинен бути розташованим праворуч від кнопки `radio`. Ви або не написали текст, або маєте друкарську помилку.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

