---
id: 5ef9b03c81a63668521804da
title: Крок 42
challengeType: 0
dashedName: step-42
---

# --description--

Використайте елемент `button`, щоб створити активну кнопку. Наприклад, `<button>Click Here</button>` створює кнопку з текстом `Click Here`.

Додайте елемент `button` із текстом `Submit` під елементом `input`. За замовчуванням натискання кнопки форми без жодних атрибутів надсилає форму до місця, вказаного в атрибуті `action`.

# --hints--

Ваш елемент `button` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Ваш елемент `button` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/button\>/));
```

Текст вашого елемента `button` повинен бути `Submit`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

Ваш елемент `button` повинен бути під елементом `input`. Вони знаходяться в неправильному порядку.

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>

--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

