---
id: 5ef9b03c81a63668521804db
title: Крок 41
challengeType: 0
dashedName: step-41
---

# --description--

Щоб користувач не надсилав форму, коли потрібна інформація відсутня, ви повинні додати атрибут `required` до елемента `input`. Немає потреби встановлювати значення для атрибута `required`. Замість цього, просто додайте слово `required` до елемента `input`, переконавшись, що між ним та іншими атрибутами є пробіл.

# --hints--

Ви або видалили свій елемент `input`, або він має недійсний синтаксис. Всі значення атрибутів повинні бути в лапках.

```js
assert($('input').length);
```

Ваш `form` повинен містити тільки елемент `input`. Видаліть будь-які додаткові елементи HTML або текст в елементі `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Ваш елемент `input` повинен мати атрибут `required`. Пам’ятайте, що ви просто додаєте слово `required` всередині теґа елемента `input`.

```js
assert($('input')[0].hasAttribute('required'));
```

Значення не повинне бути надане атрибуту `required`.

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

