---
id: 5ef9b03c81a63668521804e2
title: Крок 54
challengeType: 0
dashedName: step-54
---

# --description--

Зазвичай форми використовують прапорці для запитань, які можуть мати більше за одну відповідь. Наприклад, ось прапорець з опцією `tacos`: `<input type="checkbox"> tacos`.

Під елементом `legend`, який ви щойно додали, додайте `input` з атрибутом `type` зі значенням `checkbox`, та надайте йому опцію:

`Loving`

# --hints--

Ваш елемент `input` повинен мати початковий теґ, але не кінцевий теґ.

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

Ви повинні додати лише один елемент введення для вашого прапорця. Видаліть зайві.

```js
assert($('fieldset > input').length === 1);
```

Ваш новий елемент `input` повинен бути під елементом `legend` із текстом `What's your cat's personality?`. Вони знаходяться в неправильному порядку.

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

Ваш новий елемент `input` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового теґа.

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

Ваш новий елемент `input` повинен мати атрибут `type` зі значенням `checkbox`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

Хоча ви встановили атрибут `type` нового елемента `input` як `checkbox`, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Текст `Loving` повинен бути розташованим праворуч від вашого прапорця. Переконайтеся, що між елементом та текстом є пробіл. Ви або не написали текст, або маєте друкарську помилку.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

