---
id: 5ef9b03c81a63668521804d0
title: Крок 25
challengeType: 0
dashedName: step-25
---

# --description--

Відзначте слово `love` в елементі `figcaption`, обернувши його в елемент виділення `em`.

# --hints--

Ваш елемент виділення `em` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Ваш елемент виділення `em` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/em\>/));
```

Ви або видалили елемент `figcaption`, або в ньому відсутній початковий чи кінцевий теґ.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Ваш елемент виділення `em` повинен оточувати текст `love`. Ви або не написали текст, або маєте друкарську помилку.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Текст `figcaption` повинен бути `Cats love lasagna`. Перевірте наявність друкарських помилок та необхідних пробілів навколо початкового та кінцевого теґів елемента `em`.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

