---
id: 5efb2c990dc218d6c85f89b2
title: Крок 43
challengeType: 0
dashedName: step-43
---

# --description--

Хоча ви додали свою кнопку під полем введення тексту, на сторінці вони зображаються один біля одного. Це тому що обидва елементи `input` та `button` є <dfn>вбудованими елементами</dfn>, які не з’являються в нових рядках.

Кнопка, яку ви додали, надішле форму за замовчуванням. Однак, якщо покладатись на поведінку за замовчуванням, може виникнути плутанина. Додайте атрибут `type` зі значенням `submit` до `button`, щоб було зрозуміло, що це кнопка відправки.

# --hints--

Ваш елемент `button` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Ваш елемент `button` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/button\>/));
```

Ваш елемент `button` не має атрибута `type`. Перевірте, чи є пробіл після назви початкового теґа.

```js
assert($('button')[0].hasAttribute('type'));
```

Ваш елемент `button` повинен мати атрибут `type` зі значенням `submit`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

Хоча ви встановили атрибут `type` елемента `button` як `submit`, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

