---
id: 5efb23e70dc218d6c85f89b1
title: Крок 38
challengeType: 0
dashedName: step-38
---

# --description--

Є багато видів введення, які можна створити з допомогою атрибута `type`. Ви можете легко створити поле пароля, кнопку скидання або елемент керування, щоб користувачі могли вибрати файл зі свого комп’ютера.

Створіть текстове поле для отримання тексту від користувача, додавши атрибут `type` зі значенням `text` до елемента `input`.

# --hints--

Ви або видалили свій елемент `input`, або він має недійсний синтаксис. Якщо ви додали атрибути, переконайтеся, що їхні значення написані в лапках.

```js
assert($('input').length);
```

Ваш елемент `form` повинен містити тільки елемент `input`. Видаліть будь-які зайві елементи HTML або текст між теґами елемента `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Ваш елемент `input` не має атрибута `type` зі значенням `text`. Перевірте, чи є пробіл після назви початкового теґа та/або є пробіли перед усіма назвами атрибутів.

```js
assert($('input')[0].hasAttribute('type'));
```

Ваш елемент `input` повинен мати атрибут `type` зі значенням `text`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

Хоча ви встановили атрибут `type` елемента `input` як `text`, краще завжди писати значення атрибута в лапках.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

