---
id: 5efc518e8d6a74d05e68af75
title: Крок 57
challengeType: 0
dashedName: step-57
---

# --description--

Додайте атрибут `name` зі значенням `personality` до елемента `input` прапорця.

Хоча ви не помітите цього в браузері, це полегшує обробку вашої вебформи для сервера, особливо якщо є декілька прапорців.

# --hints--

Переконайтеся, що прапорець досі присутній.

```js
assert($('input[type="checkbox"]')[0]);
```

Елемент `input` прапорця не має атрибута `name`. Перевірте, чи є пробіл після назви початкового теґа.

```js
assert($('input[type="checkbox"]')[0].hasAttribute('name'));
```

Ваш елемент `input` прапорця повинен мати атрибут `name` зі значенням `personality`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
assert(
  $('input[type="checkbox"]')[0]
    .getAttribute('name')
    .match(/^personality$/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

