---
id: 5efc4f528d6a74d05e68af74
title: Крок 55
challengeType: 0
dashedName: step-55
---

# --description--

Існує інший спосіб пов’язати текст елемента `input` з самим елементом. Ви можете вкласти текст в елемент `label` та додати атрибут `for` з таким самим значенням, що й атрибут `id` елемента `input`.

Пов’яжіть текст `Loving` з прапорцем, вклавши текст `Loving` в елемент `label` та розмістивши його праворуч від елемента `input` прапорця.

# --hints--

Переконайтеся, що прапорець досі присутній.

```js
assert($('input[type="checkbox"]')[0]);
```

Ваш прапорець досі повинен мати атрибут `id` зі значенням `loving`. Можливо, ви видалили атрибут або змінили його значення.

```js
assert($('input[type="checkbox"]')[0].id === 'loving');
```

Текст `Loving` вже не повинен розташовуватися праворуч від вашого прапорця. Його слід загорнути в елемент `label`. Переконайтеся, що між двома елементами є пробіл.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  !checkboxInputElem.nextSibling.nodeValue
    .replace(/\s+/g, ' ')
    .match(/ Loving/i)
);
```

Вам потрібно буде додати новий елемент `label`, в який буде вкладено текст `Loving`. Переконайтеся, що він має початковий та кінцевий теґи.

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

Новий елемент `label` має розташовуватися праворуч від вашого прапорця. Переконайтеся, що між двома елементами є пробіл.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

Новий елемент `label` не має атрибуту `for`. Перевірте, чи є пробіл після назви початкового теґа.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

Новий елемент `label` повинен мати атрибут `for` зі значенням `loving`. Ви або не написали значення, або маєте друкарську помилку. Пам'ятайте, що значення атрибутів повинні бути в лапках.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.getAttribute('for').match(/^loving$/));
```

Текст `Loving` повинен бути вкладеним в новий елемент `label`. Ви або не написали текст, або маєте друкарську помилку.

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

