---
id: 5f05a1d8e233dff4a68508d8
title: Крок 46
challengeType: 0
dashedName: step-46
---

# --description--

Створіть ще одну радіокнопку під першою. Вкладіть її в межі елемента `label` з `Outdoor` як текст `label`. Надайте радіокнопці атрибут `id` зі значенням `outdoor`.

# --hints--

`input` лише початкової радіокнопки `Indoor` повинен мати `id` зі значенням `indoor`. Якщо потрібно, ви можете перезапустити крок, щоб отримати початковий HTML.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Ви не повинні вносити жодних змін до радіокнопки `Indoor`. Якщо потрібно, ви можете перезапустити крок, щоб отримати початковий HTML.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Ви повинні додати один новий елемент `input`, вкладений в новий елемент `label`. Переконайтеся, що ваш новий `label` має початковий та кінцевий теґи.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Ви не повинні додавати жодних атрибутів до будь-якого початкового теґу `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

Ваш новий `input` повинен мати атрибут `id` зі значенням `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

Ваш новий `input` повинен мати атрибут `type` зі значенням `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Ви не повинні додавати жодних нових елементів, крім `input`, вкладеного в `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Зліва від нового `input` не повинно бути тексту.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Текст мітки нової радіокнопки повинен бути `Outdoor`. Сюди входять великі літери.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Ваша нова радіокнопка та мітка мають бути безпосередньо під/після радіокнопки та мітки `Indoor`. Між ними не повинно бути інших теґів.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

