---
id: 5f07be6ef7412fbad0c5626b
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Перш ніж додавати будь-який новий вміст, вам слід використати елемент `section`, щоб відокремити вміст фотографій котів від майбутнього вмісту.

Візьміть елементи `h2`, коментар, `p` та елемент прив'язки (`a`) і вкладіть їх в елемент `section`.

# --hints--

Ваш елемент `section` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Ваш елемент `section` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

Весь елемент `section` повинен бути між початковим та кінцевим теґами елемента `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Наявні `h2`, коментар, елемент `p` та елемент прив'язки (`a`) повинні бути між початковим та кінцевим теґами елемента `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

Елемент `h1` не повинен бути вкладеним в елементі `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

