---
id: 5f07be6ef7412fbad0c5626b
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Перш ніж додавати будь-який новий вміст, вам слід використати елемент `section`, щоб відокремити вміст фотографій котів від майбутнього вмісту.

Візьміть усі елементи, які зараз знаходяться в елементі `main`, та вкладіть їх в елемент `section`.

# --hints--

Ваш елемент `section` повинен мати початковий теґ. Початкові теґи мають такий синтаксис: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Ваш елемент `section` повинен мати кінцевий теґ. Кінцеві теґи мають `/` відразу після символу `<`.

```js
assert(code.match(/<\/section\>/));
```

Весь елемент `section` повинен бути між початковим та кінцевим теґами елемента `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Наявні `h2`, коментар, елемент `p` та елемент anchor (`a`) повинні бути між початковим та кінцевим теґами елемента `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElems = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElems.length === 3);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
    </main>
--fcc-editable-region--
  </body>
</html>
```

