---
id: 62dabe2ef403a12d5d295273
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Ви перетворили текст `cat photos` всередині елемента `p` на посилання, тому вам не потрібне друге посилання під елементом `p`. Видаліть весь елемент прив’язки під елементом `p`.

# --hints--

Your code should only have one paragraph (`p`) element.

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

Ваш код повинен мати лише один елемент прив'язки (`a`).

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

Do not make any changes to the elements you created in the previous step.

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

Під елементом `p` не повинно бути тексту `cat photos`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
