---
id: 62dabe2ef403a12d5d295273
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Ви перетворили текст `cat photos` всередині елемента `p` на посилання, тому вам не потрібне друге посилання під елементом `p`. Видаліть весь елемент прив’язки під елементом `p`.

# --hints--

Ваш код повинен мати лише один елемент прив'язки (`a`).

```js
assert(document.querySelectorAll('a').length === 1);
```

Елемент `p` досі повинен мати елемент прив'язки (`a`) всередині себе.

```js
assert($('p > a').length);
```

Значення посилання `href` повинне бути `https://freecatphotoapp.com`. Ви або не написали значення `href`, або маєте друкарську помилку.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Текстом посилання повинен бути `cat photos`. Ви або не написали текст, або маєте друкарську помилку.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Видимим вмістом елемента `p` повинен бути `Click here to view more cat photos.` Двічі перевірте текст, пробіли або пунктуацію обох елементів `p` та вкладеного елемента прив'язки.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

Під елементом `p` не повинно бути тексту `cat photos`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more <a href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
