---
id: 635090f47eb6d9563a6fed05
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Il prossimo passo è trovare il numero al centro della lista. Se la lista ha un numero dispari di numeri, il numero centrale è la mediana. Se la lista ha un numero pari di numeri, la mediana è la media dei due numeri centrali.

Puoi controllare se un numero è pari o dispari con l'operatore <dfn>resto</dfn>, che è rappresentato dal simbolo `%`. Questo operatore restituisce il resto della divisione di due numeri. Se il resto è `0`, il numero è pari. Se il resto è `1`, il numero è dispari:

```js
array.length % 2 === 0;
```

Dichiara una variabile `median`. Usando l'operatore ternario, controlla se la lunghezza di `array` è pari. Se la lunghezza di `array` è pari, trova i due numeri centrali e calcola la loro media. Se la lunghezza di `array` è dispari, trova il numero centrale e assegnalo alla variabile `median`.

# --hints--

La funzione `getMedian` dovrebbe avere una variabile `median`.

```js
assert.match(getMedian.toString(), /median\s*=/);
```

La variabile `median` dovrebbe usare un operatore ternario per verificare se `array.length` è pari.

```js
assert.match(getMedian.toString(), /median\s*=\s*array\.length\s*%\s*2\s*===\s*0\s*\?/);
```

Se `array.length` è pari, la variabile `median` dovrebbe utilizzare la funzione `getMean` per calcolare la media dei due numeri centrali. Il primo argomento dovrebbe essere il valore di `sorted` ad `array.length / 2` e il secondo ad `array.length / 2 - 1`.

```js
assert.match(getMedian.toString(), /median\s*=\s*array\.length\s*%\s*2\s*===\s*0\s*\?\s*getMean\(\s*\[sorted\[array\.length\s*\/\s*2\]\s*,\s*sorted\[\s*array\.length\s*\/\s*2\s*-\s*1\]\]\)\s*\:/);
```

Se `array.length` è dispari, la variabile `median` dovrebbe assegnare il numero centrale alla variabile `median`. Il numero centrale è il valore di `sorted` a `Math.floor(array.length / 2)`.

```js
assert.match(getMedian.toString(), /median\s*=\s*array\.length\s*%\s*2\s*===\s*0\s*\?\s*getMean\(\s*\[sorted\[array\.length\s*\/\s*2\]\s*,\s*sorted\[\s*array\.length\s*\/\s*2\s*-\s*1\]\]\)\s*\:\s*sorted\[\s*Math\.floor\(\s*array\.length\s*\/\s*2\)\];/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort((a, b) => a - b);
}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));

  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
