---
id: 635060a5c03c950f46174cb5
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

La statistica è un modo di usare la matematica per dare senso ai dati. Ci aiuta a capire i modelli e le tendenze dei dati, in modo da poter fare previsioni e prendere decisioni sulla base di tali informazioni.

In questa sfida, costruirai una calcolatrice statistica che prende una serie di numeri e restituisce media, mediana, moda, deviazione standard e varianza.

L'HTML e il CSS ti sono stati forniti. Sentiti libero di esplorare il codice – potresti notare che la funzione `calculate` viene chiamata quando il modulo viene inviato. Quando sei pronto, dichiara una variabile `calculate` e assegnale una funzione vuota nel file `script.js`.

# --hints--

Dovresti usare la parola chiave `const` per dichiarare la variabile `calculate`.

```js
assert.match(code, /const\s*calculate/);
```

La variabile `calculate` dovrebbe essere una funzione.

```js
assert.isFunction(calculate);
```

La funzione `calculate` dovrebbe essere vuota.

```js
assert.match(calculate.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
