---
id: 646c48df8674cf2b91020ecc
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Add a `link` element within your `head` element. For that `link` element, set the `rel` attribute to `stylesheet` and the `href` to `./styles.css`.

# --hints--

Your should have a `link` element.

```js
assert.match(code, /<link/)
```

Your `link` element should have a `rel` attribute set to `stylesheet`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1/)
```

Your `link` element should have an `href` attribute with the value `./styles.css`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1\s+href=\1\.\/styles\.css\1\s*\/?>/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
