---
id: 63b60aaaa65f8922bfce6b7e
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Erstelle in deinem `span`-Element ein `label`-Element für einen `entry-dropdown` und weise ihm den Text `Add food or exercise:` zu. Then create a `select` element with the `id` set to `entry-dropdown` and a `name` set to `options`. Füge unter dieses ein `button`-Element mit der `id` auf `add-entry` gesetzt und dem Text `Add Entry` hinzu.

Gib deinem `button`-Element ein `type`-Attribut, das auf `button` eingestellt ist, um das automatische Absenden des Formulars zu unterbinden.

# --hints--

Du solltest ein `label`-Element zu deinem `span`-Element hinzufügen.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Dein neues `label`-Element sollte ein `for`-Attribut enthalten, das auf `entry-dropdown` gesetzt ist.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Dein neues `label`-Element sollte den Text `Add food or exercise:` enthalten.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Du solltest deinem `span`-Element ein `select`-Element hinzufügen.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Dein `select`-Element sollte nach deinem `label`-Element stehen.

```js
assert(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName === 'LABEL');
```

Dein neues `select`-Element sollte ein `id`-Attribut auf `entry-dropdown` gesetzt haben.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Dein neues `select`-Element sollte ein `name`-Attribut auf `options` gesetzt haben.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Your `select` element should be empty and have a closing `</select>` tag.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

Du solltest ein `button`-Element zu deinem `span`-Element hinzufügen.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Dein `button`-Element sollte nach deinem `select`-Element stehen.

```js
assert(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName === 'SELECT');
```

Dein neues `button`-Element sollte ein `id`-Attribut auf `add-entry` gesetzt haben.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Dein neues `button`-Element sollte den Text `Add Entry` haben.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
