---
id: 63b607af6fcdb119aae9b16a
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

In your form, users should have the capability to add various meal types along with their calorie counts.

Erstelle ein `fieldset`-Element mit der `id` auf `breakfast` gesetzt.

Erstelle innerhalb dieses Elements eine `legend` mit dem Text `Breakfast` und einen leeren `div` mit der `class` auf `input-container` gesetzt.

# --hints--

Du solltest ein `fieldset`-Element in deiner `form` erstellen.

```js
assert.exists(document.querySelector('form fieldset'));
```

Dein neues `fieldset`-Element sollte nach deinem `input`-Element stehen.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

Dein `fieldset`-Element sollte eine `id` auf `breakfast` gesetzt haben.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Du solltest ein `legend`-Element innerhalb deines `fieldset` erstellen.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

Dein `legend`-Element sollte den Text `Breakfast` haben.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Du solltest ein `div`-Element innerhalb deines `fieldset` erstellen.

```js
assert.exists(document.querySelector('form fieldset div'));
```

Dein `div`-Element sollte nach deinem `legend`-Element stehen.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

Dein `div`-Element sollte eine `class` auf `input-container` gesetzt haben.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
