---
id: 67ed03ac474c48692f41749e
title: Build a Hash Table
challengeType: 23
dashedName: build-a-hash-table
---

# --description--

In this lab, you will build a hash table from scratch. A hash table is a data structure that stores key-value pairs. A hash table works by taking the key as an input and then hashing this key according to a specific hashing function. 

For the purpose of this lab, the hashing function will be simple: it will sum the Unicode values of each character in the key. The hash value will then be used as the actual key to store the associated value. The same hash value would also be used to retrieve, update, or delete the value associated with the key.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should define a class named `HashTable` with a `collection` property initialized as an empty dictionary inside its `__init__` method. The `collection` dictionary should store key-value pairs based on the hashed value of the key.

2. The `HashTable` class should have four instance methods: `hash`, `add`, `remove`, and `lookup`.

3. The `hash` method should:

    - Take `string` as a parameter.
    - Return a hashed value computed as the sum of the Unicode (ASCII) values of each character in the string. You can use the `ord` function for this computation.

4. The `add` method should:

    - Take a key-value pair -`key`, `val` as parameters and compute the hash of the key.
    - Use the computed hashed value to store the key-value pair inside the `collection` dictionary. 
    - If multiple keys produce the same hash value, their key-value pairs should be stored in a nested dictionary under the same hash value.

5. The `remove` method should:

    - Take `key` as a parameter and compute its hash.
    - Remove the corresponding key-value pair from the hash table.

6. The `lookup` method should:

    - Take `key` as a parameter.
    - Compute the hash of the key, and return the corresponding value stored inside the hash table. 
    - If the key does not exist in the collection, it should return `None`.

# --hints--

You should define the `HashTable` class.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

The `HashTable` should initialize a `collection` property as an empty dictionary within the `__init__` method.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("HashTable").find_function("__init__").has_stmt("self.collection={}")
    `))
})
```

The `HashTable` class should have a `hash` method.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

The `hash` method should take `string` as a parameter.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("HashTable").find_function("hash").has_args("self,string")`)) })
```

The `hash` method should return the sum of the Unicode values of each character in the string.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

The `HashTable` class should have an `add` method.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

The `add` method should take `key` and `value` as parameters.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("HashTable").find_function("add").has_args("self,key,val")`)) })
```

The `HashTable` class should have a `remove` method.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

The `remove` method should take `key` as a parameter.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("HashTable").find_function("remove").has_args("self,key")`)) })
```

The `HashTable` class should have a `lookup` method.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

The `lookup` method should take `key` as the parameter.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("HashTable").find_function("lookup").has_args("self,key")`)) })
```

`HashTable.hash("golf")` should return `424`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable.add("golf", "sport")` should add the key-value pair to the collection at index `424`.

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable.add("dear", "friend")` and `HashTable.add("read", "book")` should add both the key-value pairs to the collection at index `412` as a nested dictionary.

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

When the `"golf", "sport"` key-value pair exists in the hash table, `HashTable.remove("golf")` should remove the key-value pair from the collection.

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

When the `"golf", "sport"` key-value pair exists in the hash table, `HashTable.lookup("golf")` should return `"sport"`.

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

When the `"golf", "sport"` key-value pair does not exist in the collection, `HashTable.lookup("golf")` should return `None`.

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.remove("golf")

        assert ht.lookup("golf") is None
`) })

```


# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
