---
id: 66f1805aac05f00a7028fbcc
title: Build a Hash Table Class
challengeType: 14
dashedName: build-a-hash-table-class
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should define a class named `HashTable` with a `collection` property initialized as an empty object inside its `constructor` method. The `collection` object should store key-value pairs based on the hashed value of the key.
1. The `HashTable` class should have four instance methods, namely `hash`, `add`, `remove`, and `lookup`.
1. The `hash` method should take a string as a parameter and return the hashed value computed as the sum of the UTF-16 code unit of each character in the string. Use the `charCodeAt` method for that.
1. The `add` method should take a key-value pair as the parameters, and hash the key. The hashed value should be used to store the key-value pair inside the hash table (`collection` object). When the same hash value is produced by multiple different keys, each key-value pair should be stored using the same hash value.
1. The `remove` method should take a key as a parameter and remove the corresponding key-value pair from the hash table.
1. The `lookup` method should take a key as a parameter and return the corresponding value stored inside the hash table, or `null` if not present.

# --hints--

You should have a `HashTable` class.

```js
assert.isFunction(HashTable);
assert.isObject(new HashTable());
```

Your `HashTable` class should have a `collection` property initialized as an empty object.

```js
const table = new HashTable();
assert.isObject(table.collection);
assert.isEmpty(table.collection);
```

Your `HashTable` class should have a `hash` method.

```js
const table = new HashTable();
assert.isFunction(table.hash);
```

`new HashTable().hash("freeCodeCamp")` should return `1182`.

```js
const table = new HashTable();
assert.strictEqual(table.hash("freeCodeCamp"), 1182);
```

The `hash` method should take a string as a parameter and return the hashed value computed as the sum of the UTF-16 code unit of each character in the string using the `charCodeAt` method.

```js
const table = new HashTable();
assert.strictEqual(table.hash("Hey you"), 675);
```

Your `HashTable` class should have an `add` method.

```js
const table = new HashTable();
assert.isFunction(table.add);
```

After calling `hashTable.add("FCC", "freeCodeCamp")`, `hashTable.collection` should become `{ 204: { FCC: 'freeCodeCamp' } }`.

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
assert.deepEqual(table.collection, { 204: { FCC: 'freeCodeCamp' } });
```

After calling `hashTable.add("FCC", "freeCodeCamp")` and `hashTable.add("CFC", "chloroflourocarbon")`, `hashTable.collection` should become `{ 204: { FCC: 'freeCodeCamp', CFC: 'chloroflourocarbon' } }`.

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.add("CFC", "chloroflourocarbon");
assert.deepEqual(table.collection, { 204: { FCC: 'freeCodeCamp', CFC: 'chloroflourocarbon' } });
```

Items should be added using the hash function.

```js

```

The `add` method should take a key-value pair as the parameters, hash the key and store the hashed key-value pair inside `collection` object.

```js
const table = new HashTable();
table.add("", "");
assert.deepEqual(table.collection, { : { : '' } });
```

When the same hash value is produced by multiple different keys, the `add` method should store each key-value pair using the same hash value.

```js

```

Your `HashTable` class should have a `remove` method.

```js
const table = new HashTable();
assert.isFunction(table.remove);
```

`hashTable.remove("fCC")` should remove the `fCC` key from the `collection` object.

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.remove("FCC");
assert.isEmpty(table.collection);
```

The `remove` method should accept a key as input and should remove the associated key value pair.

```js
const table = new HashTable();
table.add("FCC", "freeCodeCamp");
table.remove("FCC");
assert.isEmpty(table.collection);
```

The `remove` method should only remove the correct key value pair.

```js
```

Your `HashTable` class should have a `lookup` method.

```js
const table = new HashTable();
assert.isFunction(table.lookup);
```

`hashTable.lookup("fCC")` should return the value of the `fCC` key stored within the `collection` object.

```js

```

`hashTable.lookup("fCC")` should return `null` if there's no `fCC` key stored within the `collection` object.

```js

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class HashTable {
    constructor() {
        this.collection = {};
    }

    hash(string) {
        let hashed = 0;
        for (let i = 0; i < string.length; i++) {
            hashed += string.charCodeAt(i);
        }
        return hashed
    }

    add(key, val) {
        const theHash = this.hash(key);
        if (!this.collection.hasOwnProperty(theHash)) {
            this.collection[theHash] = {};
        }
        this.collection[theHash][key] = val;
    }

    remove(key) {
        const theHash = this.hash(key);
        const hashedObj = this.collection[theHash];
        if (hashedObj.hasOwnProperty(key)) {
            delete hashedObj[key];
        }
        if (!Object.keys(hashedObj).length) {
            delete this.collection[theHash];
        }
    }

    lookup(key) {
        const theHash = this.hash(key);
        if (this.collection.hasOwnProperty(theHash)) {
            return this.collection[theHash][key];
        }
        return null
    }
}
```
