---
id: 67172b43f84bcd2dec238a3d
title: Build a Book Organizer
challengeType: 14
dashedName: build-a-book-organizer
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have an array of objects named `books` where each object in the array should have a string `title`, another string `authorName`, and a number `releaseYear`. 

2. Your `books` array should have a minimum of three objects.

3. You should have a callback function named `sortByYear` that accepts two books, `book1`, and `book2` as parameter for sorting the array.

4. The `sortByYear` function should return `-1` if the `releaseYear` of `book1` object is smaller than that of the `book2` object.

5. The `sortByYear` function should return `1` if the `releaseYear` of `book1` object is bigger than that of the `book2` object.

6. The `sortByYear` function should return `0` for all other scenarios.

7. You should filter out books written after a certain year such as 1950 from the `books` array and save the filtered array in a new array named `filteredBooks`.

8. You should sort the books in the `books` array according to their `releaseYear` in ascending order and save the sorted array in a new array named `filteredBooksSorted`.

9. You should iterate over the `filteredBooksSorted` array and print out the details of each book following the `<BOOK_TITLE> by <BOOK_AUTHOR> (<releaseYear>)` format.

# --hints--

You should have an array `books` in your code.

```js
assert.isArray(books)
```

Your `books` array should have at least three book objects.

```js
assert.isAtLeast(books.length, 3)
```

Your `books` array should contain objects each with a string `title`, another string `authorName`, and a number `releaseYear`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

You should have an array `filteredBooks` in your code.

```js
assert.isArray(filteredBooks)
```

You should have an array `filteredBooksSorted` in your code.

```js
assert.isArray(filteredBooksSorted)
```

You should call the `sort` higher order function by passing the `sortByYear` callback function on the `filteredBooks` array.

```js
assert.isTrue(filteredBooksSorted.every((book, index, array) => index === 0 || book.releaseYear >= array[index - 1].releaseYear))
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

let filteredBooksSorted = filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooksSorted.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
