---
id: 67172b43f84bcd2dec238a3d
title: Build a Book Organizer
challengeType: 0
dashedName: build-a-book-organizer
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a callback function named `sortByYear` that accepts two books, `book1`, and `book2` as parameter for sorting the array.

2. The `sortByYear` function should return `-1` if the `release_year` of `book1` object is smaller than that of the `book2` object.

3. The `sortByYear` function should return `1` if the `release_year` of `book1` object is bigger than that of the `book2` object.

4. The `sortByYear` function should return `0` for all other scenarios.

5. You should have an array of objects named `books` where each object in the array should have a string `title`, another string `author_name`, and a number `release_year`. 

6. Your `books` array should have a minimum of three objects.

7. You should call the `filter` higher-order function on the `books` array with a callback to filter out books written after a certain year such as 1950 and save the filtered array in a new array named `filteredBooks`.

8. You should call the `sort` higher order function by passing the `sortByYear` callback function on the `filteredBooks` array to sort the books according to their `release_year` in ascending order and save the sorted array in a new array named `filteredBooksSorted`.

9. You should call the `forEach` higher order function on the `filteredBooksSorted` array and print out the details of each book following the `<BOOK_TITLE> by <BOOK_AUTHOR> (<RELEASE_YEAR>)` format.

# --hints--

You should have a function `sortByYear` in your code.

```js
assert.isFunction(sortByYear);
```

Your `sortByYear` function should return `-1` if the `release_year` of `book1` object is smaller than that of the `book2` object, `1` if the `release_year` of `book1` object is larger than that of the `book2` object, and `0` in all other scenarios.

```js
assert.equal(sortByValue({release_year: 1913}, {release_year: 1925}), -1);
assert.equal(sortByValue({release_year: 1925}, {release_year: 1913}), 1);
assert.equal(sortByValue({release_year: 1925}, {release_year: 1925}), 0);
```

You should have an array `books` in your code.

```js
assert.isArray(books)
```

Your `books` array should have at least three book objects.

```js
assert.lenghtOf(books, 3)
```

Your `books` array should contain objects each with a string `title`, another string `author_name`, and a number `release_year`.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "author_name", "release_year"]);
  assert.isString(book.title);
  assert.isString(book.author_name);
  assert.isNumber(book.release_year);
})
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.release_year < book2.release_year) {
    return -1;
  }

  if (book1.release_year > book2.release_year) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    author_name: "F. Scott Fitzgerald",
    release_year: 1925,
  },
  {
    title: "In Search of Lost Time",
    author_name: "Marcel Proust",
    release_year: 1913,
  },
  {
    title: "Ulysses",
    author_name: "James Joyce",
    release_year: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    author_name: "Gabriel García Márquez",
    release_year: 1967,
  },
  {
    title: "The Catcher in the Rye",
    author_name: "J. D. Salinger",
    release_year: 1951,
  },
];

let filteredBooks = books.filter((book) => book.release_year < 1950);

let filteredBooksSorted = filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooksSorted.forEach((book) => {
  console.log(`${book.title} by ${book.author_name} (${book.release_year})`);
});
```
