---
id: 6564883669b5af4b69f794cc
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Your `all_characters` variable is a string formed by all lowercase and uppercase letters, all the 10 digits and several special characters. 

Just before it, add a comment saying `Combine all characters`.

# --hints--

You should add the comment just above your `all_characters` variable.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
