---
id: 656477845006313fbfea0ad1
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Although the effect might seem equal to `random.choice()`, `secrets` ensure you the most secure source of randomness that your operating system can provide.

Now, delete your two `print()` calls.

# --hints--

You should delete your two `print()` calls.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
