---
id: 67ebd3fe712a39661dc28a3a
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

To account for the setter function, you will have to attach a `value` attribute and an `onChange` event to the `select` and a `value` attribute to each of the `option` elements.

That's a long process, right? So the best way to handle this is to create an array for the values of the options, map through that array, and use each element of the array as the values of the options.

To do that, a `powerSourceOptions` array has been provided for you, so, take a look at it.

For now, remove all the options inside the `select` element apart from the first one.


# --hints--

You should create a `powerSourceOptions` array.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*/)
```

Your `powerSourceOptions` array should have the string `"Bitten by a strange creature"` as its first element.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*('|")Bitten\s+by\s+a\s+strange\s+creature("|')/)
```

Your `powerSourceOptions` array should have the string `"Radioactive exposure"` as its second element.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*('|")Bitten\s+by\s+a\s+strange\s+creature("|'),\s*('|")Radioactive\s+exposure('|")/)
```

Your `powerSourceOptions` array should have the string `"Science experiment"` as its third element.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*('|")Bitten\s+by\s+a\s+strange\s+creature("|'),\s*('|")Radioactive\s+exposure('|"),\s*('|")Science\s+experiment("|')/)
```

Your `powerSourceOptions` array should have the string `"Alien heritage"` as its fourth element.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*('|")Bitten\s+by\s+a\s+strange\s+creature("|'),\s*('|")Radioactive\s+exposure('|"),\s*('|")Science\s+experiment("|'),\s*('|")Alien\s+heritage('|")/)
```

Your `powerSourceOptions` array should have the string `"Ancient artifact discovery"` as its fifth element.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*('|")Bitten\s+by\s+a\s+strange\s+creature("|'),\s*('|")Radioactive\s+exposure('|"),\s*('|")Science\s+experiment("|'),\s*('|")Alien\s+heritage('|"),\s*('|")Ancient\s+artifact\s+discovery("|')/)
```

Your `powerSourceOptions` array should have the string `"Other"` as its sixth element.

```js
assert.match(code, /(const|let)\s+powerSourceOptions\s*=\s*\[\s*('|")Bitten\s+by\s+a\s+strange\s+creature("|'),\s*('|")Radioactive\s+exposure('|"),\s*('|")Science\s+experiment("|'),\s*('|")Alien\s+heritage('|"),\s*('|")Ancient\s+artifact\s+discovery("|'),\s*('|")Other("|')\s*\];?/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.0/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.0/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const powerSourceOptions = [
    'Bitten by a strange creature',
    'Radioactive exposure',
    'Science experiment',
    'Alien heritage',
    'Ancient artifact discovery',
    'Other'
  ];

  const [heroName, setHeroName] = React.useState('');
  const [realName, setRealName] = React.useState('');
  const [powerSource, setPowerSource] = React.useState('');
  const [powers, setPowers] = React.useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
        <label>
          Hero Name
          <input
            type='text'
            value={heroName}
            onChange={e => setHeroName(e.target.value)}
          />
        </label>
        <label>
          Real Name
          <input
            type='password'
            value={realName}
            onChange={e => setRealName(e.target.value)}
          />
        </label>
        </div>
        <label className='section column'>
          How did you get your powers?
          <select>
            <option value=''>
              Select one...
            </option>
            --fcc-editable-region--
            <option>
              Bitten by a strange creature
            </option>
            <option>
              Radioactive exposure
            </option>
            <option>
              Science experiment
            </option>
            <option>
              Alien heritage
            </option>
            <option>
              Ancient artifact discovery
            </option>
            <option>
              Other
            </option>
            --fcc-editable-region--
          </select>
        </label>
      </form>
    </div>
  )
};
```
