---
id: 67ebc0e1c183cb2d743e513b
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Time to create the labels and inputs. Create a `label` with the text `Hero Name` and an `input` element of type `text` inside it.

Connect the `Hero Name` `input` to the `heroName` state variable by giving it a `value` of attribute set to `heroName`.

# --hints--

You should create a `label` element.

```js
assert.exists(document.querySelector("label"));
```

Your `label` element should have `Hero Name` as its text.

```js
const labelEl = document.querySelector("label")
assert.equal(labelEl.textContent, "Hero Name");
```

You should create an `input` of type `text` element inside your `label` element.

```js
const inputEl = document.querySelector("label > input");
assert.equal(inputEl.getAttribute("type"), "text");
```

Your `input` element should have its value set to the `heroName` state variable.

```js
assert.match(code, /value=\{heroName\}/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.0/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.0/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = React.useState('');
  const [realName, setRealName] = React.useState('');
  const [powerSource, setPowerSource] = React.useState('');
  const [powers, setPowers] = React.useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
        --fcc-editable-region--

        --fcc-editable-region--
        </div>
      </form>
    </div>
  )
};
```
