---
id: 6a1c46843b36624be487988d
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now you need to handle potential problems. What if the user clicks the "Search" button without typing anything into the input field?

Right now, the function takes that empty string and returns `" is available to borrow!"`. This is not very helpful.

Before your success `return` statement, add an `if` statement. Check if the `title` parameter is strictly equal to an empty string (`""`).

For now, leave the body of the `if` statement empty after adding it.

# --hints--

You should create an `if` statement.

```js
assert.match(code, /if\s*\(/);
```

Your `if` statement should check if `title` is strictly equal to an empty string.

```js
assert.match(code, /if\s*\(\s*title\s*===\s*(["'])\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css"/>
</head>

<body>
    <div>
        <h1>Find book title</h1>
        <input type="text" id="bookTitleInput" placeholder="Enter book title">
        <button id="search">Search</button>
        <p id="statusMessage"></p>
    </div>
    <script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
--fcc-editable-region--

--fcc-editable-region--
  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  const result = findBook(bookInput.value.trim());
  statusDisplay.textContent = result;
});
```
