---
id: 6a1c49deb109227f4020e0e9
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Cool! But, as a developer, you also need to see the full error object to debug effectively. 

Inside your `catch` block, right below where you update the `textContent`, add a `console.error()` statement and pass it your `err` object.

# --hints--

You should call the `console.error()` method inside your `catch` block.

```js
assert.match(code, /console\.error\s*\(/);
```

You should pass the `err` object to your `console.error()` call.

```js
assert.match(code, /console\.error\s*\(\s*err\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8" />
        <meta name="viewport" content="width=device-width, initial-scale=1.0" />
        <title>Library Book Search</title>
        <link rel="stylesheet" href="styles.css" />
    </head>
    <body>
        <div>
            <h1>Find book title</h1>
            <input type="text" id="bookTitleInput" placeholder="Enter book title">
            <button id="search">Search</button>
            <p id="statusMessage"></p>
            <p id="scanStatus"></p>
        </div>
        <script src="script.js"></script>
    </body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  try {
    const result = findBook(bookInput.value.trim());
    statusDisplay.textContent = result;
  } catch (err) {
    statusDisplay.textContent = err.message;
--fcc-editable-region--

--fcc-editable-region--
  }
});
```
