---
id: 6a1c4a2773f60894ce905694
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

For advanced debugging, you can use the `debugger` statement to pause your JavaScript execution and inspect your program. 

Inside your `catch` block, right below your `console.error()` statement, add a `debugger;` statement.

*Note: If you have your browser's Developer Tools open when you run the tests, your browser will actually pause on this line!*

# --hints--

You should add a `debugger;` statement inside your `catch` block.

```js
assert.match(code.replace(/\s+/g, ''), /console\.error\(err\);debugger;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8" />
        <meta name="viewport" content="width=device-width, initial-scale=1.0" />
        <title>Library Book Search</title>
        <link rel="stylesheet" href="styles.css" />
    </head>
    <body>
        <div>
            <h1>Find book title</h1>
            <input type="text" id="bookTitleInput" placeholder="Enter book title">
            <button id="search">Search</button>
            <p id="statusMessage"></p>
            <p id="scanStatus"></p>
        </div>
        <script src="script.js"></script>
    </body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  try {
    const result = findBook(bookInput.value.trim());
    statusDisplay.textContent = result;
  } catch (err) {
    statusDisplay.textContent = err.message;
    console.error(err);
--fcc-editable-region--

--fcc-editable-region--
  }
});
```
