---
id: 6a1c48f85f99dbe8f8a41e9e
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Inside your second `if` statement, you need to throw another error to stop the execution.

Use the `throw` keyword to throw a `new Error()`. Pass the string `title + " not found!"` as an argument to the Error constructor.

# --hints--

Inside your second 'if' statement, you should use the 'throw' keyword to throw a 'new Error()' with the exact string: title + " not found!".

```js
assert.throws(() => findBook("Nonexistent Book"), Error, "Nonexistent Book not found!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Find book title</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter book title">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
}

input {
  padding: 0.5rem;
  width: 200px;
}

button {
  padding: 0.5rem 1rem;
  cursor: pointer;
  background-color: #000000;
  color: white;
}

#statusMessage {
  margin-top: 1.5rem;
  font-weight: bold;
  height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
  --fcc-editable-region--
    
  --fcc-editable-region--
  }
  
  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  const result = findBook(bookInput.value.trim());
  statusDisplay.textContent = result;
});
```
