---
id: 6a1c49231334592cc5a4c6b0
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Great! However, currently, if your `find_my_book` function throws an error, it will completely crash the script because nothing is there to catch it. 

To handle these errors gracefully, you need to wrap the code that might fail inside a `try` block. 

Look at the `searchBtn` event listener at the bottom of the file. Wrap all the existing code inside that callback function in a `try` block.

*Note: Ignore the red error line below the curly braces, you'll fix that in the next step.*
# --hints--

You should add a `try` block inside your event listener callback.

```js
assert.match(code, /try\s*\{/);
```

Your `try` block should wrap the existing `const result` declaration and the `if(result)` statement.

```js
assert.match(code.replace(/\s+/g, ''), /try\{constresult=find_my_book\(bookInput\.value\);if\(result\)\{statusDisplay\.textContent=result;?\}\}\}\);/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css"/>
</head>

<body>
    <div>
        <h1>Find book title</h1>
        <input type="text" id="bookTitleInput" placeholder="Enter book title">
        <button id="search">Search</button>
        <p id="statusMessage"></p>
    </div>
    <script src="script.js"></script>
</body>
</html>
```

```css
body{
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input{
    padding: 0.5rem;
    width: 200px;
}

button{
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage{
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const find_my_book = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
--fcc-editable-region--
  const result = find_my_book(bookInput.value);
  if (result) {
    statusDisplay.textContent = result;
  }
--fcc-editable-region--
});
```
