---
id: 6a1c48b465a40113ecf3766f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Your function will now stop running and throw an error if the user tries to search with an empty field.

But what if the user searches for a book that is not in your `library` array? You need to handle that exception, too.

Below your first `if` statement, create a second `if` statement. Inside the condition, use the logical NOT operator (`!`) and the `.includes()` method to check if the `library` array does not include the `title` passed into the function.

# --hints--

You should create an `if` statement that checks if `library` does not include `title`.

```js
assert.match(code, /if\s*\(\s*!library\.includes\(\s*title\s*\)\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8" />
        <meta name="viewport" content="width=device-width, initial-scale=1.0" />
        <title>Library Book Search</title>
        <link rel="stylesheet" href="styles.css" />
    </head>
    <body>
        <div>
            <h1>Find book title</h1>
            <input type="text" id="bookTitleInput" placeholder="Enter book title">
            <button id="search">Search</button>
            <p id="statusMessage"></p>
            <p id="scanStatus"></p>
        </div>
        <script src="script.js"></script>
    </body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  const result = findBook(bookInput.value.trim());
  statusDisplay.textContent = result;
});
```
