---
id: 6a1c4995ca0a594713d129e1
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

To capture and handle the errors from your `try` block, you need to add a `catch` block. You can pass a parameter to this block to inspect the actual error object that was thrown.

Directly below your `try` block, add a `catch` block and pass it the parameter `err`. Inside the block, update the DOM by setting the text content of `statusDisplay` to `err.message`.

# --hints--

You should add a `catch` block that takes `err` as a parameter. Inside the block, you should set `statusDisplay.textContent` to `err.message`.

```js
document.getElementById('bookTitleInput').value = "Nonexistent Book";
document.getElementById('search').click();
assert.strictEqual(document.getElementById('statusMessage').textContent, "Nonexistent Book not found!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Find book title</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter book title">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
}

input {
  padding: 0.5rem;
  width: 200px;
}

button {
  padding: 0.5rem 1rem;
  cursor: pointer;
  background-color: #000000;
  color: white;
}

#statusMessage {
  margin-top: 1.5rem;
  font-weight: bold;
  height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  try {
    const result = findBook(bookInput.value.trim());
    statusDisplay.textContent = result;
--fcc-editable-region--
  }
--fcc-editable-region--
});
```
