---
id: 6a1c47779b9594c071919b18
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now it is time to handle the empty string.

While you can throw any value, it is best practice to throw an `Error` object.

Inside your `if` statement, use the `throw` keyword to throw a `new Error()`. Pass the exact string `"Search field cannot be empty"` into the `Error` constructor.

*(Note: If you click the Search button after writing this, nothing will appear on the screen! The error halts the script silently because nothing is catching it yet. You will fix this in a few steps. However, you can click on the console for now to check if your code works.)*

# --hints--

Inside your `if` statement, you should use the `throw` keyword to throw a `new Error()` with the exact string "Search field cannot be empty".

```js
assert.throws(() => findBook(""), Error, "Search field cannot be empty");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Find book title</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter book title">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
}

input {
  padding: 0.5rem;
  width: 200px;
}

button {
  padding: 0.5rem 1rem;
  cursor: pointer;
  background-color: #000000;
  color: white;
}

#statusMessage {
  margin-top: 1.5rem;
  font-weight: bold;
  height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
  if (title === "") {
  --fcc-editable-region--
    
  --fcc-editable-region--
  }
  
  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  const result = findBook(bookInput.value.trim());
  statusDisplay.textContent = result;
});
```
