---
id: 6a1c45e3d23c5c908748b790
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now that your `findBook` function is set up, you can create the "happy path". This is the expected behavior of the code when everything goes right.

When a user searches for a book, the function should return a success message. This message will eventually be displayed on the screen.

Inside your `findBook` function, use the `return` keyword to return the exact string: `title + " is available to borrow!"`.

You can use standard string concatenation or template literals.

# --hints--

You should use the `return` keyword inside your function.

```js
assert.match(code, /return\s+/);
```

You should return the correct success string using string concatenation or template literals.

```js
assert.match(code, /return\s+(?:title\s*\+\s*["'] is available to borrow!["']|`\$\{title\} is available to borrow!`)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css"/>
</head>

<body>
    <div>
        <h1>Find book title</h1>
        <input type="text" id="bookTitleInput" placeholder="Enter book title">
        <button id="search">Search</button>
        <p id="statusMessage"></p>
    </div>
    <script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

const findBook = (title) => {
--fcc-editable-region--

--fcc-editable-region--
}

searchBtn.addEventListener('click', () => {
  const result = findBook(bookInput.value.trim());
  statusDisplay.textContent = result;
});
```
