---
id: 6a1d3032bdcd0ab7f47a251c
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Now it is time to put your new element to work!

Inside your `finally` block, right underneath where you clear the input, set the `textContent` of your `scanStatusDisplay` variable to exactly `"Library scan finished."`.

Try running a few successful and failed searches in the preview, and observe how the UI responds. Congratulations on completing your Library Book Search!

# --hints--

You should set `scanStatusDisplay.textContent` to `"Library scan finished."` inside your `finally` block.

```js
assert.match(code.replace(/\s+/g, ''), /finally\{[^}]*scanStatusDisplay\.textContent=["'`]Libraryscanfinished\.["'`]/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div id="app">
      <h1>Library Book Search</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter a book title...">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');
const scanStatusDisplay = document.getElementById('scanStatus')

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  try {
    const result = findBook(bookInput.value.trim());
    statusDisplay.textContent = result;
  } catch (err) {
    statusDisplay.textContent = err.message;
    console.error(err);
  } finally {
    bookInput.value = "";
  --fcc-editable-region--

  --fcc-editable-region--
  }
});

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div id="app">
      <h1>Library Book Search</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter a book title...">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
}

input {
    padding: 0.5rem;
    width: 200px;
}

button {
    padding: 0.5rem 1rem;
    cursor: pointer;
    background-color: #000000;
    color: white;
}

#statusMessage {
    margin-top: 1.5rem;
    font-weight: bold;
    height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');
const scanStatusDisplay = document.getElementById('scanStatus');

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
}

searchBtn.addEventListener('click', () => {
  try {
    const result = findBook(bookInput.value.trim());
    statusDisplay.textContent = result;
  } catch (err) {
    statusDisplay.textContent = err.message;
    console.error(err);
  } finally {
    bookInput.value = "";
    scanStatusDisplay.textContent = "Library scan finished.";
  }
});
```
