---
id: 61967e74a8e3690ab6292daa
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Seleziona l'elemento `body` e dagli un `background` con un gradiente lineare angolato di 45 gradi in senso orario, con colore di partenza `rgb(118, 201, 255)` e colore finale `rgb(247, 255, 222)`.

# --hints--

Dovresti usare un selettore `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Dovresti usare la proprietà `background` nel selettore `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Dovresti dare alla proprietà `background` un il valore `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
