---
id: 61968df2acd5550bf1616c34
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Normalizza la dimensione della pagina rimuovendo `margin` e `padding` dell'elemento `body`.

# --hints--

Dovresti dare all'elemento `body` una proprietà `margin` con il valore `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Dovresti dare all'elemento `body` una proprietà `padding` con il valore `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.padding, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));

}
--fcc-editable-region--
```
