---
id: 61969e7451455614217e901b
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Per rendere la montagna più realistica, puoi utilizzare la funzione di trasformazione `skew`, che richiede due argomenti. Il primo è un angolo di inclinazione dell'asse x e il secondo è un angolo di inclinazionazione dell'asse y.

Usa la proprietà `transform` per inclinare la montagna di `0deg` nell'asse x e `44deg` nell'asse y.

# --hints--

Dovresti dare all'elemento `.left-mountain` una proprietà `transform`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform);
```

Dovresti usare la funzione `skew` su `transform`.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform, 'skew');
```

Dovresti dare all'elemento `.left-mountain` una proprietà `transform` con il valore `skew(0deg, 44deg)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.getPropVal('transform', true), 'skew(0deg,44deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
