---
id: 619665c9abd72906f3ad30f9
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Creerai un felice pinguino che saluta, ed esplorerai ulteriormente le trasformazioni e animazioni in CSS nel processo.

Comincia con il tuo boilerplate HTML di base. Includi la dichiarazione `DOCTYPE`, l'elemento `html`, i tag `meta` appropriati, gli elementi `head`, `body` e `title`. Inoltre, collega il tuo foglio di stile alla pagina.

# --hints--

Il tuo codice dovrebbe avere una dichiarazione `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Il tuo codice dovrebbe avere un elemento `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Il tuo codice dovrebbe avere un elemento `head` all'interno dell'elemento `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Il tuo codice dovrebbe avere un elemento `body` all'interno dell'elemento `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Il tuo elemento `head` dovrebbe venire prima del tuo elemento `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Un elemento `meta` dovrebbe avere un attributo `name` con valore di `viewport`, e un attributo `content` con valore di `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

L'altro elemento `meta` dovrebbe avere l'attributo `charset` con valore di `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

L'elemento `title` dovrebbe avere del testo.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Dovrebbe esserci un elemento `link`.

```js
assert(/<link/.test(code))
```

Il tuo elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con un valore di `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Il tuo elemento `link` dovrebbe avere un attributo `type` con un valore di `text/css`.

```js
assert.match(code, /<link[\s\S]*?type=('|"|`)text\/css\1/)
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con un valore di `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
