---
id: 6197cff995d03905b0cca8ad
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

Il pinguino sarà composto da due sezioni principali: la testa e il corpo.

All'interno di `.penguin`, aggiungi due nuovi elementi `div`. Il primo con l'attributo `class` impostato su `penguin-head`, e il secondo con l'attributo `class` impostato su `penguin-body`.

# --hints--

Dovresti aggiungere due nuovi elementi `div` a `.penguin`. Dovresti avere `--fcc-expected--` elementi `.penguin > div` invece di `--fcc-actual--`.

```js
assert.equal(document.querySelectorAll('.penguin > div')?.length, 2);
```

Dovresti dare al primo `div` un attributo `class` con il valore `penguin-head`.

```js
assert.include(document.querySelector('.penguin > div:nth-of-type(1)')?.className, 'penguin-head');
```

Dovresti dare al secondo `div` un attributo `class` con il valore `penguin-body`.

```js
assert.include(document.querySelector('.penguin > div:nth-of-type(2)')?.className, 'penguin-body');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
--fcc-editable-region--
    <div class="penguin">

    </div>
--fcc-editable-region--
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
