---
id: 616965351e74d4689eb6de30
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Finalmente, utiliza una etiqueta `viewport` `<meta>` para asegurarse de que tu página se vea igual en cualquier dispositivo.

Anida un elemento `meta` de autocierre dentro del elemento `head`. Dale un atributo `name` con el valor `viewport` y un atributo `content` con el valor `width=device-width, initial-scale=1.0`.

# --hints--

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Tu nuevo elemento `meta` deben ser un elemento de autocierre.

```js
assert(code.match(/<\/meta>/i) === null);
```

Tu nuevo elemento `meta` debe tener un atributo `name` con el valor`viewport` y un atributo `content` con el valor `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
