---
id: 61695ab9f6ffe951c16d03dd
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Anida un elemento `head` dentro del elemento `html`. Después del elemento `head` añade un elemento `body`.

# --hints--

Debes tener una etiqueta `head` de apertura.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta `head` de cierre.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta `body` de apertura.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta `body` de cierre.

```js
assert(code.match(/<\/body\s*>/i));
```

Solo debe haber una etiqueta `head` de apertura.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Solo debe haber una etiqueta `head` de cierre.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Solo debe haber una etiqueta `body` de apertura.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Solo debe haber una etiqueta `body` de cierre.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Tu elemento `head` debe estar vacío.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Tu elemento `body` debe estar vacío.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Tu elemento `body` debe estar después tu elemento `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Tu elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Tu elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
