---
id: 61695d1fbc003856628bf561
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Para decirle a los navegadores cómo codificar los caracteres en tu página, dale a `charset` el valor `utf-8`. `utf-8` es un conjunto de caracteres universales que incluye casi todos los caracteres de todos los idiomas creados por el humano.

Dentro del elemento `head`, anida un elemento `meta` con un atributo `charset` con el valor `utf-8`. Recuerda que los elementos `meta` son de autocierre y no necesitan etiqueta de cierre.

# --hints--

Debes tener un elemento `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

Tu elemento `meta` debe ser un elemento de autocierre.

```js
assert(code.match(/<\/meta>/i) === null);
```

Tu elemento `meta` debe tener un atributo `charset` con el valor`utf-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
