---
id: 616971b3cd990186b66c99a1
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Ahora que tu archivo CSS externo está configurado, puedes comenzar a darle estilo a la página.

Como recordatorio, aquí está cómo apuntar un elemento de párrafo y alinearlo a la derecha:

```css
p {
  text-align: right;
}
```

Crea una nueva regla CSS que apunta al elemento `h1`, y establece su propiedad `text-align` al `center`.

# --hints--

Debes utilizar un selector `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

La regla CSS de tu selector `h1` debe tener la propiedad `text-align` con el valor `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
