---
id: 61695197ac34f0407e339882
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Como has visto en los proyectos anteriores, las páginas web deben comenzar con la declaración `DOCTYPE html`, seguida de un elemento `html`.

Añade una declaración `DOCTYPE html` en la parte superior del documento y un elemento `html` después de la declaración. Dale al elemento `html` un atributo `lang` con el valor `en`.

# --hints--

Tu código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debes definir que el tipo de documento es `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de apertura con un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/gi));
```

Tu declaración `DOCTYPE` debe estar al inicio de tu código HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
