---
id: 61695c4aad56f95497c19583
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Recuerda que el elemento `title` le da a los motores de búsqueda información extra sobre la página. También le dice a los navegadores qué texto mostrar en la pestaña de la página y en la barra de título cuando la página está abierta.

Dentro del elemento `head`, anida un elemento `title` con el texto `Colored Markers`.

# --hints--

Debes tener una etiqueta `title` de apertura.

```js
assert(code.match(/<title\s*>/i));
```

Debes tener una etiqueta `title` de cierre.

```js
assert(code.match(/<\/title\s*>/i));
```

Tu proyecto debe tener el título `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el texto del título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
