---
id: 61696ef7ac756c829f9e4048
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

En este proyecto trabajar치s con un archivo CSS externo para darle estilo a la p치gina. Ya creamos un archivo `styles.css` para ti. Pero antes de que puedas usarlo, necesitar치s enlazarlo a la p치gina.

Nest a `link` element within the `head` element. Dale un atributo `rel` con el valor `stylesheet` y un atributo `href` con el valor `styles.css`.

# --hints--

Debes tener un elemento `link` de autocierre.

```js
assert(document.querySelectorAll('link').length === 1);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
