---
id: 616968c2c94c8071b349c146
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Ahora estás listo para empezar a añadir contenido a la página.

Dentro del elemento `body`, anida un elemento `h1` con el texto `CSS Color Markers`.

# --hints--

Tu código debe tener un elemento `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Debes tener una etiqueta `h1` de apertura.

```js
assert(code.match(/<h1\s*>/i));
```

Debes tener una etiqueta `h1` de cierre.

```js
assert(code.match(/<\/h1\s*>/i));
```

Tu elemento `h1` debe estar dentro del elemento `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Tu elemento `h1` debe tener el texto `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el texto del elemento `h1`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
