---
id: 616965351e74d4689eb6de30
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Puedes tener múltiples elementos `meta` de autocierre en una página web. Cada elemento `meta` añade información sobre la página que no puede ser expresada por otros elementos HTML.

Añade otro elemento `meta` de autocierre dentro del `head`. Añade un atributo `name` con el valor `viewport` y un atributo `content` con el valor `width=device-width, initial-scale=1.0` de modo que la página se vea igual en cualquier dispositivo.

# --hints--

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Tu nuevo elemento `meta` deben ser un elemento de autocierre.

```js
assert(code.match(/<\/meta>/i) === null);
```

Tu nuevo elemento `meta` debe tener un atributo `name` con el valor`viewport` y un atributo `content` con el valor `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
