---
id: 616d3a67ccf800ad94ec89ae
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ahora añadirás algunos elementos, a los que darás forma de marcadores o plumones en los siguientes pasos.

First, within the `body` element, add a `div` element and set its `class` attribute to `container`. Asegúrate que el elemento `div` esté debajo del elemento `h1`.

# --hints--

Debes tener una etiqueta `div` de apertura.

```js
assert(code.match(/<div.*>/i));
```

Debes tener una etiqueta `div` de cierre.

```js
assert(code.match(/<\/div\s*>/i));
```

Tu elemento `div` debe estar dentro de tu elemento `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Tu elemento `div` debe estar después del elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Tu elemento `div` debe tener un atributo `class` con el valor `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
