---
id: 616d525007b8c5d8b3308b1c
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Tu marcador se vería mejor si estuviera centrado en la página. Una forma fácil de hacerlo, es con la `margin` <dfn>propiedad abreviada</dfn>.

En el proyecto anterior, estableciste el margen de los elementos por separado, con propiedades como `margin-top` (margen superior) y `margin-left` (margen-izquierdo). La propiedad abreviada `margin` facilita el establecer múltiples áreas del un margen al mismo tiempo.

Para centrar tu marcador en la página, dale el valor `auto` a la propiedad `margin`. Esto les da el valor `auto` a `margin-top`, `margin-right`, `margin-bottom` y `margin-left`.

# --hints--

Tu regla CSS `.marker` debe tener una propiedad `margin` con un valor `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
