---
id: 616d55bd160a95e22453a081
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Ya que tienes un marcador centrado y con color, es hora de añadir los otros marcadores.

En él `div` con la clase `container`, añade otros dos elementos `div` y dales una clase con el valor `marker` a cada uno.

# --hints--

El primer elemento `div` que acabas de añadir debe tener una etiqueta de apertura.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

El primer elemento `div` que acabas de añadir debe tener una etiqueta de cierre.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

El segundo elemento `div` que acabas de añadir debe tener una etiqueta de apertura.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

El segundo elemento `div` que acabas de añadir debe tener una etiqueta de cierre.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Tus nuevos elementos `div` deben estar dentro del `div` que tiene la clase `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

Tus nuevos elementos `div` deben tener una clase `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
