---
id: 616d4a84b756d9c4b8255093
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Es hora de agregar algo de color al marcador. Recuerda que una forma de añadir color a un elemento es utilizando el nombre del <dfn>color</dfn> en inglés, como `black` (negro), `cyan` (cian) o `yellow` (amarillo).

Como recordatorio, aquí está cómo vincular la clase `freecodecamp`:

```css
.freecodecamp {

}
```

Crea una nueva regla CSS que apunta a la clase `marker`y establezca la `background-color` a `red`.

**Nota:** No verás ningún cambio después de agregar el CSS.

# --hints--

Debes crear un selector de clase para orientar la clase `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

La regla CSS `.marker` debe tener una propiedad `background-color` establecida en `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
