---
id: 616d50b93ba424d6282c99cf
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Puedes ver que tu marcador parece no tener color. El color de fondo (background color) se aplicó, pero como el elemento `div` del marcador está vacío y no tiene ninguna altura asignada por defecto.

En tu regla CSS `.marker`, agrega una propiedad `width` (ancho) con un valor de `200px` y una propiedad `height` (altura) con un valor dé `25px`.

# --hints--

Tu regla CSS `.marker` debe tener una propiedad `width` con un valor dé `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

Tu regla CSS `.marker` debe tener una propiedad `height` con un valor dé `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
