---
id: 61764c602bee6974e7790f35
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Para dar a los marcadores diferentes colores, necesitará añadir una clase única a cada uno. Múltiples clases pueden ser añadidas a un elemento listándolas en el atributo `class` y separándolas con un espacio. Por ejemplo, lo siguiente añade las clases `animal` y `dog` a un elemento `div`.

```html
<div class="animal dog">
```

If you add multiple classes to an HTML element, the styles of the first classes you list may be overridden by later classes.

To begin, add the class `one` to the first marker `div` element.


# --hints--

You should add the class `one` to the first marker `div` element.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Your first marker `div` should have the classes `marker` and `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
