---
id: 61764f7e5240eb7ccc7f6a0a
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

Luego, crea una nueva regla CSS que apunta a la clase `one` y establece su propiedad `background-color` a `red`.

# --hints--

Debes usar un selector de clases para apuntar a la clase `one`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one'));
```

Tu regla CSS `.one` debe tener una propiedad `background-color` establecida a `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
