---
id: 617b65579ce424bf5f02ca73
title: Paso 21
challengeType: 0
dashedName: step-21
---

# --description--

Hay dos modelos principales de colores: el modelo <dfn>aditivo</dfn> RGB (rojo, verde, azul) utilizado en dispositivos electrónicos, y el modelo <dfn>sustractivo</dfn> CMYK (cian, magenta, amarillo, negro) utilizado para impresión.

En este proyecto, trabajarás con el modelo RGB. Esto significa que los colores comienzan como negros y cambian a medida que se introducen diferentes niveles de rojo, verde y azul. Una forma fácil de ver esto, es con la función CSS `rgb`.

Crea una nueva regla CSS que apunte a la clase `container` y establece su `background-color` a negro con `rgb(0, 0, 0)`.

# --hints--

Debes utilizar un selector de clase para apuntar a la clase `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Tu regla CSS `.container` debe contener una propiedad `background-color` establecida a `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
