---
id: 617bd6ec666b1da2587e4e37
title: Paso 34
challengeType: 0
dashedName: step-34
---

# --description--

Existen tres colores terciarios más: verde chartreuse (verde + amarillo), azur (azul + cyan), y rosa (rojo + magenta).

Para crear verde chartreuse, actualiza la función `rgb` de la clase `.one` de manera tal que el rojo esté en `127` y el verde en su maximo valor.

Para el azur, actualiza la función `rgb` de la clase `.two` de manera tal que el verde esté en `127` y el azul en su valor máximo.

Finalmente, para el rosa, que a veces es llamado rosa brillante, actualiza la función `rgb` de la clase `.three` para que el azul esté en `127` y el rojo es su valor máximo.

# --hints--

Tu regla CSS `.one` debe contener una propiedad `background-color` establecida a `rgb(127, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

Tu regla CSS `.two` debe contener una propiedad `background-color` establecida a `rgb(0, 127, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

Tu regla CSS `.three` debe contener una propiedad `background-color` establecida a `rgb(255, 0, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
