---
id: 617b954d9f4f6217a749380e
title: Paso 25
challengeType: 0
dashedName: step-25
---

# --description--

Ahora, aumenta un poco más el espacio vertical entre tus marcadores y los bordes del elemento `container` en el que se encuentran.

Dentro del selector `.container` de CSS, usa la propiedad de `padding` para agregar `10px` de relleno superior e inferior y `0` de relleno a izquierda y derecha. Esta funciona de forma similar a la propiedad `margin` que usaste anteriormente.

# --hints--

No deberías de eliminar la propiedad `background-color` y su valor del selector CSS `.container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

Tu regla CSS `.container` debería tener una propiedad `padding` con el valor `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
