---
id: 617b8e0d93a8d10d9a90e720
title: Paso 23
challengeType: 0
dashedName: step-23
---

# --description--

Observa que el `background-color` (color de fondo) de tu marcador sigue siendo rojo. Esto es porque has establecido el rojo de la función `rgb` en su valor máximo de `255`, o 100% rojo, y tanto el verde como el azul con valores de `0`.

Ahora utiliza la función `rgb` para establecer los otros colores.

En la regla CSS `.two`, utiliza la función `rgb` para establecer el `background-color` al valor máximo para el verde, y a `0` para los demás valores. Y en la regla CSS `.three`, utiliza la función `rgb` para establecer el `background-color` al valor máximo para el azul y a `0` para los demás valores.

# --hints--

Tu regla CSS `.two` no debería usar la palabra clave para color `green` para indicar la propiedad `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

Tu regla CSS `.two` debería tener la propiedad `background-color` establecida en `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

Tu regla CSS `.three` no debería usar la palabra clave para color `blue` para indicar la propiedad `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

Tu regla CSS `.three` debería tener la propiedad `background-color` establecia en `rgb(0, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
