---
id: 617b8b38f32bf2080a140675
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Una función es una pieza de código que puede tomar una entrada y realizar una acción específica. La función CSS `rgb` acepta valores, o <dfn>argumentos</dfn>, para rojo, verde y azul y produce un color:

```css
rgb(red, green, blue);
```

Cada valor rojo, verde y azul es un número de `0` a `255`. `0` significa que hay 0% de ese color, y es negro. `255` significa que hay 100% de ese color.

En la regla CSS `.one`, reemplaza la palabra clave `red` con la función `rgb`. Para la función `rgb`, establezca el valor para rojo a `255`, el valor para verde a `0`, y el valor para azul a `0`.

# --hints--

Tu regla de CSS `.one` no debe usar la palabra clave de color `red` para establecer la propiedad `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

Tu regla CSS `.one` debe tener una propiedad `background-color` establecida a `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
