---
id: 618a0927005553b74bfa05ae
title: Paso 36
challengeType: 0
dashedName: step-36
---

# --description--

Una rueda de color es un círculo donde colores similares, o <dfn>hues</dfn>, están cerca unos de otros, y los diferentes están más separados. Por ejemplo, el rojo puro está entre los tonos rosa y naranja.

Dos colores opuestos entre sí en la rueda de color se llaman <dfn>complementary colors</dfn>. Si se combinan dos colores complementarios, producen gris. Pero cuando se colocan lado a lado, estos colores producen un fuerte contraste visual y parecen más brillantes.

En la función `rgb` para la `.one` regla CSS, establezca el valor rojo al máximo de `255` para producir rojo puro. En la función `rgb` para el `.two` regla CSS, establezca para verde y azul el máximo de `255` para producir cian.

# --hints--

La regla CSS `.one` debe tener la propiedad `background-color` establecida a `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

Tu regla CSS `.two` debe contener una propiedad `background-color` establecida en `rgb(0, 255, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
