---
id: 618a00ed1ca871a2b3aca0cb
title: Paso 35
challengeType: 0
dashedName: step-35
---

# --description--

Una vez estudiados los colores primarios, secundarios y terciarios sobre un círculo cromático, será más sencillo comprender otros conceptos de teoría del color y qué impacto pueden tener en el diseño.

First, in the CSS rules `.one`, `.two`, and `.three`, adjust the values in the `rgb` function so that the `background-color` of each element is set to pure black. Recuérdese que la función `rgb` utiliza el modelo aditivo, en el que los colores comienzan como negro y cambian a medida que los valores de rojo, verde y azul aumentan.

# --hints--

Tu regla CSS `.one` debe tener una propiedad `background-color` establecida a `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

Tu regla CSS `.two` debe tener una propiedad `background-color` establecida a `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

Tu regla CSS `.three` debe tener una propiedad `background-color` establecida a `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
