---
id: 619b761916dac02643940022
title: Paso 46
challengeType: 0
dashedName: step-46
---

# --description--

Es práctica muy común aplicar color a un elemento CSS con valores <dfn>hexadecimales</dfn> (hex). Puede sonar complicado, pero no es más que otra forma de expresar colores RGB.

Los valores hexadecimales empiezan con el carácter `#`, siguiendo seis caracteres entre 0-9 y A-F. El primer par de caracteres representa el rojo, el segundo el verde y el tercero el azul. Pr ejemplo, `#4B5320`.

In the `.green` class selector, set the `background-color` property to a hex color code with the values `00` for red, `FF` for green, and `00` blue.

# --hints--

Tu regla CSS `.green` debe tener la propiedad `background-color` establecidad a `#00FF00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
