---
id: 619b74fa777a2b2473c94f82
title: Paso 45
challengeType: 0
dashedName: step-45
---

# --description--

Actualiza el selector de clase CSS `.two` de manera que apunte a la nueva clase `green`. And update the `.three` class selector so it targets the new `blue` class.

# --hints--

Tu código no debe tener el selector de clase `.two`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

Debes usar un selector de clase para seleccionar la clase `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

Tu regla CSS `.green` debe tener la propiedad `background-color` establecida a `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Tu código no debe tener el selector de clase `.three`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

Debes usar un selector de clase para seleccionar la clase `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Tu regla CSS `.blue` debe tener la propiedad `background-color` establecida a `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
