---
id: 619b7fd56aa2253778bcf5f7
title: Paso 48
challengeType: 0
dashedName: step-48
---

# --description--

Otra forma de representar colores es el modelo <dfn>HSL</dfn>, siglas de hue, saturation y lightness (matiz, saturación y luminosidad).

La función de CSS `hsl` acepta tres valores: un número entre 0 y 360 para el matiz (hue), un porcentaje para la saturación (saturation), y otro porcentaje para la luminosidad (lightness).

Sobre un círculo cromático, el matiz (hue) rojo está en 0 grados, el verde en 120 y el azul en 240.

La saturación (saturation) es la intensidad de un color desde 0%, o gris, hasta 100% para el color puro.

La luminosidad (lightness) es lo brillante que parece un color, desde 0% o completamente negro, hasta 100%, completamente blanco, siendo neutro el 50%.

En la regla CSS `.blue`, utiliza la función `hsl` para cambiar la propiedad `background-color` a azul puro. Establece el tono a `240`, la saturación a `100%` y la luminosidad a `50%`.

# --hints--

La regla CSS `.blue` debe tener la propiedad `background-color` establecida a `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
