---
id: 61a489b8579e87364b2d2cdb
title: Paso 49
challengeType: 0
dashedName: step-49
---

# --description--

Hemos visto unas cuantas maneras de establecer colores planos en CSS, pero también se pueden aplicar transiciones de color, o <dfn>gradientes</dfn>, sobre un elemento.

Se llama gradiente, o degradado, a la transición progresiva de un color a otro. La función CSS `linear-gradient` permite controlar la dirección de la transición a lo largo de una línea, y qué colores participan en el degradado.

Una cosa a tener en cuenta es que la función `linear-gradient` lo que en realidad crea es un elemento `image`, y se asocia normalmente con la propiedad `background`, la cual puede aceptar una imagen como valor.

En la regla CSS `.red`, sustituye la propiedad `background-color` por `background`.

# --hints--

La regla CSS `.red` debe tener la propiedad `background` con el valor `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
