---
id: 61a493ead935c148d2b76312
title: Paso 50
challengeType: 0
dashedName: step-50
---

# --description--

La función `linear-gradient` es muy flexible -- aquí hay una sintaxis básica que usaras en este tutorial:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` es la dirección de la línea a lo largo de la cual tendrá lugar la transición. `color1` y `color2` son argumentos de color, representando los colores usados en la transición. Se pueden pasar a la función en cualquier formato, ya sea el nombre del color, hex, `rgb` o `hsl`.

Vamos a aplicar al primer marcador un gradiente de rojo a verde a lo largo de una línea de 90 grados.

Primero, en la regla CSS `.red`, establece la propiedad `background` a `linear-gradient()`, y al argumento `gradientDirection` pásale el valor `90deg`.

# --hints--

Tu regla CSS `.red` debe tener la propiedad `background` con el valor `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
