---
id: 61addaf7e83988b59a7d18ff
title: Paso 82
challengeType: 0
dashedName: step-82
---

# --description--

La propiedad abreviada `border-left` permite establecer el grosor, el estilo y el color del borde izquierdo de forma simultánea.

La sintaxis sería:

```css
border-left: width style color;
```

En la regla CSS `.sleeve`, reemplaza las propiedades `border-left-width`, `border-left-style` y `border-left-color` por la propiedad abreviada `border-left`. Los valores de grosor, estilo y color del borde izquierdo deben ser los mismos.

# --hints--

La regla CSS `.sleeve` no debe tener la propiedad `border-left-width`.

```js
assert(!/border-left-width/g.test(code));
```

La regla CSS `.sleeve` no debe tener la propiedad `border-left-style`.

```js
assert(!/border-left-style/g.test(code));
```

La regla CSS `.sleeve` no debe tener la propiedad `border-left-color`.

```js
assert(!/border-left-color/g.test(code));
```

La regla CSS `.sleeve` debe tener la propiedad abreviada `border-left` con el valor `10px solid black`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left-width: 10px;
  border-left-style: solid;
  border-left-color: black;
}
--fcc-editable-region--

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
