---
id: 61b093219e7fc020b43b1bb4
title: Paso 74
challengeType: 0
dashedName: step-74
---

# --description--

Otra forma de establecer la opacidad de un elemento es con el canal <dfn>alpha</dfn>. Similar a la propiedad `opacity`, el canal alfa controla qué tan transparente u opaco es un color.

Ya has establecido la opacidad de la carátula con un color con nombre y la propiedad `opacity`, pero puedes agregar un canal alfa a las otras propiedades de color CSS.

Dentro de la regla `.sleeve`, quite la propiedad y el valor `opacity`.

# --hints--

La regla CSS `.sleeve` no debe tener una propiedad y un valor `opacity`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
  opacity: 0.5;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
