---
id: 61b0936d9e7fc020b43b1bb8
title: Paso 78
challengeType: 0
dashedName: step-78
---

# --description--

Parece que tu manga desapareció, pero no te preocupes, todavía está allí. Lo que ha ocurrido es que tu nuevo `div` de tapa está ocupando todo el ancho del marcador, y está empujando el mango hacia la siguiente línea.

Esto se debe a que la propiedad predeterminada `display` para los elementos `div` es `block`. Entonces, cuando dos elementos `block` están uno al lado del otro, se apilan como bloques reales. Por ejemplo, los elementos del marcador están apilados uno encima del otro.

Para colocar dos elementos `div` en la misma línea, establezca sus propiedades `display` en `inline-block`.

Cree una nueva regla para dirigirse a las clases `cap` y `sleeve` y establezca `display` en `inline-block`.

# --hints--

Debe usar un selector de clase para orientar las clases `cap` y `sleeve`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

La regla CSS debe tener una propiedad `display` establecida en `inline-block`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
