---
id: 61b0a44a6b865738ba49b9fb
title: Paso 86
challengeType: 0
dashedName: step-86
---

# --description--

La última cosa que haremos será añadir un ligero sombreado a cada marcador para hacerlos parecer aún más realistas.

La propiedad `box-shadow` permite aplicar una o más sombras alrededor de un elemento. Esta sería la sintaxis básica:

```css
box-shadow: offsetX offsetY color;
```

Comienza añadiendo una sombra simple al marcador rojo.

En la regla CSS `.red`, añade la propiedad `box-shadow` con los valores `5px` para `offsetX`, `5px` para `offsetY` y `red` para `color`.

# --hints--

La regla CSS `.red` debe tener la propiedad abreviada `box-shadow` con el valor `5px 5px red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
