---
id: 61b30995968123ceb6b76167
title: Paso 88
challengeType: 0
dashedName: step-88
---

# --description--

Observa que los límites de la sombra son muy bruscos. En la propiedad `box-shadow`, se encuentra el valor opcional `blurRadius`:

```css
box-shadow: offsetX offsetY blurRadius color;
```

Si no se indica un valor para `blurRadius`, será `0` por defecto, produciendo sombras con los límites bruscos. Cuanto mayor sea el valor de `blurRadius`, mayor será el efecto de difuminado de la sombra.

En la regla CSS `.green`, añade la propiedad `box-shadow` con los valores `5px` para `offsetX`, `5px` para `offsetY`, `5px` para `blurRadius` y `green` para `color`.

# --hints--

La regla CSS `.green` debe tener la propiedad abreviada `box-shadow` con el valor `5px 5px 5px green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow === 'green 5px 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

--fcc-editable-region--
.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
