---
id: 5f3313e74582ad9d063e3a38
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`html` 要素内に `head` 要素を追加して、そこに `title` 要素を追加してください。 `title` 要素のテキストは `Cafe Menu` としてください。

# --hints--

`<head>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head>/i));
```

`</head>` の終了タグが 1 つ必要です。

```js
assert(code.match(/<head>/i));
```

`<title>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<title>/i));
```

`</title>` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/title>/i));
```

`title` 要素は `head` 要素の中にネストされている必要があります。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

`title` 要素は `Cafe Menu` というテキストをもつ必要があります。 スペルが正しいか確認してください。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
