---
id: 5f33310c1851c6c4da013250
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

メニューの中身を追加しましょう。 `main` 要素を既存の `header` 要素の下に追加してください。 後ほど、カフェが提供するコーヒーとデザートの値段をここに入れます。

# --hints--

`<main>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<main>/i));
```

終了タグ `</main>` が 1 つ必要です。

```js
assert(code.match(/<\/main>/i));
```

`header` 要素を変更しないでください。 誤って終了タグを削除していないか確認してください。

```js
assert($('header').length === 1);
```

`main` タグを `header` タグの後に置く必要があります。

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
