---
id: 5f33294a6af5e9188dbdb8f3
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

メニューの中身を追加しましょう。 `main` 要素を、既存の `body` 要素の中に追加してください。 後ほど、カフェが提供するコーヒーとデザートの値段をここに入れます。

# --hints--

`<main>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<main>/i));
```

終了タグ `</main>` が 1 つ必要です。

```js
assert(code.match(/<\/main>/i));
```

`body` 要素を変更しないでください。 誤って終了タグを削除していないか確認してください。

```js
assert($('body').length === 1);
```

`main` タグは `body` タグの中に置く必要があります。

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
