---
id: 5f332b23c2045fb843337579
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

カフェが 2020 年にオープンしたということを訪問者に知らせるために、`p` 要素を `h1` 要素の下に追加して、そのテキストを `Est. 2020` にしてください。

# --hints--

`<p>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<p>/i));
```

終了タグ `</p>` が 1 つ必要です。

```js
assert(code.match(/<\/p>/i));
```

既存の `h1` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert($('h1').length === 1);
```

`p` 要素は `h1` 要素の下に置く必要があります。

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

`p` 要素は `Est. 2020` というテキストをもつ必要があります。

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

