---
id: 5f33071498eb2472b87ddee4
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

猫の写真アプリの最後の数ステップで学んだように、ウェブページを作り始めるのに必要な基本構造があります。

`<!DOCTYPE html>` タグと、 `lang` 属性に `en`を設定した `html` 要素を追加してください。

# --hints--

`DOCTYPE` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

`lang` 属性に `en` を設定した開始タグ `<html>` が必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`<html>` の終了タグが 1 つ必要です。 終了タグは初めの山括弧 `<` の後に `/` があることを忘れないようにしましょう。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
