---
id: 5f3326b143638ee1a09ff1e3
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

実際に表示されるコンテンツを作成する準備として、`head` 要素の下に `body` 要素を追加してください。

# --hints--

`<body>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body>/i));
```

終了タグ `</body>` が 1 つ必要です。

```js
assert(code.match(/<\/body>/i));
```

`head` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

`body` 要素を `head` 要素の後に置く必要があります。

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

