---
id: 5f33294a6af5e9188dbdb8f3
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

カフェの名前は `CAMPER CAFE` です。 `h1` 要素を `body` 要素内に 1 つ追加してください。 そこにカフェの名前を目立つように大文字で入れてください。

# --hints--

`<h1>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<h1>/i));
```

終了タグ `</h1>` が 1 つ必要です。

```js
assert(code.match(/<\/h1>/i));
```

既存の `body` 要素を変更しないでください。

```js
assert($('body').length === 1);
```

`h1` 要素は `body` 要素の中にネストされている必要があります。

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

`h1` 要素は `CAMPER CAFE` というテキストをもつ必要があります。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
