---
id: 5f33071498eb2472b87ddee4
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS は、HTML 文書のスタイル設定に使われる言語です。 CSS を使って、HTML 要素が画面にどのように表示されるべきかを記述します。

猫の写真アプリの最後の数ステップで学んだように、ウェブページを作り始めるのに必要な基本構造があります。 HTML 文書には `DOCTYPE` 宣言と `html` 要素が必要です。 `DOCTYPE` は、この文書がどのバージョンの HTML で書かれているかをブラウザに伝えます。 そして `html` 要素は、他すべての要素を内包するルート要素を表します。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

`<!DOCTYPE html>` タグと、 `lang` 属性に `en`を設定した `html` 要素を追加してください。

# --hints--

`<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

言語に英語を指定した開始タグ `<html lang="en">` が必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

終了タグ `</html>` が必要です。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
