---
id: 5f332b23c2045fb843337579
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

前回のステップで追加した `p` 要素はカフェについての補足情報を提供しているので、`h1` 要素と `p` 要素の両方を `header` 要素の中にネストしましょう。

# --hints--

`<header>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<header>/i));
```

終了タグ `</header>` が 1 つ必要です。

```js
assert(code.match(/<\/header>/i));
```

`h1` 要素は `header` 要素の中にネストされている必要があります。

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

`p` 要素は `header` 要素の中にネストされている必要があります。

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
<html>
```

