---
id: 5f3477ae9675db8bb7655b30
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

前のステップで、<dfn>要素型セレクター</dfn>を使用して `h1` 要素にスタイルを適用しました。 次は、それぞれ新しく要素型セレクターを追加して、`h2` 要素と `p` 要素を中央に配置してください。

# --hints--

既存の `h1` セレクターを変更しないでください。

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

新しい `h2` セレクターを追加してください。

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

新しい `p` セレクターを追加してください。

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

`h1` 要素の `text-align` を `center` に設定する必要があります。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

`h2` 要素の `text-align` を `center` に設定する必要があります。

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

`p` 要素の `text-align` を `center` に設定する必要があります。

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
    <main>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
<html>
```

