---
id: 5f3477cb303c5cb61b43aa9b
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

テキストが再び中央に配置されたので、CSS ファイルへのリンクは機能しています。 このファイルにもう一つ、`body` 要素の `background-color` プロパティを `brown` に変更するスタイルを追加してください。

# --hints--

`body` セレクターが必要です。

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

`background-color` プロパティを `brown` に設定する必要があります。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

`body` 要素の背景色は `brown` である必要があります。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

