---
id: 5f3477cbcb6ba47918c1da92
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

ページのスタイルを、モバイル機器でもデスクトップやラップトップと同様の表示にするには、特別な `content` 属性をもつ `meta` 要素を追加する必要があります。

`head` 要素内に以下を追加してください:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

コードには 2 つの `meta` 要素が必要です。

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

`meta` 要素には `name` 属性があり、値が `viewport` に設定されている必要があります。

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

`meta` 要素には `content` 属性があり、値が `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
