---
id: 5f344f9c805cd193c33d829c
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

`style` 要素の中で、次のように要素を指定してプロパティを設定することで、その要素にスタイルを適用できます:

```css
element {
 property: value;
}
```

`h1` 要素の `text-align` プロパティを `center` に設定して、中央揃えにしてください。

# --hints--

`style` 要素内に `h1` セレクターが必要です。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

`text-align` プロパティの値を `center` に設定する必要があります。

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

`h1` セレクターの `text-align` プロパティを `center` に設定する必要があります。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

