---
id: 5f3477ae8466a9a3d2cc953c
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

これで CSS は `styles.css` ファイル内にあるので、`style` 要素とその中身をすべて削除してください。 それを削除すると、中央に配置されていたテキストが再び左側に戻ります。

# --hints--

コード内に `style` タグが無いようにしてください。

```js
assert(!code.match(/style/i));
```

HTML ファイル内に CSS セレクターが無いようにしてください。

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

