---
id: 5f344fbc22624a2976425065
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`h2` 要素を `section` 要素内に作成し、そのテキストを `Coffee` にしてください。

# --hints--

`<h2>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<h2\s*>/i));
```

終了タグ `</h2>` が 1 つ必要です。

```js
assert(code.match(/<\/h2\s*>/i));
```

既存の `section` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert($('section').length === 1);
```

`h2` 要素は `section` 要素の中にある必要があります。

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

`h2` 要素のテキストは `Coffee` にする必要があります。

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
