---
id: 5f344fc1520b6719f2e35605
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

メニューには 2 つのセクションを作ります。ひとつはコーヒー用でもうひとつはデザート用です。 `section` 要素を `main` 要素内に 1 つ追加し、注文できるコーヒーの一覧を配置する場所を用意します。

# --hints--

`<section>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<section\s*>/i));
```

`</section>` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/section\s*>/i));
```

既存の `main` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert($('main').length === 1);
```

`section` 要素は `main` 要素の中にある必要があります。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
