---
id: 5f3477aefa51bfc29327200b
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

`style` タグの中に CSS を記述することで、3 つの要素にスタイルを適用しました。 これでも機能しますが、これからもっと多くのスタイルを使用するので、すべてのスタイルを別のファイルに記述してそれにリンクするのが最善です。

既に別のファイル `styles.css` を作成して、エディタ画面をそのファイルに切り替えてあります。 エディタ上部のタブを使用してファイルの間を移動できます。

まず、作成したスタイルを `styles.css` ファイルの中に書き写すことから始めましょう。 `style` の開始タグと終了タグは除外してください。

# --hints--

`styles.css` ファイルには `h1, h2, p` の要素型セレクターが必要です。

```js
(getUserInput) => {
  assert(getUserInput('editableContents').replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
}
```

セレクターの `text-align` プロパティは `center` に設定されている必要があります。

```js
(getUserInput) => {
  assert(getUserInput('editableContents').match(/text-align:\s*center;?/));
}
```

コードに `<style>` タグを含めないでください。

```js
  assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
