---
id: 5f3477cb2e27333b1ab2b955
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

スタイルを再度適用するためには `styles.css` ファイルをリンクする必要があります。 自己終了要素の `link` を 1 つ `head` 要素内にネストしてください。 その `rel` 属性を `stylesheet` に、`href` 属性を `styles.css` に設定してください。

# --hints--

コードには `link` 要素が 1 つ必要です。

```js
const link = document.querySelector('link');
assert(link);
```

既存の `head` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert($('head').length === 1);
```

自己終了要素 `link` が 1 つ必要です。

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
const link = document.querySelector('head > link');
assert(link);
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
