---
id: 5f3477cb2e27333b1ab2b955
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

スタイルを再度適用するためには `styles.css` ファイルをリンクする必要があります。 自己閉じ要素の `link` を 1 つ `head` 要素内にネストしてください。 それに、`rel` 属性値 `stylesheet` 、`type` 属性値 `text/css`、`href` 属性値 `styles.css` を加えてください。

# --hints--

コードには `link` 要素が 1 つ必要です。

```js
// link is removed -> if exists, replaced with style
const link = document.querySelector('head > style');
assert(link);
```

既存の `head` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert($('head').length === 1);
```

`link` 要素は自己閉じ要素である必要があります。

```js
assert(code.match(/<link[\w\W\s]+\/?>/i));
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
assert(code.match(/rel\s*=\s*('|")stylesheet\1/i));
```

`link` 要素には `type` 属性があり、値が `text/css` に設定されている必要があります。

```js
assert(code.match(/type\s*=\s*('|")text\/css\1/i));
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
assert(code.match(/href\s*=\s*('|")styles.css\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
    <main>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
<html>
```

```css
h1, h2, p {
  text-align: center;
}
```
