---
id: 5f3477ae34c1239cafe128be
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

ここでは全く同じスタイルを適用する要素型セレクターが 3 つあります。 セレクターのリストを作成することで、複数の要素に同じスタイルのグループを適用できます。 各セレクターは次のようにコンマで区切ります:

```css
selector1, selector2 {
  property: value;
}
```

既存の 3 つの型セレクターを削除し、`h1`、`h2`、`p` の各要素のテキストを中央揃えにする 1 つのセレクター リストに置き換えてください。

# --hints--

3 つの要素すべてを選択するために `h1, h2, p` というひとつの要素型セレクターを使用する必要があります。 必ずこの順番で記載してください。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

`style` 要素内に入れるセレクターは 1 つだけです。

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

