---
id: 5f356ed69db0a491745e2bb6
title: ステップ 27
challengeType: 0
dashedName: step-27
---

# --description--

クラスのスタイルを `div` 要素に適用するため、`div` 要素の開始タグから `id` 属性を削除して `class` 属性を追加しましょう。 クラスの値は `menu` としてください。

# --hints--

`div` が今まで通り表示されている必要があります。 `<div>` タグが不正な形式になっていないか確認してください。

```js
assert($('div').length === 1);
```

`div` 要素に `menu` クラスが必要です。

```js
assert($('div').attr('class').includes('menu'));
```

`div` 要素が `menu` という ID を持たないようにしてください。

```js
assert(!$('div#menu').length);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

