---
id: 5f356ed63c7807a4f1e6d054
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

次の目標は、`div` がページの幅全体に広がらないようにすることです。 そうするには CSS の `width` プロパティが最適です。

`id` 属性を持つ特定の要素を選択するには、`id` セレクターが使用できます。 `id` セレクターは、要素の `id` の値の前に `#` の記号を付けることで定義されます。 例えば、要素が `cat` という `id` を持っている場合、その要素を選択するには次のようにします:

```css 
#cat {
  width: 250px;
}
```

では `#menu` セレクターを使って、要素の幅を `300px` に設定しましょう。

# --hints--

`#menu` セレクターが必要です。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

`width` プロパティを `300px` に設定する必要があります。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

`div` の幅は 300px である必要があります。

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

