---
id: 5f356ed6199b0cdef1d2be8f
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

ここまでは、要素型セレクターと ID セレクターを使用して要素にスタイルを適用してきました。 しかし、一般的に要素のスタイル設定にはそれとは違うセレクターが使われます。

<dfn>クラスセレクター</dfn>は、次のようにクラス名の前にドットをつけることで定義されます:

```css
.class-name {
  styles
}
```

`#menu` を `.menu` という名前のクラスと置き換えて、既存の `#menu` セレクターをクラスセレクターに変更してください。

# --hints--

`.menu` クラスセレクターが必要です。

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

`#menu` セレクターが無いようにしてください。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

