---
id: 5f35e5c4321f818cdc4bed30
title: ステップ 29
challengeType: 0
dashedName: step-29
---

# --description--

良さそうですね。 ではメニュー項目を追加し始めましょう。 空の `article` 要素を 1 つ、`Coffee` という見出しの下に追加してください。 ここにはカフェで提供しているコーヒー各種のフレーバーと価格を含める予定です。

# --hints--

`<article>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<article>/i));
```

終了タグ `</article>` が 1 つ必要です。

```js
assert(code.match(/<\/article>/i));
```

既存の `h2` 要素を変更しないようにしてください。

```js
assert($('h2').length === 1);
```

`article` 要素を `h2` 要素の後に置く必要があります。

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

