---
id: 5f35e5c44359872a137bd98f
title: ステップ 28
challengeType: 0
dashedName: step-28
---

# --description--

カフェの主な商品はコーヒーなので、ページの背景にコーヒー豆の画像を使いましょう。

`body` 要素型セレクター内のコメントと中身を削除してください。 次に、`background-image` プロパティを追加し、値を `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)` に設定してください。

# --hints--

コメントアウトされている `background-color` プロパティを削除する必要があります。

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

`body` セレクターにはコメントを含まないようにしてください。

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

`background-image` プロパティを `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)` に設定する必要があります。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

`body` 要素はコーヒー豆の背景画像をもつ必要があります。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

