---
id: 5f3c866de7a5b784048f94b1
title: ステップ 36
challengeType: 0
dashedName: step-36
---

# --description--

望んでいた形に近いものになりましたが、ここでフレーバーと価格が一直線上にあるとよさそうです。 `p` 要素は <dfn>ブロックレベル</dfn>要素なので、親要素の幅全体に広がっています。

これを一直線上に配置するには、`p` 要素にとあるスタイルを適用して<dfn>インライン</dfn>要素のように動作させる必要があります。 `class` 属性と値 `item` を、`Coffee` という見出しの下にある最初の `article` 要素に追加してください。

# --hints--

`article` 要素に `item` クラスを適用する必要があります。

```js
assert(code.match(/<article\s*class=('|")item\1\s*>/i))
```

`item` クラスをもつ要素が 1 つだけ必要です。

```js
assert($('.item').length === 1);
```

最初の `article` 要素に `item` クラスが必要です。

```js
assert($('article')[0].className === 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
