---
id: 5f3c866daec9a49519871816
title: ステップ 30
challengeType: 0
dashedName: step-30
---

# --description--

`article` 要素は一般的に、関連する情報を持つ複数の要素を含みます。 今回はコーヒーのフレーバーとその価格を含みます。 2 つの `p` 要素を `article` 要素内にネストしてください。 その最初の要素には `French Vanilla` というテキストを、また 2 つ目の要素には `3.00` というテキストを入れてください。

# --hints--

既存の `article` 要素を変更しないでください。

```js
assert($('article').length === 1);
```

`article` 要素には `p` 要素が 2 つ必要です。

```js
assert($('article').children('p').length === 2);
```

最初の `p` 要素には `French Vanilla` というテキストが必要です。

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

2 つ目の `p` 要素には `3.00` というテキストが必要です。

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
