---
id: 5f3ef6e0a81099d9a697b550
title: ステップ 64
challengeType: 0
dashedName: step-64
---

# --description--

`footer` の中に、`p` 要素を 1 つ追加してください。 次に、リンク先が `https://www.freecodecamp.org` でテキストが `Visit our website` のアンカー (`a`) 要素を、`p` 要素の中にネストしてください。

# --hints--

既存の `footer` 要素を変更しないでください。

```js
assert(document.querySelectorAll("footer").length === 1);
```

新しい `p` 要素は `footer` 要素の中にネストされている必要があります。 `p` 要素は 1 つだけ必要です。

```js
assert(document.querySelectorAll("footer > p").length === 1);
assert(document.querySelectorAll("footer p").length === 1);
```

新しい `a` 要素は新しく作成した `p` 要素の中にネストされている必要があります。 `a` 要素は 1 つだけ必要です。

```js
assert(document.querySelectorAll("footer > p > a").length === 1);
assert(document.querySelectorAll("footer a").length === 1);
```

新しい `a` 要素には `Visit our website` というテキストが必要です。

```js
assert(document.querySelector("footer > p > a")?.innerText === "Visit our website");
```

新しい `a` 要素は `https://www.freecodecamp.org` にリンクされている必要があります。 `a` 要素は `href` 属性を使用してリンクを作成することを思い出してください。

```js
assert(document.querySelector("footer > p > a")?.href === "https://www.freecodecamp.org/");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
--fcc-editable-region--
      <footer>

      </footer>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

