---
id: 61695197ac34f0407e339882
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

As you've learned in the lectures, webpages should start with a `DOCTYPE html` declaration, followed by an `html` element.

Add a `DOCTYPE html` declaration at the top of the document, and an `html` element after that. Give the `html` element a `lang` attribute with `en` as its value.

# --hints--

Your code should contain the `DOCTYPE` reference.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` reference.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

You should define the document type to be `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

You should close the `DOCTYPE` declaration with a `>` after the type.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Your `html` element should have an opening tag with a `lang` attribute of `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Your `html` element should have a closing tag.

```js
assert.match(code, (/<\/html\s*>/gi);
```

Your `DOCTYPE` declaration should be at the beginning of your HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
