---
id: 60a3e3396c7b40068ad69986
title: الخطوة 28
challengeType: 0
dashedName: step-28
---

# --description--

قم بإنشاء `div` جديد مع `class` بقيمة `three` اسفل عنصر `.two`. سيكون هذا المستطيل الثالث.

# --hints--

يجب عدم تغيير عنصري `.one` و `.two` الحاليين.

```js
assert(document.querySelectorAll('.one').length === 1);
assert(document.querySelectorAll('.two').length === 1);
```

يجب أن يكون `div` الجديد متداخلاً في عنصر `.canvas` الخاص بك.

```js
assert(document.querySelector('.canvas').children[2].tagName === 'DIV');
```

يجب أن يأتي `div` الجديد بعد عنصر `.two` الخاص بك.

```js
assert(document.querySelector('.two').nextElementSibling.tagName === 'DIV');
```

يجب أن يحتوي عنصر `div` الجديد على `class` بقيمة `three`.

```js
assert(document.querySelector('.canvas').children[2].className.split(' ').includes('three'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
