---
id: 60a3e3396c7b40068ad69970
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

قبل أن تبدأ في styling (تصميم) `div` التي قمت بإضافتها، تحتاج إلى ربط CSS الخاص بك بـ HTML.

قم بإضافة عنصر `link` لربط ملف `styles.css` الخاص بك. قم بتعيين `href` إلى `./styles.css`، وتذكر تعيين سمة `rel` إلى `stylesheet`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
assert(/<link/.test(code))
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `rel` بقيمة `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
