---
id: 60a3e3396c7b40068ad6996a
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

في الوقت الحالي، يجب أن تكون على دراية بالعناصر الأساسية التي يجب أن تحتوي عليها صفحة HTML.

عيّن كودك مع تعريف `DOCTYPE` و عنصر `html` يحدد نوع لغة الصفحة بالإنجليزية، وعنصر `head`، وعنصر `body`.

# --hints--

يجب أن يحتوي الكود على إعلان `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يكون لديك علامة `<html>` فتح مع سمة `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي الكود الخاص بك على عنصر `head` داخل عنصر `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

يجب أن يحتوي الكود الخاص بك على عنصر `body` داخل عنصر `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

عنصر `head` الخاص بك يجب أن يأتي قبل عنصر `body` الخاص بك.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
