---
id: 60a3e3396c7b40068ad69973
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

قم بتغيير لون `background-color` للـ canvas إلى `#4d0f00`.

# --hints--

يجب عليك تعيين خاصية `background-color` إلى `#4d0f00`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

يجب أن يحتوي عنصر `.canvas` على `background-color` من `#4d0f00`.

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
