---
id: 60a3e3396c7b40068ad6997e
title: الخطوة 21
challengeType: 0
dashedName: step-21
---

# --description--

الآن يتم توسيط `.one` أفقيًا ، لكن الـ margin top يتجاوز الـ canvas وحدود الإطار، مما يؤدي إلى إزاحة الـ canvas بأكمله لأسفل بمقدار 20 بكسل.

قم بإضافة `padding` من `1px` إلى عنصر `.canvas` لإعطاء عنصر `.one` شيئًا صلبًا لدفعه بعيدًا.

# --hints--

يجب عليك تعيين خاصية `padding` إلى `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

يجب أن يكون للعنصر `.canvas` الخاص بك `padding` بقيمة `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
