---
id: 60a3e3396c7b40068ad6996e
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

الـ margin هو المنطقة الواقعة خارج الصندوق، ويمكن استخدامه للتحكم في المساحة بين الصناديق أو العناصر الأخرى.

هنا العنصر السفلي لديه margin top أكبر، يدفعه إلى أسفل الصفحة.

الآن بعد أن فهمت CSS box model، لنبدأ في لوحة Rothko.

قم بإزالة عنصر `<img>`.

# --hints--

يجب ألا يكون لديك عنصر `img` في الكود الخاص بك.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
