---
id: 60a3e3396c7b40068ad6998b
title: الخطوة 33
challengeType: 0
dashedName: step-33
---

# --description--

من المفيد أن تدفع الـ margins الخاصة بك في اتجاه واحد.

في هذه الحالة، الـ bottom margin لعنصر `.one` يدفع `.two` إلى الأسفل بـ 20 بكسل.

في منتقي `.two` ، استخدم خاصية `margin` لتعيين top margin إلى `0`، و horizontal margin إلى `auto`، و bottom margin إلى `20px`. سيؤدي هذا إلى إزالة الـ top margin، وتوسيط العنصر أفقيًا، وضبط الـ bottom margin الي 20 بكسل.

# --hints--

يجب عليك تعيين خاصية `margin` إلى `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

يجب أن يكون لعنصر `.two` الخاص بك `margin` من `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
