---
id: 60a3e3396c7b40068ad69972
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

أضف خاصية `height` مع القيمة `600px` إلى قاعدة `.canvas` الخاصة بك.

# --hints--

يجب عليك تعيين خاصية `height` إلى `600px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

يجب أن يكون لعنصر `.canvas` الخاص بك `height` من `600px`.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
