---
id: 60a3e3396c7b40068ad6996f
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

أضف عنصر `div` في `body`.

عيّن سمة `class` بقيمة `canvas`. على سبيل المثال، `<div class="my-div">`.

هذا سيعمل كـ canvas للوحتك.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

يجب أن يحتوي عنصر `div` على `class` مع القيمة `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
