---
id: 60a3e3396c7b40068ad6997f
title: الخطوة 22
challengeType: 0
dashedName: step-22
---

# --description--

إضافة 1 بكسل من الـ padding إلى الأعلى، الأسفل، اليسار، واليمين للـ canvas، غير أبعاده إلى 502 pixels في 602 pixels.

استبدل خاصية `padding` بـ `overflow` بقيمة `hidden` - مما يعيد الـ canvas مرة أخرى إلى أبعادها الأصلية.

# --hints--

يجب عليك إزالة خاصية `padding` من منتقى `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

يجب عليك تعيين خاصية `overflow` إلى `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

يجب أن يكون للعنصر `.canvas` الخاص بك `overflow` بقيمة `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
