---
id: 60a3e3396c7b40068ad69979
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

أضف عنصر `div` جديد داخل عنصر `.canvas` الخاص بك.

منح `div` الجديد سمة `class` بقيمة `one`. سيكون هذا المستطيل الأول لك.

# --hints--

يجب عليك إضافة عنصر `div` جديد.

```js
assert(document.querySelectorAll('div').length === 3);
```

يجب عليك دمج عنصر `div` الجديد داخل عنصر `.canvas` الخاص بك.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

يجب أن يحتوي `div` الجديد على سمة `class` ذات قيمة `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
