---
id: 60a3e3396c7b40068ad69978
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

قم باستخدام margins لضبط المسافات خارج عنصر ما.

باستخدام خاصية `margin`، أعطى لعنصر `.frame` الـ margin الرأسي بقيمة `20px` و الـ margin الأفقي بقيمة `auto`. هذا سيحرك الـ frame إلى أسفل 20 بكسل وسيركزه أفقياً في الصفحة.

# --hints--

يجب عليك تعيين خاصية `margin` إلى `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

يجب أن يكون لعنصر `.frame` الخاص بك `margin` من `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
