---
id: 60a3e3396c7b40068ad69970
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

قبل أن تبدأ في تصميم `div` الذي إضافته، تحتاج إلى ربط ملفك من CSS بملفك من HTML.

قم بإضافة عنصر `link` لربط ملف `styles.css` الخاص بك. عيّن `href` بقيمة `styles.css`، وتذكر أن تعيين سمة (attribute) من `rel` بقيمة `stylesheet`.

# --hints--

يجب أن يحتوي كودك على عنصر `link`.

```js
assert(/<link/.test(code))
```

يجب أن يحتوي عنصرك `link` على سمة `rel` بقيمة `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

يجب أن يحتوي عنصرك `link` على سمة `href` بقيمة `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
