---
id: 60a3e3396c7b40068ad6996c
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في CSS box model يتم التعامل مع كل عنصر HTML كصندوق بأربع مناطق.

تخيل أنك تتلقى صندوقاً من بائع التجزئة المفضل لديك على الإنترنت -- المحتوى هو العنصر في الصندوق، أو في حالتنا، header أو paragraph أو عنصر image.

قم بتغيير سمة `src` في `<img>` من `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` إلى `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

يجب أن يحتوي عنصر `img` الجديد على `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

