---
id: 60a3e3396c7b40068ad6996b
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل عنصر `head` ، أضف وسم `meta` الذي يعين `charset` إلى `UTF-8`، وعنصر `title` بقيمة `Rothko Painting`.

ضمن عنصر `body` ، أضف عنصر `img` مع `src` من `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

يجب أن يحتوي الكود الخاص بك علي وسم `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

يجب أن يقوم وسم `meta` بتعيين `charset` إلى `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` يجب أن يكون `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

الكود الخاص بك يجب أن يحتوي على عنصر `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

يجب أن يحتوي عنصر `img` الجديد على `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
