---
id: 60a3e3396c7b40068ad69971
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

حان الوقت لـ CSS.

على الرغم من أن `<div>` الخاص بك لا يحتوي على نص، فإنه لا يزال يعامل كمربع يحتوي على محتوى. اكتب قاعدة CSS تستخدم منتقي الصف `.canvas` وتعين `width` إلى 500 بكسل. إليك قاعدة CSS تعين عرض class الـ `card` إلى 300 بكسل:

```css
.card {
  width: 300px;
}
```

# --hints--

الكود الخاص بك يجب أن يحتوي على منتقي `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

يجب عليك تعيين خاصية `width` إلى `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

يجب أن يكون لعنصر `.canvas` الخاص بك `width` من `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
