---
id: 60a3e3396c7b40068ad69977
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

استخدم padding لضبط المسافات داخل عنصر ما.

في `.frame`، استخدم خاصية `padding` لزيادة المساحة بين عناصر `.frame` و `.canvas` بـ `50px`. وسيؤدي padding إلى زيادة المساحة في الأعلى، الأسفل، اليسار، واليمين لحدود العنصر والـ canvas داخله.

# --hints--

يجب عليك تعيين خاصية `padding` إلى `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

يجب أن يكون لعنصر `.frame` الخاص بك `padding` من `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
