---
id: 68e4276483b0c71658aab1d3
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

As you recall from earlier lessons and workshops, you can extract portions of a string by using the `slice()` method.

Here is a reminder:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Change the value of the `camelCasedVersion` from an empty string to the result of using the `slice()` method on the `lowercaseWord` variable. Pass in the number `0` and `5` which represent the start and end indices for the `slice()` method.

Now you should see the word `"camel"` in the console.

# --hints--

Your `camelCasedVersion` variable should not be an empty string.

```js
assert.isNotEmpty(camelCasedVersion);
```

You should assign `lowercaseWord.slice(0, 5)` to the `camelCasedVersion` variable.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

You should NOT assign the literal string `"camel"` to the `camelCasedVersion`. Please reread through the instructions for the correct approach.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
