---
id: 68e4276483b0c71658aab1d2
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

There will be times where you will need to use multiple methods to create a new string. Over the next few steps, you will use some of the methods you have learned so far to create a new string that uses the camelCase naming convention.

Start by creating a variable called `lowercaseWord` and assign it the string `"camelcase"`. 

Below that `lowercaseWord` variable, create another variable called `camelCasedVersion` and assign it an empty string for now.

Below that `camelCasedVersion` variable, add a `console.log()` with the string `"Camel cased version:"` followed by another `console.log()` for the `camelCasedVersion` variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `lowercaseWord`.

```js
assert.exists(lowercaseWord);
```

Your `lowercaseWord` should be a string.

```js
assert.isString(lowercaseWord);
```

You should assign the string `"camelcase"` to your `lowercaseWord` variable.

```js
assert.equal(lowercaseWord, "camelcase");
```

You should have a `camelCasedVersion` variable.

```js
assert.exists(camelCasedVersion);
```

Your `camelCasedVersion` variable should be a string.

```js
assert.isString(camelCasedVersion);
```

You should assign an empty string to your `camelCasedVersion` variable.

```js
assert.isEmpty(camelCasedVersion);
```

You should log the string `"Camel cased version:"` to the console.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

You should log the `camelCasedVersion` variable to the console.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
