---
id: 68e4276483b0c71658aab1cd
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Now it is time to see the result.

Start by adding a `console.log()` with the string `"Result of trimming whitespace from both ends:"`.

Then below that `console.log()`, add another `console.log()` with the `cleanedInput` variable. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the string `"Result of trimming whitespace from both ends:"` to the console.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

You should log the `cleanedInput` variable to the console.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
