---
id: 68e4276483b0c71658aab1d1
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

If you need to create a new string where all of the letters are in lowercase, you can use the `toLowerCase()` method.  

Start by creating a variable called `lowerCaseInput` and assign it the result of using the `toLowerCase()` method on the `cleanedInput` variable. 

Then log the string `"Result of using the toLowerCase() method:"`.

Finally, add another `console.log()` for the `lowerCaseInput` variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a `lowerCaseInput` variable.

```js
assert.exists(lowerCaseInput);
```

Your `lowerCaseInput` variable should be a string.

```js
assert.isString(lowerCaseInput);
```

You should assign `cleanedInput.toLowerCase()` to your `lowerCaseInput` variable.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

You should log the string `"Result of using the toLowerCase() method:"` to the console.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

You should log the `lowerCaseInput` variable to the console.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
