---
id: 68e4276483b0c71658aab1d0
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

There will be times where you will need to create a new string with a different casing. 

To create a new string with all of the letters converted to uppercase, you can use the `toUpperCase()` method.

Here is an example:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Create a variable called `upperCaseInput` and assign it the result of using the `toUpperCase()` method on the `cleanedInput` variable.

# --hints--

You should have a `upperCaseInput` variable.

```js
assert.exists(upperCaseInput);
```

Your `upperCaseInput` variable should be a string.

```js
assert.isString(upperCaseInput);
```

You should assign `cleanedInput.toUpperCase()` to your `upperCaseInput` variable.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
