---
id: 68f1179eaea3bb59ba434386
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

To trim whitespace at the end of a string, you can use the `trimEnd()` method.

Here is an example:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Create a variable called `trimmedEnd` and assign it the result of using the `trimEnd()` method on the `userInput` variable. 

# --hints--

You should have a `trimmedEnd` variable.

```js
assert.exists(trimmedEnd);
```

Your `trimmedEnd` variable should be a string.

```js
assert.isString(trimmedEnd);
```

You should assign `userInput.trimEnd()` to your `trimmedEnd` variable.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
