---
id: 68f3f5ca1d2a490c97f14f03
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now it is time to see the results of trimming whitespace at the end of the string.

Below your `trimmedEnd` variable, add a `console.log()` with the string `"After using the trimEnd() method, trailing spaces removed:"`. 

Then below that `console.log()`, add another `console.log()` with the `trimmedEnd` variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the string `"After using the trimEnd() method, trailing spaces removed:"` to the console.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

You should log the `trimmedEnd` variable to the console.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--
const trimmedEnd = userInput.trimEnd();

--fcc-editable-region--
```
