---
id: 68f3f3d65d51d80ac0b412fc
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now it is time to see the results of trimming whitespace at the beginning of the string.

Below your `trimmedStart` variable, add a `console.log()` with the string `"After using the trimStart() method, leading spaces removed:"`. 

Then below that `console.log()`, add another `console.log()` with the `trimmedStart` variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the string `"After using the trimStart() method, leading spaces removed:"` to the console.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

You should log the `trimmedStart` variable to the console.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

--fcc-editable-region--
const trimmedStart = userInput.trimStart();

--fcc-editable-region--
```
