---
id: 68f3eeaa65614107ecd5f1f2
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Below your first `console.log()`, add another `console.log()` to log the `userInput` variable to the console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should log the `userInput` variable to the console.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
