---
id: 68f405c8c9df44132c8fe1bb
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

As you recall from earlier lessons, the camelCase naming convention starts with a lowercase letter for the first word, and each subsequent word begins with an uppercase letter followed by lowercase letters.

Here are some reminder examples: 

```md
freeCodeCamp
camelCase
learningIsFun
```

The second word in the `lowercaseWord` variable is `"case"`. To access the `c` in that word, you can use `lowercaseWord[5]`. 

Use the `+` operator to concatenate `lowercaseWord.slice(0, 5)` with the result of using the correct method for converting strings to uppercase on `lowercaseWord[5]`.

Now you should see `camelC` in the console.

# --hints--

You should concatenate `lowercaseWord.slice(0, 5)` with `lowercaseWord[5].toUpperCase()`.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

You should NOT assign the literal string `"camelC"` to the `camelCasedVersion`. Please re read through the instructions for the correct approach.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

