---
id: 69c918841987bf16c1d1babb
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now you will create the main function that does all the work.

Create a function named `shortestPath` that takes three parameters: `matrix`, `startNode`, and `targetNode`.

You will be writing the code for the next couple of steps inside the function, but for now, leave the body empty.

# --hints--

You should create a function named `shortestPath`.

```js
assert.isFunction(shortestPath);
```

Your `shortestPath` function should have the parameters `matrix`, `startNode`, and `targetNode`.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
