---
id: 69c918841987bf16c1d1bac1
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Now you'll add the main loop that drives the algorithm. It runs once for each node in the graph, selecting the closest unvisited node each time.

Inside `shortestPath`, add a `for` loop that runs `n` times using `i` as the loop variable.

At the start of the loop body, using `let` or `var`, declare two variables `minDistance` and `current` set to `INF`, and `-1` respectively. These will track the closest unvisited node found in each iteration.

# --hints--

You should have a `for` loop inside `shortestPath` that runs `n` times. Don't forget to set `i` with `let` or `var` inside it.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

You should have a `minDistance` variable set to `INF` inside the loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

You should have a `current` variable set to `-1` inside the loop.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
