---
id: 69c918841987bf16c1d1bac0
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

As the algorithm runs, you need to track which nodes have already been processed so you don't revisit them.

Create a variable named `visited` and assign it `new Array(n).fill(false)`.

# --hints--

You should create a variable named `visited` inside `shortestPath`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

Your `visited` variable should be set to `new Array(n).fill(false)`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);

}
--fcc-editable-region--
```
