---
id: 6140d263016325162fd076fe
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

創建一個 `.cabin` 選擇器。 設置 `background-color` 爲 `red`, `width` 爲 `20%`，`height`爲 `20%`。

# --hints--

你應該有一個 `.cabin` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin'));
```

你的 `.cabin`選擇器應該有一個 `background-color` 屬性，設置爲 `red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin')?.backgroundColor === 'red');
```

你的 `.cabin` 選擇器應該有一個 `width` 屬性設置爲 `20%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin')?.width === '20%');
```

你的 `.cabin` 選擇器應該有一個 `height` 屬性設置爲 `20%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin')?.height === '20%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
