---
id: 6140c9d35015ae0ba0c250e8
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `body` 元素中添加一個 `div`，並給它一個 `wheel` 的 `class`。

在你的新 `div` 中，添加六個 `span` 元素，其中 `class` 設置爲 `line`，以及六個 `class` 設置爲 `cabin` 的 `div` 元素。

# --hints--

你應該在 `body` 元素中創建一個新的 `div`。

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

`body` 元素中的 `div` 應該有一個 `wheel` 的 `class`。

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

你的 `.wheel` 元素應包含六個 `span` 元素。

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

`.wheel` 元素中的六個 `span` 元素應具有 `line` 的 `class`。

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

你的 `.wheel` 元素應該有六個 `div` 元素。

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

`.wheel` 元素中的六個 `div` 元素應該有 `cabin` 的 `class`。

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
