---
id: 6140cbeec34e970dfe75e710
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

爲你的 `.wheel` 元素創建一個選擇器。 首先將 `border` 設置爲 `2px solid black`，將 `border-radius` 設置爲 `50%`，並將 `margin-left` 設置爲 `50px`。

# --hints--

你應該有一個 `.wheel` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

你的 `.wheel` 選擇器應該有一個 `border` 屬性設置爲 `2px solid black`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

你的 `.wheel` 選擇器應該將 `border-radius` 屬性設置爲 `50%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

你的 `.wheel` 選擇器應該有一個 `margin-left` 屬性設置爲 `50px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
