---
id: 6140d0069049f5139d78da40
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

`transform-origin`屬性用於設置應用CSS變換點。 例如，當執行 `rotate`時（你將在本項目的後面進行），`transform-origin`決定了元素圍繞哪一點旋轉。

給 `.line` 選擇器一個 `0% 0%` 的 `transform-origin` 屬性。 這將使原點從左邊偏移 `0%`，從上面偏移 `0%`，把它設置在元素的左上角。

# --hints--

你的 `.line` 選擇器應該有一個 `transform-origin` 屬性，設置爲 `0% 0%`。

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
