---
id: 6140cdebd39d6a101e747529
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

爲你的 `.line` 元素創建一個選擇器。 首先將 `background-color` 設置爲 `black`，將 `width` 設置爲 `50%`，並將 `height` 設置爲 `2px`。

# --hints--

你應該有一個 `.line` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

你的 `.line` 選擇器應該將 `background-color` 屬性設置爲 `black`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

你的 `.line` 選擇器應該有一個 `width` 屬性值爲 `50%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

你的 `.line` 選擇器應該有一個 `height` 屬性值爲 `2px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
