---
id: 6140cbeec34e970dfe75e710
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstelle einen Selektor für dein `.wheel`-Element. Beginne, indem du `border` auf `2px solid black` setzt, `border-radius` auf `50%` und `margin-left` auf `50px`.

# --hints--

Du solltest einen `.wheel`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Dein `.wheel`-Selektor sollte eine, auf `2px solid black` gesetzte, `border`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Dein `.wheel`-Selektor sollte eine, auf `50%` gesetzte, `border-radius`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Dein `.wheel`-Selektor sollte eine, auf `50px` gesetzte, `margin-left`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
