---
id: 6140cd32d018ed0f600eefce
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Setze die `position`-Eigenschaft des `.wheel`-Selektors auf `absolute`. Setze sowohl die `height` als auch die `width` auf `55vw`.

# --hints--

Dein `.wheel`-Selektor sollte eine `position`-Eigenschaft von `absolute` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

Dein `.wheel`-Selektor sollte eine `height`-Eigenschaft von `55vw` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

Dein `.wheel`-Selektor sollte eine `width`-Eigenschaft von `55vw` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
