---
id: 6140d4bc9db3c81c51a09ab7
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Führe das Muster fort, wende folgende, spezifische Regeln auf die `.cabin`-Elemente an:

- Die zweite `.cabin` sollte eine, auf `17%` gesetzte, `right`-Eigenschaft sowie eine, auf `93.5%` gesetzte, `top`-Eigenschaft haben.
- Die dritte `.cabin` sollte eine, auf `67%` gesetzte, `right`-Eigenschaft sowie eine, auf `93.5%` gesetzte, `top`-Eigenschaft haben.
- Die vierte `.cabin` sollte eine, auf `-8.5%` gesetzte, `left`-Eigenschaft sowie eine, auf `50%` gesetzte, `top`-Eigenschaft haben.
- Die fünfte `.cabin` sollte eine, auf `17%` gesetzte, `left`-Eigenschaft sowie eine, auf `7%` gesetzte, `top`-Eigenschaft haben.
- Die sechste`.cabin` sollte eine, auf `17%` gesetzte, `right`-Eigenschaft sowie eine, auf `7%` gesetzte, `top`-Eigenschaft haben.

# --hints--

Du solltest einen `.cabin:nth-of-type(2)`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

Dein `.cabin:nth-of-type(2)`-Selektor sollte eine, auf `17%` gesetzte, `right`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right === '17%');
```

Dein `.cabin:nth-of-type(2)`-Selektor sollte eine, auf `93.5%` gesetzte, `top`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top === '93.5%');
```

Du solltest einen `.cabin:nth-of-type(3)`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

Dein `.cabin:nth-of-type(3)`-Selektor sollte eine, auf `67%` gesetzte, `right`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right === '67%');
```

Dein `.cabin:nth-of-type(3)`-Selektor sollte eine, auf `93.5%` gesetzte, `top`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top === '93.5%');
```

Du solltest einen `.cabin:nth-of-type(4)`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

Dein `.cabin:nth-of-type(4)`-Selektor sollte eine, auf `-8.5%` gesetzte, `left`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left === '-8.5%');
```

Dein `.cabin:nth-of-type(4)`-Selektor sollte eine, auf `50%` gesetzte, `top`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top === '50%');
```

Du solltest einen `.cabin:nth-of-type(5)`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

Dein `.cabin:nth-of-type(5)`-Selektor sollte eine, auf `17%` gesetzte, `left`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left === '17%');
```

Dein `.cabin:nth-of-type(5)`-Selektor sollte eine, auf `7%` gesetzte, `top`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top === '7%');
```

Du solltest einen `.cabin:nth-of-type(6)`-Selektor haben.

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

Dein `.cabin:nth-of-type(6)`-Selektor sollte eine, auf `17%` gesetzte, `right`-Eigenschaft haben.

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

Dein `.cabin:nth-of-type(6)`-Selektor sollte eine, auf `7%` gesetzte, `top`-Eigenschaft haben.

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
