---
id: 6140df547f09402144e40b92
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

The `animation-name` property is used to link a `@keyframes` rule to a CSS selector. Der Wert dieser Eigenschaft sollte mit dem Namen der `@keyframes`-Regel übereinstimmen. Give your `.wheel` selector an `animation-name` property set to `wheel`.

The `animation-duration` property is used to set how long the animation should sequence to complete. Die Zeit sollte entweder in Sekunden (`s`) oder Millisekunden (`ms`) angegeben werden. Die `animation-duration`-Eigenschaft deines `.wheel`-Selektors sollte `10s` betragen.

# --hints--

Your `.wheel` selector should have an `animation-name` property set to `wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationName === 'wheel');
```

Your `.wheel` selector should have an `animation-duration` property set to `10s`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationDuration === '10s');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}
--fcc-editable-region--

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
   0% {
     transform: rotate(0deg);
   }
   100% {
     transform: rotate(360deg);
   }
}
```
