---
id: 6140cdebd39d6a101e747529
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Erstelle einen Selektor für deine `.line`-Elemente. Beginne, indem du `background-color` auf `black` setzt, `width` auf `50%` und `height` auf `2px`.

# --hints--

Du solltest einen `.line`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

Dein `.line`-Selektor sollte eine `background-color`-Eigenschaft von `black` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

Dein `.line`-Selektor sollte eine `width`-Eigenschaft von `50%` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

Dein `.line`-Selektor sollte eine `height`-Eigenschaft von `2px` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
