---
id: 6140c9d35015ae0ba0c250e8
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge ein `div` innerhalb deines `body`-Elements ein und gib ihm eine `class` von `wheel`.

Füge in deinem neuen `div` sechs `span`-Elemente mit einer `class` von `line` sowie sechs `div`-Elemente mit einer `class` von `cabin` ein.

# --hints--

Du solltest ein neues `div` innerhalb deines `body`-Elements hinzufügen.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

Dein `div` innerhalb deines `body`-Elements sollte eine `class` von `wheel` haben.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

Dein `.wheel`-Element sollte sechs `span`-Elemente enthalten.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

Deine sechs `span`-Elemente innerhalb deines `.wheel`-Elements sollten eine `class` von `line` haben.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

Dein `.wheel`-Element sollte sechs `div`-Elemente enthalten.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

Deine sechs `div`-Elemente innerhalb deines `.wheel`-Elements sollten eine `class` von `cabin` haben.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
