---
id: 6140d0069049f5139d78da40
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Die `transform-origin`-Eigenschaft wird zur Festlegung des Punkts, an welchem eine CSS-Transformation angewendet wird, genutzt. Möchtest du beispielsweise ein `rotate` durchführen (was du im späteren Projektverlauf machen wirst), legt `transform-origin` den Punkt fest, um welchen das Element rotiert.

Gib dem `.line`-Selektor eine `transform-origin`-Eigenschaft von `0% 0%`. Hierdurch wird der Ausgangspunkt um `0%` von links und um `0%` von oben verschoben, wodurch dieser auf die obere linke Ecke des Elements verschoben wird.

# --hints--

Dein `.line`-Selektor sollte eine `transform-origin`-Eigenschaft von `0% 0%` haben.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
