---
id: 6140d10d50636e14695013b2
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Erstelle einen, auf dein zweites `.line`-Element abzielenden, Selektor. Setze die `transform`-Eigenschaft auf `rotate(60deg)`.

Denk daran, dass die `transform`-Eigenschaft es dir erlaubt, die Form eines Elements zu verändern. In diesem Fall dreht der `rotate(60deg)`-Wert das Element um 60 Grad im Uhrzeigersinn, um seinen `transform-origin`-Punkt herum.

# --hints--

Du solltest einen `.line:nth-of-type(2)`-Selektor erstellen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Dein `.line:nth-of-type(2)`-Selektor sollte eine, auf `rotate(60deg)` gesetzte, `transform`-Eigenschaft haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
