---
id: 6140c7e645d8e905819f1dd4
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Beginne mit dem standardmäßigen Boilerplate-Code. Füge deine `DOCTYPE`-Deklaration, dein `html`-Element mit der Sprache auf Englisch gesetzt, deine `head` und `body`-Elemente hinzu.

Füge dein `meta`-Element dem richtigen `charset` hinzu, dein `title`-Element, und ein `link`-Element für die `./styles.css`-Datei.

Setze den `title` auf `Ferris Wheel`.

# --hints--

Dein Code sollte die `DOCTYPE` Referenz beinhalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest nach der `DOCTYPE` Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein abschließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest einen einleitenden `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest einen abschließenden `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest einen öffnenden `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest einen abschließenden `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head` und `body`- Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte innerhalb des `html`-Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Dein Code sollte ein `meta`-Element haben.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Dein `meta`-Element sollte ein `charset`-Attribut mit dem Wert `UTF-8` haben.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Dein Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein Projekt sollte einen Titel namens `Ferris Wheel` haben.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

Dein Code sollte ein `link`-Element haben.

```js
assert.match(code, /<link/)
```

Du solltest ein selbstschließendes `link`-Element haben.

```js
assert(document.querySelectorAll('link').length === 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
