---
id: 6140d1a351e88f159ed54fca
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Erstelle mit demselben Muster einen separaten Selektor für den dritten `.line`, den vierten `.line`, den fünften `.line` und den sechsten `.line`.

Setze die `transform`-Eigenschaft des dritten `.line` auf `rotate(120deg)`, des vierten auf `rotate(180deg)`, des fünften auf `rotate(240deg)` und des sechsten auf `rotate(300deg)`.

# --hints--

Du solltest einen `.line:nth-of-type(3)`-Selektor erstellen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Dein `.line:nth-of-type(3)`-Selektor sollte eine `transform`-Eigenschaft von `rotate(120deg)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform === 'rotate(120deg)');
```

Du solltest einen `.line:nth-of-type(4)`-Selektor erstellen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Dein `.line:nth-of-type(4)`-Selektor sollte eine `transform`-Eigenschaft von `rotate(180deg)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform === 'rotate(180deg)');
```

Du solltest einen `.line:nth-of-type(5)`-Selektor erstellen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Dein `.line:nth-of-type(5)`-Selektor sollte eine `transform`-Eigenschaft von `rotate(240deg)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform === 'rotate(240deg)');
```

Du solltest einen `.line:nth-of-type(6)`-Selektor erstellen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Dein `.line:nth-of-type(6)`-Selektor sollte eine `transform`-Eigenschaft von `rotate(300deg)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform === 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
