---
id: 617ace7d831f9c16a569b38a
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Gib deinem `.wheel`-Selektor eine `max-height` und `max-width`-Eigenschaft, beide auf `500px` gesetzt.

# --hints--

Dein `.wheel`-Selektor sollte eine `max-height`-Eigenschaft von `500px` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight === '500px');
```

Dein `.wheel`-Selektor sollte eine `max-width`-Eigenschaft von `500px` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth === '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
}
--fcc-editable-region--
```
