---
id: 613297a923965e0703b64796
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

你可能已经熟悉 `meta` 元素；它用于指定有关页面的信息，例如标题、描述、关键词和作者。

给你的页面一个 `meta` 元素，然后给它的属性 `charset` 一个合适的值。

`charset` 属性指定了页面的字符编码，`UTF-8` 是现在大多数浏览器支持的唯一编码。

# --hints--

你应该在 `head` 元素中创建一个 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你应该给 `meta` 元素添加一个 `charset` 属性，值为 `UTF-8`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
