---
id: 61329b210dac0b08047fd6ab
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

继续学习 `meta` 元素，`viewport` 定义告诉浏览器如何呈现页面。 包括改善移动设备上的视觉可访问性，并改进 _SEO_（搜索引擎优化）。

添加带有 `content` 属性的 `viewport` 定义，详细说明页面的 `width` 和 `initial-scale`。

# --hints--

你应该在 `head` 中创建另一个 `meta` 元素。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

你应该给 `meta` 元素添加一个 `name` 属性，值为 `viewport`。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

你应该给 `meta` 添加一个 `content` 属性，值为 `width=device-width, initial-scale=1`。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
