---
id: 61329d52e5010e08d9b9d66b
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

另一个关于可访问性和 SEO 的重要 `meta` 元素是 `description` 定义。 `content` 属性的值被搜索引擎用来提供你的页面的描述。

添加一个 `meta` 元素，并将 `name` 属性设置为 `description`，并为其添加一个有用的 `content` 属性。

# --hints--

你应该给 `head` 添加一个新的 `meta` 元素。

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

你应该给 `meta` 添加一个值为 `description` 的 `name` 属性。

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

你应该给 `meta` 添加一个 `content` 属性。

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

`content` 属性的值不应超过 165 个字符。 _这是 Google 接受的最大描述长度。_

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
