---
id: 6133acc353338c0bba9cb553
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

最后，在 `head` 元素中，`title` 元素对于屏幕阅读器理解页面内容很有用。 此外，它是 _SEO_ 的重要组成部分。

给你的页面一个描述性的简洁的 `title`。

# --hints--

你应该给 `head` 元素添加一个 `title` 元素。

```js
assert.exists(document.querySelector('head > title'));
```

`title` 的文本不应超过 60 个字符。

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

尝试让你的 `title` 元素的文本作更多描述性说明。 _提示：至少 15 个字符_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
