---
id: 6133d11ef548f51f876149e3
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

导航是无障碍的核心部分，屏幕阅读器依靠你提供你的页面结构。 这是通过语义 HTML 元素完成的。

将 `header` 元素和 `main` 元素添加到你的页面。

`header` 元素将用于介绍页面，以及提供导航菜单。

`main` 元素将包含页面的核心内容。

# --hints--

你应该将给 `body` 元素添加一个 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你应该给 `body` 元素添加一个 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

`header` 元素应该在 `main` 元素之前。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
