---
id: 613e2546d0594208229ada50
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在 `header` 中，通过嵌套一个 `img` 元素、`h1` 元素和 `nav` 元素来提供有关页面的上下文。

`img` 元素应该指向 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`，并且有一个 `id` 为 `logo`。

`h1` 元素应该包含文本 `HTML/CSS Quiz`。

# --hints--

你应该将 `img` 元素添加到 `header` 元素。

```js
assert.exists(document.querySelector('header > img'));
```

你应该将 `h1` 元素添加到 `header` 元素。

```js
assert.exists(document.querySelector('header > h1'));
```

你应该将 `nav` 元素添加到 `header` 元素。

```js
assert.exists(document.querySelector('header > nav'));
```

你应该按照依次添加 `img` 元素、`h1` 元素和 `nav` 元素。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

你应该给 `img` 元素添加一个 `src` 属性，值为 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

你应该给 `img` 元素添加一个 `id` 属性，值为 `logo`。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

你应该给 `h1` 元素添加文本 `HTML/CSS Quiz`。

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
