---
id: 613e275749ebd008e74bb62e
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

_SVG_（可缩放矢量图形）的一个有用之处是它包含一个 `path` 属性，该属性允许在不影响图像分辨率的情况下缩放图像。

`img` 当前是默认尺寸，这个尺寸太大。 可以使用它的 `id` 作为选择器来缩放图像，并将 `width` 设置为 `max(100px, 18vw)`。

# --hints--

你应该使用 `#logo` 选择器来定位 `img` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

你应该给 `img` 元素添加一个值为 `max(100px, 18vw)` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
