---
id: 6140883f74224508174794c0
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

使 `header` 占用其父容器的全宽，设置它的 `height` 为 `50px`，并将属性 `background-color` 设置为 `#1b1b32`。 然后设置 `display` 属性时使用 _Flexbox_。

# --hints--

你应该使用 `header` 元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

你应该把 `header` 元素 `width` 属性的值设置为 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

你应该把 `header` 的 `height` 属性设置为 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

你应该把 `header` 元素 `background-color` 属性的值设置为 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

你应该把 `header` 元素 `display` 属性设置为 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
