---
id: 62a1166ed9a56d439c0770e7
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Crie três elementos `span` dentro do elemento `#stats`. Dê a cada uma deles a classe `stat`. Em seguida, dê ao primeiro o texto `XP: 0`, ao segundo o texto `Health: 100` e ao terceiro o texto `Gold: 50`.

# --hints--

Você precisa de três elementos `span` dentro do elemento `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Você deve dar aos três novos elementos `span` uma classe chamada `stat`.

```js
assert.exists(document.querySelectorAll('#stats > .stat')?.length, 3);
```

O primeiro elemento `.stat` deve ter o texto `XP: 0` fornecido.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

O segundo elemento `.stat` deve ter o texto `Health: 100` fornecido.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

O terceiro elemento `.stat` deve ter o texto `Gold: 50` fornecido.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
--fcc-editable-region--
<body>
    <div id="game">
        <div id="stats">

        </div>
        <div id="controls"></div>
        <div id="monsterStats"></div>
        <div id="text"></div>
    </div>
</body>
--fcc-editable-region--
</html>
```
