---
id: 62a115879a6d51422652cbfc
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Crie quatro elementos `div` dentro do elemento `#game`. Dê a eles os seguintes valores de `id`, na ordem: `stats`, `controls`, `monsterStats` e `text`.

# --hints--

Você deve criar quatro elementos `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Você deve dar a um dos novos elementos `div` um `id` de `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Você deve dar a um dos novos elementos `div` um `id` de `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Você deve dar a um dos novos elementos `div` um `id` de `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Você deve dar a um dos novos elementos `div` um `id` de `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Você deve colocar os novos elementos `div` na ordem correta.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Você deve colocar os novos elementos `div` dentro do elemento `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
--fcc-editable-region--
<body>
  <div id="game">
  </div>
</body>
--fcc-editable-region--
</html>
```
