---
id: 62a2509ba163e020bb9d84ea
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Agora, você pode começar a escrever o JavaScript. Comece criando um elemento `script`. Este elemento é usado para carregar o JavaScript no arquivo HTML.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Você deve ter apenas um elemento `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

O elemento `script` deve ter uma tag de abertura.

```js
assert.match(code, /<script\s*>/i);
```

O elemento `script` deve ter uma tag de fechamento.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
