---
id: 62a24190868ca51c0b6e83c7
title: Etapa 26
challengeType: 0
dashedName: step-26
---

# --description--

Por enquanto, oculte o seu elemento `#monsterStats` com a propriedade `display`. Não mude nenhum outro estilo.

# --hints--

Você deve ter um seletor `#monsterStats`.

```js
const monsterStats = new __helpers.CSSHelp(document).getStyle('#monsterStats');
assert.exists(monsterStats);
```

O seletor `#monsterStats` deve ter a propriedade `display` de `none`.

```js
const display = new __helpers.CSSHelp(document).getStyle('#monsterStats')?.getPropertyValue('display');
assert.equal(display, 'none');
```

O seletor `#monsterStats` deve ter uma `border` de `1px solid #0a0a23`.

```js
const border = new __helpers.CSSHelp(document).getStyle('#monsterStats')?.getPropertyValue('border');
assert.equal(border, '1px solid rgb(10, 10, 35)');
```

O seletor `#monsterStats` deve ter `5px` de preenchimento.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#monsterStats')?.getPropertyValue('padding');
assert.equal(padding, '5px');
```

O seletor `#monsterStats` deve ter uma `color` de `#ffffff`.

```js
const color = new __helpers.CSSHelp(document).getStyle('#monsterStats')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

O seletor `#monsterStats` deve ter um `background-color` de `#c70d0d`.

```js
const background = new __helpers.CSSHelp(document).getStyle('#monsterStats')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(199, 13, 13)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">Experiência: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Vida: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Ouro: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Vá a loja</button>
        <button id="button2">Vá a caverna</button>
        <button id="button3">Lute com um dragão</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Nome do Monstro: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Vida: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Seja bem-vindo ao Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

--fcc-editable-region--
#monsterStats {

  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}
--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
