---
id: 62a23d1c5f1c93161f3582ae
title: Etapa 19
challengeType: 0
dashedName: step-19
---

# --description--

Semelhante ao elemento `#stats`, o elemento `#monsterStats` necessita de dois elementos `span`. Dê a eles a classe `stat` e dê ao primeiro elemento o texto `Monster Name:` e ao segundo o texto `Health:`. Após o texto de cada um, adicione um elemento `strong` com um elemento `span` vazio aninhado. Dê ao primeiro elemento `span` interno um `id` de `monsterName` e ao segundo elemento `span` interno um `id` de `monsterHealth`.

# --hints--

O elemento `#monsterStats` deve ter dois elementos `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

Os novos elementos `span` devem ter um valor de `class` de `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

O primeiro elemento `span` deve ter um elemento `strong` com um elemento `span` vazio aninhado.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

O segundo elemento `span` deve ter um elemento `strong` com um elemento `span` vazio aninhado.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

O primeiro elemento `span` deve ter o texto `Monster Name:`. Certifique-se de que o espaçamento está correto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

O segundo elemento `span` deve ter o texto `Health:`. Certifique-se de que o espaçamento está correto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

Os elementos `strong` e `span` devem vir após o texto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

O primeiro elemento `span` interno deve ter um `id` de `monsterName`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

O segundo elemento `span` interno deve ter um `id` de `monsterHealth`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
