---
id: 62a240c67f3dbb1a1e6d95ee
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Dê ao `#game` uma largura máxima de `500px` e uma altura máxima de `400px`. Defina a `background-color` como `lightgray` e a `color` como `white`. Use margens para centralizá-lo definindo as margens superior e inferior como `0` e a margem à esquerda e à direita como `auto`. Por fim, dê a ele `10px` de espaçamento (padding) em todos os quatro lados.

# --hints--

Você deve ter um seletor `#game`.

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

O seletor `#game` deve ter um `max-width` com o valor de `500px`.

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

O seletor `#game` deve ter um `max-height` com o valor de `400px`.

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

O elemento `#game` deve ter a `background-color` com o valor de `lightgray`.

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'lightgray');
```

O seletor `#game` deve ter o atributo `color` com o valor `white`.

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'white');
```

O seletor `#game` deve ter uma `margin` definida como `0 auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin');
assert.equal(margin, '0px auto');
```

O seletor `#game` deve ter `10px` de espaçamento em todos os lados.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
