---
id: 62a23cb9bc467a147516b500
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Para o elemento `#controls`, crie três elementos `button`. O primeiro deve ter o `id` definido como `button1` e o texto `Go to store`. O segundo deve ter o `id` definido como `button2` e o texto `Go to cave`. O terceiro deve ter o `id` definido como `button3` e o texto `Fight dragon`.

# --hints--

Você precisa adicionar três elementos `button` ao elemento `#controls`.

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

O primeiro botão deve ter o atributo `id` definido como `button1`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

O primeiro botão deve ter o texto `Go to store`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

O segundo botão deve ter o atributo `id` definido como `button2`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

O segundo botão deve ter o texto `Go to cave`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

O terceiro botão deve ter o atributo `id` definido como `button3`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

O terceiro botão deve ter o texto `Fight dragon`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
--fcc-editable-region--
        <div id="controls">

        </div>
--fcc-editable-region--
        <div id="monsterStats"></div>
        <div id="text"></div>
    </div>
</body>
</html>
```
