---
id: 62a255dae245b52317da824a
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Uma das ferramentas mais poderosas é o console do desenvolvedor. Dependendo do navegador, ele pode ser aberto pressionando `F12` ou `Ctrl+Shift+I`. No Mac, você pode pressionar `Option + ⌘ + C` e selecionar "Console". Você também pode clicar no botão "Console" acima da janela de pré-visualização para ver o console embutido.

O console do desenvolvedor incluirá erros produzidos pelo código, mas você também pode usá-lo para ver valores de variáveis no código, o que é útil para depuração.

Adicione uma linha `console.log("Hello World");` entre as tags `script`. Em seguida, clique no botão "Console" para abrir o console. You should see the text `"Hello World"`.


# --hints--

Você deve adicionar uma linha `console.log("Hello World");` no código.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

A linha `console.log("Hello World");` deve estar entre as tags `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
