---
id: 62a257659d0d1e2456f24ba2
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Antes de começar a escrever o código do projeto, você deve movê-lo para um arquivo próprio para manter as coisas organizadas.

Remova a linha `console.log("Hello World");`. Em seguida, dê ao elemento `script` vazio um atributo `src` definido como `./script.js`.

# --hints--

Você não deve ter uma linha `console.log("Hello World");` no código.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

O elemento `script` deve ter o atributo `src` definido como `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
