---
id: 62a3c2fccf186146b59c6e96
title: Passo 36
challengeType: 0
dashedName: step-36
---

# --description--

`button1` representa o primeiro elemento `button`. Esses elementos têm uma propriedade especial chamada `onclick`, que você pode usar para determinar o que acontece quando alguém clica nesse botão.

Você pode acessar as propriedades em JavaScript de algumas maneiras diferentes. A primeira é com a <dfn>notação de ponto</dfn>. Acessar a propriedade `onclick` de um botão teria essa aparência:

```js
button.onclick
```

Use a notação de ponto para definir a propriedade `onclick` de `button1` para a variável `goStore`. Essa variável será algo que você escreverá mais tarde. Observe que `button1` já foi declarado. Portanto, você não precisa usar `let` ou `const`.

# --hints--

Você deve usar a notação do ponto para acessar a propriedade `onclick` de `button1`.

```js
assert.match(code, /button1\.onclick/);
```

Você não deve usar `let` ou `const`.

```js
assert.notMatch(code, /(let|const)\s+button1\.onclick/);
```

Você deve definir a propriedade `onclick` de `button1` para que seja a variável `goStore`.

```js
assert.match(code, /button1\.onclick\s*=\s*goStore/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");

--fcc-editable-region--
// initialize buttons

--fcc-editable-region--
```
