---
id: 62a3a0a3c0a4b32915d26a6e
title: 6° passo
challengeType: 0
dashedName: step-6
---

# --description--

Remova a linha `console.log("Hello World");` para começar a escrever o código do projeto.

Em JavaScript, uma <dfn>variável</dfn> é usada para armazenar um valor. Para usar uma variável, primeiro você deve <dfn>declarar</dfn>. Por exemplo, para declarar uma variável chamada `camperbot`, você escreve:

```js
let camperbot;
```

A palavra-chave `let` informa ao JavaScript que você está declarando uma variável. Declare uma variável chamada `xp`.

# --hints--

Você não deve ter uma linha `console.log("Hello World");` no código.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Você deve usar a palavra-chave `let` para declarar a variável.

```js
assert.match(code, /let/);
```

Você deve declarar uma variável chamada `xp`.

```js
assert.match(code, /let\s+xp/);
```

Você não deve atribuir um valor à variável.

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

Não se esqueça do ponto e vírgula no final da linha.

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
