---
id: 62a3b41c9494f937560640ab
title: Passo 10
challengeType: 0
dashedName: step-11
---

# --description--

Declare mais duas variáveis chamadas `monsterHealth` e `inventory`, mas não as inicialize.

# --hints--

Você deve usar `let` para declarar a variável `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Você deve usar camelCase para nomear a variável `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` não deve ter um valor.

```js
assert.isUndefined(monsterHealth);
```

Você não deve atribuir um valor à variável `monsterHealth`. Lembre-se do ponto e vírgula.

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

Você deve usar `let` para declarar a variável `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` não deve ter um valor.

```js
assert.isUndefined(inventory);
```

Você não deve atribuir um valor à variável `inventory`. Lembre-se do ponto e vírgula.

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
