---
id: 62a3a75d8466a12e009eff76
title: 7° passo
challengeType: 0
dashedName: step-7
---

# --description--

Inicialize outra variável chamada `health` com um valor de `100`, e uma variável chamada `gold` com um valor de `50`.

# --hints--

Você deve usar `let` para declarar uma variável chamada `health`.

```js
assert.match(code, /let\s+health/);
```

Você deve inicializar uma variável chamada `health` com um valor de `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

Você deve usar `let` para declarar uma variável chamada `gold`.

```js
assert.match(code, /let\s+gold/);
```

Você deve inicializar uma variável chamada `gold` com um valor de `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` deve ter um valor de `100`.

```js
assert.equal(health, 100);
```

`gold` deve ter um valor de `50`.

```js
assert.equal(gold, 50);
```

`xp` ainda deve ter um valor de `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
