---
id: 62a3a75d8466a12e009eff76
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Inicialize outra variável chamada `health` com um valor de `100`, e uma variável chamada `gold` com um valor de `50`.

# --hints--

Você deve usar `var` para declarar uma variável chamada `health`.

```js
assert.match(code, /var health/);
```

Você deve inicializar uma variável chamada `health` com um valor de `100`.

```js
assert.match(code, /var health\s?=\s?100/);
```

Você deve usar `var` para declarar uma variável chamada `gold`.

```js
assert.match(code, /var gold/);
```

Você deve inicializar uma variável chamada `gold` com um valor de `50`.

```js
assert.match(code, /var gold\s?=\s?50/);
```

`health` deve ter um valor de `100`.

```js
assert.equal(health, 100);
```

`gold` deve ter um valor de `50`.

```js
assert.equal(gold, 50);
```

`xp` ainda deve ter um valor de `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
var xp = 0;
--fcc-editable-region--
```
