---
id: 62a3a488b24fb32b91155d56
title: 7° passo
challengeType: 0
dashedName: step-7
---

# --description--

É possível <dfn>atribuir</dfn> um valor às variáveis. Quando você faz isso enquanto declara, isso se chama <dfn>inicialização</dfn>. Por exemplo:

```js
let age = 32;
```

This would initialize the `age` variable with a value of `32`, a number.

Inicialize a variável `xp` para que tenha um valor de `0`, um número.

# --hints--

`xp` deve ter um valor de `0`.

```js
assert.equal(xp, 0);
```

Você deve inicializar a variável `xp` como `0`. Não se esqueça do ponto e vírgula no final da linha.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp;
--fcc-editable-region--
```
