---
id: 62a3b49686792938718b90d3
title: Etapa 12
challengeType: 0
dashedName: step-12
---

# --description--

Todas as variáveis que você atribuiu têm valores que são números. O JavaScript tem vários tipos de dados diferentes. A próxima que você vai usar é a <dfn>string</dfn>. As strings são usadas para armazenar coisas como palavras ou texto. Strings são envolvidas por aspas duplas, aspas simples ou acentos graves/crases (em inglês, backticks). Aqui está um exemplo de declaração de uma variável com uma string:

```js
let developer = "Naomi";
```

Atribua a variável `inventory` para que tenha o valor de `stick`.

# --hints--

Você deve definir `inventory` com a string `stick`.

```js
assert.equal(inventory, "stick");
```

Você deve inicializar a variável `inventory` com a string `stick`.

```js
assert.match(code, /let\s+inventory\s*=\s*('|"|`)stick\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
