---
id: 62a3bec30ea7f941412512dc
title: Passo 33
challengeType: 0
dashedName: step-33
---

# --description--

Assim como você fez com os botões, crie variáveis para os seguintes `id`s e use `querySelector()` para dar a eles o elemento como um valor:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats` e `monsterName`.

Lembre-se de declarar estes com a palavra-chave `const` e nomeie as variáveis para que correspondam aos `id`s.

# --hints--

Você deve declarar uma variável `text` com `const`.

```js
assert.match(code, /const text/);
```

A variável `text` ainda deve ter o valor do elemento `#text`.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

Você deve declarar uma variável `xpText` com `const`.

```js
assert.match(code, /const xpText/);
```

A variável `xpText` ainda deve ter o valor do elemento `#xpText`.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

Você deve declarar uma variável `healthText` com `const`.

```js
assert.match(code, /const healthText/);
```

A variável `healthText` ainda deve ter o valor do elemento `#healthText`.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

Você deve declarar uma variável `goldText` com `const`.

```js
assert.match(code, /const goldText/);
```

A variável `goldText` ainda deve ter o valor do elemento `#goldText`.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

Você deve declarar uma variável `monsterStats` com `const`.

```js
assert.match(code, /const monsterStats/);
```

A variável `monsterStats` ainda deve ter o valor do elemento `#monsterStats`.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

Você deve declarar uma variável `monsterName` com `const`.

```js
assert.match(code, /const monsterName/);
```

A variável `monsterName` ainda deve ter o valor do elemento `#monsterName`.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
