---
id: 62a3a7e4f1060e2fc5ffb34b
title: 9° passo
challengeType: 0
dashedName: step-9
---

# --description--

Create another variable called `currentWeapon` and set it to `0`.

Quando um nome de variável tem várias palavras, a convenção em JavaScript é usar o que chamamos <dfn>camelCase</dfn>. A primeira palavra é minúscula, e a primeira letra de cada palavra seguinte é maiúscula.

```js
let thisIsCamelCase;
```

# --hints--

You should use `let` to declare a variable called `currentWeapon`.

```js
assert.match(code, /let\s+currentWeapon/i);
```

Você deve usar CamelCase para nomear a variável.

```js
assert.match(code, /currentWeapon/);
```

A variável `currentWeapon` deve ser definido como `0`.

```js
assert.equal(currentWeapon, 0);
```

Você deve inicializar a variável como `0`.

```js
assert.match(code, /let\s+currentWeapon\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
