---
id: 62a3b506dbaead396f58a701
title: Etapa 13
challengeType: 0
dashedName: step-13
---

# --description--

O inventário do jogador poderá guardar vários itens. Você precisará usar um tipo de dados que possa fazer isso. Um <dfn>array</dfn> pode ser usado para manter vários valores. Por exemplo:

```js
let order = ["first", "second", "third"];
```

Este é um array que tem três valores. Observe como os valores são separados por vírgulas.

Change your `inventory` variable to be an array with the strings `"stick"`, `"dagger"`, and `"sword"`.

# --hints--

A variável `inventory` deve ser um array.

```js
assert.isArray(inventory);
```

A variável `inventory` deve ter três valores.

```js
assert.lengthOf(inventory, 3);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

Your `inventory` variable should include the string `"dagger"`.

```js
assert.include(inventory, "dagger");
```

Your `inventory` variable should include the string `"sword"`.

```js
assert.include(inventory, "sword");
```

A variável `inventory` deve ter os valores na ordem correta.

```js
assert.deepEqual(inventory, ["stick", "dagger", "sword"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = "stick";
--fcc-editable-region--
```
