---
id: 62a3a7e4f1060e2fc5ffb34b
title: Passo 22
challengeType: 0
dashedName: step-22
---

# --description--

Crie outra variável chamada `currentWeapon` e defina-a como 0.

Quando um nome de variável tem várias palavras, a convenção em JavaScript é usar o que chamamos <dfn>camelCase</dfn>. A primeira palavra é minúscula, e a primeira letra de cada palavra seguinte é maiúscula.

# --hints--

Você deve usar `var` para declarar uma variável chamada `currentWeapon`.

```js
assert.match(code, /var currentweapon/i);
```

Você deve usar CamelCase para nomear a variável.

```js
assert.match(code, /currentWeapon/);
```

A variável `currentWeapon` deve ser definido como `0`.

```js
assert.equal(currentWeapon, 0);
```

Você deve inicializar a variável como `0`.

```js
assert.match(code, /var currentWeapon\s?=\s?0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
var xp = 0;
var health = 100;
var gold = 50;
--fcc-editable-region--
```
