---
id: 62a3b79d520a7f3d0e25afd6
title: Passo 29
challengeType: 0
dashedName: step-29
---

# --description--

O JavaScript interage com o HTML usando o <dfn>Modelo de Objeto de Documento</dfn>, ou DOM. O DOM é uma árvore de objetos que representa o HTML. Você pode acessar o HTML usando o objeto `document`, que representa todo o documento em HTML.

Um método para encontrar elementos específicos no HTML é usando o método `querySelector()`. O método `querySelector()` recebe um seletor de CSS como argumento e retorna o primeiro elemento que corresponde ao seletor. Por exemplo, para encontrar o elemento `<h1>` no HTML, você escreveria:

```js
let h1 = document.querySelector("h1");
```

Observe que `h1` é uma string e corresponde ao seletor CSS que você usaria.

Crie uma variável `button1` e use `querySelector()` para atribuir ao elemento o `id` de `button1`. Lembre-se de que os seletores de CSS `id` são prefixados com um `#`.

# --hints--

Você deve usar `let` para declarar a variável `button1`.

```js
assert.match(code, /let button1/);
```

Você deve usar `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Você deve usar o seletor `#button1`.

```js
assert.match(code, /querySelector\(('|")#button1\1\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
