---
id: 62a3b506dbaead396f58a701
title: Passo 27
challengeType: 0
dashedName: step-27
---

# --description--

O inventário do jogador poderá guardar vários itens. Você precisará usar um tipo de dados que possa fazer isso. Um <dfn>array</dfn> pode ser usado para manter vários valores. Por exemplo:

```js
let order = ["first", "second", "third"];
```

Este é um array que tem três valores. Observe como os valores são separados por vírgulas.

Altere a variável `inventory` para que seja um array com as strings `stick`, `dagger` e `sword`.

# --hints--

A variável `inventory` deve ser um array.

```js
assert.isArray(inventory);
```

A variável `inventory` deve ter três valores.

```js
assert.lengthOf(inventory, 3);
```

A variável `inventory` deve incluir a string `stick`.

```js
assert.include(inventory, "stick");
```

A variável `inventory` deve incluir a string `dagger`.

```js
assert.include(inventory, "dagger");
```

A variável `inventory` deve incluir a string `sword`.

```js
assert.include(inventory, "sword");
```

A variável `inventory` deve ter os valores na ordem correta.

```js
assert.deepEqual(inventory, ["stick", "dagger", "sword"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = "stick";
--fcc-editable-region--
```
