---
id: 62a3b5843544ce3a77459c27
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

For now, you want the player to start with just the `"stick"`. Change the `inventory` array to have `"stick"` as its only value.

# --hints--

A variável `inventory` ainda deve ser um array.

```js
assert.isArray(inventory);
```

A variável `inventory` deve ter somente um valor.

```js
assert.lengthOf(inventory, 1);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick", "dagger", "sword"];
--fcc-editable-region--
```
