---
id: 62a39f5a5790eb27c1e5d4bf
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

A visualização foi alterada para o novo arquivo `script.js`. Lembre-se de que você pode usar as abas acima para alternar entre arquivos.

Adicione a linha `console.log("Hello World");` a este arquivo e veja-o aparecer no console.

# --hints--

Você deve ter uma linha `console.log("Hello World");` no código.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
