---
id: 62a3b3eab50e193608c19fc6
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Declare uma variável chamada `fighting` mas não a inicialize com um valor. Lembre-se de terminar a linha com ponto e vírgula.

# --hints--

Você deve usar `let` para declarar uma variável `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

A variável `fighting` não deve ter um valor.

```js
assert.isUndefined(fighting);
```

Você não deve atribuir um valor à variável `fighting`. Não se esqueça do ponto e vírgula no final da linha.

```js
assert.match(code, /let\s+fighting\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
--fcc-editable-region--
```
