---
id: 62a3b79d520a7f3d0e25afd6
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

O JavaScript interage com o HTML usando o <dfn>Modelo de Objeto de Documento</dfn>, ou DOM. O DOM é uma árvore de objetos que representa o HTML. Você pode acessar o HTML usando o objeto `document`, que representa todo o documento em HTML.

Um método para encontrar elementos específicos no HTML é usando o método `querySelector()`. O método `querySelector()` recebe um seletor de CSS como argumento e retorna o primeiro elemento que corresponde ao seletor. Por exemplo, para encontrar o elemento `<h1>` no HTML, você escreveria:

```js
let h1 = document.querySelector("h1");
```

Observe que `h1` é uma string e corresponde ao seletor CSS que você usaria.

Crie uma variável `button1` e use `querySelector()` para atribuir ao elemento o `id` de `button1`. Lembre-se de que os seletores de CSS `id` são prefixados com um `#`.

# --hints--

Você deve usar `let` para declarar a variável `button1`.

```js
assert.match(code, /let\s+button1/);
```

Você deve usar `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

You should use the `"#button1"` selector.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
