---
id: 62a7bf06d2ad9d1c5024e833
title: Passo 44
challengeType: 0
dashedName: step-44
---

# --description--

Você também precisará atualizar as funções que são executadas quando os botões forem clicados novamente.

Na função `goStore()`, atualize a propriedade `onclick` para cada botão para executar `buyHealth`, `buyWeapon` e `goTown`, respectivamente.

# --hints--

Você deve usar a notação do ponto para acessar a propriedade `onclick` de `button1`.

```js
assert.match(code, /button1\.onclick/);
```

Você não deve usar `let` ou `const` para acessar a propriedade `onclick` de `button1`.

```js
assert.notMatch(code, /(let|const)\s+button1.onclick/);
```

Você deve definir a propriedade `onclick` de `button1` para que seja `buyHealth`.

```js
assert.match(code, /button1\.onclick\s*=\s*buyHealth/);
```

Você deve definir a propriedade `onclick` de `button1` na função `goStore`.

```js
assert.match(goStore.toString(), /button1\.onclick\s*=\s*buyHealth/);
```

Você deve usar a notação do ponto para acessar a propriedade `onclick` de `button2`.

```js
assert.match(code, /button2\.onclick/);
```

Você não deve usar `let` ou `const` para acessar a propriedade `onclick` de `button2`.

```js
assert.notMatch(code, /(let|const)\s+button2.onclick/);
```

Você deve definir a propriedade `onclick` de `button2` para que seja `buyWeapon`.

```js
assert.match(code, /button2\.onclick\s*=\s*buyWeapon/);
```

Você deve definir a propriedade `onclick` de `button2` na função `goStore`.

```js
assert.match(goStore.toString(), /button2\.onclick\s*=\s*buyWeapon/);
```

Você deve usar a notação do ponto para acessar a propriedade `onclick` de `button3`.

```js
assert.match(code, /button3\.onclick/);
```

Você não deve usar `let` ou `const` para acessar a propriedade `onclick` de `button3`.

```js
assert.notMatch(code, /(let|const)\s+button3.onclick/);
```

Você deve definir a propriedade `onclick` de `button3` para que seja `goTown`.

```js
assert.match(code, /button3\.onclick\s*=\s*goTown/);
```

Você deve definir a propriedade `onclick` de `button3` na função `goStore`.

```js
assert.match(goStore.toString(), /button3\.onclick\s*=\s*goTown/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
