---
id: 60f0286404aefb0562a4fdf9
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Add a `title` and a `meta` element to the `head`. Give your project a title of `Registration Form`, and give a `charset` attribute with a value of `UTF-8` to your `meta` element.

# --hints--

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

The `title` element should be within the `head` element.

```js
assert.exists(document.querySelector('head > title'));
```

Your project should have a title of `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Remember, the casing and spelling matters for the title.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
