---
id: 60f81616cff80508badf9ad5
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

At this point, you should be able to submit the form. However, you might notice not much happens.

To make the form more interactive, add the `required` attribute to the `input` elements in the first `fieldset`.

Now, if you try to submit the form without filling in the required fields, you will see an error message.

# --hints--

You should give the first `input` element a `required` attribute.

```js
assert.equal(document.querySelector('input')?.required, true);
```

You should give the second `input` element a `required` attribute.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.required, true);
```

You should give the third `input` element a `required` attribute.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.required, true);
```

You should give the fourth `input` element a `required` attribute.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.required, true);
```

You should not give the `submit` `input` a `required` attribute.

```js
assert.equal(document.querySelector('input[type="submit"]')?.required, false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
