---
id: 63541ef4f96cd82e8e6c788a
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

The `method` attribute specifies how to send form-data to page specified in the `action` attribute. The form-data can be sent as URL variables (with `method="get"`) or as HTTP post transaction (with `method="post"`).

Add `post` value to `method` attribute in `form`.

# --hints--

You should add `post` value to `method` attribute.

```js
assert.equal(document.querySelector('form')?.method, 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'></form>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
