---
id: a6e40f1041b06c996f7b2406
title: Build a First Element Finder
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

In this lab, you will create a function that looks through an array and returns the first element in it that passes a 'truth test'.

Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have a function named `findElement` that accepts an array and a function as arguments.
2. The function should return the first item in the array that passes a 'truth test'. This means that, calling the passed in function `func`, given an element `x`, the 'truth test' is passed if `func(x)` is `true`.
3. If no element passes the test, return `undefined`.

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` should return `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` should return `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
