---
id: 6976d8a7cd2a1a5ffcd7f9c4
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Each course has a cost. You need to store these amounts in variables to use them later. Since these amounts include cents, you will use the float type, which is used to represent decimal numbers. Here's an example of a variable with a float value:

```py
change = 2.35
```

Create four variables: `benjis_bill` set to `37.89`, `mortys_bill` set to `57.34`, `marcys_bill` set to `39.39`, and `my_bill` set to `64.21`. 

# --hints--

You should define a variable named `benjis_bill` and assign it the value `37.89`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('benjis_bill').is_equivalent('benjis_bill = 37.89')
    `))
})
```

You should define a variable named `mortys_bill` and assign it the value `57.34`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('mortys_bill').is_equivalent('mortys_bill = 57.34')
    `))
})
```

You should define a variable named `marcys_bill` and assign it the value `39.39`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('marcys_bill').is_equivalent('marcys_bill = 39.39')
    `))
})
```

You should define a variable named `my_bill` and assign it the value `64.21`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('my_bill').is_equivalent('my_bill = 64.21')
    `))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

--fcc-editable-region--

--fcc-editable-region--
```
