---
id: 6977a8bbcbdd557886bd0718
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

With the tip now included, you have the final amount for the entire group. You have to determine how much each person owes by dividing the total bill by the number of friends. 

In Python, you use the forward slash `/` to perform division. For example:

```py
half = 10 / 2
```

Create a variable named `final_bill` and assign it the result of dividing `running_total` by `num_of_friends`. 

Finally, use the `print()` function to display the string `Bill per person:` followed by a space and the value of `final_bill`.

Note: You might notice that your output has many decimal places (like 198.8299...). This is a common quirk in computer math, and we will learn how to clean it up in the next step.

# --hints--

You should define a variable named `final_bill`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('final_bill')
    `))
})
```

You should use the `/` operator to divide `running_total` by `num_of_friends` and assign the result to your `final_bill` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('final_bill').is_equivalent('final_bill = running_total / num_of_friends')
    `))
})
```

You should print the string `Bill per person:` followed by a space and the `final_bill` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Bill per person:', final_bill)") or _Node(_code).has_call("print(f'Bill per person: {final_bill}')")`))
})
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

benjis_bill = 37.89
mortys_bill = 57.34
marcys_bill = 39.39
my_bill = 64.21

running_total = benjis_bill + mortys_bill + marcys_bill + my_bill
print('Total bill so far:', running_total)

tip = running_total * 0.25
print('Tip amount:', tip)

running_total += tip
print('Total with tip:', running_total)

--fcc-editable-region--

--fcc-editable-region--
```
