---
id: 658255d5f955175b270f251d
title: Step 33
challengeType: 20
dashedName: step-33
---

# --description--

Inside your `if` statement, create a variable named `category` to store the expense category. Assign it a call to `input()` and use the `'Enter category: '` as the argument.

# --hints--

You should declare a variable named `category` in your `if` statement.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("category")`)) })
```

You should assign `input('Enter category: ')` to your `category` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("category").is_equivalent("category = input('Enter category: ')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--       
        if choice == '1':
            amount = float(input('Enter amount: '))
            
--fcc-editable-region--
```
