---
id: 6582267ef39f5047411c1743
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Next, define a function named `print_expenses` that takes one parameter `expenses`. This function will later be used to display each expense in your list.

Fill the body of your new function with a `pass` statement.

# --hints--

You should define a new function named `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Your `print_expenses` function should take a single parameter `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
