---
id: 5dc174fcf86c76b9248c6eb2
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML 元素中擁有開始標籤，如 `<h1>`，和結束標籤，如 `</h1>`。

元素的文本在元素的開始標籤和結束標籤之間。

找到 `h1` 元素，並將其文本更改爲：

`CatPhotoApp`

# --hints--

代碼中的文本應該是 `CatPhotoApp`。 你可能想要檢查你的拼寫。

```js
assert(code.match(/catphotoapp/i));
```

你的 `h1` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('h1'));
```

你的 `h1` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/h1\>/));
```

你應該只有一個 `h1` 元素。 刪除多餘的 `h1` 元素。

```js
assert(document.querySelectorAll('h1').length === 1);
```

你的 `h1` 元素中的文本應該是 `CatPhotoApp`。 你可能忽略了此文本，或者有拼寫錯誤，或沒有把它放在 `h1` 元素的開始和結束標籤之間。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

你似乎使用了一個瀏覽器擴展，正在修改頁面。 一定要關閉所有瀏覽器擴展。

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

