---
id: 5dc17dc8f86c76b9248c6eb5
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

註釋讓你在不影響瀏覽器顯示內容的情況下，留下信息。 它也能讓你的代碼失效。 在 HTML 中，一個註釋以 `<!--` 開始，包含數行文本，並以 `-->` 結束。 例如此註釋 `<!-- TODO: Remove h1 -->` 包含着文本 `TODO: Remove h1`。

使用以下文本在 `p` 元素上方添加註釋：

`TODO: Add link to cat photos`

# --hints--

你的註釋應該以 `<!--` 開始。 你在定義註釋的開頭時缺少了一個或多個字符。

```js
assert(code.match(/<!--/));
```

你的註釋應以 `-->` 結束。  你在定義註釋的結尾時缺少了一個或多個字符。

```js
assert(code.match(/-->/));
```

你的代碼不應該有多餘的開始/結束註釋的字符。 你有一個多餘的 `<!--` 或 `-->` 顯示在瀏覽器中。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

你的註釋應該包含文本 `TODO: Add link to cat photos`。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

你的註釋應該在 `p` 元素之上。 你把順序寫錯了。

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

