---
id: 5dc17d3bf86c76b9248c6eb4
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 元素用於在網站上創建一段文本。 在你的 `h2` 元素下方創建一個 `p` 元素，併爲其提供以下文本：

`See more cat photos in our gallery.`

# --hints--

你的 `p` 元素應該有一個開始標籤。 起始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('p'));
```

你的 `p` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/p\>/));
```

你的 `p` 元素的文本應該是 `See more cat photos in our gallery.` 你可能忽略了文本或有拼寫錯誤。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

你的 `p` 元素應該在 `h2` 元素的下方。 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

