---
id: 5dc1798ff86c76b9248c6eb3
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 到 `h6` 標題元素用於表明其下方內容的重要性。 數字越低，重要性越高，所以 `h2` 元素所具有的重要性要低於 `h1` 元素。 每頁只使用一個 `h1` 元素，並將重要性較低的標題放在重要性較高的標題之下。

在 `h1` 元素下方，添加帶有以下文本的 `h2` 元素：

`Cat Photos`

# --hints--

你的 `h1` 元素應該有一個開始標籤。 開始標籤的語法是：`<elementName>`。

```js
assert(document.querySelector('h1'));
```

你的 `h1` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/h1\>/));
```

你應該只有一個 `h1` 元素。 刪除多餘的。

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

你的 `h1` 元素的文本應該是 “CatPhotoApp”。 你要麼省略了文本，要麼有錯字。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

你的 `h2` 元素應該有一個開始標籤。 開始標籤的語法是：`<elementName>`。

```js
assert(document.querySelector('h2'));
```

你的 `h2` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/h2\>/));
```

你的 `h2` 元素中的文本應該是 “Cat Photos”。 僅將文本 `Cat Photos` 放在開始和結束 `h2` 標籤之間。

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

你的 `h2` 元素應該放在 `h1` 元素下面。 `h1` 元素更重要，必須在 `h2` 元素之上。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

