---
id: 5dc2385ff86c76b9248c6eb7
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 有些元素能夠指明不同的內容區域。 這些元素能讓你的 HTML 易於閱讀，並有助於搜索引擎優化（SEO）和無障礙訪問。

Identify the main section of this page by adding a `<main>` opening tag before the `h1` element, and a `</main>` closing tag after the `p` element.

# --hints--

你的 `main` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('main'));
```

你的 `main` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/main\>/));
```

Your `main` element's opening tag should be below the `body` element. 你把順序寫錯了。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Your `main` element's opening tag should be above the `h1` element. 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

你的 `main` 元素的開始標籤應該在 `p` 元素的下方。 你把順序寫錯了。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/click here to view more cat photos/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
--fcc-editable-region--
  </body>
</html>
```

