---
id: 5dc23991f86c76b9248c6eb8
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在上一步中，你將 `h2` 元素、註釋和 `p` 元素嵌套在 `main` 元素中。 嵌套元素比其父元素多縮進兩個空格，可提高可讀性：

```html
<ul> 
  <li> This `li` element indented </li>
  <li> This `li` element is also indented </li>
</ul>
```

在 `h2`、comment 和 `p` 元素前面再添加兩個空格，這樣你的 HTML 就更具可讀性。

# --hints--

你的代碼應該有一個 `h2` 元素，其中包含 `Cat Photos` 文本。 你可能意外地刪除了它，或者它缺少了開始和結束標籤，或者是文本的內容被改變了。

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

你不應添加示例中的 `ul` 或 `li` 元素。

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

你的 `h2` 元素應位於 `main` 元素的開始標籤下方，並且其開始標籤應在距離行首 6 個空格處的位置。

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s{6}<h2>/));
```

你的代碼應該有註釋。 你從之前的步驟中刪除了註釋。

```js
assert(code.match(/<!--.*-->/));
```

評論的文本應該是 `TODO: Add link to cat photos`。 不要更改註釋的文本或間距。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

你的註釋應位於 `h2` 元素下方，並從距離行首 6 個空格處開始寫。

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

你的代碼應該有一個 `p` 元素。 你刪除了前一步驟中的 `p` 元素。

```js
assert(document.querySelector('p'));
```

`p` 元素的文本應該是 `Click here to view more cat photos.`。 不要更改 `p` 元素的文本、間距或標點符號。

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

你的 `p` 元素應該在註釋的下方，並且它的起始標籤距離行首應有 6 個空格。

```js
assert(code.toLowerCase().match(/-->\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

