---
id: 5dc2385ff86c76b9248c6eb7
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 有些元素能夠指明不同的內容區域。 這些元素能讓你的 HTML 易於閱讀，並有助於搜索引擎優化（SEO）和無障礙訪問。

指明一個主要區域，可通過在 `h1` 元素後面添加一個 `<main>` 開始標籤，並在 `p` 元素後面添加一個 `</main>` 結束標籤來實現。

# --hints--

你的 `main` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('main'));
```

你的 `main` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/main\>/));
```

你的 `main` 元素的開始標籤應該在 `h1` 元素的下方。 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('MAIN'));
```

你的 `main` 元素的開始標籤應該在 `h2` 元素上方。 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('main,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H2'));
```

你的 `main` 元素的開始標籤應該在 `p` 元素的下方。 你把順序寫錯了。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/click here to view more cat photos/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
--fcc-editable-region--
  </body>
</html>
```

