---
id: 5dc24614f86c76b9248c6ebd
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以使用錨點元素（`a`）鏈接到另一個頁面。 例如 `<a href='https://freecodecamp.org'></a>` 會鏈接到 `freecodecamp.org`。

在段落下方添加一個錨點元素並鏈接到 `https://freecatphotoapp.com`。 此時，鏈接不會出現在預覽中。

# --hints--

你的錨點元素（`a`）應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的錨點元素（`a`）應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的錨點元素（`a`）應該在 `p` 元素的下方。 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

你的錨點元素（`a`）缺少 `href` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert(document.querySelector('a').hasAttribute('href'));
```

你的錨點元素（`a`）應該鏈接到 `https://freecatphotoapp.com`。 你可能省略了這個 URL 或者有拼寫錯誤。

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

雖然你將錨點元素（a）的 `href` 屬性指向了正確的鏈接，但建議始終以引號來包裹屬性的值。

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

