---
id: 5dc2385ff86c76b9248c6eb7
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 有些元素能夠指明不同的內容區域。 這些元素能讓你的 HTML 易於閱讀，並有助於搜索引擎優化（SEO）和無障礙訪問。

通過在 `h1` 元素之前添加 `<main>` 開始標籤和在 `p` 元素之後添加 `</main>` 結束標籤來標識此頁面的主要部分。

# --hints--

你的 `main` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('main'));
```

你的 `main` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/main\>/));
```

你的 `main` 元素的開始標籤應該在 `body` 元素的下方。 你把順序寫錯了。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

你的 `main` 元素的開始標籤應該在 `h1` 元素上方。 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

你的 `main` 元素的結束標籤應該在 `p` 元素的下方。 你把順序寫錯了。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/see more cat photos in our gallery/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

