---
id: 5dc24073f86c76b9248c6ebb
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>屬性</dfn>是寫在一個元素的開始標籤中的特殊詞彙，它們用於控制這個元素的行爲。 在 `img` 元素中的 `src` 屬性明確了一個圖片的 URL（圖片所在的位置）。

這是一個 `img` 元素的示例，其中包含一個 `src` 屬性指向 freeCodeCamp logo：

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

在現有的 `img` 元素中，添加一個具有以下 URL 的 `src` 屬性：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

你的代碼應該有一個 `img` 元素。 你可能已經刪除了 `img` 元素，或者你沒有用引號將 `src` 屬性的值括起來。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素應該有一個 `src` 屬性。 你可能忽略了屬性或者有拼寫錯誤。 請確保元素名稱和屬性名稱之間有一個空格。

```js
assert(document.querySelector('img').src);
```

你的 `img` 元素的 `src` 屬性應設置爲 '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'。 你可能忽略了 URL 或者有拼寫錯誤。 URL 的大小寫很重要。

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

儘管你已將 `img` 元素的 `src` 設置爲正確的 URL，但建議始終用引號將屬性的值括起來。

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

