---
id: 5dc23991f86c76b9248c6eb8
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在上一步中，你將 `h1` 元素、`h2` 元素、註釋和 `p` 元素嵌套在 `main` 元素中。 這叫作 *嵌套*。 嵌套的元素（子元素）的位置應該放在被他們嵌套的元素（父元素）中新的一行並且右側保留兩個空格的間距的後面。 這個間距被稱爲縮進，它被用來使 HTML 更容易讀。

`h1` 元素、`h2` 元素和註釋比下面代碼中的 `main` 元素多縮進了兩個空格。 使用鍵盤上的空格鍵在 `p` 元素前面再添加兩個空格，使其縮進得當。

# --hints--

你的代碼應該有一個 `h2` 元素，其中包含 `Cat Photos` 文本。 你可能意外地刪除了它，或者它缺少了開始和結束標籤，或者是文本的內容被改變了。

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

你不應添加示例中的 `ul` 或 `li` 元素。

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

不應更改 `h2` 元素行上的縮進。 它的開始標籤前應該有 6 個空格。 可以重啓步驟來恢復原始縮進。

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

你的代碼應該有註釋。 你從之前的步驟中刪除了註釋。

```js
assert(code.match(/<!--.*-->/));
```

評論的文本應該是 `TODO: Add link to cat photos`。 不要更改註釋的文本或間距。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

你不應更改註釋元素行上的縮進。 它的開始標籤前應該有 6 個空格。 可以重啓步驟來恢復原始縮進。

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

你的代碼應該有一個 `p` 元素。 你刪除了前一步驟中的 `p` 元素。

```js
assert(document.querySelector('p'));
```

`p` 元素的文本應該是 `See more cat photos in our gallery.` 不要更改 `p` 元素的文本、間距或標點符號。

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/see\s+more\s+cat\s+photos\s+in\s+our\s+gallery\.?$/)
);
```

開始 `p` 標籤縮進應該和 `h2` 和註釋元素一致。 它的開始標籤前應該有 6 個空格。

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

