---
id: 5dc24165f86c76b9248c6ebc
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

所有的 `img` 元素都應該有一個 `alt` 屬性。 `alt` 屬性的文本（值）有兩個作用，第一個作用是讓屏幕閱讀器可以知曉圖片的內容，這會對網頁的可訪問性有很大提升；另一個作用是當圖片無法加載時，頁面需要顯示的替代文本。 例如 `<img src="cat.jpg" alt="A cat">` 有一個 `alt` 屬性，其文本爲 `A cat`。

在 `img` 元素中，添加一個帶有以下文本的 `alt` 屬性：

`A cute orange cat lying on its back`

# --hints--

你的代碼應該有一個 `img` 元素。 你刪除了前一步驟中的 `img` 元素。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素應該有一個 `alt` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

你的 `img` 元素的 `alt` 屬性值被寫成別的了，不再是 “A cute orange cat lying on its back” 了。 確保將 `alt` 屬性的值放在引號中。

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

