---
id: 5ddbd81294d8ddc1510a8e56
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

鏈接的文本必須放置在錨點元素（`a`）的起始和閉合標籤之間。 例如，`<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` 就是一個包含着文本 `click here to go to freeCodeCamp.org` 的鏈接。

將錨點文本 `link to cat pictures` 添加到錨點元素中。 這將成爲鏈接的文本。

# --hints--

你的錨點元素（`a`）應該有一個起始標籤。 起始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的錨點元素（`a`）應該有一個閉合標籤。 閉合標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的錨點元素（`a`）中的文本應該是 `link to cat pictures`。 確保將鏈接的文本放置在錨點元素（`a`）的起始和閉合標籤之間。

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

