---
id: 5dfa2407b521be39a3de7be1
title: 步驟 14
challengeType: 0
dashedName: step-14
---

# --description--

要在新標籤頁中打開鏈接，你可以在錨元素（`a`）中使用 `target` 屬性。

`target` 屬性指定了在哪裏打開鏈接文檔。 `target="_blank"` 意思是在新標籤頁或窗口中打開鏈接的文檔。

這是一個帶有 `target` 屬性的 `a` 元素的基本語法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

向錨元素（`a`）的開始標籤中添加值爲 `_blank` 的 `target` 屬性，用以在新的標籤頁中打開鏈接。

# --hints--

你的 `p` 元素應該嵌套一個文本爲 `cat photos` 的錨元素（`a`）。 你可能刪除了它或有拼寫錯誤。

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

你的錨元素（`a`）缺少一個 `target` 屬性。 請檢查在起始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert(document.querySelector('a').hasAttribute('target'));
```

`target` 屬性的值應爲 `_blank`。 你可能省略了這個值或者有拼寫錯誤。 別忘了應該用引號包裹屬性的值。

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

