---
id: 5dfa22d1b521be39a3de7be0
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

你現在可以在預覽中看到圖像旁邊帶下劃線的字樣 `cat photos`。 那是你的鏈接，點擊它一下。

在你的 `p` 元素的文本中，將單詞 `cat photos` 轉換爲指向 `https://freecatphotoapp.com` 的鏈接，將這些單詞放到新創建的錨點元素內。

# --hints--

你應該在 `p` 元素中嵌套一個新的錨點（`a`）元素。

```js
assert($('p > a').length);
```

鏈接的 href 值應爲 `https://freecatphotoapp.com`。 你可能忽略了 href 值或者有拼寫錯誤。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

鏈接的文本應該是 `cat photos`。 你可能忽略了文本或有拼寫錯誤。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

在嵌套了錨點元素（`a`）之後，瀏覽器中顯示的 `p` 元素應該僅爲 `Click here to view more cat photos.`。再次檢查 `p` 和嵌套其中的錨點元素的文本、間距或標點。

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

