---
id: 5dfb6250eacea3f48c6300b2
title: 步驟 22
challengeType: 0
dashedName: step-22
---

# --description--

在無序列表之後，添加一個新圖像，它具有 `src` 屬性，值爲：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

而它的 `alt` 屬性值爲：

`A slice of lasagna on a plate.`

# --hints--

在 `</ul>` 結束標籤之後應該有一個 `img` 元素。

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

新圖像沒有 `alt` 屬性。 檢查開始標籤名稱後是否有空格和/或所有屬性名稱前是否有空格。

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

新圖像的 `alt` 值應爲 `A slice of lasagna on a plate.`。確保將 `alt` 屬性的值放在引號中。

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

新圖像沒有 `src` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

新圖像的 `src` 值應爲 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`。 確保將 `src` 屬性的值放在引號中。

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

儘管你已將新圖像的 `src` 設置爲正確的 URL，但建議始終將屬性值放在引號中。

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

