---
id: 5dfb6a35eacea3f48c6300b4
title: 步驟 24
challengeType: 0
dashedName: step-24
---

# --description--

圖像標題（`figcaption`）元素用於添加標題以描述 `figure` 元素中包含的圖像。 例如，`<figcaption>A cute cat</figcaption>` 添加標題 `A cute cat`。

在嵌套在 `figure` 元素中的圖像之後，添加一個 `figcaption` 元素，其文本爲：

`Cats love lasagna.`

# --hints--

你的 `figcaption` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('figcaption'));
```

你的 `figcaption` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert(code.match(/<\/figcaption\>/));
```

`figcaption` 元素應該嵌套在 `figure` 元素中。

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

Lasagna `img` 元素應該嵌套在 `figure` 元素中。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

嵌套在 `figure` 元素中的 `figcaption` 元素應位於 `img` 元素下方。 你的順序是錯誤的。

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

你的 `figcaption` 元素的文本應該是 `Cats love lasagna.`。 你可能忽略了文本或有拼寫錯誤。

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

