---
id: 5ef9b03c81a63668521804d9
title: 步驟 40
challengeType: 0
dashedName: step-40
---

# --description--

佔位符文本用於提示人們在輸入框中輸入什麼樣的信息。 例如，`<input type="text" placeholder="Email address">`。

將佔位符文本 `cat photo URL` 添加到你的 `input` 元素中。

# --hints--

你要麼刪除了 `input` 元素，要麼它的語法無效。 所有屬性的值都應該放在引號中。

```js
assert($('input').length);
```

你的 `form` 應該只包含 `input` 元素。 刪除 `form` 元素中的任何其他 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

你的 `input` 元素沒有 `placeholder` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert($('input')[0].hasAttribute('placeholder'));
```

你的 `input` 元素應該有一個 `placeholder` 屬性，其值爲 `cat photo URL`。  你可能省略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert(
  $('input')[0]
    .getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

儘管你已將 `input` 元素的 `placeholder` 屬性設置爲 `cat photo URL`，但建議始終將屬性值放在引號中。

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

