---
id: 5ef9b03c81a63668521804e5
title: 步驟 61
challengeType: 0
dashedName: step-61
---

# --description--

爲了使複選框或單選按鈕默認被選中，你需要爲其添加 `checked` 屬性。 無需爲 `checked` 屬性設置值。 只需將單詞 `checked` 添加到 `input` 元素，確保它和其他屬性之間有空格。

默認選中第一個單選按鈕和第一個複選框。

# --hints--

確保仍然有兩個單選按鈕和三個複選框嵌套在它們各自的 `fieldset` 元素中。

```js
assert(
  $('input[type="radio"]').length === 2 &&
    $('fieldset > input[type="checkbox"]').length === 3
);
```

你的第一個單選按鈕，`id` 爲 `indoor`，應該有 `checked` 屬性。

```js
assert($('input[type="radio"]')[0].hasAttribute('checked'));
```

第二個單選按鈕不應具有 `checked` 屬性。

```js
assert(!$('input[type="radio"]')[1].hasAttribute('checked'));
```

第一個複選框缺少 `checked` 屬性。

```js
assert($('input[type="checkbox"]')[0].hasAttribute('checked'));
```

第二個複選框不應具有 `checked` 屬性。

```js
assert(!$('input[type="checkbox"]')[1].hasAttribute('checked'));
```

第三個複選框不應具有 `checked` 屬性。

```js
assert(!$('input[type="checkbox"]')[2].hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

