---
id: 5ef9b03c81a63668521804db
title: 步驟 41
challengeType: 0
dashedName: step-41
---

# --description--

爲了防止用戶在缺少所需信息時提交你的表單，你需要將 `required` 屬性添加到 `input` 元素。 無需爲 `required` 屬性設置值。 只需將單詞 `required` 添加到 `input` 元素，確保它和其他屬性之間有空格。

# --hints--

你要麼刪除了 `input` 元素，要麼它的語法無效。 所有屬性的值都應該放在引號中。

```js
assert($('input').length);
```

你的 `form` 應該只包含 `input` 元素。 刪除 `form` 元素中的任何其他 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

你的 `input` 元素應有 `required` 屬性。 請記住，你只需在 `input` 元素的標籤中添加單詞 `required`。

```js
assert($('input')[0].hasAttribute('required'));
```

不應爲 `required` 屬性賦予值。

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

