---
id: 5ef9b03c81a63668521804d2
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

有序列表（`ol`）的代碼類似於無序列表，但有序列表中的列表項在顯示時會被編號。

在第二個 `section` 元素的最後一個 `h3` 元素之後，添加一個包含以下三個列表項的有序列表：

`flea treatment` `thunder` `other cats`

# --hints--

你的 `ol` 元素應該有一個開始標籤。 開始標籤的語法是：`<elementName>`。

```js
assert(document.querySelector('ol'));
```

你的 `ol` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/ol>/));
```

`ol` 元素應位於第二個 `section` 元素的結束標籤之上。 你把順序寫錯了。

```js
assert($('main > section')[1].lastElementChild.nodeName === 'OL');
```

三個 `li` 元素應該嵌套在 `ol` 元素中。

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

你應該有三個 `li` 元素，其中包含文本 `flea treatment`、`thunder` 和 `other cats`，順序不限。

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

你應該只有一個 `ol` 元素。

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

