---
id: 5ef9b03c81a63668521804d4
title: 步驟 32
challengeType: 0
dashedName: step-32
---

# --description--

`strong` 元素用於指示某些文本非常重要或緊急。

在你剛剛添加的 `figcaption` 中，通過將 `hate` 放在 `strong` 元素中來表明它非常重要。

# --hints--

你的 `strong` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('strong'));
```

你的 `strong` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/strong\>/));
```

你的 `strong` 元素應圍繞 `Cats hate other cats.` 文本中的單詞 `hate`。 你可能忽略了文本或有拼寫錯誤。

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.toLowerCase() === 'hate'
);
```

`figcaption` 的文本應該是 `Cats hate other cats.`。 檢查是否有拼寫錯誤，以及 `strong` 元素的開始和結束標籤周圍是否存在必要的空格。

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

