---
id: 5ef9b03c81a63668521804d6
title: 步驟 35
challengeType: 0
dashedName: step-35
---

# --description--

現在你將添加一個 Web 表單來收集來自用戶的信息。

在 `Cat Form` 標題之後，添加一個 `form` 元素。

# --hints--

你的 `form` 元素應該有一個開始標籤和一個結束標籤。 你可能缺少一個或兩個必需的標籤，或者它們的順序錯誤。

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

你的 `form` 元素標籤的順序不正確。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

嵌套在最後一個 `section` 元素中的 `form` 元素應位於 `h2` 元素之下。 你的 `h2` 元素和 `form` 元素的順序錯誤。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

`form` 元素應該沒有內容。 刪除 `form` 元素標籤之間的任何 HTML 元素或文本。

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

