---
id: 5ef9b03c81a63668521804dc
title: 步驟 44
challengeType: 0
dashedName: step-44
---

# --description--

對於需要從多個選項中獲得一個答案的問題，你可以使用單選按鈕。

下面是一個帶有 `cat` 選項的單選按鈕示例：`<input type="radio"> cat`。 請記住，`input` 元素是<dfn>自閉合的</dfn>。

在文本輸入之前，添加一個具有如下選項的單選按鈕：

`Indoor`

# --hints--

你應該爲單選按鈕創建一個 `input` 元素。 請檢查語法。

```js
assert($('form > input').length >= 2);
```

你的 `input` 元素應該有一個開始標籤，而不是一個結束標籤。

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

你應該只爲單選按鈕添加一個 `input` 元素。 刪除多餘的。

```js
assert($('form > input').length === 2);
```

你的新 `input` 元素應位於現有 `input` 之上，其中 `type` 屬性設置爲 `text`。 你把順序寫錯了。

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

你的新 `input` 元素沒有 `type` 屬性。 檢查確認開始標籤名稱後應該有一個空格。

```js
assert($('input')[0].hasAttribute('type'));
```

你的新 `input` 元素應該只有一個 `type` 屬性。 刪除多餘的。

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

你的新 `input` 元素應該有一個 `type` 屬性，其值爲 `radio`。 你可能落下了這個值或者有拼寫錯誤。 請記住，應該把屬性值放在引號中。

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

儘管你已將新的 `input` 元素的 `type` 屬性設置爲 `radio`，但建議始終將屬性的值放在引號中。

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

`radio` 按鈕的 `Indoor` 文本應該位於它之後而不是之前。

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

`Indoor` 文本應直接位於 `radio` 按鈕的右側。 你可能落下了文本或有拼寫有誤。

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

