---
id: 5ef9b03c81a63668521804df
title: 步驟 46
challengeType: 0
dashedName: step-46
---

# --description--

`id` 屬性用於標識特定的 HTML 元素。 每個 `id` 屬性的值必須不同於整個頁面的所有其他 `id` 值。

將值爲 `indoor` 的 `id` 屬性添加到單選按鈕。 當元素具有多個屬性時，屬性的順序無關緊要。

# --hints--

你的單選按鈕仍應位於 `label` 元素的開始標籤和結束標籤之間。

```js
const labelChildNodes = [...$('form > label')[0].childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

你的單選按鈕應該有一個 `id` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert($('input')[0].hasAttribute('id'));
```

單選按鈕都應該有 `type` 屬性。 請確保在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert($('input')[0].hasAttribute('type'));
```

單選按鈕應該有一個 `id` 屬性，其值爲 `indoor`。 你可能忽略了這個值或者有拼寫錯誤。 所有屬性的值都應該放在引號中。

```js
assert($('input')[0].id.match(/^indoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

