---
id: 5ef9b03c81a63668521804d8
title: 步驟 37
challengeType: 0
dashedName: step-37
---

# --description--

`input` 元素允許你通過多種方式從 Web 表單中收集數據。 與 `img` 元素一樣，`input` 元素是<dfn>自閉合的</dfn>，並且不需要結束標籤。

在 `form` 元素中嵌套一個 `input` 元素。

# --hints--

你的 `form` 元素應具有正確順序的開始標籤和結束標籤。 你可能缺少一個或兩個必需的標籤，或者它們的順序錯誤。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

你的 `form` 元素的開始標籤應該只有一個 `action` 屬性。 刪除你可能輸入的任何其他內容。

```js
assert([...document.querySelector('form').attributes].length < 2);
```

你應該創建一個 `input` 元素。 請檢查語法。

```js
assert(document.querySelector('input'));
```

你的 `input` 元素應該有一個開始標籤，而不是一個結束標籤。

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

你的 `input` 元素應該嵌套在 `form` 元素中。

```js
assert(document.querySelector('form > input'));
```

你的 `form` 應該只包含 `input` 元素。 刪除 `form` 元素標籤之間的任何 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">

        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

