---
id: 5ef9b03c81a63668521804e2
title: 步驟 54
challengeType: 0
dashedName: step-54
---

# --description--

對於可能有多個答案的問題，表單通常使用複選框。 例如，這是一個帶有 `tacos` 選項的複選框：`<input type="checkbox"> tacos`。

在你剛剛添加的 `legend` 元素下，添加一個 `input`，將其 `type` 屬性設置爲 `checkbox`，併爲其指定選項：

`Loving`

# --hints--

複選框的 `input` 元素應該有一個開始標籤，而不是一個結束標籤。

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

你應該只爲你的複選框添加一個 input 元素。 刪除多餘的。

```js
assert($('fieldset > input').length === 1);
```

你的新 `input` 元素應位於 `legend` 元素下方，並帶有文本 `What's your cat's personality?`。 你把順序寫錯了。

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

你的新 `input` 元素沒有 `type` 屬性。 檢查確認開始標籤名稱後應該有一個空格。

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

你的新 `input` 元素應該有一個 `type` 屬性，其值爲 `checkbox`。 你可能忽略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

儘管你已將新的 `input` 元素的 `type` 屬性設置爲 `checkbox`，但建議始終將屬性的值放在引號中。

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

文本 `Loving` 應直接位於複選框的右側。 請確保元素和文本之間有一個空格。 你可能忽略了文本或有拼寫錯誤。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

