---
id: 5ef9b03c81a63668521804d7
title: 步驟 36
challengeType: 0
dashedName: step-36
---

# --description--

`action` 屬性指示應該將表單數據發送到哪裏。 例如，`<form action="/submit-url"></form>` 告訴瀏覽器應該將表單數據發送到路徑`/submit-url`。

將值爲 `https://freecatphotoapp.com/submit-cat-photo` 的 `action` 屬性添加到 `form` 元素。

# --hints--

你的 `form` 元素應具有正確順序的開始標籤和結束標籤。 你可能缺少一個或兩個必需的標籤，或者它們的順序錯誤。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

嵌套在最後一個 `section` 元素中的 `form` 元素應位於 `h2` 元素下方。 你的 `h2` 元素和 `form` 元素的順序錯誤。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

你的 `form` 元素應該沒有內容。 刪除 `form` 元素標籤之間的任何 HTML 元素或文本。

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

你的 `form` 元素沒有 `action` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

你的 `form` 元素應該有一個 `action` 屬性，其值爲 `https://freecatphotoapp.com/submit-cat-photo`。

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

儘管你已將 `action` 屬性設置爲正確的 URL，但建議始終將屬性的值放在引號中。

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

