---
id: 5efb23e70dc218d6c85f89b1
title: 步驟 38
challengeType: 0
dashedName: step-38
---

# --description--

你可以使用 `type` 屬性創建多種輸入。 你可以輕鬆創建密碼字段、重置按鈕或控件，讓用戶從他們的計算機中選擇文件。

通過將值爲 `text` 的 `type` 屬性添加到 `input` 元素，創建一個文本字段以從用戶那裏獲取文本輸入。

# --hints--

你要麼刪除了 `input` 元素，要麼它的語法無效。 如果你添加了屬性，請確保將它們的值放在引號中。

```js
assert($('input').length);
```

你的 `form` 元素應該只包含 `input` 元素。 刪除 `form` 元素標籤之間的任何額外 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

你的 `input` 元素沒有值爲 `text` 的 `type` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert($('input')[0].hasAttribute('type'));
```

你的 `input` 元素應該有一個 `type` 屬性，其值爲 `text`。 你可能忽略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

儘管你已將 `input` 元素的 `type` 屬性設置爲 `text`，但建議始終將屬性的值放在引號中。

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

