---
id: 5efb2c990dc218d6c85f89b2
title: 第 43 步
challengeType: 0
dashedName: step-43
---

# --description--

即使你在文本輸入下方添加了按鈕，它們也會在頁面上彼此相鄰。 這是因爲 `input` 和 `button` 元素都是<dfn>內聯元素</dfn>，它們不會出現在新的行上。

你添加的按鈕將默認提交表單。 然而，依賴默認行爲可能造成混亂。 將值爲 `submit` 的 `type` 屬性添加到 `button` 以明確它是一個提交按鈕。

# --hints--

你的 `button` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('button'));
```

你的 `button` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/button\>/));
```

你的 `button` 元素沒有 `type` 屬性。 檢查確認開始標籤名稱後應該有一個空格。

```js
assert($('button')[0].hasAttribute('type'));
```

你的 `button` 元素應該有一個 `type` 屬性，其值爲 `submit`。 你可能忽略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

儘管你已將 `button` 元素的 `type` 屬性設置爲 `submit`，但建議始終將屬性的值放在引號中。

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

