---
id: 5efc4f528d6a74d05e68af74
title: 步驟 56
challengeType: 0
dashedName: step-56
---

# --description--

還有另一種方法可以將 `input` 元素的文本與元素本身相關聯。 你可以將文本嵌套在 `label` 元素中，並添加與 `input` 元素的 `id` 具有相同值的 `for` 屬性。

通過僅給 `label` 元素添加文本 `Loving`，並給它添加適當的 `for` 屬性值，將文本 `Loving` 與複選框相關聯。

# --hints--

你應該確保該複選框仍然存在。

```js
assert($('input[type="checkbox"]')[0]);
```

你的複選框應該仍然有一個 `id` 屬性，其值爲 `loving`。  應該是 `--fcc-expected--`，而不是 `--fcc-actual--`。

```js
assert.equal($('input[type="checkbox"]')[0].id, 'loving');
```

文本 `Loving` 不應再直接位於複選框的右側。 它應該被包裹在 `label` 元素中。 確保兩個元素之間有空格。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  !checkboxInputElem.nextSibling.nodeValue
    .replace(/\s+/g, ' ')
    .match(/ Loving/i)
);
```

你將需要添加一個新的 `label` 元素，在其中嵌套文本 `Loving`。 確保它有一個開始標籤和一個結束標籤。

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

新的 `label` 元素應直接位於複選框的右側。 確保兩個元素之間有空格。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

新的 `label` 元素沒有 `for` 屬性。 檢查確認開始標籤名稱後應該有一個空格。

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

新的 `label` 元素應該有一個 `for` 屬性，其值爲 `loving`。  應該是 `--fcc-expected--`，而不是 `--fcc-actual--`。

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert.equal(labelElem.getAttribute('for'), 'loving');
```

文本 `Loving` 應該嵌套在新的 `label` 元素中。 你可能忽略了文本或有拼寫錯誤。

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

