---
id: 5efc575c8d6a74d05e68af77
title: 步驟 59
challengeType: 0
dashedName: step-59
---

# --description--

在前一個複選框之後添加最後一個複選框，其 `id` 屬性值爲 `energetic`。 `name` 屬性應該與前面的複選框相同。

還要在包含文本 `Energetic` 的新複選框的右側添加一個 `label` 元素。 確保將 `label` 元素與新複選框相關聯。

# --hints--

你需要添加一個新複選框。

```js
assert($('input[type="checkbox"]').length === 3);
```

你的新複選框應該有一個值爲 `energetic` 的 `id` 屬性和一個值爲 `personality` 的 `name` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
const checkboxes = [...$('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'energetic' &&
      checkbox.getAttribute('name') === 'personality'
  )
);
```

你的新複選框應該在第二個複選框之後。 你把順序寫錯了。

```js
const checkboxes = [...$('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('lazy') < checkboxes.indexOf('energetic'));
```

在新複選框的右側，應該有 `label` 元素和文本 `Energetic`。

```js
const nextElementSibling = $('input[type="checkbox"]')[2].nextElementSibling;
assert(
  nextElementSibling.nodeName === 'LABEL' &&
    nextElementSibling.innerText.replace(/\s+/g, '').match(/^Energetic$/i)
);
```

新的 `label` 應該有一個 `for` 屬性，其值與新複選框的 `id` 屬性值相同。 你可能忽略了這個值或者有拼寫錯誤。

```js
assert(
  $('input[type="checkbox"]')[2].nextElementSibling.getAttribute('for') ===
    'energetic'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality"> <label for="lazy">Lazy</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

