---
id: 5efc518e8d6a74d05e68af75
title: 步驟 57
challengeType: 0
dashedName: step-57
---

# --description--

將值爲 `personality` 的 `name` 屬性添加到複選框 `input` 元素。

雖然你不會在瀏覽器中注意到這一點，但這樣做會使服務器更容易處理你的 Web 表單，尤其是當有多個複選框時。

# --hints--

你應該確保該複選框仍然存在。

```js
assert($('input[type="checkbox"]')[0]);
```

複選框 `input` 元素沒有 `name` 屬性。 檢查確認開始標籤名稱後應該有一個空格。

```js
assert($('input[type="checkbox"]')[0].hasAttribute('name'));
```

複選框 `input` 元素應該有一個 `name` 屬性，其值爲 `personality`。 你可能忽略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert(
  $('input[type="checkbox"]')[0]
    .getAttribute('name')
    .match(/^personality$/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

