---
id: 5f05a1d8e233dff4a68508d8
title: 步驟 47
challengeType: 0
dashedName: step-47
---

# --description--

在第一個單選按鈕下方創建另一個單選按鈕。 將其放在 `label` 元素內部，`label` 內的文字爲 `Outdoor`。 單選按鈕的 `id` 屬性爲 `outdoor`。

# --hints--

應該只有原來的 `Indoor` 單選按鈕 `input` 的 `id` 爲 `indoor`。 如果需要，可以重新啓動步驟來恢復原始的 HTML。

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

不應該修改 `Indoor` 單選按鈕。 如果需要，可以重新啓動步驟來恢復原始的 HTML。

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

應該在新創建的 `label` 元素內添加一個新的 `input` 元素。 請確保新的 `label` 同時有一個開始和結束標籤。

```js
assert(document.querySelectorAll('label input').length === 2);
```

`label` 開始標籤不應該有任何屬性。

```js
assert(code.match(/<label\s*>/g).length === 2);
```

你的 `input` 應該有一個 `id` 屬性，其值爲 `outdoor`。

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

你的新 `input` 元素應該有一個 `type` 屬性，其值爲 `radio`。

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

除了 `input` 嵌套在 `label` 之外，不應該添加任何新的元素。

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

新的 `input` 的左邊不應該有文本。

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

新的單選按鈕的標籤文本必須是 `Outdoor`。 首字母要大寫。

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

新的單選按鈕和標籤應該在 `Indoor` 單選按鈕和標籤下面/後面。 兩者之間不應有其他標籤。

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

