---
id: 5f05a1d8e233dff4a68508d8
title: 步驟 46
challengeType: 0
dashedName: step-46
---

# --description--

在新的 `label` 元素中嵌套另一個帶有選項 `Outdoor` 的單選按鈕。 新的單選按鈕應放在第一個單選按鈕之後。 此外，將其 `id` 屬性值設置爲：

`outdoor`

# --hints--

你將需要添加一個新的 `label` 元素來嵌套新的單選按鈕。 確保它有一個開始標籤和一個結束標籤。

```js
assert(
  document.querySelectorAll('label').length === 2 &&
    code.match(/<\/label\>/g).length === 2
);
```

`Outdoor` 文本應直接位於新 `radio` 按鈕的右側。 請確保元素和文本之間有一個空格。 你可能忽略了文本或有拼寫錯誤。

```js
const radioButtons = [...$('input')];
assert(
  radioButtons.filter((btn) =>
    btn.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Outdoor/i)
  ).length
);
```

你的新單選按鈕和相關標籤應位於第一個的下方。 你把順序寫錯了。

```js
const collection = [
  ...document.querySelectorAll('input[type="radio"]')
].map((node) => node.nextSibling.nodeValue.replace(/\s+/g, ''));
assert(collection.indexOf('Indoor') < collection.indexOf('Outdoor'));
```

你的新單選按鈕應該有一個 `id` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且/或所有的屬性名稱前面也要有一個空格。

```js
assert($('input')[1].hasAttribute('id'));
```

你的新單選按鈕應該有一個 `id` 屬性，其值爲 `outdoor`。 你可能省略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert($('input')[1].id.match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

