---
id: 5f07c98cdb9413cbd4b16750
title: 步驟 17
challengeType: 0
dashedName: step-17
---

# --description--

是時候添加一個新部分了。 在現有 `section` 元素下方添加第二個 `section` 元素。

# --hints--

你的 `section` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert(document.querySelectorAll('section').length >= 2);
```

你應該只添加一個 `section` 開始標籤。 請刪除多餘的。

```js
assert(document.querySelectorAll('section').length === 2);
```

你的 `section` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/section>/g).length >= 2);
```

你應該只添加一個 `section` 結束標籤。 請刪除多餘的。

```js
assert(code.match(/<\/section>/g).length === 2);
```

第二個 `section` 元素不應嵌套在第一個 `section` 元素中。

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

兩個 `section` 元素都應該在 `main` 元素的開始標籤和結束標籤之間。

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

