---
id: 5f07be6ef7412fbad0c5626b
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

在添加任何新內容之前，你應該使用 `section` 元素將貓咪照片內容與未來的內容分開。

將你的 `h2` 元素、註釋、`p` 元素和錨點（`a`）元素嵌套在一個 `section` 元素中。

# --hints--

你的 `section` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert(document.querySelector('section'));
```

你的 `section` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert(code.match(/<\/section\s*>/i));
```

整個 `section` 元素應該在 `main` 元素的開始標籤和結束標籤之間。

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

現有的 `h2` 元素、註釋、`p` 元素和錨（`a`）元素應位於 `section` 元素的開始標籤和結束標籤之間。

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

`h1` 元素不應該被嵌套在 `section` 元素中。

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

