---
id: 62dabe2ef403a12d5d295273
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

現在將 `p` 元素中的文本 `cat photos` 變成了鏈接， 已經不需要 `p` 元素下面的第二個鏈接了。 刪除 `p` 元素下面的錨點元素。

# --hints--

你的代碼應該只包含一個錨點元素（`a`）。

```js
assert(document.querySelectorAll('a').length === 1);
```

`p` 元素內應該還有一個錨點(`a`) 元素。

```js
assert($('p > a').length);
```

鏈接的 `href` 值應爲 `https://freecatphotoapp.com`。 你可能落下了 `href` 值或者有拼寫有誤。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

鏈接的文本應該是 `cat photos`。 你可能落下了文本或有拼寫有誤。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

瀏覽器中顯示的 `p` 元素應該僅爲 `See more cat photos in our gallery.`。再次檢查 `p` 和嵌套其中的錨點元素的文本、間距或標點。

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

`p` 元素下面不應該有 `cat photos` 文本。

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
