---
id: 7cf9b03d81a65668421804c3
title: 步驟 39
challengeType: 0
dashedName: step-39
---

# --description--

爲了通過 `action` 屬性中指定的位置訪問表單的數據，你必須給文本字段一個 `name` 屬性，併爲其分配一個值來表示數據正在提交。 例如，你可以對電子郵件地址文本字段使用以下語法：`<input type="text" name="email">`。

將值爲 `catphotourl` 的 `name` 屬性添加到你的文本字段。

# --hints--

你要麼刪除了 `input` 元素，要麼它的語法無效。 所有屬性的值都應該放在引號中。

```js
assert($('input').length);
```

你的 `form` 應該只包含 `input` 元素。 刪除 `form` 元素中的任何其他 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

你的 `input` 元素沒有 `name` 屬性。 請檢查在開始標籤的名稱後面要有一個空格，且所有的屬性名稱前面也要有一個空格。

```js
assert($('input')[0].hasAttribute('name'));
```

你的 `input` 元素應該有一個 `name` 屬性，其值爲 `catphotourl`。  你可能忽略了這個值或者有拼寫錯誤。 請記住，屬性值應該放在引號中。

```js
assert(
  $('input')[0]
    .getAttribute('name')
    .match(/^catphotourl$/i)
);
```

儘管你已將 `input` 元素的 `name` 屬性設置爲 `catphotourl`，但建議始終用引號將屬性的值括起來。

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

