---
id: 65689aa3d3f2b6aad204a59e
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

If the previous condition is true, add a vertical separator(`'║'`) between the segments.

# --hints--

Use the append method to add the vertical separator to the `row_list`.

```js
assert.match(code, /row_list\.append\(\'║\'\)/)
```

Use the append method on row_list to add the vertical separator.

```js
assert.match(code, /row_list\.append/)
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []

            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
            
                if square_no != 3:
--fcc-editable-region--

--fcc-editable-region--
```
