---
id: 6568991b4d4874a4d5271337
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Initiate an inner loop to iterate over each segment (`part`) of the row.
Use the enumerate function again to get both the index (`square_no`) and the corresponding segment (`part`).

Leave the parameters empty for now.

# --hints--

The inner loop should have `square_no` as the counter and `part` as the element from iterable.

```js
assert.match(code, /for\s+square_no\s*,\s*part/)
```

You should use the `enumerate()` function, leave it empty.

```js
assert.match(code, /for\s+square_no\s*,\s*part\s+in\s+enumerate\s*\(\s*\)\:/)
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
          row_list = []
--fcc-editable-region--

--fcc-editable-region--
```
