---
id: 6568a242a3e1efc22b07274d
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Add the modified `row_empty` string to the `board_string`.

`board_string` is gradually built up as the loop iterates over each row, creating the full ASCII art representation of the sudoku board.

# --hints--

Use compound addition to add the `row_empty` string to the `board_string`.

```js
assert.match(code, /board_string\s*\+\=\s*row_empty/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
--fcc-editable-region--

--fcc-editable-region--
```
