---
id: 6568994faf481da5d37bfa40
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

note: fix test

As a row in a board is divided into three segments, divide the `line` into three parts using list slicing and add `(line[:3], line[3:6], line[6:], start = 1)` to the parameter for enumeration.

# --hints--

Divide the `line` into three parts using `list slicing (line[:3], line[3:6], line[6:])` as the parameter for enumeration.

```js
assert.match(code, /for\s+square_no\s*,\s*part\s+in\s+enumerate\s*\(\s*line\s*\[\s*:\s*3\s*\]\s*\)\:/)
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
          row_list = []
--fcc-editable-region--
          for square_no, part in enumerate():
--fcc-editable-region--
```
