---
id: 656874efd5102b81815c8ef7
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

A new instance of a class is created by using the function notation: `ClassName()`. The instantiation creates an empty object.

The `__init__` method is a special method that allows you to instantiate an object to a customized state. When a class implements an `__init__` method, `__init__` is automatically called upon instantiation.

Create an `__init__` method inside your `Board` class.

# --hints--

Your method should be named `__init__`. Don't add any parameters.

```js
({ test: () => assert(code, /def\s+__init__\s*\(\)\:/)})
``` 

# --seed--

## --seed-contents--

```py
class Board:
--fcc-editable-region--

--fcc-editable-region--
        pass
```
