---
id: 65688f93a1b6e9970f710f62
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

note: fix test

To create the top border of the board, create a variable `upper-lines` and assign it the value of `f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
`

This line creates a string called `upper_lines` that represents the top border of the sudoku board in a visually appealing ASCII art style. It uses special Unicode characters to draw the borders and intersections

# --hints--

Assign the value of `f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
` to a variable named `upper_lines`.

```js
assert.match(
  code,
  /upper_lines\s*=\s*f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'/
);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
--fcc-editable-region--

--fcc-editable-region--
```
