---
id: 6568bfd65322add674039bde
title: Step 36
challengeType: 20
dashedName: step-36
---

# --description--

Next, you will work on a method that checks if a given number can be inserted into a specified row of the sudoku board.

Create a method named `valid_in_row`. It should take three parameters:

- `self`: representing the instance of the class.
- `row`: representing the row index.
- `num`: representing the number to be checked.

Also, don't forget to add the `pass` keyword in the function body.

# --hints--

Your method should be named `valid_in_row`.

```js
({ test: () => assert(code, /def\s+valid_in_row/) });
```

The method should take three parameters: `self`, `row`, and `num`. Order matters. Don't forget to add the `pass` keyword in the function body.

```js
({
  test: () =>
    assert(code, /def\s+valid_in_row\(\s*self\s*,\s*row\s*,\s*num\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

            if index < 8:
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'
                else:
                    board_string += middle_lines
            else:
                board_string += lower_lines

        return board_string

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--

--fcc-editable-region--
```
