---
id: 6568994faf481da5d37bfa40
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now, you need to create the three line segments to pass to the `enumerate` function.

Use list slicing to create the three lists of equal length representing the `line` segment of each 3x3 square and pass them to the `enumerate()` call. Add `start = 1` to start the enumeration from `1` instead of `0`.

# --hints--

Divide the `line` into three parts using list slicing `[line[:3], line[3:6], line[6:]]` as the first parameter for enumeration. Do not forget to use `start = 1` as the second parameter to start the enumeration from 1 instead of 0.

```js
assert.match(code, /\[line\[:3\],\s*line\[3:6\],\s*line\[6:\]\]\s*,\s*start\s*=\s*1/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

--fcc-editable-region--
        for index, line in enumerate(self.board):
          row_list = []
          for square_no, part in enumerate():
--fcc-editable-region--
```
