---
id: 656896ffecbf07a2d3402a93
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Inside the loop, initialize an empty list `row_list` to store the elements of a single row in the sudoku board.

# --hints--

Your list should be named `row_list`.

```js
assert.match(code, /row_list/);
```

`row_list` should be an empty list.

```js
assert.match(code, /row_list\s*=\s*\[\]/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
--fcc-editable-region--

--fcc-editable-region--
```
