---
id: 656874efd5102b81815c8ef7
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

A constructor method is a method that is called whenever an object of the class is created. It is used to initialize the attributes (or properties) of the object. In Python, the constructor method is called __init__ and it is a reserved word.

`self` is the first parameter of the constructor, which is a reference to the instance of the class being created. It is used to access and modify the object's attributes.

Create a constructor method `__init__` of the Board class. Do not add any parameters for now.

# --hints--

Your method should be named `__init__`. Don't add any parameters.

```js
({ test: () => assert(code, /def\s+__init__\s*\(\)\:/)})
``` 

# --seed--

## --seed-contents--

```py
class Board:
--fcc-editable-region--

--fcc-editable-region--
    pass
```
