---
id: 65688f737b0ef396bf0c22d6
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Now you will move to the actual construction of the board, which is a 9x9 gird.

Define a method `__str__` within the `Board` class. Also, add the `self` parameter. This method is automatically called when you use the `str()` function on an instance of the class or when you use `print()` with the object.

# --hints--

Your method should be named `__str__`.

```js
({ test: () => assert(code, /def\s+__str__\s*\(\)\:/)})
```

You should add the parameter `self` to the method.

```js
({ test: () => assert(code, /def\s+__str__\s*\(\s*self\s*\)\:/)})
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--

--fcc-editable-region--
        pass
```
