---
id: 60a3e3396c7b40068ad69979
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Adicione um novo elemento `div` abaixo do elemento `.canvas`.

Dê ao novo elemento `div` o atributo `class` com o valor de `one`. Esse será o seu primeiro retângulo.

# --hints--

Você deve criar um elemento `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Você deve criar um elemento `div` dentro do elemento `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

O novo elemento `div` deve ter o atributo `class` definido como `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
