---
id: 60a3e3396c7b40068ad69977
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Use o preenchimento (padding) para ajustar o espaçamento dentro de um elemento.

Em `.frame`, use a propriedade abreviada `padding` para aumentar o espaço entre os elementos `.frame` e `.canvas` em `50px`. A propriedade abreviada aumentará o espaço acima, abaixo, na esquerda e na direita da borda do elemento e da tela dentro dele.

# --hints--

Você deve definir a propriedade `padding` como `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

O elemento `.frame` deve ter um valor de `padding` igual a `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
