---
id: 60a3e3396c7b40068ad69970
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Antes de começar a estilizar a `div` que você adicionou, você precisa vincular o CSS ao HTML.

Adicione um elemento `link` para vincular o arquivo `styles.css`. Defina o `href` como `styles.css` e lembre-se de definir o atributo `rel` como `stylesheet`.

# --hints--

O código deve ter um elemento `link`.

```js
assert(/<link/.test(code))
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
