---
id: 60a3e3396c7b40068ad6996f
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Adicione um elemento `div` ao `body`.

Defina o atributo `class` como `canvas`. Por exemplo, `<div class="my-div">`.

Esse elemento atuará como a tela para sua pintura.

# --hints--

O código deve ter um elemento `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

O elemento `div` deve uma `class` com o valor `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
