---
id: 60a3e3396c7b40068ad69974
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Toda pintura precisa de uma moldura.

Envolva o elemento `.canvas` em outra `div`. Dê à `div` a classe `frame`.
# --hints--

Você deve adicionar um novo elemento `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

O elemento `.canvas` deve estar dentro do novo elemento `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

O novo elemento `div` deve uma `class` com o valor `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

O novo elemento `div` deve estar dentro do elemento `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
