---
id: 60a3e3396c7b40068ad69972
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Adicione a propriedade `height` com o valor de `600px` à regra `.canvas`.

# --hints--

Você deve definir a propriedade `height` como `600px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

O elemento `.canvas` deve ter uma `height` com o valor de `600px`.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
