---
id: 60a3e3396c7b40068ad69978
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Use margens para ajustar o espaçamento fora de um elemento.

Usando a propriedade `margin`, dê ao elemento `.frame` uma margem vertical de `20px` e uma margem horizontal de `auto`. Isso moverá o frame para baixo 20 pixels e o centralizará na página na horizontal.

# --hints--

Você deve dar à propriedade `margin` o valor de `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

O elemento `.frame` deve ter um valor de `margin` igual à `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
