---
id: 60a3e3396c7b40068ad6996c
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

No modelo de caixa do CSS, todos os elementos do HTML são tratados como uma caixa com quatro áreas.

Imagine que você recebeu uma caixa de sua loja on-line favorita -- o conteúdo é o item na caixa, ou, no nosso caso, um elemento de cabeçalho, de parágrafo ou de imagem.

Altere o atributo `src` na `<img>` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` para `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

O elemento `img` deve ter para o atributo `src` o valor de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

