---
id: 60a3e3396c7b40068ad69986
title: Passo 28
challengeType: 0
dashedName: step-28
---

# --description--

Crie uma `div` com a `class` tendo o valor de `three` logo abaixo do elemento `.two`. Este será o seu terceiro retângulo.

# --hints--

Os elementos `.one` e `.two` existentes não devem ser modificados.

```js
assert(document.querySelectorAll('.one').length === 1);
assert(document.querySelectorAll('.two').length === 1);
```

O novo elemento `div` deve estar dentro do elemento `.canvas`.

```js
assert(document.querySelector('.canvas').children[2].tagName === 'DIV');
```

A nova `div` deve estar depois do elemento `.two`.

```js
assert(document.querySelector('.two').nextElementSibling.tagName === 'DIV');
```

O novo elemento `div` deve uma `class` com o valor `three`.

```js
assert(document.querySelector('.canvas').children[2].className.split(' ').includes('three'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
