---
id: 60a3e3396c7b40068ad6998b
title: Passo 33
challengeType: 0
dashedName: step-33
---

# --description--

É útil fazer com que suas margens empurrem em uma determinada direção.

Neste caso, a margem inferior do elemento `.one` empurra `.two` 20 pixels para baixo.

No seletor `.two`, use a propriedade abreviada `margin` para definir a margem superior como `0`, a margem horizontal como `auto` e a margem inferior como `20px`. Isso removerá a margem superior, centralizará o elemento na horizontal e definirá a margem inferior como 20 pixels.

# --hints--

Você deve definir a propriedade `margin` como `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

O elemento `.two` deve ter um valor de `margin` igual à `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
