---
id: 60a3e3396c7b40068ad6997e
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Agora `.one` está centralizado na horizontal, mas sua margem superior está passando os limites da tela e entrando na borda do quadro, deslocando a tela inteira 20 pixels para baixo.

Adicione um `padding` de `1px` para o elemento `.canvas` para dar ao elemento `.one` algo sólido como barreira.

# --hints--

Você deve dar à propriedade `padding` o valor de `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

O elemento `.canvas` deve ter um valor de `padding` igual a `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
