---
id: 60a3e3396c7b40068ad69971
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

É hora de usar o CSS.

Mesmo que a `<div>` não tenha texto, ela ainda é tratada como uma caixa com conteúdo. Escreva uma nova regra de CSS que use o seletor `.canvas` e defina a `width` como 500 pixels. Aqui está uma regra de CSS que define a largura da classe `card` como 300 pixels:

```css
.card {
  width: 300px;
}
```

# --hints--

Você deve ter um seletor `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Você deve definir a propriedade `width` como `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

O elemento `.canvas` deve ter um `width` com o valor de `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
