---
id: 60a3e3396c7b40068ad6996b
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do elemento `head`, adicione uma tag `meta` que define o `charset` como `UTF-8` e um elemento `title` com o valor `Rothko Painting`.

Dentro do elemento `body`, adicione um elemento `img` com um `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

O código deve ter uma tag `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

A tag `meta` deve definir o `charset` como `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

O código deve ter um elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

O `title` deve ser `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

O código deve ter um elemento `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

O elemento `img` deve ter para o atributo `src` o valor de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
