---
id: 5f332a88dc25a0fd25c7687a
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Jina la mkahawa ni `CAMPER CAFE`. Ongeza kipengele cha `h1` ndani ya kipengele chako cha `main`. Lipe jina la mkahawa kuwa herufi kubwa ili kuifanya ionekane vyema.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Hupaswi kubadilisha kipengele chako cha `main` kilichopo.

```js
assert($('main').length === 1);
```

Kipengee chako cha `h1` kinapaswa kuwekwa kwenye kipengee chako cha `main`.

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

