---
id: 5f33294a6af5e9188dbdb8f3
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Ni wakati wa kuongeza baadhi ya maudhui ya menyu. Ongeza kipengele cha `main` ndani ya kipengele cha `body` kilichopo. Hatimaye itakuwa na maelezo ya bei kuhusu kahawa na desserts zinazotolewa na mkahawa.

# --hints--

Msimbo wako unapaswa kuwa na tagi ya kufungua ya `<main>`.

```js
assert(code.match(/<main>/i));
```

Msimbo wako unapaswa kuwa na tagi ya kufunga ya `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Hupaswi kubadilisha kipengele chako cha `body` kilichopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
assert($('body').length === 1);
```

Tagi yako ya `main` inapaswa kuwa ndani ya tagi yako ya `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
