---
id: 5f332b23c2045fb843337579
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Ili kuwafahamisha wageni kuwa mkahawa huo ulianzishwa mwaka wa 2020, ongeza kipengele cha `p` chini ya kipengele cha `h1` chenye maandishi `Est. 2020`.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `<p>`.

```js
assert(code.match(/<p>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `</p>`.

```js
assert(code.match(/<\/p>/i));
```

Hupaswi kubadilisha kipengele chako cha `h1` kilichopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
assert($('h1').length === 1);
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele chako cha `h1`.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

