---
id: 5f33071498eb2472b87ddee4
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Kama ulivyojifunza katika hatua chache zilizopita za Programu ya Cat Photo, kuna muundo msingi unaohitajika ili kuanza kuunda ukurasa wako wa wavuti.

Ongeza tagi ya `<!DOCTYPE html>`, na kipengele cha `html` chenye sifa ya `lang` ya `en`.

# --hints--

Unapaswa kuwa na tamko la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Unapaswa kuwa na tagi ya ufunguzi ya `<html>` yenye sifa ya `lang` ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Unapaswa kuwa na tagi ya kufunga ya `<html>`. Kumbuka kwamba tagi za kufunga zina `/` inayofuata `<` mabano ya ufunguzi.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
