---
id: 5f331e55dfab7a896e53c3a1
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` ni mojawapo ya vipengele kadhaa vinavyotoa maelezo ya ziada yasiyoonekana kwenye ukurasa wa wavuti, lakini ni muhimu kwa injini za utafutaji au jinsi ukurasa unavyoonyeshwa.

Ndani ya kipengele cha `head`, weka kipengele cha `meta` chenye sifa inayoitwa `charset` iliyowekwa kwa thamani `utf-8` ili kuambia kivinjari jinsi ya kusimba herufi za ukurasa. Kumbuka kuwa vipengele vya `meta` vinajifunga.

# --hints--

Unapaswa kuwa na tagi ya `meta`.

```js
assert(code.match(/<meta.*?\/?>/is));
```

Tagi yako ya `meta` inapaswa kuwa na sifa ya `charset`.

```js
assert(code.match(/<meta\s+charset\s*=/i));
```

Sifa ya `charset` inapaswa kuwa na thamani ya `utf-8`.

```js
assert(code.match(/charset\s*=\s*('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
