---
id: 5f344f9c805cd193c33d829c
title: Hatua ya 11
challengeType: 0
dashedName: step-11
---

# --description--

Unaweza kuongeza mtindo kwa kipengee kwa kukibainisha katika kipengele cha `style` na kukiwekea sifa kama hii:

```css
element {
 property: value;
}
```

Weka kipengee chako cha `h1` kati kwa kuweka sifa ya `text-align` kuwa na thamani ya `center`.

# --hints--

Unapaswa kuwa na kichaguzi cha `h1` katika kipengele chako cha `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

Sifa ya `text-align` inapaswa kuwa na thamani ya `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

Kichaguzi chako cha `h1` kinapaswa kuweka sifa ya `text-align` kuwa `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

