---
id: 5f344fad8bf01691e71a30eb
title: Hatua ya 10
challengeType: 0
dashedName: step-10
---

# --description--

Hadi sasa, umewekewa vikwazo kuhusu uwasilishaji na mwonekano wa maudhui unayounda. Ili kuanza kudhibiti, ongeza kipengele cha `style` ndani ya kipengele cha `head`.

# --hints--

Msimbo wako unapaswa kuwa na tagi ya ufunguzi ya `<style>`.

```js
assert(code.match(/<style\s*>/i));
```

Msimbo wako unapaswa kuwa na tagi ya kufunga ya `</style>`.

```js
assert(code.match(/<\/style\s*>/));
```

Kipengee chako cha `style` kinapaswa kuwekwa kwenye kipengee chako cha `head`.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
