---
id: 5f3477cb2e27333b1ab2b955
title: Hatua ya 16
challengeType: 0
dashedName: step-16
---

# --description--

Sasa unahitaji kuunganisha faili ya `styles.css` ili mitindo itumike tena. Weka kipengele cha kujifunga cha `link` katika kipengele cha `head`. Ipe `rel` thamani ya sifa `stylesheet` na `href` thamani ya `styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
const link = document.querySelector('link');
assert(link);
```

Hupaswi kubadilisha kipengele chako cha `head` kilichopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
assert($('head').length === 1);
```

Unapaswa kuwa na kipengele kimoja cha kujifunga cha `link`.

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const link = document.querySelector('head > link');
assert(link);
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `rel` yenye thamani `stylesheet`.

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
