---
id: 5f3477cbcb6ba47918c1da92
title: Hatua ya 17
challengeType: 0
dashedName: step-17
---

# --description--

Ili muundo wa ukurasa uonekane sawa kwenye simu ya mkononi kama unavyofanya kwenye kompyuta ya mezani au kompyuta ya mkononi, unahitaji kuongeza kipengele cha `meta` chenye sifa maalum ya `content`.

Ongeza yafuatayo ndani ya kipengele cha `head`:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Msimbo wako unapaswa kuwa na vipengele viwili vya `meta`.

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya `name` yenye thamani ya `viewport`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya `content` yenye thamani ya `width=device-width, initial-scale=1.0`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
