---
id: 5f3477cb303c5cb61b43aa9b
title: Hatua ya 18
challengeType: 0
dashedName: step-18
---

# --description--

Maandishi yamewekwa katikati tena kwa hivyo kiungo kwa faili ya CSS kinafanya kazi. Ongeza mtindo mwingine kwenye faili unaobadilisha sifa ya `background-color` kuwa `brown` kwa kipengele cha `body`.

# --hints--

Unapaswa kutumia kichaguzi cha `body`.

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

Unapaswa kuweka sifa a `background-color` kuwa `brown`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

Kipengele chako cha `body` kinapaswa kuwa na mandharinyuma ya `brown`.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

