---
id: 5f3477aefa51bfc29327200b
title: Hatua ya 14
challengeType: 0
dashedName: step-14
---

# --description--

Umeweka mtindo kwa vipengele vitatu kwa kuandika CSS ndani ya tagi ya `style`. Hii inafanya kazi, lakini kwa kuwa kutakuwa na mitindo mingi zaidi, ni bora kuweka mitindo yote kwenye faili tofauti na kuiunganisha.

Tumekuundia faili tofauti ya `styles.css` na kubadilisha mwonekano wa kihariri hadi faili hiyo. Unaweza kubadilisha kati ya faili na vichupo vilivyo juu ya kihariri.

Anza kwa kuandika upya mitindo uliyounda kwenye faili ya `styles.css`. Hakikisha kuwa haujumuishi tagi za kufungua na kufunga za `style`.

# --hints--

Faili yako ya `styles.css` inapaswa kuwa na kiteuzi cha aina ya `h1, h2, p`.

```js
(getUserInput) => {
  assert(getUserInput('editableContents').replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
}
```

Kichaguzi chako kinapaswa kuweka sifa ya `text-align` kuwa `center`.

```js
(getUserInput) => {
  assert(getUserInput('editableContents').match(/text-align:\s*center;?/));
}
```

Msimbo wako haipaswi kuwa na tagi zozote za `<style>`.

```js
  assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
