---
id: 5f3477ae34c1239cafe128be
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Sasa una viteuzi vya aina tatu vilivyo na mtindo sawa. Unaweza kuongeza kikundi sawa cha mitindo kwa vipengele vingi kwa kuunda orodha ya viteuzi. Kila kichaguzi kimetenganishwa na koma kama hii:

```css
selector1, selector2 {
  property: value;
}
```

Futa vichaguzi vitatu vya aina vilivyopo na ubadilishe na orodha moja ya kiteuzi ambayo huweka maandishi kati kwa vipengee vya `h1`, `h2` na `p`.

# --hints--

Unapaswa kutumia kichaguzi cha aina moja kwa vipengele vyote vitatu, `h1, h2, p`. Hakikisha kutumia orodha hiyo.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Unapaswa kuwa na kichaguzi kimoja pekee katika kipengele chako cha `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

