---
id: 5f3477ae8466a9a3d2cc953c
title: Hatua ya 15
challengeType: 0
dashedName: step-15
---

# --description--

Kwa kuwa sasa una CSS katika faili ya `styles.css`, endelea na uondoe kipengele cha `style` na maudhui yake yote. Mara tu inapoondolewa, maandishi yaliyokuwa katikati yatarudishwa kushoto.

# --hints--

Hupaswi kuwa na tagi zozote za `style` katika code yako.

```js
assert(!code.match(/style/i));
```

Hupaswi kuwa na vichaguzi vyovyote vya CSS katika faili yako ya HTML.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

