---
id: 5f3477ae9675db8bb7655b30
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Katika hatua iliyotangulia, ulitumia <dfn>type selector</dfn> kuweka mtindo wa kipengele cha `h1`. Weka kati vipengee vya `h2` na `p` kwa kuongeza kiteuzi cha aina mpya kwa kila kimoja kwenye kipengele cha `style` kilichopo.

# --hints--

Hupaswi kubadilisha kichaguzi cha `h1` kilichopo.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Hupaswi kuongeza tagi mpya ya `style`. Ongeza sheria mpya za CSS kwenye tagi iliyopo ya `style`.

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

Unapaswa kuongeza kichaguzi kipya cha `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Unapaswa kuongeza kichaguzi kipya cha `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Kipengee chako cha `h1` kinapaswa kuwa na `text-align` ya `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Kipengee chako cha `h2` kinapaswa kuwa na `text-align` ya `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Kipengee chako cha `p` kinapaswa kuwa na `text-align` ya `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

