---
id: 5f356ed6199b0cdef1d2be8f
title: Hatua ya 26
challengeType: 0
dashedName: step-26
---

# --description--

Kufikia sasa umekuwa ukitumia vichaguzi vya aina kuweka mtindo kwa vipengee. <dfn>class selecto</dfn> inafafanuliwa kwa jina lililo na nukta moja kwa moja mbele yake, kama hii:

```css
.class-name {
  styles
}
```

Badilisha kichaguzi kilichopo cha `div` kuwa kiteuzi cha darasa kwa kubadilisha `div` na class inayoitwa `menu`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Hupaswi kuwa na kichaguzi cha `div`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('div');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
div {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

