---
id: 5f356ed63c7807a4f1e6d054
title: Hatua ya 21
challengeType: 0
dashedName: step-21
---

# --description--

Lengo sasa ni kufanya `div` isichukue upana mzima wa ukurasa. Sifa ya CSS `width` inafaa kwa hili. Unda kichaguzi kipya cha aina katika laha ya mtindo inayoipa kipengele chako cha `div` upana wa `300px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `div`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('div');
assert(hasDiv);
```

Unapaswa kuweka sifa ya `width` kuwa `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

`div` yako inapaswa kuwa na upana wa 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle('div')?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

