---
id: 5f35e5c4321f818cdc4bed30
title: Hatua ya 29
challengeType: 0
dashedName: step-29
---

# --description--

Inaonekana vizuri. Ni wakati wa kuanza kuongeza baadhi ya vipengee vya menyu. Ongeza kipengele tupu cha `article` chini ya kichwa cha `Coffee`. Itakuwa na ladha na bei ya kila kahawa unayotoa kwa sasa.

# --hints--

Unapaswa kuwa na tagi ya `<article>` ya kufungua.

```js
assert(code.match(/<article>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `</article>`.

```js
assert(code.match(/<\/article>/i));
```

Hupaswi kubadilisha kipengele cha `h2` kilichopo.

```js
assert($('h2').length === 1);
```

Kipengele chako cha `article` kinapaswa kuja baada ya kipengele chako cha `h2`.

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

