---
id: 5f3c866daec9a49519871816
title: Hatua ya 30
challengeType: 0
dashedName: step-30
---

# --description--

Vipengele vya `article` kwa kawaida huwa na vipengele vingi ambavyo vina taarifa zinazohusiana. Katika kesi hii, itakuwa na ladha ya kahawa na bei ya ladha hiyo. Ongeza vipengele viwili vya `p` ndani ya kipengele chako `article`. Maandishi ya wa kwanza yanapaswa kuwa `French Vanilla`, na ya pili `3.00`.

# --hints--

Hupaswi kubadilisha kipengele cha kilichopo cha `article`.

```js
assert($('article').length === 1);
```

Kipengele chako cha `article` kinapaswa kuwa na vipengele viwili vya `p`.

```js
assert($('article').children('p').length === 2);
```

Kipengele chako cha kwanza cha `p` kinapaswa kuwa na maandishi `French Vanilla`.

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

Kipengele chako cha pili cha `p` kinapaswa kuwa na maandishi `3.00`.

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
