---
id: 5f3c866d5414453fc2d7b480
title: Hatua ya 31
challengeType: 0
dashedName: step-31
---

# --description--

Kuanzia chini ya jozi iliyopo ya kahawa/bei, ongeza kahawa na bei zifuatazo kwa kutumia vipengele vya `article` vyenye vipengee viwili vya `p` vilivyowekwa ndani ya kila moja. Kama hapo awali, maandishi ya kipengele cha kwanza cha `p` yanapaswa kuwa na ladha ya kahawa na maandishi ya `p` ya pili yanapaswa kuwa na bei.

```bash
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

# --hints--

Unapaswa kuwa na vipengele vitano vya `article`.

```js
assert($('article').length === 5);
```

Kila kipengele cha `article` kinafaa kuwa na vipengele viwili vya `p`.

```js
const articles = $('article');
assert(articles[0].children.length === 2);
assert(articles[1].children.length === 2);
assert(articles[2].children.length === 2);
assert(articles[3].children.length === 2);
assert(articles[4].children.length === 2);
```

Kipengele chako cha kwanza cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `French Vanilla` na `3.00`.

```js
const children = $('article')[0].children;
assert(children[0].innerText.match(/French Vanilla/i));
assert(children[1].innerText.match(/3\.00/i));
```

Kipengele chako cha pili cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Caramel Macchiato` na `3.75`.

```js
const children = $('article')[1].children;
assert(children[0].innerText.match(/Caramel Macchiato/i));
assert(children[1].innerText.match(/3\.75/i));
```

Kipengele chako cha tatu cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Pumpkin Spice` na `3.50`.

```js
const children = $('article')[2].children;
assert(children[0].innerText.match(/Pumpkin Spice/i));
assert(children[1].innerText.match(/3\.50/i));
```

Kipengele chako cha nne cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Hazelnut` na `4.00`.

```js
const children = $('article')[3].children;
assert(children[0].innerText.match(/Hazelnut/i));
assert(children[1].innerText.match(/4\.00/i));
```

Kipengele chako cha tano cha `article` kinapaswa kuwa na vipengele vya `p` vyenye maandishi `Mocha` na `4.50`.

```js
const children = $('article')[4].children;
assert(children[0].innerText.match(/Mocha/i));
assert(children[1].innerText.match(/4\.50/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
