---
id: 5f3c866de7a5b784048f94b1
title: Hatua ya 36
challengeType: 0
dashedName: step-36
---

# --description--

Hiyo ni aina ya kile unachotaka, lakini sasa itakuwa vyema ikiwa flavor na price zingekuwa kwenye mstari mmoja. Vipengele vya `p` ni vipengee vya <dfn>block-level</dfn>, kwa hivyo huchukua upana mzima wa kipengele chao kikuu.

Ili kuziweka kwenye mstari sawa, unahitaji kutumia mtindo fulani kwenye vipengele vya `p` ili vifanye kama vipengele vya <dfn>inline</dfn>. Ili kufanya hivyo, anza kwa kuongeza sifa ya `class` yenye thamani `item` kwa kipengele cha kwanza cha `article` chini ya kichwa `Coffee`.

# --hints--

Unapaswa kutumia class ya `item` kwenye kipengele chako `article`.

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

Unapaswa kuwa na kipengele kimoja pekee chenye class ya `item`.

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

Kipengele chako cha kwanza cha `article` kinapaswa kuwa na class ya `item`.

```js
const el = document.querySelectorAll('article')[0];

assert.equal(el.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
