---
id: 5f3c866dd0d0275f01d4d847
title: Hatua ya 39
challengeType: 0
dashedName: step-39
---

# --description--

Naam, hiyo haikufanya kazi. Kuweka vipengele vya `p` kama `inline-block` na kuviweka kwenye mistari tofauti katika msimbo huunda nafasi ya ziada upande wa kulia wa kipengele cha kwanza cha `p`, na kusababisha ya pili kuhama hadi mstari unaofuata. Njia moja ya kurekebisha hili ni kufanya kila kipengele cha `p` kuwa na upana kidogo chini ya `50%`.

Badilisha thamani ya `width` iwe `49%` kwa kila darasa ili kuona kitakachotokea.

# --hints--

Unapaswa kuweka sifa ya `width` kuwa `49%` katika kiteuzi chako cha `.flavor`.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert(flavorWidth === '49%');
```

Unapaswa kuweka sifa ya `width` kuwa `49%` katika kiteuzi chako cha `.price`.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert(priceWidth === '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
--fcc-editable-region--
```

