---
id: 5f45b0731d39e15d54df4dfc
title: Hatua ya 79
challengeType: 0
dashedName: step-79
---

# --description--

Rangi chaguo-msingi ya kiungo ambacho bado hakijabofya kwa kawaida ni bluu. Rangi chaguo-msingi ya kiungo ambacho tayari kimetembelewa kutoka kwa ukurasa kwa kawaida ni zambarau.

Ili kufanya kiungo cha `footer` kiwe na rangi sawa bila kujali kiungo kimetembelewa, tumia kiteuzi cha aina kwa kipengele cha nanga (`a`) na utumie thamani `black` kwa sifa ya `color`.

# --hints--

Unapaswa kutumia kichaguzi cha `a`.

```js
const hasASelector = new __helpers.CSSHelp(document).getStyle('a');
assert(hasASelector);
```

Unapaswa kuweka sifa ya `color` kuwa `black`.

```js
const hasColor = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.color === 'black');
```

Kipengele chako cha `a` kinafaa kuwa na `color` ya `black`.

```js
const aColor = new __helpers.CSSHelp(document).getStyle('a')?.getPropertyValue('color');
assert(aColor === 'black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <hr>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <hr class="bottom-line">
      <footer>
        <p>
          <a href="https://www.freecodecamp.org" target="_blank">Visit our website</a>
        </p>
        <p>123 Free Code Camp Drive</p>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
  padding: 20px;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

hr {
  height: 2px;
  background-color: brown;
  border-color: brown;
}

.bottom-line {
  margin-top: 25px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
  margin-top: 5px;
  margin-bottom: 5px;
  font-size: 18px;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}

/* FOOTER */

footer {
  font-size: 14px;
}
--fcc-editable-region--

--fcc-editable-region--

```
