---
id: 6551fe3b1df7c9740f13f270
title: Hatua ya 11
challengeType: 20
dashedName: step-11
---

# --description--

Key aspects of variable naming in Python are:

- Some words are reserved keywords (e.g. `for`, `while`, `True`). They have a special meaning in Python, so you cannot use them for variable names.
- Variable names cannot start with a number, and they can only contain alpha-numeric characters or underscores.
- Variable names are case sensitive, i.e. `my_var` is different to `my_Var` and `MY_VAR`.
- Finally, it is a common convention to write variable names using `snake_case`, where each space is replaced by an underscore character and the words are written in lowercase letters.

Remove both calls to `print()` and declare another variable called `alphabet`. Assign the string `abcdefghijklmnopqrstuvwxyz` to this variable.

# --hints--

You should declare a variable called `alphabet`.

```js
({ test: () => assert(__userGlobals.has("alphabet")) })
```

You should assign the string `abcdefghijklmnopqrstuvwxyz` to your `alphabet` variable.

```js
({ test: () => assert.equal(__userGlobals.get("alphabet"), "abcdefghijklmnopqrstuvwxyz") })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
