---
id: 6551f6f9ec6db270eb83f3db
title: Hatua ya 7
challengeType: 20
dashedName: step-7
---

# --description--

You can see `11` printed on the terminal because `Hello World` contains 11 characters.

Another useful built-in function is `type()`, which returns the data type of a variable. Modify your `print` call to print the data type of `text`.

# --hints--

You should call the `type()` function.

```js
({ test: () => assert.match(code, /type\s*\(.*\)/) })

```

You should pass `text` to the `type()` function.

```js
({ test: () => assert.match(code, /type\s*\(\s*text\s*\)/) })

```

You should pass `type(text)` to the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
