---
id: 65520e6f2b9678799977f24d
title: Hatua ya 15
challengeType: 20
dashedName: step-15
---

# --description--

Remove the last `print()` call. Then, instead of `text[0]`, pass `text[0].lower()` to `find()` and see the output.

# --hints--

You should remove your last `print()` call.

```js
({ test: () => assert.isFalse(/print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/.test(code)) })
```

You should pass `text[0].lower()` to the `find()` method.

```js
({ test: () => assert.match(code, /index\s*=\s*alphabet\.find\s*\(\s*text\[\s*\s*0\s*\]\.lower\s*\(\s*\)\s*\)/) })

```

You should still print your `index` variable.

```js
({ test: () => assert.match(code, /print\s*\(\s*index\s*\)/) })

```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
