---
id: 6553980e0527fa115c705646
title: Hatua ya 39
challengeType: 20
dashedName: step-39
---

# --description--

Now, instead of printing `space!`, use the addition assignment operator to add the space to the current value of `encrypted_text`.

# --hints--

You should remove `print('space!')`.

```js
({ test: () => assert.isFalse(/print\s*\(\s*("|')space!\1\s*\)/.test(code)) })
```

You should use the `+=` operator to add `char` to the current value of `encrypted_text`.

```js
({ test: () => assert.match(code, /encrypted_text\s*\+=\s*char/) })
```

You should use the `+=` operator to add `char` to the current value of `encrypted_text` inside your `if` block.

```js
({ test: () => assert.match(code, /(\s*)if\s+char\s*==\s*("|')\s\2\s*:\s*?\1\1encrypted_text\s*\+=\s*char/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
