---
id: 6555df0c81300b175308557d
title: Hatua ya 84
challengeType: 20
dashedName: step-84
---

# --description--

Below the line you just modified, add another `print()` call to print `Key: python`. Again, use concatenation (`+`) to get this result.

# --hints--

You should have a `print()` call that prints `Key: python` by concatenating the string `Key:` and the value of the `custom_key` variable.

```js
({ test: () => assert.match(code, /print\s*\(\s*("|')Key:\s\1\s*\+\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
encryption = encrypt(text, custom_key)
print('Encrypted text: ' + text)

decryption = decrypt(encryption, custom_key)
print(decryption)
--fcc-editable-region--
```
