---
id: 6555db4a7b788e15795674e5
title: Hatua ya 83
challengeType: 20
dashedName: step-83
---

# --description--

Two or more strings can be concatenated by using the `+` operator. For example: `'Hello' + ' there!'` results in `'Hello there!`.

Modify `print(encryption)` to print `Encrypted text: mrttaqrhknsw ih puggrur`. Use the `+` operator to concatenate `text` to your string and pay attention to the spacing.

# --hints--

You should change your `print(encryption)` call so that it prints `Encrypted text: mrttaqrhknsw ih puggrur` using the `+` operator.

```js
({
  test: () => {
    assert.match(code, /print\s*\(\s*("|')Encrypted\stext:\s\1\s*\+\s*text\s*\)/)
    assert.isFalse(/print\s*\(\s*encryption\s*\)/.test(code))
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--    
encryption = encrypt(text, custom_key)
print(encryption)
decryption = decrypt(encryption, custom_key)
print(decryption)
--fcc-editable-region--
```
