---
id: 656886ec4e4d9b96ded3c565
title: Hatua ya 22
challengeType: 20
dashedName: step-22
---

# --description--

The iteration variable can have any valid name, but it's convenient to give it a meaningful name.

Rename your `i` variable to `char`.

# --hints--

You should replace `i` with `char`.

```js
({ test: () => assert.match(code, /for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
