---
id: 656883980318048fb11a6e3b
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

Now, remove the indentation and see what happens to the output.

# --hints--

You should remove the indentation from the line where you call `print(i)`.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s*^print\s*\(\s*i\s*\)/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
