---
id: 694648acde178bb8202d9512
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

The report card should also show whether the student is currently enrolled.

Declare a variable named `is_student` and assign it the value `True`. This value represents a yes-or-no condition.

# --hints--

You should have an `is_student` variable.

```js
assert(has_variable("is_student"))
```

The variable `is_student` should store the value `True`.

```js
assert(find_variable("is_student").is_equivalent("is_student = True"))
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
