---
id: 69461e4f98594012939bc1c6
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will practice Python data types by building a simple report card printer.

As you learned in previous lessons, variables are assigned in this way:

```py
greeting = 'Hello World'
```

Create a variable called `name` and assign it the string `Alice`. Remember to surround the value with either single or double quotes as shown in the example above.

# --hints--

You should have a `name` variable.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

The variable `name` should store the value `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
