---
id: 6998b99a8e195cd0c1f211df
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will build a simple wildlife tracker using JavaScript objects.

Here is an example of an object:

```js
const animal = {};
```

This creates an empty object.

Now create a variable named `tiger` and assign it an empty object.

# --hints--

You should create a variable named `tiger`.

```js
assert.exists(tiger);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
