---
id: 6998c3fa088be980504d13be
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now add another property to the `tiger` object called `isEndangered`.

Set the `isEndangered` property to `true`.

# --hints--

`tiger` should have an `isEndangered` property.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` should be equal to `true`.

```js
assert.strictEqual(tiger.isEndangered, true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
