---
id: 6998b99a8e195cd0c1f211df
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will build a simple wildlife tracker using JavaScript objects.

Start by creating a variable named `tiger` and assign it an object.

The object should have a property called `species` with the value `"Tiger"`.

# --hints--

You should create a variable named `tiger`.

```js
assert.exists(tiger);
```

tiger should be an object.

```js
assert.isObject(tiger);
```

tiger should have a species property with the value "Tiger".

```js
assert.strictEqual(tiger.species, "Tiger");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
