---
id: 6999faf886634b72be7a819b
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now update the `tiger` object so it includes a property called `species`.

Set its value to `"Tiger"`.

Here is an example of creating an object with a property:

```js
const animal = {
  name: "Lion"
};
```

# --hints--

tiger should be an object.

```js
assert.isObject(tiger);
```

tiger should have a species property with the value "Tiger".

```js
assert.strictEqual(tiger.species, "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
