---
id: 6999ac171928136168654774
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now create a function called `getSpecies`.

The function should take one parameter named `animal`.

Return the `species` property of the `animal` using dot notation.

Here is a small example of accessing a property from an object:

```js
const dog = {
  species: "Dog"
};

console.log(dog.species); // Dog
```

# --hints--

You should create a function named `getSpecies`.

```js
assert.exists(getSpecies);
```

getSpecies should be a function.

```js
assert.isFunction(getSpecies);
```

Calling getSpecies(tiger) should return "Tiger".

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
