---
id: 6999b0d79ea938edcb032237
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now create a function called `getProperty`.

The function should take two parameters: `animal` and `propertyName`.

Return the value of the property using bracket notation.

Here is an example of accessing a property using bracket notation:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

In this example, bracket notation allows you to access a property using a variable.

After creating the function, use `console.log` to call `getProperty(tiger, "species")` and `getProperty(elephant, "age")` so you can see the returned values in the console.

# --hints--

`getProperty` should be a function.

```js
assert.isFunction(getProperty);
```

Calling `getProperty(tiger, "species")` should return `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Calling `getProperty(elephant, "age")` should return `10`.

```js
assert.strictEqual(getProperty(elephant, "age"), 10);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(tiger, 7));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(elephant, "species"));
```
