---
id: 6999b0198d18d402c36cc5fe
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Now create a function called `removeEndangeredStatus`.

The function should take one parameter named `animal`.

Inside the function, remove the `isEndangered` property from the `animal` object using the `delete` keyword.

Return the updated `animal` object.

Here is an example of removing a property using `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly;
  return pet;
}

console.log(removeFlight(bird));
// { species: "Parrot" }
```

In this example, the `canFly` property is removed from the `bird` object.

After creating the function, use `console.log` to call `removeEndangeredStatus(tiger)` so you can see the updated object in the console.

# --hints--

`removeEndangeredStatus` should be a function.

```js
assert.isFunction(removeEndangeredStatus);
```

Calling `removeEndangeredStatus(tiger)` should remove the `isEndangered` property.

```js
removeEndangeredStatus(tiger);
assert.notProperty(tiger, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(tiger, 7));

--fcc-editable-region--

--fcc-editable-region--
```
