---
id: 6999b07c34c2cde21ee6e1d8
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Now create a function called `hasHabitat`.

The function should take one parameter named `animal`.

Return `true` if the `animal` object has a property called `"habitat"`.

Return `false` otherwise.

Use the `hasOwnProperty` method.

# --hints--

`hasHabitat` should be a function.

```js
assert.isFunction(hasHabitat);
```

Calling `hasHabitat(tiger)` should return `true` after adding a habitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Calling `hasHabitat(elephant)` should return `false` if no habitat exists.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

function getSpecies(animal) {
  return animal.species;
}

function addHabitat(animal, habitat) {
  animal.habitat = habitat;
  return animal;
}

function updateAge(animal, newAge) {
  animal.age = newAge;
  return animal;
}

function removeEndangeredStatus(animal) {
  delete animal.isEndangered;
  return animal;
}

--fcc-editable-region--

--fcc-editable-region--
```
