---
id: 6999b07c34c2cde21ee6e1d8
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Now create a function called `hasHabitat`.

The function should take one parameter named `animal`.

Return `true` if the `animal` object has a property called `"habitat"`.

Return `false` otherwise.

Use the `hasOwnProperty` method.

Here is an example of using `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color");
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

In this example, the `hasColor` function uses `hasOwnProperty` to check whether the `color` property exists on an object.

After creating the function, use `console.log` to call `hasHabitat(tiger)` and `hasHabitat(elephant)` so you can see both results in the console.

# --hints--

`hasHabitat` should be a function.

```js
assert.isFunction(hasHabitat);
```

Calling `hasHabitat(tiger)` should return `true` after adding a habitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Calling `hasHabitat(elephant)` should return `false` if no habitat exists.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(tiger, 7));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
