---
id: 6999adc72c02812a28c01e31
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now create a function called `updateAge`.

The function should take two parameters: `animal` and `newAge`.

Inside the function, update the `age` property of the `animal` object to `newAge`.

Return the updated `animal` object.

# --hints--

`updateAge` should be a function.

```js
assert.isFunction(updateAge);
```

Calling `updateAge(tiger, 7)` should update the age property.

```js
updateAge(tiger, 7);
assert.strictEqual(tiger.age, 7);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

function getSpecies(animal) {
  return animal.species;
}

function addHabitat(animal, habitat) {
  animal.habitat = habitat;
  return animal;
}

--fcc-editable-region--

--fcc-editable-region--
```
