---
id: 6999ac171928136168654774
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Now create a function called `getSpecies`.

The function should take one parameter named `animal`.

Return the `species` property of the `animal` using dot notation.

Here is a small example of accessing a property from an object:

```js
const dog = {
  species: "Dog"
};

const getAnimalSpecies = (pet) => {
  return pet.species;
};

console.log(getAnimalSpecies(dog)); // Dog
```

In this example, the function takes an object as a parameter and returns its `species` property.

After creating the function, use `console.log` to call `getSpecies(tiger)` so you can see the result in the console.

# --hints--

You should create a function named `getSpecies`.

```js
assert.exists(getSpecies);
```

`getSpecies` should be a function.

```js
assert.isFunction(getSpecies);
```

Calling `getSpecies(tiger)` should return `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
