---
id: 6999b07c34c2cde21ee6e1d8
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

In this step, you will create a function that checks whether an object has a specific property.

Here is an example of using `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

In this example, the `hasColor` function uses `hasOwnProperty` to check whether the `color` property exists on an object.

Now create a function called `hasHabitat`.

The function should take one parameter named `animal`.

Return `true` if the `animal` object has a property called `"habitat"`.

Return `false` otherwise.

Use the `hasOwnProperty` method.

After creating the function, use `console.log` to call `hasHabitat(tiger)` and `hasHabitat(elephant)` so you can see both results in the console.

# --hints--

You should create a function named `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

The `hasHabitat` function should have only one parameter.

```js
assert.lengthOf(hasHabitat, 1);
```

`hasHabitat` should use the `hasOwnProperty` method to check for the `habitat` property.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

You should log `hasHabitat(tiger)` to the console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

You should log `hasHabitat(elephant)` to the console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Calling `hasHabitat(tiger)` should return `true` after adding a habitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Calling `hasHabitat(elephant)` should return `false` if no habitat exists.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` should use the function parameters and work with any object.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
