---
id: 6999adc72c02812a28c01e31
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Now create a function called `updateAge`.

The function should take two parameters: `animal` and `newAge`.

Inside the function, update the `age` property of the `animal` object to `newAge`.

Return the updated `animal` object.

Here is an example of updating a property inside a function:

```js
const dog = {
  age: 4
};

function changeAge(pet, updatedAge) {
  pet.age = updatedAge;
  return pet;
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

In this example, the `age` property is updated to a new value.

After creating the function, use `console.log` to call `updateAge(tiger, 7)` so you can see the updated `tiger` object in the console.

# --hints--

`updateAge` should be a function.

```js
assert.isFunction(updateAge);
```

Calling `updateAge(tiger, 7)` should update the age property.

```js
updateAge(tiger, 7);
assert.strictEqual(tiger.age, 7);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
