---
id: 69999ff7f7862fdfb885ea5d
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now update your `tiger` object by adding a new property called `age`.

Set the `age` property to the value `5`.

# --hints--

`tiger` should have an `age` property.

```js
assert.property(tiger, "age");
```

`tiger.age` should be equal to `5`.

```js
assert.strictEqual(tiger.age, 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
