---
id: 6999adc72c02812a28c01e31
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

In this step, you will create a function that updates a property of an object.

Here is an example of updating a property inside a function:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

In this example, the `age` property is updated to a new value.

Now create a function called `updateAge`. The function should take two parameters: `animal` and `newAge`.

Inside the function, update the `age` property of the `animal` object to `newAge`. Return the updated `animal` object.

After creating the function, use `console.log` to call `updateAge(elephant, 12)` so you can see the updated `elephant` object in the console.

# --hints--

You should create a function named `updateAge`.

```js
assert.isFunction(updateAge);
```

The `updateAge` function should have two parameters: `animal` and `newAge`.


`updateAge` should use dot notation to update the `age` property.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

The `updateAge` function should return the updated `animal` object.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

You should log `updateAge(elephant, 12)` to the console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Calling `updateAge(elephant, 12)` should update the age property.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` should use the function parameters and work with any object.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
