---
id: 6999b0198d18d402c36cc5fe
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now create a function called `removeEndangeredStatus`.

The function should take one parameter named `animal`.

Inside the function, remove the `isEndangered` property from the `animal` object using the `delete` keyword.

Return the updated `animal` object.

# --hints--

`removeEndangeredStatus` should be a function.

```js
assert.isFunction(removeEndangeredStatus);
```

Calling `removeEndangeredStatus(tiger)` should remove the `isEndangered` property.

```js
removeEndangeredStatus(tiger);
assert.notProperty(tiger, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

function getSpecies(animal) {
  return animal.species;
}

function addHabitat(animal, habitat) {
  animal.habitat = habitat;
  return animal;
}

function updateAge(animal, newAge) {
  animal.age = newAge;
  return animal;
}

--fcc-editable-region--

--fcc-editable-region--
```
