---
id: 699a00de9f564bb0effaa14d
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

In this step, you will create a function to access the `age` property of an object.

Here is an example:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age;
}

console.log(getAnimalAge(cat)); // 3
```

In this example, the function takes an object as a parameter and returns its `age` property.

Now create a function called `getAge`.

The function should take one parameter named `animal`.

Return the `age` property of the `animal` using dot notation.

After creating the function, use `console.log` to call `getAge(tiger)` so you can see the result in the console.


# --hints--

You should create a function named `getAge`.

```js
assert.exists(getAge)
```

`getAge` should be a function.

```js
assert.isFunction(getAge);
```

The `getAge` function should have only one parameter.

```js
assert.lengthOf(getAge, 1);
```

Calling `getAge(tiger)` should return `5`.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
