---
id: 699a00de9f564bb0effaa14d
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now create one more function called `getAge`.

The function should take one parameter named `animal`.

Return the `age` property of the `animal` using dot notation.

Here is an example:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age;
}

console.log(getAnimalAge(cat)); // 3
```

After creating the function, use `console.log` to call `getAge(tiger)` so you can see the result in the console.


# --hints--

You should create a function named `getAge`.

```js
assert.exists(getAge)
```

`getAge` should be a function.

```js
assert.isFunction(getAge);
```

Calling `getAge(tiger)` should return `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
