---
id: 615f378014c2da526a109c81
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

La fuente es un poco pequeña. Cree un selector `html` y configure la fuente para que tenga un tamaño de `16px`.

# --hints--

Debe tener un selector `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

Su selector `html` debe tener una propiedad `font-size` establecida en `16px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.fontSize === '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

body {
  font-family: 'Open Sans', sans-serif;
}
```
