---
id: 615f4f9e4a40566b776a8f38
title: Paso 27
challengeType: 0
dashedName: step-27
---

# --description--

Cree un nuevo selector `.large` y asígnele una propiedad `height` establecida en `10px`. También cree un selector `.large, .medium` y establezca la propiedad `background-color` en `black`.

# --hints--

Debería tener un nuevo selector `.large`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.large'));
```

Su selector `.large` debe tener una propiedad `height` establecida en `10px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.large')?.height === '10px');
```

Debería tener un nuevo selector `.large, .medium`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.large, .medium'));
```

Su selector `.large, .medium` debe tener una propiedad `background-color` establecida en `black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.large, .medium')?.backgroundColor === 'black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

--fcc-editable-region--

--fcc-editable-region--
```
