---
id: 615f357957e370510f21ea16
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Cree un selector `body` y asígnele una `font-family` establecida en `Open Sans` con un respaldo de `sans-serif`.

Recuerda que las fuentes con espacios en el nombre deben estar entre comillas para CSS.

# --hints--

Deberías tener un selector `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Su selector `body` debe tener una propiedad `font-family` establecida en `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
