---
id: 615f34948891834dd77655a6
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Agregue un segundo elemento `p` con el texto `Serving size 2/3 cup (55g)`.

# --hints--

Deberías tener un segundo elemento `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Su segundo elemento `p` debe estar dentro de su elemento `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Tu segundo elemento `p` debe ir después de tu elemento `p` existente.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Su segundo elemento `p` debe tener el texto `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
