---
id: 615f2d4150fe0d4cbd0f2628
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Debajo del elemento `h1`, agregue un elemento `p` con el texto `8 servings per container`.

# --hints--

Debe agregar un nuevo elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

El elemento `p` debe estar dentro del elemento `body`.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

El elemento `p` debe venir después del elemento `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Su elemento `p` debe tener el texto `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
