---
id: 615f3b091162165948e1cb33
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Si inspecciona su elemento `.label` con las herramientas de desarrollo de su navegador, puede notar que en realidad tiene 288 píxeles de ancho en lugar de 270. Esto se debe a que, de forma predeterminada, el navegador incluye el borde y el relleno al determinar el tamaño de un elemento.

Para resolver esto, restablezca el modelo de caja creando un selector `*` y dándole una propiedad `box-sizing` de `border-box`.

# --hints--

Debe crear un selector `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

El selector `*` debe tener una propiedad `box-sizing` establecida en `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
