---
id: 615f34ecc1091b4fd5a8a484
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro de su elemento `head`, agregue un elemento `link` con el atributo `rel` establecido en `stylesheet` y el `href` Atributo establecido en `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Esto importará la familia de fuentes `Open Sans`, con los valores de peso de fuente `400`, `700` y `800`.

También agregue un elemento `link` para vincular su archivo `styles.css`.

# --hints--

Su código debe tener dos elementos `link` de cierre automático.

```js
assert(document.querySelectorAll('link').length === 2);
```

Ambos elementos `link` deben tener el atributo `rel` establecido en `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno de sus elementos `link` debe tener un atributo `href` establecido en `./styles.css`.

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Uno de sus elementos `link` debe tener un atributo `href` establecido en `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
