---
id: 615f2abbe7d18d49a1e0e1c8
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Hemos proporcionado una plantilla HTML básica para usted.

Cree un elemento `h1` dentro de su elemento `body` y asígnele el texto `Nutrition Facts`.

# --hints--

Deberías añadir un nuevo elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Tu elemento `h1` debe estar dentro de tu elemento `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Tu elemento `h1` debe tener el texto `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
