---
id: 615f3cafd794015aa9547a6d
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Recuerde que el uso de `h1`, `h2` y etiquetas similares determinan la estructura semántica de su HTML. Sin embargo, puede ajustar el CSS de estos elementos para controlar el flujo visual y la jerarquía.

Cree una regla `h1` y establezca la propiedad `font-weight` en `800`. Esto hará que su texto `h1` sea más audaz.

# --hints--

Debe crear un selector `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

El selector `h1` debe tener una propiedad `font-weight` establecida en `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
