---
id: 3c008a051e5b4e91afec1e39
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Now that your `greet_customer` function is defined, you need to call it and store the result. When you call a function, you pass arguments inside the parentheses. The function executes and returns a value that you can save in a variable.

Call `greet_customer` with the argument `"Alice"` and assign the result to a variable named `greeting`. Then, use `print()` to display the value of `greeting`.

# --hints--

You should have a variable named `greeting`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('greeting')
    `))
})
```

You should assign `greet_customer('Alice')` to your `greeting` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('greeting').is_equivalent("greeting = greet_customer('Alice')")
    `))
})
```

You should print the `greeting` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call('print(greeting)')
    `))
})
```

# --seed--

## --seed-contents--

```py
tax_rate = 0.08

def greet_customer(name):
    return "Welcome, " + name + "!"

--fcc-editable-region--

--fcc-editable-region--
```
