---
id: 5d2faf6fe1f246c8ac8015bc
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Now you'll create your first function. As you learned in the lesson, you use the `def` keyword, followed by the function name, a pair of parentheses with any parameters, and a colon. The code inside the function must be indented.

Here's a quick reminder:

```py
def greet():
    print('Hello!')
```

Define a function named `greet_customer` that takes one parameter called `name`. Inside the function body, return the string `"Welcome, " + name + "!"`.

# --hints--

You should define a function named `greet_customer`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function('greet_customer')`)) })
```

Your `greet_customer` function should return the string `"Welcome, " + name + "!"`.

```js
({
    test: () => runPython(`
exec(_code)
assert greet_customer('Alice') == 'Welcome, Alice!'
assert greet_customer('Bob') == 'Welcome, Bob!'
`)
})
```

# --seed--

## --seed-contents--

```py
tax_rate = 0.08

--fcc-editable-region--

--fcc-editable-region--
```
