---
id: 609236672e304b56b87dc939
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Next, you need a function to calculate the subtotal for an order. This function will take two parameters: the `price` of a single item and the `quantity` of items ordered.

Define a function named `calculate_subtotal` that takes two parameters: `price` and `quantity`. Inside the function, return the product of `price` and `quantity`.

# --hints--

You should define a function named `calculate_subtotal`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function('calculate_subtotal')`)) })
```

Your `calculate_subtotal` function should return the product of `price` and `quantity`.

```js
({
    test: () => runPython(`
exec(_code)
assert calculate_subtotal(5, 3) == 15
assert calculate_subtotal(2, 4) == 8
`)
})
```

# --seed--

## --seed-contents--

```py
tax_rate = 0.08

def greet_customer(name):
    return "Welcome, " + name + "!"

greeting = greet_customer("Alice")
print(greeting)

--fcc-editable-region--

--fcc-editable-region--
```
