---
id: 62de12b9f6c04cbba0540246
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Now you need a function that calculates the total amount including tax. This function will demonstrate how functions can access global variables.

In Python, a function can read a global variable without needing to pass it as a parameter. The `tax_rate` variable you defined earlier is in the global scope, so it can be accessed from inside any function.

Define a function named `calculate_total` that takes one parameter called `subtotal`. Inside the function, return the result of `subtotal * (1 + tax_rate)`. This uses the global `tax_rate` variable to calculate the total with tax included.

# --hints--

You should define a function named `calculate_total`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function('calculate_total')`)) })
```

Your `calculate_total` function should return the subtotal multiplied by `(1 + tax_rate)`.

```js
({
    test: () => runPython(`
exec(_code)
assert calculate_total(10) == 10 * (1 + 0.08)
assert calculate_total(100) == 100 * (1 + 0.08)
`)
})
```

# --seed--

## --seed-contents--

```py
tax_rate = 0.08

def greet_customer(name):
    return "Welcome, " + name + "!"

greeting = greet_customer("Alice")
print(greeting)

def calculate_subtotal(price, quantity):
    return price * quantity

subtotal = calculate_subtotal(5, 3)
print("Subtotal:", subtotal)

--fcc-editable-region--

--fcc-editable-region--
```
