---
id: 66017b97bad249648c771ac4
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will practice defining functions and working with scope by building a coffee shop order calculator. This tool will calculate order subtotals, apply a tax rate, and display a summary for customers.

Start by creating a global variable that stores the tax rate for the coffee shop. This variable will be accessible throughout your entire program, including from inside the functions you will write later.

Create a variable named `tax_rate` and assign it the value `0.08`.

# --hints--

You should have a variable named `tax_rate`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('tax_rate')
    `))
})
```

You should assign the float `0.08` to your `tax_rate` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('tax_rate').is_equivalent('tax_rate = 0.08')
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
