---
id: 9bb6dc4df2764148b27df17d
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Now it's time to use your `calculate_subtotal` function. Imagine a customer orders 3 drinks that cost $5 each.

Call `calculate_subtotal` with `5` and `3` as arguments, and assign the result to a variable named `subtotal`. Then, use `print()` to display the string `"Subtotal: "` followed by a space and the value of `subtotal`.

# --hints--

You should have a variable named `subtotal`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('subtotal')
    `))
})
```

You should assign `calculate_subtotal(5, 3)` to your `subtotal` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('subtotal').is_equivalent('subtotal = calculate_subtotal(5, 3)')
    `))
})
```

You should print the string `"Subtotal: "` followed by a space and the `subtotal` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Subtotal:', subtotal)") or _Node(_code).has_call("print(f'Subtotal: {subtotal}')")
    `))
})
```

# --seed--

## --seed-contents--

```py
tax_rate = 0.08

def greet_customer(name):
    return "Welcome, " + name + "!"

greeting = greet_customer("Alice")
print(greeting)

def calculate_subtotal(price, quantity):
    return price * quantity

--fcc-editable-region--

--fcc-editable-region--
```
