---
id: 9e6b22901f194cd1ae0a3b08
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

You have all the pieces in place. Now it's time to use your `calculate_total` function to compute the final total, calculate the tax amount, and display a complete order summary.

Call `calculate_total` with your `subtotal` variable as the argument, and assign the result to a variable named `total`. Then, create a variable named `tax` and assign it the value of `total - subtotal` to find just the tax amount.

Finally, print the string `"Tax: "` followed by a space and the `tax` variable, and then print the string `"Total: "` followed by a space and the `total` variable. With that, the coffee shop order calculator is complete.

# --hints--

You should have a variable named `total`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_variable('total')
    `))
})
```

You should assign `calculate_total(subtotal)` to your `total` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('total').is_equivalent('total = calculate_total(subtotal)')
    `))
})
```

You should have a variable named `tax` that equals `total - subtotal`.

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_variable('tax').is_equivalent('tax = total - subtotal')
    `))
})
```

You should print the string `"Tax: "` followed by a space and the `tax` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Tax:', tax)") or _Node(_code).has_call("print(f'Tax: {tax}')")
    `))
})
```

You should print the string `"Total: "` followed by a space and the `total` variable.

```js
({
    test: () => assert(runPython(`
    _Node(_code).has_call("print('Total:', total)") or _Node(_code).has_call("print(f'Total: {total}')")
    `))
})
```

# --seed--

## --seed-contents--

```py
tax_rate = 0.08

def greet_customer(name):
    return "Welcome, " + name + "!"

greeting = greet_customer("Alice")
print(greeting)

def calculate_subtotal(price, quantity):
    return price * quantity

subtotal = calculate_subtotal(5, 3)
print("Subtotal:", subtotal)

def calculate_total(subtotal):
    return subtotal * (1 + tax_rate)

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
tax_rate = 0.08

def greet_customer(name):
    return "Welcome, " + name + "!"

greeting = greet_customer("Alice")
print(greeting)

def calculate_subtotal(price, quantity):
    return price * quantity

subtotal = calculate_subtotal(5, 3)
print("Subtotal:", subtotal)

def calculate_total(subtotal):
    return subtotal * (1 + tax_rate)

total = calculate_total(subtotal)
tax = total - subtotal
print("Tax:", tax)
print("Total:", total)
```
