---
id: 68b94f23f7b3ee14078c89c5
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will use the `iframe` element to display a video. The basic HTML boilerplate has been prepared for you.

Begin by creating an `h1` element with the text `iframe Video Display`.

# --hints--

You should have an `h1` element.

```js
assert.exists(document.querySelector('h1'))
```

Your `h1` element should have the text `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
