---
id: 68ba9c4f1688914d72876098
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

The next attribute you'll add is `referrerpolicy`. It is the rule that determines how much detail you share when your page connects to another page.

Add the `referrerpolicy` attribute and set it to `strict-origin-when-cross-origin`. This shares the full address on the same site, only the site name on other sites, and nothing on insecure sites.

# --hints--

Your `iframe` element should have a `referrerpolicy` attribute set to `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
