---
id: 68ba9c4f1688914d72876097
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Add `encrypted-media`, `gyroscope`, and `web-share` to the existing values in the `allow` attribute.

These three will allow the use of encrypted media extensions to protect the video, let the iframe pop out into picture-in-picture mode when needed, and allow sharing the iframe content through the device's native share dialogs. 

# --hints--

The `allow` attribute of your `iframe` element should have `encrypted-media` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const values = iframeElAllowAttr.trim().split(/[;\s]+/).map(v => v.trim()).filter(v => v)
assert.strictEqual(values[3], 'encrypted-media')
```

The `allow` attribute of your `iframe` element should have `gyroscope` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const values = iframeElAllowAttr.trim().split(/[;\s]+/).map(v => v.trim()).filter(v => v)
assert.strictEqual(values[4], 'gyroscope')
```

The `allow` attribute of your `iframe` element should have `web-share` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const values = iframeElAllowAttr.trim().split(/[;\s]+/).map(v => v.trim()).filter(v => v)
assert.strictEqual(values[5], 'web-share')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write"
    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```
