---
id: 68ba9c4f1688914d72876096
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

One of the attributes is `allow`. It's like a permission list that tells the browser what features the `iframe` is allowed to use.

Add the `allow` attribute with the value `accelerometer`, `autoplay`, and `clipboard-write`.

`accelerometer` lets the `iframe` use motion sensors so it can detect things like device tilting and rotation. `autoplay` lets the video start playing automatically, and `clipboard-write` lets the iframe write data to the user’s clipboard.

# --hints--

Your `iframe` element should have an `allow` attribute set to `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const accelerometer = iframeElAllowAttr.trim().split(' ')[0]
assert.strictEqual(accelerometer, 'accelerometer')
```

The `allow` attribute of your `iframe` element should have `autoplay` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const autoplay = iframeElAllowAttr.trim().split(' ')[1]
assert.strictEqual(autoplay, 'autoplay')
```

The `allow` attribute of your `iframe` element should have `clipboard-write` as one of its values.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
const clipboardWrite = iframeElAllowAttr.trim().split(' ')[2]
assert.strictEqual(clipboardWrite, 'clipboard-write')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <br />
--fcc-editable-region--
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"

    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```
