---
id: 662bdb24a63307df57b3dffe
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

As said in the previous step, `Equation` is the class you want to act as interface, but at the moment is just a regular class.

Nothing prevents you from creating an instance of `Equation`. Verify it by instantiating an `Equation` object and assigning it to a variable `eq`.

# --hints--

You should declare a variable `eq` and assign it an instance of `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self): pass
        
    def analyze(self): pass
        

class LinearEquation(Equation): pass

--fcc-editable-region--
lin_eq = LinearEquation()

--fcc-editable-region--
```
