---
id: 662fab95a3f903a7aa028550
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Decorators can be combined with each other to obtain enhanced functionalities. Turn the `grade` method into a property.

For that you'll need to use the `property` decorator in both classes and as the first decorator within the `Equation` class.

# --hints--

You should decorate with `@property` and `@abstractmethod` the `grade` method within the `Equation` class. The order matters.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("grade").has_decorators("property", "abstractmethod")`)) })
```

You should decorate with `@property` the `grade` method within the `LinearEquation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("grade").has_decorators("property")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self): pass
    
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    def grade(self):
        return 1
--fcc-editable-region--
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation()

```
