---
id: 662fa2e2cf27c09f21f4f5d0
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Inside the `Equation` class, define a new abstract method `degree`. Then, define it within the `LinearEquation` class to avoid errors.

Later on, you'll use this method as part of the validation of the arguments passed to instantiate the equation objects.

# --hints--

You should define a method named `degree` within the `Equation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("degree")`)) })
```

Your `degree` method should take one parameter, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("degree").has_args("self")`)) })
```

You should decorate with `@abstractmethod` the `degree` method within the `Equation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("degree").has_decorators("abstractmethod")`)) })
```

You should define a method named `degree` within the `LinearEquation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_function("degree")`)) })
```

Your `degree` method should take one parameter, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("degree").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self): pass
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    def solve(self): pass
    
    def analyze(self): pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
