---
id: 662fba5aa58abab470a2e6b8
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

It's time to go back to the `__init__` method. Depending on the equation type, you'll need to pass a variable number of arguments during the instantiation.

Add a second parameter `args` to the method and use the `*` operator to make it accept a variable number of arguments.

# --hints--

Your `__init__` method should take two parameters, `self`, and `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self): pass
--fcc-editable-region--
    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation()

```
