---
id: 662fa7a2e834a7a3530db75b
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

The `grade` method will return a number indicating the grade of the equation, which is the exponent of the highest `x` term. Therefore, within the `LinearEquation` class, return `1` from the `grade` method.

# --hints--

Your `grade` method should return `1`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("grade").has_return("1")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    def __init__(self): pass
    
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
--fcc-editable-region--
class LinearEquation(Equation):
    def grade(self): pass
--fcc-editable-region--    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation()
```
