---
id: 6639fdcc701833a54c364211
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

After validating the coefficients, you need to store them in an instance attribute. Use a dictionary comprehension to create a dictionary in which the key is the grade of the coefficient and the corresponding value is the coefficient, and assign it to an attribute named `coefficients`.

For example, a `LinearEquation` object instantiated with `2` and `4` should have the folowing `coefficients` attribute: `{0: 4, 1: 2}`, because `4` corresponds to zero-th grade of `x` and `2` corresponds to the first grade of `x`.

For this, you may find useful the `reversed()` function, which takes a sequence as its argument and returns a reverse iterator.

# --hints--

You should declare an attribute named `coefficients` within your `__init__` method.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_variable("self.coefficients")`)) })
```

Your `coefficients` attribute should be a dictionary containing key-value pairs in the form grade-coefficient.

```js
({ test: () => assert(runPython(`
LinearEquation(1, 6).coefficients == {0: 6, 1: 1} and LinearEquation(-3.5, 0).coefficients == {0: 0, 1: -3.5}
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.grade + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.grade + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
    
    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
--fcc-editable-region--
```
