---
id: 6639f947d3a1818c9322c64a
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

The last step of validating the coefficients is checking that the highest grade coefficient is different from zero. Remember that highest grade coefficient should be passed as the first argument when instantiating the object.

Add an `if` statement for that and raise a `ValueError` and use the following string to provide a custom message `"Highest order coefficient must be different from zero"`.

# --hints--

You should create an `if` statement that checks if the first coefficient passed to instantiate the equation is equal to zero.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]")
`)) })
```

You should raise a `ValueError` within the new `if` statement and use the provided string to return a custom error message.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[1].find_bodies()[0].has_stmt("raise ValueError('Highest order coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.grade + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.grade + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
    
    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
--fcc-editable-region--
```
