---
id: 663a2ecc3bc690f17d2c636f
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

Now the root is always displayed with three decimal digits, regardless of their values.

Remove both the `lin_eq.solve()` and `print(lin_eq.results)` calls from your code.

# --hints--

You should remove both your `lin_eq.solve()` and `print(lin_eq.results)` calls.

```js
({ test: () => runPython(`
assert not _Node(_code).has_call("lin_eq.solve()")
assert not _Node(_code).has_call("print(lin_eq.results)")
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}
        
        
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')
            
        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        self.results.append(f'x = {x:.3f}')
        return x
    
    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)
lin_eq.solve()
print(lin_eq.results)
--fcc-editable-region--
```
