---
id: 663a2dd1901cbeecc28748bd
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

It's time to test the `solve` method. Call it on `lin_eq`. Then, print the `results` attribute of `lin_eq`.

# --hints--

You should call the `solve` method of you `lin_eq` object.

```js
({ test: () => assert(runPython(`
_Node(_code).has_call("lin_eq.solve()") or _Node(_code).has_call("print(lin_eq.solve())")
`)) })
```

You should print the `results` attribute of `lin_eq` after calling the `solve` method.

```js
({ test: () => assert(runPython(`
_Node(_code).is_ordered("lin_eq.solve()", "print(lin_eq.results)") or _Node(_code).is_ordered("print(lin_eq.solve())", "print(lin_eq.results)")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []
        
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')
            
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f"x = {round(x, 3)}")
        return x
    
    def analyze(self): pass
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)
--fcc-editable-region--
```
