---
id: 663a08e4207660b6684a2d4a
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Now, complete your `__init__` method by creating two attributes, `results` and `details` and assign both of them an empty list. You'll use them later on in your project.

# --hints--

You should declare a `results` attribute and assign it an empty list.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_variable("self.results").is_equivalent("self.results = []")`)) })
```

You should declare a `details` attribute and assign it an empty list.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_variable("self.details").is_equivalent("self.details = []")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
--fcc-editable-region--
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        
--fcc-editable-region--
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
```
