---
id: 663a26512edd9fdb5e10f28b
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

In a previous step, you declared an attribute named `results`. You'll use it to store strings containing the results of the equation and then display them to have a nice output. Create an f-string that has the form `x =` followed by a space and the result rounded to the third decimal digit, and append it to the `results` list.

To round your result, use the `round()` function. It takes the number to round as its first argument, and the number of decimal digits as the second argument.

# --hints--

You should append a string containing the solution of your linear equation in the form `x = <result>` to the `results` attribute. Remember to round the result to the third decimal digit.

```js
({ test: () => runPython(`
eq = LinearEquation(2.2, 1.5)
eq.solve()
assert eq.results[0] == 'x = -0.682', "Expected different result" 
assert len(eq.results) == 1, "Expected different 'results' length"
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []
        
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')
            
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
--fcc-editable-region--
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        
--fcc-editable-region--
        return x

    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
