---
id: 663b7fefd437bd984e091cbf
title: Step 32
challengeType: 20
dashedName: step-32
---

# --description--

Next, create a new class named `QuadraticEquation` and make it inherit from `Equation`. You'll use this new class to represent quadratic equations.

Inside your new class, define a `degree` property which returns the degree of the equation. Also, define the `solve` and `analyze` methods. You will take care of the implementation in the following steps.

# --hints--

You should create a new class named `QuadraticEquation` and make it inherit from the `Equation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").inherits_from("Equation")`)) })
```

You should define a `solve` method within the `QuadraticEquation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("solve")`)) })
```

Your `solve` method should take a single parameter, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("solve").has_args("self")`)) })
```

You should define an `analyze` method within the `QuadraticEquation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("analyze")`)) })
```

Your `analyze` method should take a single parameter, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("analyze").has_args("self")`)) })
```

You should define a `degree` method within the `QuadraticEquation` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_function("degree")`)) })
```

Your `degree` method should take a single parameter, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("degree").has_args("self")`)) })
```

You should decorate with `@property` the `degree` method.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").find_function("degree").has_decorators("property")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []
        
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')
            
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return x

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}
--fcc-editable-region--

--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
