---
id: 663b95d65caeb3ca04c5fef4
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

If the coefficient has non-zero value, you can have 3 cases:

- `n == 0`: the term is made by the coefficient itself
- `n == 1`: the term is the coefficient times `x`
- `n > 1`: the term is the coefficient times `x**n`

Inside your `if` statement, create an `if` statement for the first case and append a string containing the coefficient to the `terms` list. Use an f-string for that.

# --hints--

You should create an `if` statement to check if `n` is equal to `0` inside your existing `if` statement.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("n==0")`)) })
```

You should append `f'{coefficient}'` to the `terms` list within your new `if` statement.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("terms.append(f'{coefficient}')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.grade + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.grade + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                pass
        equation_string = ' '.join(terms)
        return equation_string
--fcc-editable-region--
    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
