---
id: 663c981b9b06922e13a97fe9
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

Create an `elif` clause for the case `n == 1`. Within the `elif` clause, create an f-string cointaing the coefficient directly followed by a lowercase `x`. This time, instead of appending the string to the `terms` list, use the `.insert()` method to insert the string at the beginning of the list.

For that you need to pass two arguments to `.insert()`: the index at which insert the string, and the string to insert.

# --hints--

You should create an `elif` clause to check if `n` is equal to `1`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("n==1")`)) })
```

You should insert `f'{coefficient}x'` at the beginning of the `terms` list within your new `elif` clause.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__str__").find_for_loops()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].is_equivalent("terms.insert(0, f'{coefficient}x')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []
--fcc-editable-region--
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient}')
        equation_string = ' '.join(terms)
        return equation_string
--fcc-editable-region--
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
