---
id: 664cb04a16fe6938708967ef
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

After joining the terms, concatenate the string `' = 0'` to `equation_string` to display the complete equation.

Also, to refine the output, remove any leading `+` sign from `equation_string`.

# --hints--

The `__str__` method should return a different string representation.

```js
({ test: () => assert(runPython(`
eq1 = LinearEquation(4, 2)
str(eq1) == '4x +2 = 0'
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.grade + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.grade + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
--fcc-editable-region--
        equation_string = ' '.join(terms)
        
        return equation_string
--fcc-editable-region--
    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)
```
