---
id: 664c670069bae45fd060c25d
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

Next, print your `lin_eq` instance.

# --hints--

You should print `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(lin_eq)")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.grade + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.grade + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.coefficients = []
    
    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self): pass
    
    def analyze(self): pass
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
--fcc-editable-region--
```
