---
id: 664eec7f38234443b42c206f
title: Step 35
challengeType: 20
dashedName: step-35
---

# --description--

Your equation is currently represented as `11x**2 -1x +1 = 0`, but it would be nice not to display the coefficient multiplying \\( x \\) when it's equal to one. So that the equation is represented as `11x**2 -x +1 = 0`.

Import the `re` module. You are going to use a regular expression to substitute the coefficients for this case during the next steps.

# --hints--

You should import the `re` module.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import re")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--

--fcc-editable-region--
class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                else:
                    terms.insert(0, f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'

        return equation_string.strip('+')
        
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return x

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    def degree(self):
        return 2
        
    def solve(self):
        pass
        
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
