---
id: 6650f037c017aa6855a608e3
title: Step 46
challengeType: 20
dashedName: step-46
---

# --description--

The graph of any quadratic equation has a parabolic shape. The x coordinate of the vertex of the parabola can be found at \\( x = - \frac{b}{2a} \\).

From the `analyze` method, return  dictionary containing two keys, `'x'`, and `'y'`, and the corresponding values of vertex x and y coordinates, respectively.

Use the relation above to find the x coordinate. Then, use the x coordinate to calculate the y coordinate.

# --hints--

You should return a dictionary containing two keys, `'x'`, and `'y'`, and the corresponding values of vertex x and y coordinates, respectively.

```js
({ test: () => runPython(`
eq = QuadraticEquation(16, 2, 1)
assert eq.analyze() == {'x': -0.0625, 'y': 0.9375}
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):   
    def __init__(self, *args):
        if (self.grade + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.grade + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest order coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                else:
                    terms.insert(0, f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

    @property
    @abstractmethod
    def grade(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def grade(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f"x = {round(x, 3)}")
        return [x]

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    def grade(self):
        return 2

    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
        if self.delta == 0:
            self.results.append(f'x = {round(x1, 3)}')
            return [x1]

        self.results.extend([f'x1 = {round(x1, 3)}', f'x2 = {round(x2, 3)}'])
        return [x1, x2]
--fcc-editable-region--
    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(1, 2, 1)
print(quadr_eq)
print(quadr_eq.solve())
print(quadr_eq.results)
```
