---
id: 6650583d9d9a194714da47f0
title: Step 41
challengeType: 20
dashedName: step-41
---

# --description--

The roots of quadratic equations can be found by applying the formula:

\\[ x = \frac{-b \pm \sqrt{b^2 - 4ac}}{2a} \\]

After the `if` statement, create two variables: assign them the root obtained by taking the plus sign between \\( b \\) and the square root of the discriminant in the formula above, and the root obtained by taking the minus sign. Then, return a list containing your two new variables.

# --hints--

Your `solve` method should return a list containing the correct solutions.

```js
({ test: () => runPython(`
eq = QuadraticEquation(-1, 2, 3)
assert eq.solve() == [-1, 3] or eq.solve() == [3, -1]
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                else:
                    terms.insert(0, f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return x

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    def degree(self):
        return 2
--fcc-editable-region--
    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
--fcc-editable-region--
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(11, -1, 1)
print(quadr_eq)

```
