---
id: 6650e88cc500673ec881c9ca
title: Step 42
challengeType: 20
dashedName: step-42
---

# --description--

In this case, the discriminant is zero and the two roots are coincidents.

Create an `if` statement to check if the discriminant is zero. This time, append a single string in the form `x = <root>` to the `results` attribute. As before, round the root to the third decimal digit. Then, return a list containing the root.

# --hints--

You should create an `if` statement to check if the `delta` attribute is equal to zero.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("QuadraticEquation").find_function("solve").find_ifs()[1].find_conditions()[0]
node.is_equivalent("self.delta == 0") or node.is_equivalent("not self.delta")
`)) })
```

You should append the string `x = <root>`, where the root is rounded to the third decimal digit, to the `results` attribute within your new `if` statement.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("QuadraticEquation").find_function("solve").find_ifs()[1].find_bodies()[0]
node.has_stmt("self.results.append(f'x = {x1:.3f}')") or node.has_stmt("self.results.append(f'x = {x2:.3f}')")
`)) })
```

You should return a list containing the root within your new `if` statement.

```js
({ test: () => runPython(`
eq = QuadraticEquation(4, 4, 1)
assert eq.solve() == [-0.5]
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}
        

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        self.results.append(f'x = {x:.3f}')
        return x

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    degree = 2

    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
--fcc-editable-region--
        
--fcc-editable-region--
        self.results.extend([f'x1 = {x1:.3f}', f'x2 = {x2:.3f}'])
        return [x1, x2]

    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
quadr_eq = QuadraticEquation(1, 2, 1)
print(quadr_eq)
print(quadr_eq.solve())
print(quadr_eq.results)
```
