---
id: 6655b8f52c08b6d812d58425
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

The `round()` function rounds your result to the specified number of decimal digits. Although you specified `3` as the number of digits, your output contains a single decimal digits. This happens because trailing decimal zeros are neglected from the string representation.

To include a fixed number of decimal digits in your string, you can use the format specifier `.nf`, where `n` is the number of decimal digits to display.

Delete the `round()` call from your string. Within the curly braces, write `x` followed by a colon and the format specifier needed to display 3 decimal digits.

# --hints--

You should append `f'x = {x:.3f}'` to the `results` attribute.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("self.results.append(f'x = {x:.3f}')")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}
        
        
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')
            
        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
--fcc-editable-region--
        self.results.append(f'x = {round(x, 3)}')
--fcc-editable-region--
        return x
    
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
lin_eq.solve()
print(lin_eq.results)
```
