---
id: 6656fc80df65eeacd3b16d5a
title: Step 64
challengeType: 20
dashedName: step-64
---

# --description--

As a last step, modify the string appended to the `details` attribute of your quadratic equation so that the text placed after the equal sign is right-aligned for each line. Your final output should look like this:

```py

---Quadratic Equation---

    x**2 +2x +1 = 0     

-------Solutions--------

       x = -1.000       

--------Details---------

concavity =      upwards
min =    (-1.000, 0.000)

```

With that, the project is complete!

# --hints--

The `solver` function should return a different string.

```js
({ test: () => runPython(`
eq = QuadraticEquation(-4, 3, 2)
expected = """
---Quadratic Equation---

   -4x**2 +3x +2 = 0    

-------Solutions--------

      x1 = -0.425       
       x2 = 1.175       

--------Details---------

concavity =    downwards
max =     (0.375, 2.562)
"""
assert solver(eq) == expected
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                else:
                    terms.insert(0, f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

    @property
    @abstractmethod
    def degree(self): pass
    
    @property
    @abstractmethod
    def type(self): pass
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    @property
    def type(self):
        return "Linear Equation"
        
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return [x]

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]:>16}\ny-intercept = {self.coefficients[0]:>10}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    def degree(self):
        return 2
        
    @property
    def type(self):
        return "Quadratic Equation"

    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
        if self.delta == 0:
            self.results.append(f'x = {x1:.3f}')
            return [x1]

        self.results.extend([f'x1 = {x1:.3f}', f'x2 = {x2:.3f}'])
        return [x1, x2]

    def analyze(self):
        x = -self.coefficients[1] / (2 * self.coefficients[2])
        y = (
            self.coefficients[2] * x**2
            + self.coefficients[1] * x
            + self.coefficients[0]
        )
        vertex = {'x': x, 'y': y}
        if self.coefficients[2] > 0:
            concavity = 'upwards'
            vertex['m'] = 'min'
        else:
            concavity = "downwards"
            vertex['m'] = 'max'
        coord = f'({x:.3f}, {y:.3f})'
--fcc-editable-region--
        self.details.append(f'concavity = {concavity}\n{vertex["m"]} = {coord}')
--fcc-editable-region--
        return {'vertex': vertex, 'concavity': concavity}

def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError('Argument must be an Equation object')
    if not equation.results:
        equation.solve()
    if not equation.details:
        equation.analyze()
    output_string = f'\n{equation.type:-^24}'
    output_string += f'\n\n{str(equation):^24}\n\n'
    output_string += f'{"Solutions":-^24}\n\n'
    for result in equation.results:
        output_string += f'{result:^24}\n'
    output_string += f'\n{"Details":-^24}\n\n'
    for detail in equation.details:
        output_string += f'{detail}\n'
    return output_string

lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)
print(solver(quadr_eq))

```

# --solutions--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                else:
                    terms.insert(0, f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

    @property
    @abstractmethod
    def degree(self): pass
    
    @property
    @abstractmethod
    def type(self): pass
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    @property
    def type(self):
        return "Linear Equation"
        
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return [x]

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]:>16}\ny-intercept = {self.coefficients[0]:>10}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    def degree(self):
        return 2
        
    @property
    def type(self):
        return "Quadratic Equation"

    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
        if self.delta == 0:
            self.results.append(f'x = {x1:.3f}')
            return [x1]

        self.results.extend([f'x1 = {x1:.3f}', f'x2 = {x2:.3f}'])
        return [x1, x2]

    def analyze(self):
        x = -self.coefficients[1] / (2 * self.coefficients[2])
        y = (
            self.coefficients[2] * x**2
            + self.coefficients[1] * x
            + self.coefficients[0]
        )
        vertex = {'x': x, 'y': y}
        if self.coefficients[2] > 0:
            concavity = 'upwards'
            vertex['m'] = 'min'
        else:
            concavity = "downwards"
            vertex['m'] = 'max'
        coord = f'({x:.3f}, {y:.3f})'
        self.details.append(f'concavity = {concavity:>12}\n{vertex["m"]} = {coord:>18}')
        return {'vertex': vertex, 'concavity': concavity}

def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError('Argument must be an Equation object')
    if not equation.results:
        equation.solve()
    if not equation.details:
        equation.analyze()
    output_string = f'\n{equation.type:-^24}'
    output_string += f'\n\n{str(equation):^24}\n\n'
    output_string += f'{"Solutions":-^24}\n\n'
    for result in equation.results:
        output_string += f'{result:^24}\n'
    output_string += f'\n{"Details":-^24}\n\n'
    for detail in equation.details:
        output_string += f'{detail}\n'
    return output_string

lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)
print(solver(quadr_eq))

```
