---
id: 6656e3f3c987505def490b81
title: Step 51
challengeType: 20
dashedName: step-51
---

# --description--

From now on, you'll keep building the output by concatenating strings to `output_string`.

Create a string containing the string representation of your equation centered in a width of 24 characters. Make the string begin and end with two newline characters, and add your new string to the current value of `output_string`.

# --hints--

The `solver` function should return a different string.

```js
({ test: () => runPython(`
eq = LinearEquation(4, 3)
expected = """
----Linear Equation-----

       4x +3 = 0        

"""
assert solver(eq) == expected
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}
        

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

    
    
    @property
    @abstractmethod
    def type(self): pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    @property
    def type(self):
        return "Linear Equation"
        
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        self.results.append(f'x = {x:.3f}')
        return [x]

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    degree = 2
        
    @property
    def type(self):
        return "Quadratic Equation"

    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
        if self.delta == 0:
            self.results.append(f'x = {x1:.3f}')
            return [x1]

        self.results.extend([f'x1 = {x1:.3f}', f'x2 = {x2:.3f}'])
        return [x1, x2]

    def analyze(self):
        x = -self.coefficients[1] / (2 * self.coefficients[2])
        y = (
            self.coefficients[2] * x**2
            + self.coefficients[1] * x
            + self.coefficients[0]
        )
        vertex = {'x': x, 'y': y}
        if self.coefficients[2] > 0:
            concavity = 'upwards'
            vertex['m'] = 'min'
        else:
            concavity = "downwards"
            vertex['m'] = 'max'
        coord = f'({x:.3f}, {y:.3f})'
        self.details.append(f'concavity = {concavity}\n{vertex["m"]} = {coord}')
        return {'vertex': vertex, 'concavity': concavity}

def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError('Argument must be an Equation object')
    if not equation.results:
        equation.solve()
    if not equation.details:
        equation.analyze()
--fcc-editable-region--
    output_string = f'\n{equation.type:-^24}'

--fcc-editable-region--
    return output_string

lin_eq = LinearEquation(2, 3)
print(solver(lin_eq))
quadr_eq = QuadraticEquation(1, 2, 1)
```
