---
id: 6656262ec602ff688c50c8a9
title: Step 52
challengeType: 20
dashedName: step-52
---

# --description--

After the `if` statement you wrote in the previous step, ensure that `solve()` and `analyze()` methods have not been called yet, otherwise you would keep adding the same items to the `results` and `details` lists.

Check if the `results` and `details` attributes of `equation` contain elements. If not, call the `solve` method, and the `analyze` method, respectively.

# --hints--

You should check if `equation.results` does not contain elements and call `equation.solve()` if it is empty.

```js
({ test: () => runPython(`
eq1 = LinearEquation(4, 3)
assert eq1.results == []
eq1.analyze()
solver(eq1)
assert eq1.results == ['x = -0.750']
assert eq1.details == ['slope = 4\\ny-intercept = 3']
eq2 = QuadraticEquation(-5, 2, 1)
assert eq2.results == []
solver(eq2)
assert eq2.results == ['x1 = -0.290', 'x2 = 0.690']
solver(eq2)
assert eq2.results == ['x1 = -0.290', 'x2 = 0.690']
`) })
```

You should check if `equation.details` does not contain elements and call `equation.analyze()` if it is empty.

```js
({ test: () => runPython(`
eq1 = LinearEquation(4, 3)
assert eq1.details == []
eq1.solve()
solver(eq1)
assert eq1.details == ['slope = 4\\ny-intercept = 3']
assert eq1.results == ['x = -0.750']
eq2 = QuadraticEquation(-5, 2, 1)
assert eq2.details == []
solver(eq2)
assert eq2.details == ['concavity = downwards\\nmax = (0.200, 1.200)']
solver(eq2)
assert eq2.details == ['concavity = downwards\\nmax = (0.200, 1.200)']
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}
        

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))

        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        self.results.append(f'x = {x:.3f}')
        return [x]

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    degree = 2

    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
        if self.delta == 0:
            self.results.append(f'x = {x1:.3f}')
            return [x1]

        self.results.extend([f'x1 = {x1:.3f}', f'x2 = {x2:.3f}'])
        return [x1, x2]

    def analyze(self):
        x = -self.coefficients[1] / (2 * self.coefficients[2])
        y = (
            self.coefficients[2] * x**2
            + self.coefficients[1] * x
            + self.coefficients[0]
        )
        vertex = {'x': x, 'y': y}
        if self.coefficients[2] > 0:
            concavity = 'upwards'
            vertex['m'] = 'min'
        else:
            concavity = "downwards"
            vertex['m'] = 'max'
        coord = f'({x:.3f}, {y:.3f})'
        self.details.append(f'concavity = {concavity}\n{vertex["m"]} = {coord}')
        return {'vertex': vertex, 'concavity': concavity}
--fcc-editable-region--
def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError('Argument must be an Equation object')
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)
```
