---
id: 66562f71937f877c66123bbe
title: Step 54
challengeType: 20
dashedName: step-54
---

# --description--

The first thing to display at the top of the output will be the equation type. Add an abstract property named `type` to the `Equation` class. Then, add the new property to the `LinearEquation` class and to the `QuadraticEquation` class. Make them return the string `'Linear Equation'` and `'Quadratic Equation'`, respectively.

# --hints--

You should define an abstract property named `type` within the `Equation` class.

```js
({ test: () => runPython(`
foo = _Node(_code).find_class("Equation").find_function("type")
assert foo.has_args("self")
assert foo.has_decorators("property", "abstractmethod")
`) })
```

The `type` property of the `LinearEquation` class should return the string `'Linear Equation'`.

```js
({ test: () => runPython(`
foo = _Node(_code).find_class("LinearEquation").find_function("type")
assert foo.has_args("self")
assert foo.has_decorators("property")
assert foo.has_return("'Linear Equation'")
`) })
```

The `type` property of the `QuadraticEquation` class should return the string `'Quadratic Equation'`.

```js
({ test: () => runPython(`
foo = _Node(_code).find_class("QuadraticEquation").find_function("type")
assert foo.has_args("self")
assert foo.has_decorators("property")
assert foo.has_return("'Quadratic Equation'")
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                else:
                    terms.insert(0, f'{coefficient:+}x**{n}')

        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r'(?<!\d)1(?=x)', '', equation_string.strip('+'))
--fcc-editable-region--
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return [x]

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}

class QuadraticEquation(Equation):
    def __init__(self, *args):
        super().__init__(*args)
        self.delta = (
            self.coefficients[1] ** 2 - 4 * self.coefficients[2] * self.coefficients[0]
        )
    @property
    def degree(self):
        return 2
--fcc-editable-region--
    def solve(self):
        if self.delta < 0:
            self.results.append('No real roots')
            return []
            
        x1 = (-self.coefficients[1] + (self.delta) ** 0.5) / (2 * self.coefficients[2])
        x2 = (-self.coefficients[1] - (self.delta) ** 0.5) / (2 * self.coefficients[2])
        if self.delta == 0:
            self.results.append(f'x = {x1:.3f}')
            return [x1]

        self.results.extend([f'x1 = {x1:.3f}', f'x2 = {x2:.3f}'])
        return [x1, x2]

    def analyze(self):
        x = -self.coefficients[1] / (2 * self.coefficients[2])
        y = (
            self.coefficients[2] * x**2
            + self.coefficients[1] * x
            + self.coefficients[0]
        )
        vertex = {'x': x, 'y': y}
        if self.coefficients[2] > 0:
            concavity = 'upwards'
            vertex['m'] = 'min'
        else:
            concavity = "downwards"
            vertex['m'] = 'max'
        coord = f'({x:.3f}, {y:.3f})'
        self.details.append(f'concavity = {concavity}\n{vertex["m"]} = {coord}')
        return {'vertex': vertex, 'concavity': concavity}

def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError('Argument must be an Equation object')
    if not equation.results:
        equation.solve()
    if not equation.details:
        equation.analyze()

lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)
```
