---
id: 665ee783d35cb68875c626d4
title: Step 31
challengeType: 20
dashedName: step-31
---

# --description--

Now, remove both the `lin_eq.analyze()` and `print(lin_eq.details)` calls from your code.

# --hints--

You should remove both your `lin_eq.analyze()` and `print(lin_eq.details)` calls.

```js
({ test: () => runPython(`
assert not _Node(_code).has_call("lin_eq.analyze()")
assert not _Node(_code).has_call("print(lin_eq.details)")
`) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):   
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        for arg in args:
            if not isinstance(arg, (int, float)):
                raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {n: arg for n, arg in enumerate(reversed(args))}
        self.results = []
        self.details = []
        
    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if coefficient:
                if n == 0:
                    terms.append(f'{coefficient:+}')
                elif n == 1:
                    terms.insert(0, f'{coefficient:+}x')
                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')
            
    @property
    @abstractmethod
    def degree(self): pass    
    
    @abstractmethod
    def solve(self): pass
        
    @abstractmethod
    def analyze(self): pass
        
class LinearEquation(Equation):
    @property
    def degree(self):
        return 1
    
    def solve(self):
        x = -self.coefficients[0] / self.coefficients[1]
        self.results.append(f'x = {x:.3f}')
        return x

    def analyze(self):
        self.details.append(
            f'slope = {self.coefficients[1]}\ny-intercept = {self.coefficients[0]}'
        )
        return {'slope': self.coefficients[1], 'intercept': self.coefficients[0]}
--fcc-editable-region--
lin_eq = LinearEquation(2, 3)
print(lin_eq)
lin_eq.analyze()
print(lin_eq.details)
--fcc-editable-region--
```
