---
id: 685523c05fb41d0848c4b52f
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside the first `div` element, create a `span` element with the text `Soundflow`, a `h1` element with the text `Discover New Music`, and a `p` element with the text `Stream your favorite tracks and discover new artists.`.

# --hints--

You should create a `span` element with the text `Soundflow` inside the first `div` element.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

You should create a `h1` element with the text `Discover New Music` inside the first `div` element.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

You should create a `p` element with the text `Stream your favorite tracks and discover new artists.` inside the first `div` element.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--

      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
