---
id: 5d822fd413a79914d39e98c9
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Karibu kwenye mradi wa CSS Variables Skyline! Anza kwa kuongeza tamko la `!DOCTYPE html` juu ya hati ili kivinjari kijue ni aina gani ya hati inayosoma.

# --hints--

Msimbo wako unapaswa kuwa na rejeleo la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Unapaswa kujumuisha nafasi baada ya rejeleo la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Unapaswa kufafanua aina ya hati kuwa `html`.

```js
assert(code.match(/html/gi));
```

Unapaswa kufunga tamko la `DOCTYPE` kwa `>` baada ya aina.

```js
assert(code.match(/html\s*>/gi));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
