---
id: 5d822fd413a79914d39e98d0
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

Tengeneza kipengele cha `div` ndani ya `body` na kiwe na class ya `background-buildings`. Hii itakuwa chombo cha kikundi cha majengo.

# --hints--

Unapaswa kuunda kipengele cha `div`.

```js
assert.exists(document.querySelector('div'));
```

Kipengele cha `div` kinafaa kuwa ndani ya `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Kipengele chako cha `div` kinapaswa kuwa na class ya `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

