---
id: 5d822fd413a79914d39e98d1
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Kipe kipengele chako cha `.background-buildings` `width` na `height` ya `100%` ili iwe na upana na urefu kamili wa mzazi wake, `body`.

# --hints--

Unapaswa kutumia darasa la `background-buildings` ili kuchagua kipengele sahihi.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Kipengele chako cha `.background-buildings` kinapaswa kuwa na `width` ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Kipengele chako cha `.background-buildings` kinapaswa kuwa na `height` ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

