---
id: 5d822fd413a79914d39e98ca
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ongeza tagi za `html` za kufungua na kufunga chini ya `DOCTYPE` ili uwe na mahali pa kuanza kuweka msimbo. Hakikisha umeweka lugha kwa Kiingereza.

# --hints--

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Kipengele chako cha `html` kinafaa kuwa na tagi ya ufunguzi yenye sifa ya `lang` ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na tagi ya kufunga.

```js
assert(code.match(/<\/html\s*>/));
```

Tagi zako `html` zinapaswa kuwa katika mpangilio sahihi.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Unapaswa kuwa na kipengele kimoja cha `html` tu.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
