---
id: 5d822fd413a79914d39e98cd
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Katika CSS, unaweza kulenga kila kitu kwa kutumia alama ya nyota (*). Ongeza mpaka kwenye kila kitu kwa kutumia kichaguzi cha `*`, na kuipa `border` ya `1px solid black`. Hii ni mbinu inayosaidia kuona vizuri mahali pa vipengele na ukubwa wake. Baadaye utaondoa hii.

# --hints--

Unapaswa kutumia kichaguzi cha `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Unapaswa kutumia sifa ya `border` kuunda mtindo wa vipengele vyote.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Vipengele vyote vinapaswa kuwa na mpaka wa `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
