---
id: 5d822fd413a79914d39e98cf
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Unaweza kuona `body` (ni sanduku la ndani kabisa kwenye ukurasa wako); sanduku lililozunguka ni kipengele cha `html`. Fanya `body` yako kujaza skrini yote kwa kuipa sifa ya `height` ya `100vh`. Ondoa `margin` ya chaguo-msingi kutoka kwenye kipengele cha `body` kwa kumwekea `margin` ya `0`. Hatimaye, weka sifa ya `overflow` kuwa `hidden` ili kuficha pau zozote za kusogeza zinazoonekana wakati kitu kinapita nyuma ya kituo cha kutazama.

# --hints--

Unapaswa kutumia kichaguzi cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` yako inapaswa kuwa na `height` ya `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` yako inapaswa kuwa na `margin` ya `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body` yako inapaswa kuwa na sifa ya `overflow` iliyowekwa kuwa `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

