---
id: 5d822fd413a79914d39e98d3
title: Hatua ya 11
challengeType: 0
dashedName: step-11
---

# --description--

Wekeza vipengele vinne vya `div` ndani ya sanduku la `.bb1`. Zipe darasa za `bb1a`, `bb1b`, `bb1c`, na `bb1d` kwa utaratibu huo. Jengo hili litakuwa na sehemu nne.

# --hints--

Unapaswa kuunda vipengele vinne vya vipya vya `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Unapaswa kukipa kipengele kimoja kipya cha `div` class ya `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Unapaswa kukipa kipengele kimoja kipya cha `div` class ya `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Unapaswa kukipa kipengele kimoja kipya cha `div` class ya `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Unapaswa kukipa kipengele kimoja kipya cha `div` class ya `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Unapaswa kuweka vipengele vipya vya `div` kwa utaratibu sahihi.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Unapaswa kuweka kipengele cha `div` ndani ya kipengele cha `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

