---
id: 5d822fd413a79914d39e98d2
title: Hatua ya 10
challengeType: 0
dashedName: step-10
---

# --description--

Weka `div` yenye class ya `bb1` katika kontena la majengo ya usuli (background-buildings). Fungua faili yako ya `styles.css`, kisha weka `.bb1` kuwa na `width` ya `10%` na `height` ya `70%`. "bb" inasimamia "background building", hili litakuwa jengo lako la kwanza.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Unapaswa kuipa `div` mpya class ya `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Unapaswa kutumia kichaguzi cha class ya `.bb1` ili kuunda mtindo wa kipengele.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Unapaswa kukipa kipengele cha `.bb1` `width` ya `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Unapaswa kukipa kipengele cha `.bb1` `height` ya `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

