---
id: 5d5a813321b9e3db6c106a46
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript è un linguaggio molto potente che ti permette di realizzare siti web interattivi.

Per iniziare crea il tuo boilerplate HTML standard con `DOCTYPE`, `html`, `head` e `body`, poi aggiungi un tag `meta` per il `charset`. Aggiungi un elemento `title` usando il testo `RPG - Dragon Repeller`. Includi un tag `link` per il foglio di stile per collegare il file `styles.css`.

Infine crea un elemento `div` con l'`id` impostato su `game` all'interno del `body`.

# --hints--

Il codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un carattere `>`.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di apertura. Non dimenticare l'attributo `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio dell'HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag `head` di apertura.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag `head` di chiusura.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag `body` di apertura.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag `body` di chiusura.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe essere dentro l'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere dentro l''elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Il codice dovrebbe avere un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

L'elemento `meta` dovrebbe avere un attributo `charset` con il valore `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il codice dovrebbe avere un elemento `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Dovresti avere un elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

L'elemento `div` dovrebbe avere un attributo `id` con il valore `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

L'elemento `div` dovrebbe trovarsi all'interno dell'elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
