---
id: 62a1166ed9a56d439c0770e7
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Crea tre elementi `span` all'interno dell'elemento `#stats`. Assegna a ognuno di loro la classe `stat`. Poi dai al primo il testo `XP: 0`, al secondo il testo `Health: 100` e al testo il testo `Gold: 50`.

# --hints--

Dovresti avere tre elementi `span` all'interno dell'elemento `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Dovresti dare ai tre nuovi elementi `span` la classe `stat`.

```js
assert.exists(document.querySelectorAll('#stats > .stat')?.length, 3);
```

Il primo elemento `.stat` dovrebbe avere il testo `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Il secondo elemento `.stat` dovrebbe avere il testo `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Il terzo elemento `.stat` dovrebbe avere il testo `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
--fcc-editable-region--
<body>
    <div id="game">
        <div id="stats">

        </div>
        <div id="controls"></div>
        <div id="monsterStats"></div>
        <div id="text"></div>
    </div>
</body>
--fcc-editable-region--
</html>
```
