---
id: 62a255dae245b52317da824a
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Uno degli strumenti più potenti è la tua console per sviluppatori. A seconda del browser, potrebbe aprirsi premendo `F12` o `Ctrl+Shift+I`. Su Mac puoi premere `Option + ⌘ + C` e selezionare "Console". Puoi anche cliccare sul pulsante "Console" sopra la finestra di anteprima per vedere la nostra console integrata.

La console per sviluppatori include gli errori generati dal tuo codice, ma puoi anche usarla per vedere i valori delle variabili nel codice, che è molto utile per il debugging.

Aggiungi la riga `console.log("Hello World");` tra i tag `script`. Quindi clicca sul pulsante "Console" per aprire la console. You should see the text `"Hello World"`.


# --hints--

Dovresti aggiungere la riga `console.log("Hello World");` nel codice.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

La riga `console.log("Hello World");` dovrebbe trovarsi tra i tag `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
