---
id: 62a23d1c5f1c93161f3582ae
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In modo simile all'elemento `#stats`, l'elemento `#monsterStats` ha bisogno di due elementi `span`. Assegna loro la classe `stat` e dai al primo elemento il testo `Monster Name:` e al secondo il testo `Health:`. In ognuno di essi, dopo il testo aggiungi un elemento `strong` con annidato un elemento `span` vuoto.

# --hints--

L'elemento `#monsterStats` dovrebbe avere due elementi `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

I nuovi elementi `span` dovrebbero avere entrambi un attributo `class` con il valore `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Il primo elemento `span` dovrebbe avere il testo `Monster Name:`. Assicurati che la spaziatura sia corretta.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].innerText, 'Monster Name: ');
```

Il secondo elemento `span` dovrebbe avere il testo `Health:`. Assicurati che la spaziatura sia corretta.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[1].textContent, 'Health: ');
```

Il primo elemento `span` dovrebbe avere un elemento `strong` con annidato un elemento `span` vuoto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Il secondo elemento `span` dovrebbe avere un elemento `strong` con annidato un elemento `span` vuoto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Gli elementi `strong` e `span` dovrebbero trovarsi dopo il testo.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name: <strong>/);
assert.match(spans[1].innerHTML, /Health: <strong>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
--fcc-editable-region--
        <div id="monsterStats">

        </div>
--fcc-editable-region--
        <div id="text"></div>
    </div>
</body>
</html>
```
