---
id: 62a3cfc8328d3351b95d4f61
title: Step 42
challengeType: 0
dashedName: step-42
---

# --description--

La proprietà `innerText` controlla il testo che appare in un elemento HTML. Ad esempio:

```js
const info = document.querySelector("#info");
info.innerText = "Hello World";
```

Questo codice cambia l'elemento assegnato alla variabile `info` in modo che abbia il testo `Hello World`.

Quando un giocatore clicca sul pulsante `Go to store`, vuoi che cambino i pulsanti e il testo. Rimuovi il codice nella funzione `goStore` e aggiungi una riga che aggiorna il testo di `button1` in `Buy 10 health (10 gold)`.

# --hints--

Non dovresti avere la riga `console.log("Going to store.");` nel tuo codice.

```js
assert.notMatch(code, /console.log\(('|")Going to store\.\1\)/);
```

Dovresti usare la dot notation per accedere alla proprietà `innerText` di `button1`.

```js
assert.match(code, /button1\.innerText/);
```

Non dovresti usare `let` o `const` per accedere alla proprietà `innerText` di `button1`.

```js
assert.notMatch(code, /(let|const)\s+button1.innerText/);
```

Dovresti aggiornare la proprietà `innerText` di `button1` in modo che sia `Buy 10 health (10 gold)`.

```js
assert.match(code, /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

Dovresti aggiornare la proprietà `innerText` nella funzione `goStore`.

```js
assert.match(goStore.toString(), /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    border: 1px solid black;
    padding: 5px;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
