---
id: 62a39f5a5790eb27c1e5d4bf
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

La tua visuale adesso è passata al nuovo file `script.js`. Ricorda che puoi usare le schede in alto per passare da un file all'altro.

Aggiungi la riga `console.log("Hello World");` in questo file e guarda come appare sullo schermo.

# --hints--

Dovresti avere la riga `console.log("Hello World");` nel tuo codice.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
