---
id: 62a3a488b24fb32b91155d56
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Alle variabili può essere <dfn>assegnato</dfn> un valore. Se lo fai quando le dichiari, si parla di <dfn>inizializzazione</dfn>. Ad esempio:

```js
let age = 32;
```

This would initialize the `age` variable with a value of `32`, a number.

Inizializza la variabile `xp` con il valore `0`, un numero.

# --hints--

`xp` dovrebbe avere il valore `0`.

```js
assert.equal(xp, 0);
```

Dovresti inizializzare la variabile `xp` a `0`. Non dimenticare il punto e virgola alla fine della riga.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp;
--fcc-editable-region--
```
