---
id: 62a3b41c9494f937560640ab
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Dichiara altre due variabili chiamate `monsterHealth` e `inventory`, ma non inizializzarle.

# --hints--

Dovresti usare la parola chiave `let` per dichiarare la variabile `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Dovresti usare il camelCase per il nome della variabile `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` non dovrebbe avere un valore.

```js
assert.isUndefined(monsterHealth);
```

Non dovresti assegnare un valore alla variabile `monsterHealth`. Ricorda il punto e virgola.

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

Dovresti usare la parola chiave `let` per dichiarare la variabile `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` non dovrebbe avere un valore.

```js
assert.isUndefined(inventory);
```

Non dovresti assegnare un valore alla variabile `inventory`. Ricorda il punto e virgola.

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
