---
id: 62a3a7e4f1060e2fc5ffb34b
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Create another variable called `currentWeapon` and set it to `0`.

Quando il nome di una variabile contiene più parole, la convenzione in JavaScript è di usare ciò che è chiamato <dfn>camelCase</dfn>. La prima parola è tutta minuscola e la prima lettera di ogni parola successiva è maiuscola.

```js
let thisIsCamelCase;
```

# --hints--

You should use `let` to declare a variable called `currentWeapon`.

```js
assert.match(code, /let\s+currentWeapon/i);
```

Dovresti usare il camelCase per il nome della variabile.

```js
assert.match(code, /currentWeapon/);
```

La variabile `currentWeapon` dovrebbe essere impostata a `0`.

```js
assert.equal(currentWeapon, 0);
```

Dovresti inizializzare la variabile a `0`.

```js
assert.match(code, /let\s+currentWeapon\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
