---
id: 62a3a0a3c0a4b32915d26a6e
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Rimuovi la riga `console.log("Hello World");` per iniziare a scrivere il codice del progetto.

In JavaScript, una <dfn>variabile</dfn> viene usata per conservare un valore. Per usare una variabile, la devi prima <dfn>dichiarare</dfn>. Ad esempio, per dichiarare una variabile chiamata `camperbot`, dovrai scrivere:

```js
let camperbot;
```

The `let` keyword tells JavaScript you are declaring a variable. Dichiara una variabile chiamata `xp`.

# --hints--

Non dovresti avere la riga `console.log("Hello World");` nel tuo codice.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

You should use the `let` keyword to declare your variable.

```js
assert.match(code, /let/);
```

Dovresti dichiarare una variabile chiamata `xp`.

```js
assert.match(code, /let\s+xp/);
```

Non dovresti assegnare un valore alla variabile.

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

Non dimenticare il punto e virgola alla fine della riga.

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
