---
id: 62a3b41c9494f937560640ab
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Dichiara altre due variabili chiamate `monsterHealth` e `inventory`, ma non inizializzarle.

# --hints--

Dovresti usare la parola chiave `let` per dichiarare la variabile `monsterHealth`.

```js
assert.match(code, /let monsterHealth/i);
```

Dovresti usare il camelCase per il nome della variabile `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` non dovrebbe avere un valore.

```js
assert.isUndefined(monsterHealth);
```

Non dovresti assegnare un valore alla variabile `monsterHealth`. Ricorda il punto e virgola.

```js
assert.match(code, /let monsterHealth;/);
```

Dovresti usare la parola chiave `let` per dichiarare la variabile `inventory`.

```js
assert.match(code, /let inventory/i);
```

`inventory` non dovrebbe avere un valore.

```js
assert.isUndefined(inventory);
```

Non dovresti assegnare un valore alla variabile `inventory`. Ricorda il punto e virgola.

```js
assert.match(code, /let inventory;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
