---
id: 62a3bec30ea7f941412512dc
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

Proprio come hai fatto con i pulsanti, crea delle variabili per i seguenti `id` e usa `querySelector()` per dar loro l'elemento come valore:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats` e `monsterName`.

Ricorda di dichiarare le variabili con la parola chiave `const` e di dar loro nomi corrispondenti agli `id`.

# --hints--

Dovresti dichiarare una variabile `text` con `const`.

```js
assert.match(code, /const text/);
```

La variabile `text` dovrebbe avere il valore dell'elemento `#text`.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

Dovresti dichiarare una variabile `xpText` con `const`.

```js
assert.match(code, /const xpText/);
```

La variabile `xpText` dovrebbe avere il valore dell'elemento `#xpText`.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

Dovresti dichiarare una variabile `healthText` con `const`.

```js
assert.match(code, /const healthText/);
```

La variabile `healthText` dovrebbe avere il valore dell'elemento `#healthText`.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

Dovresti dichiarare una variabile `goldText` con `const`.

```js
assert.match(code, /const goldText/);
```

La variabile `goldText` dovrebbe avere il valore dell'elemento `#goldText`.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

Dovresti dichiarare una variabile `monsterStats` con `const`.

```js
assert.match(code, /const monsterStats/);
```

La variabile `monsterStats` dovrebbe avere il valore dell'elemento `#monsterStats`.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

Dovresti dichiarare una variabile `monsterName` con `const`.

```js
assert.match(code, /const monsterName/);
```

La variabile `monsterName` dovrebbe avere il valore dell'elemento `#monsterName`.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
