---
id: 62a3b3eab50e193608c19fc6
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Declare a variable called `fighting` but do not initialize it with a value. Ricordati di terminare la riga con un punto e virgola.

# --hints--

Dovresti usare `let` per dichiarare una variabile chiamata `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

La variabile `fighting` non dovrebbe avere un valore.

```js
assert.isUndefined(fighting);
```

Non dovresti assegnare un valore alla variabile `fighting`. Non dimenticare il punto e virgola alla fine della riga.

```js
assert.match(code, /let\s+fighting\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
--fcc-editable-region--
```
