---
id: 62a3b49686792938718b90d3
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

I valori delle variabili che hai assegnato finora sono tutti dei numeri. JavaScript ha diversi tipi di dati. Il prossimo che userai è la <dfn>stringa</dfn>. Le stringhe vengono usate per conservare cose come parole o testo. Le stringhe sono racchiuse da virgolette doppie, singole o backtick. Ecco un esempio di dichiarazione di una variabile con una stringa:

```js
let developer = "Naomi";
```

Assegna la variabile `inventory` in modo che il suo valore sia `stick`.

# --hints--

Dovresti impostare `inventory` con la stringa `stick`.

```js
assert.equal(inventory, "stick");
```

Dovresti inizializzare la variabile `inventory` con la stringa `stick`.

```js
assert.match(code, /let\s+inventory\s*=\s*('|"|`)stick\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
