---
id: 62a3b506dbaead396f58a701
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

L'inventario del giocatore nel tuo gioco potrà contenere più oggetti. Dovrai usare un tipo di dato adatto per questo scopo. Un <dfn>array</dfn> può essere usato per contenere valori multipli. Ad esempio:

```js
let order = ["first", "second", "third"];
```

Questo è un array che contiene tre valori. Nota che i valori sono separati da una virgola.

Cambia la variabile `inventory` in modo che sia un array con le stringhe `stick`, `dagger` e `sword`.

# --hints--

La variabile `inventory` dovrebbe essere un array.

```js
assert.isArray(inventory);
```

La variabile `inventory` dovrebbe contenere tre valori.

```js
assert.lengthOf(inventory, 3);
```

La variabile `inventory` dovrebbe includere la stringa `stick`.

```js
assert.include(inventory, "stick");
```

La variabile `inventory` dovrebbe includere la stringa `dagger`.

```js
assert.include(inventory, "dagger");
```

La variabile `inventory` dovrebbe includere la stringa `sword`.

```js
assert.include(inventory, "sword");
```

La variabile `inventory` dovrebbe avere i valori nell'ordine corretto.

```js
assert.deepEqual(inventory, ["stick", "dagger", "sword"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = "stick";
--fcc-editable-region--
```
