---
id: 62a3c668afc43b4a134cca81
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Le funzioni sono degli strumenti speciali che ti permettono di eseguire sezioni di codice in momenti specifici. Puoi dichiarare funzioni con la parola chiave `function`. Ecco un esempio di una funzione chiamata `functionName` - nota le parentesi graffe di apertura e chiusura. Indicano la porzione di codice all'interno della funzione.

```js
function functionName() {

}
```

Crea una funzione vuota chiamata `goStore`. Corrisponderà alla variabile `goStore` usata precedentemente.

# --hints--

Dovresti dichiarare `goStore` con la parola chiave `function`.

```js
assert.match(code, /function\s*goStore()/);
```

`goStore` dovrebbe essere definita.

```js
assert.isDefined(goStore);
```

`goStore` dovrebbe essere una funzione.

```js
assert.isFunction(goStore);
```

`goStore` dovrebbe essere una funzione vuota.

```js
assert.equal(goStore.toString(), "function goStore() {}");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--

--fcc-editable-region--
```
