---
id: 62a8b0b5053f16111b0b6b5f
title: Step 59
challengeType: 0
dashedName: step-59
---

# --description--

Ora puoi consolidare una parte del tuo codice. Inizia copiando il codice dall'interno della funzione `goTown` e incollandolo nella funzione `update`. Poi, rimuovi tutto il codice dalle funzioni `goTown` e `goStore`.

# --hints--

La funzione `update` dovrebbe impostare `button1.innerText` su `Go to store`.

```js
assert.match(update.toString(), /button1\.innerText\s*=\s*('|")Go to store\1/);
```

La funzione `update` dovrebbe impostare `button2.innerText` su `Go to cave`.

```js
assert.match(update.toString(), /button2\.innerText\s*=\s*('|")Go to cave\1/);
```

La funzione `update` dovrebbe impostare `button3.innerText` su `Fight dragon`.

```js
assert.match(update.toString(), /button3\.innerText\s*=\s*('|")Fight dragon\1/);
```

La funzione `update` dovrebbe impostare `button1.onclick` su `goStore`.

```js
assert.match(update.toString(), /button1\.onclick\s*=\s*goStore/);
```

La funzione `update` dovrebbe impostare `button2.onclick` su `goCave`.

```js
assert.match(update.toString(), /button2\.onclick\s*=\s*goCave/);
```

La funzione `update` dovrebbe impostare `button3.onclick` su `fightDragon`.

```js
assert.match(update.toString(), /button3\.onclick\s*=\s*fightDragon/);
```

La funzione `update` dovrebbe impostare `text.innerText` su `You are in the town square. You see a sign that says \"Store\".`.

```js
assert.match(update.toString(), /text\.innerText\s*=\s*"You are in the town square. You see a sign that says \\"Store\\"\."/);
```

La funzione `goTown` dovrebbe essere vuota.

```js
assert.match(goTown.toString(), /function goTown\(\) \{\}/);
```

La funzione `goStore` dovrebbe essere vuota.

```js
assert.match(goStore.toString(), /function goStore\(\) \{\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");
const locations = [
    {
        name: "town square",
        "button text": ["Go to store", "Go to cave", "Fight dragon"],
        "button functions": [goStore, goCave, fightDragon],
        text: "You are in the town square. You see a sign that says \"Store\"."
    },
    {
        name: "store",
        "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
        "button functions": [buyHealth, buyWeapon, goTown],
        text: "You enter the store."
    }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function update(location) {

}

function goTown() {
  button1.innerText = "Go to store";
  button2.innerText = "Go to cave";
  button3.innerText = "Fight dragon";
  button1.onclick = goStore;
  button2.onclick = goCave;
  button3.onclick = fightDragon;
  text.innerText = "You are in the town square. You see a sign that says \"Store\".";
}

function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
  button1.onclick = buyHealth;
  button2.onclick = buyWeapon;
  button3.onclick = goTown;
  text.innerText = "You enter the store.";
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}
```
