---
id: 62fc20387ef88d1d1998aac5
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Usa `querySelector()` per ottenere gli altri due elementi `button` usando i loro `id`: `button2` e `button3`. Conservali in variabili chiamate `button2` e `button3`. Ricorda di usare `const`.

# --hints--

Dovresti dichiarare una variabile `button2` con `const`.

```js
assert.match(code, /const\s+button2/);
```

Your `button2` variable should have the value of your `"#button2"` element.

```js
assert.deepEqual(button2, document.querySelector('#button2'));
```

Dovresti dichiarare una variabile `button3` con `const`.

```js
assert.match(code, /const\s+button3/);
```

Your `button3` variable should have the value of your `"#button3"` element.

```js
assert.deepEqual(button3, document.querySelector('#button3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
