---
id: 69eefea5a952ae85871df0dc
title: Data Sanitizer Sweep
challengeType: 26
dashedName: data-sanitizer-sweep
---

# --description--

In this lab, you will clean semi-structured log files by stripping disallowed tokens and tracking repair statistics.

Here are some examples:

`
  "DEBUG: Server started on port 3000",
  "INFO: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%",
  "//INFO: Cache cleared successfully",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
`

For example, if the disallowed tokens are `e, o, r` and the replacement character is `#`, sanitizing the first line changes:

DEBUG: Server started on port 3000 to DEBUG: S##v## sta#t#d #n p##t 3000

Fatal markers are characters that should stop parsing immediately. In this example, `$` is a fatal marker. If it appears (like in the second-to-last log line), the program should stop processing additional logs.

After processing, the program should print a repair summary such as:

`Total Entry: 6 Sanitized Entry Count: 4 finished sanitize logs process`.

The sanitized entry count includes both sanitized entries and entries that did not need sanitizing. If this count is lower than total entries, parsing stopped early because a fatal marker was found. If no fatal marker appears, sanitized entry count should equal total entries.

For this lab, if a log string starts with `//` then in `cleanLogs` function this string should be skipped and count as a successfully sanitized entry count. For example, if the log array have 4 logs strings in total and all strings start with `//` then this means the result will be 

`Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process`.

If sanitization fails, the output summary should also include a reason, such as:

`Sanitizing stopped Failure Reason: Fatal character occurred`.

For this lab the fatal marker will be `$`.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Story:** 

1. You should have a variable called logs which is an array of strings representing log entries.

2. You should implement sanitizeLog(entry, blacklist) that loops over characters and replaces banned patterns. -- entry: The individual strings in the array -- blacklist: Currently an array of characters

3. You should implement cleanLogs(logs, blacklist) that uses continue to skip informational lines and break on fatal markers. -- logs: The array of log strings -- blacklist: Currently an array of characters

4. You should track counts of sanitized vs untouched entries inside an accumulator object. This object stores running totals as you loop through logs, like `{ sanitized: 0, untouched: 0 }`.

5. You should output a summary showing total entries, sanitized entries, and failure reasons.

# --before-each--

# --hints--

You should have a function called `sanitizeLog`.

```js
assert.isFunction(sanitizeLog)
```

You should have a function called `cleanLogs`.

```js
assert.isFunction(cleanLogs)
```

`sanitizeLog("DEBUG: Server started on port 3000",["a", "A", "c", "T"])` should return `DEBUG: Server st#rted on port 3000`

```js
assert.strictEqual(sanitizeLog("DEBUG: Server started on port 3000",["a", "A", "c", "T"]),"DEBUG: Server st#rted on port 3000");
```

`sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "c", "T"])` should return `W#RN: Disk us#ge #t 92%, threshold=90%`

```js
assert.strictEqual(sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "c", "T"]),'W#RN: Disk us#ge #t 92%, threshold=90%');
```

`cleanLogs(logs_1,["a", "A", "c", "T"])` should process all log strings without error.

```js
assert.strictEqual(cleanLogs(logs_1,["a", "A", "c", "T"]), 'Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process');
```

`cleanLogs(logs_2,["a", "A", "c", "T"])` should result in error message due to fatal character.

```js
assert.strictEqual(cleanLogs(logs_2,["a", "A", "c", "T"]), 'Sanitizing stopped Failure Reason: Fatal character occurred');
```

`cleanLogs(logs_3,["a", "A", "c", "T"])` should process all log strings without error.

```js
assert.strictEqual(cleanLogs(logs_3,["a", "A", "c", "T"]), 'Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process');
```

# --seed--

## --seed-contents--

```js
function printCustomMessage(entryCount, sanitizedCount, failureReason) {
  if (failureReason == "") {
    return(
      "Total Entry: " +
        entryCount +
        " Sanitized Entry Count: " +
        sanitizedCount +
        " finished sanitize logs process."
    )
  } else {
    return("Sanitizing stopped Failure Reason: " + failureReason);
  }
};

var logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

var logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

var logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --solutions--

```js
function sanitizeLog(entry, blacklist) {
  var sanitizedEntry = [];

  for (let i = 0; i < entry.length; i++) {
    if (blacklist.includes(entry[i])) {
      sanitizedEntry.push("#");
    } else {
      sanitizedEntry.push(entry[i]);
    }
  }

  return sanitizedEntry.join("");
};

function cleanLogs(logs, blacklist) {
  let accumulator = 0;
  let fatalError = "$";
  let informationalIndicator = "//";

  let sanitizedLogs = [];

  for (let i = 0; i < logs.length; i++) {
    if (logs[i].includes(fatalError)) {
      printCustomMessage(logs.length, accumulator, "Fatal character occurred");
      break;
    } else if (logs[i].includes(informationalIndicator)) {
      accumulator += 1;
      continue;
    } else {
      tmpSanitized = sanitizeLog(logs[i], blacklist);
      sanitizedLogs.push(tmpSanitized);

      accumulator += 1;
    }
  }

  printCustomMessage(logs.length, accumulator, "");
};
```
