---
id: 6823f9df49cc206af5471a30
title: Debug Camperbot's Profile Page
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot just started learning how to code and is trying to build a profile page. They asked a friend to look over their code so far and their friend said it had tons of errors.

Your job is to fix all of Camperbot's errors so they can continue to build their profile page. Complete the items in the user stories below and click "Run the Tests" to see if you fixed all the errors.

**User Stories:**

1. Camperbot is not following best practices and has multiple `h1` elements on the page. Remove all `h1` elements except for the first one.
2. Camperbot is trying to use a `heading2` element. But that element does not exist in HTML. Fix that mistake so it uses the correct second level heading element. 
3. Camperbot is trying to add some paragraphs but is using the wrong element. Fix these errors to use the correct HTML element for paragraphs.
4. Camperbot is using an `h3` element for the `Background and Interests` subheading but it has a syntax error. Spot the issue and resolve it. 

# --hints--

You should only have one `h1` on the page.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

You should completely the text `Welcome!` from the page.

```js
assert.isTrue(code.indexOf("Welcome") === -1); 
```

You should completely the text `this is fun` from the page.

```js
assert.isTrue(code.indexOf("this is fun") === -1); 
```

You should completely the text `Coding is awesome!` from the page.

```js
assert.isTrue(code.indexOf("Coding is awesome!") === -1); 
```

You should not remove the `<h1>Hello from Camperbot!</h1>` from the page.

```js
assert.strictEqual(document.querySelector("h1:first-of-type")?.innerText, "Hello from Camperbot!");
```

You should have a valid `h2` element on the page.

```js
assert.exists(document.querySelector("h2"));
```

You should not change the `About` text for the `h2` element.

```js
assert.strictEqual(document.querySelector("h2")?.innerText, "About");
```

Your first paragraph element should have a valid opening tag. Remember that `p` elements have opening tags using this syntax: `<elementName>`.

```js
assert.match(code, /<p>\s*My\s+name/);
```

Your first paragraph element should have a valid closing tag. Remember that `p` elements have closing tags using this syntax: `</elementName>`.

```js
assert.match(code, /new\s+things\.<\/p>/);
```

Your second paragraph element should have a valid opening tag. Remember that `p` elements have opening tags using this syntax: `<elementName>`.

```js
assert.match(code, /<p>\s*I\s+enjoy/);
```

Your second paragraph element should have a valid closing tag. Remember that `p` elements have closing tags using this syntax: `</elementName>`.

```js
assert.match(code, /solving\s+puzzles\.<\/p>/);
```

Your first valid paragraph element should have the text `My name is Camperbot and I love reading and learning new things.`. Double check your spelling.

```js
assert.strictEqual(document.querySelector("p:first-of-type")?.innerText, "My name is Camperbot and I love reading and learning new things.");
```

Your `h3` element should have a valid closing tag. Remember that `h3` elements have closing tags using this syntax: `</elementName>`.

```js
assert.match(code, /<\/h3\>/);
```

Your second valid paragraph element should have the text `I enjoy solving puzzles.`. Double check your spelling.

```js
assert.strictEqual(document.querySelector("h3 + p")?.innerText, "I enjoy solving puzzles.");
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<h1>Welcome!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love reading and learning new things.<pp>

<h3>Background and Interests<h3>
<h1>This is fun</h1>
<pp>I enjoy solving puzzles.<pp>

<h1>Coding is awesome!</h1>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love reading and learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
