---
id: 669fc7e141e4703748c558bf
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside the `html` element, add a `head` element.

Inside the `head` element, add a `title` element. 

The `title` element's text should be `Mr. Whiskers' Blog`.

# --hints--

You should have an opening `head` tag.

```js
assert.match(code, /<head>/i);
```

You should have a closing `head` tag.

```js
assert.match(code, /<\/head>/i);
```

You should have an opening `title` tag.

```js
assert.match(code, /<title>/i);
```

You should have a closing `title` tag.

```js
assert.match(code, /<\/title>/i);
```

Your `title` element should be nested in your `head` element.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Your `title` element should have the text `Mr. Whiskers' Blog`. You may need to check your spelling.

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Mr. Whiskers' Blog");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
