---
id: 669fd003cb89ee3c2402e041
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

In this introductory content, you will want to show an image of Mr. Whiskers with a caption.

Below the `h1` element, start by adding a `figure` element.

# --hints--

Your `figure` element should have an opening tag. 

```js
assert.match(code, /<figure\>/);
```

Your `figure` element should have a closing tag. 

```js
assert.match(code, /<\/figure\>/);
```

Your `figure` element should be within your `header` element.

```js
const figureElement = document.querySelector('figure');
assert(figureElement.parentElement.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>

    </header>
  </body>
--fcc-editable-region--
</html>
```
