---
id: 669fdc11c9b0263fe0814a7a
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Inside each of the `li` elements, you will need to have an anchor element.

For the first anchor element, the text should be `About` and the `href` attribute value should be `"#about"`. The hash symbol in front of the `about` represents an `id` name which be added later on in the project.

For the second anchor element, the text should be `Posts` and the `href` attribute value should be `"#posts"`.

For the third anchor element, the text should be `Contact` and the `href` attribute value should be `"#contact"`.

# --hints--

Your first anchor element should have the text of `About` inside your first `li` element.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement.innerText, "About");
```

Your first anchor element should have an `href` attribute set to `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

Your second anchor element should have the text of `Posts` inside your second `li` element.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement.innerText, "Posts");
```

Your second anchor element should have an `href` attribute set to `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

Your third anchor element should have the text of `Contact` inside your third `li` element.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement.innerText, "Contact");
```

Your third anchor element should have an `href` attribute set to `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
--fcc-editable-region--
      <nav>
        <ul>
          <li></li>
          <li></li>
          <li></li>
        </ul>
      </nav>
--fcc-editable-region--
    </header>
  </body>
</html>
```
