---
id: 669fc938d38e6e38ace9251e
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `"UTF-8"`.

# --hints--

You should have a `meta` element.

```js
assert(document.querySelector("meta"));
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Your `meta` tag should have a `charset` attribute.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Your `charset` attribute should have a value of `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>

  </head>
  --fcc-editable-region--
</html>
```
