---
id: 669fcb5106eeb13ab7afcd09
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

The first section you will build out is the page <dfn>header</dfn>. 

The `header` element is used represent introductory content like page navigation and other introductory information. 

Here is an example using the `header` element:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Inside the `body` element, add a `header` element.

Inside the `header` element, add an `h1` with the text of `Welcome to Mr. Whiskers' Blog Page!`.


# --hints--

Your code should have an opening `header` tag.

```js
assert.match(code, /<header>/i);
```

Your code should have a closing `header` tag.

```js
assert.match(code, /<\/header>/i);
```

Your `header` element should be within your `body` element.

```js
const headerElement = document.querySelector('header');
assert.strictEqual(headerElement.parentElement.tagName, 'BODY');
```

Your code should have an opening `h1` tag.

```js
assert.match(code, /<h1>/i);
```

Your code should have a closing `h1` tag.

```js
assert.match(code, /<\/h1>/i);
```

Your `h1` element should have the text of `Welcome to Mr. Whiskers' Blog Page!`. Double check for any spelling errors.

```js
const h1Text = document.querySelector("h1").innerText
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

Your `h1` element should be within your `header` element.

```js
const h1Element = document.querySelector('h1');
assert.strictEqual(h1Element.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  --fcc-editable-region--
  <body>
    
  </body>
  --fcc-editable-region--
</html>
```
