---
id: 66a4a467acf8759c6455e61a
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Inside the `figure` element, add an `img` element.

The `img` `src` attribute should have a value of `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` and the `alt` text should have a value of `"a cat peacefully sleeping"`.

Below your `img` element, add a `figcaption` with the text `Mr. Whiskers Sleeping`.

# --hints--

You should have an `img` element. 

```js
assert(document.querySelector("img"));
```

Your `img` element should be within your `figure` element.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement.parentElement.tagName, 'FIGURE');
```

Your `img` element should have a `src` attribute set to `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Your `img` element should have an `alt` attribute set to `"a cat peacefully sleeping"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement.getAttribute('alt');

assert.strictEqual(alt, 'a cat peacefully sleeping');
```

Your `figcaption` element should have an opening tag. 

```js
assert.match(code, /<figcaption\>/);
```

Your `figcaption` element should have a closing tag. 

```js
assert.match(code, /<\/figcaption\>/);
```

Your `figcaption` element should be within your `figure` element.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement.parentElement.tagName, 'FIGURE');
```

Your `figcaption` element should have a the text of `Mr. Whiskers Sleeping`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement.innerText, "Mr. Whiskers Sleeping")
```

Your `figcaption` should be below the `img` element. You have them in the wrong order.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement.lastElementChild.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      --fcc-editable-region--
      <figure>

      </figure>
      --fcc-editable-region--
    </header>
  </body>
</html>
```
