---
id: 67d1edfde28a4133ba64ad26
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Now, enhance the phone number by making it clickable so users can tap and call directly from their devices.

You previously learned how the `tel` makes phone numbers clickable:

```html
<a href="tel:2345678912">234-567-8912</a>
```

Now, locate the first `p` element inside the `address` element. Wrap the phone number in an `a` anchor tag with `href="tel:5555555555"`, ensuring it remains properly formatted within the `p` element.

# --hints--

You should wrap the phone number with an `a` anchor tag.

```js
const firstPElement = document.querySelector('footer #contact address p:nth-of-type(1)');
assert.match(firstPElement.innerHTML, /Phone:\s*<a\b[^>]*>.*?<\/a>/i);
```

You should include the `href` attribute inside the `a` anchor tag.

```js
const firstPElement = document.querySelector('footer #contact address p:nth-of-type(1)');
assert.match(firstPElement.innerHTML, /<a\s+href=["'][^"']+["']>.*?<\/a>/i);
```

Make sure the `href` attribute inside the `a` anchor tag starts with `tel:`, contains only numbers, and does not include dashes `(-)`.

```js
const firstPElement = document.querySelector('footer #contact address p:nth-of-type(1)');
assert.match(firstPElement.innerHTML,  /<a\s+href=["']tel:\+?[0-9]+["']>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Bath</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Birthday Party</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
      </section>
    </main>
--fcc-editable-region--
    <footer>
        <section id="contact">
            <h2>Contact</h2>
            <address>
                <p>Phone: 555-555-5555</p>
                <p>Email: fake@email.com</p>
            </address>
        </section>
    </footer>
--fcc-editable-region--
  </body>
</html>
```
