---
id: 5dc174fcf86c76b9248c6eb2
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML 元素有一个开始和结束标签，内容在标签之间。

以下是基本语法：

```html
<openingTag>content</closingTag>
```

你将学习的第一个元素是 `h1` 元素。 `h1` 元素是一个标题元素，用于网页的主标题。

```html
<h1>This is a main heading</h1>
```

将下面 `h1` 元素的文本从 `Hello World` 更改为 `CatPhotoApp`，并观察浏览器预览中的变化。

完成后，按“检查你的代码”按钮查看代码是否正确。

# --hints--

代码中应包含文本 `CatPhotoApp`。 你可能想要检查你的拼写。

```js
assert(code.match(/catphotoapp/i));
```

你的 `h1` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('h1'));
```

你的 `h1` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/h1\>/));
```

你有不止一个 `h1` 元素。 删除多余的 `h1` 元素。

```js
assert(document.querySelectorAll('h1').length === 1);
```

你的 `h1` 元素的文本应为 `CatPhotoApp`。 你可能忽略了此文本，或者有拼写错误，或没有把它放在 `h1` 元素的开始和结束标签之间。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

你似乎使用了一个浏览器扩展，正在修改页面。 一定要关闭所有浏览器扩展。

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

