---
id: 5dc17d3bf86c76b9248c6eb4
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 元素用于在网站上创建一段文本。 在你的 `h2` 元素下方创建一个 `p` 元素，并为其提供以下文本：

`See more cat photos in our gallery.`

# --hints--

你的 `p` 元素应该有一个开始标签。 起始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('p'));
```

你的 `p` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/p\>/));
```

你的 `p` 元素的文本应该是 `See more cat photos in our gallery.` 你可能忽略了文本或有拼写错误。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

你的 `p` 元素应该在 `h2` 元素的下方。 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

