---
id: 5dc174fcf86c76b9248c6eb2
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML 元素中拥有开始标签，如 `<h1>`，和结束标签，如 `</h1>`。

元素的文本在元素的开始标签和结束标签之间。

找到 `h1` 元素，并将其文本更改为：

`CatPhotoApp`

# --hints--

代码中的文本应该是 `CatPhotoApp`。 你可能想要检查你的拼写。

```js
assert(code.match(/catphotoapp/i));
```

你的 `h1` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('h1'));
```

你的 `h1` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/h1\>/));
```

你应该只有一个 `h1` 元素。 删除多余的 `h1` 元素。

```js
assert(document.querySelectorAll('h1').length === 1);
```

你的 `h1` 元素中的文本应该是 `CatPhotoApp`。 你可能忽略了此文本，或者有拼写错误，或没有把它放在 `h1` 元素的开始和结束标签之间。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

你似乎使用了一个浏览器扩展，正在修改页面。 一定要关闭所有浏览器扩展。

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

