---
id: 5dc24073f86c76b9248c6ebb
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>属性</dfn>是写在一个元素的开始标签中的特殊词汇，它们用于控制这个元素的行为。 在 `img` 元素中的 `src` 属性明确了一个图片的 URL（图片所在的位置）。 使用 `src` 属性的 `img` 元素示例：`<img src="https://www.example.com/the-image.jpg">`。

在现有的 `img` 元素中，添加具有以下 URL 的 `src` 属性：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

你的代码应该有一个 `img` 元素。 你可能已经删除了 `img` 元素，或者你没有用引号将 `src` 属性的值括起来。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素应该有一个 `src` 属性。 你要么省略了该属性，要么有错字。 请确保元素名称和属性名称之间有一个空格。

```js
assert(document.querySelector('img').src);
```

你的 `img` 元素的 `src` 属性应设置为 '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'。 你可能省略了这个 URL 或者有拼写错误。 URL 的大小写很重要。

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

尽管你已将 `img` 元素的 `src` 设置为正确的 URL，但建议始终用引号将属性的值括起来。

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

