---
id: 5dc24614f86c76b9248c6ebd
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以使用锚点元素（`a`）链接到另一个页面。 例如 `<a href='https://freecodecamp.org'></a>` 会链接到 `freecodecamp.org`。

在段落下方添加一个锚点元素并链接到 `https://freecatphotoapp.com`。 此时，链接不会出现在预览中。

# --hints--

你的锚点元素（`a`）应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的锚点元素（`a`）应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的锚点元素（`a`）应该在 `p` 元素的下方。 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

你的锚点元素（`a`）缺少 `href` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert(document.querySelector('a').hasAttribute('href'));
```

你的锚点元素（`a`）应该链接到 `https://freecatphotoapp.com`。 你可能省略了这个 URL 或者有拼写错误。

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

虽然你将锚点元素（a）的 `href` 属性指向了正确的链接，但建议始终以引号来包裹属性的值。

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

