---
id: 5dc23991f86c76b9248c6eb8
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在上一步中，你将 `h2` 元素、注释和 `p` 元素嵌套在 `main` 元素中。 嵌套元素比其父元素多缩进两个空格，可提高可读性：

```html
<ul> 
  <li> This `li` element indented </li>
  <li> This `li` element is also indented </li>
</ul>
```

在 `h2`、comment 和 `p` 元素前面再添加两个空格，这样你的 HTML 就更具可读性。

# --hints--

你的代码应该有一个 `h2` 元素，其中包含 `Cat Photos` 文本。 你可能意外地删除了它，或者它缺少了开始和结束标签，或者是文本的内容被改变了。

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

你不应添加示例中的 `ul` 或 `li` 元素。

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

你的 `h2` 元素应位于 `main` 元素的开始标签下方，并且其开始标签应在距离行首 6 个空格处的位置。

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s{6}<h2>/));
```

你的代码应该有注释。 你从之前的步骤中删除了注释。

```js
assert(code.match(/<!--.*-->/));
```

评论的文本应该是 `TODO: Add link to cat photos`。 不要更改注释的文本或间距。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

你的注释应位于 `h2` 元素下方，并从距离行首 6 个空格处开始写。

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

你的代码应该有一个 `p` 元素。 你删除了前一步骤中的 `p` 元素。

```js
assert(document.querySelector('p'));
```

`p` 元素的文本应该是 `Click here to view more cat photos.`。 不要更改 `p` 元素的文本、间距或标点符号。

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

你的 `p` 元素应该在注释的下方，并且它的起始标签距离行首应有 6 个空格。

```js
assert(code.toLowerCase().match(/-->\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

