---
id: 5dc23f9bf86c76b9248c6eba
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

你可以通过使用 `img` 元素来为你的网站添加图片。 `img` 元素只有一个开始标签，没有结束标签。 一个没有结束标签的元素，它的标签被称为<dfn>自闭合标签</dfn>。

在 `p` 元素的下方添加一个 `img` 元素。 此时，没有图片会出现在浏览器中。

# --hints--

你的 `img` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('img'));
```

你的 `img` 应该没有结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(!code.match(/<\/img\>/));
```

应该只有一个 `img` 元素。 删除多余的。

```js
assert(document.querySelectorAll('img').length === 1);
```

你的 `img` 元素应该在 `p` 元素的下方。 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

