---
id: 5dc24165f86c76b9248c6ebc
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

所有的 `img` 元素都应该有一个 `alt` 属性。 `alt` 属性的文本（值）有两个作用，第一个作用是让屏幕阅读器可以知晓图片的内容，这会对网页的可访问性有很大提升；另一个作用是当图片无法加载时，页面需要显示的替代文本。 例如 `<img src="cat.jpg" alt="A cat">` 有一个 `alt` 属性，其文本为 `A cat`。

在 `img` 元素中，添加一个带有以下文本的 `alt` 属性：

`A cute orange cat lying on its back`

# --hints--

你的代码应该有一个 `img` 元素。 你删除了前一步骤中的 `img` 元素。

```js
assert(document.querySelector('img'));
```

你的 `img` 元素应该有一个 `alt` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

你的 `img` 元素的 `alt` 属性值被写成别的了，不再是 “A cute orange cat lying on its back” 了。 确保将 `alt` 属性的值放在引号中。

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

