---
id: 5dc2385ff86c76b9248c6eb7
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 有些元素能够指明不同的内容区域。 这些元素能让你的 HTML 易于阅读，并有助于搜索引擎优化（SEO）和无障碍访问。

指明一个主要区域，可通过在 `h1` 元素后面添加一个 `<main>` 开始标签，并在 `p` 元素后面添加一个 `</main>` 结束标签来实现。

# --hints--

你的 `main` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('main'));
```

你的 `main` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/main\>/));
```

你的 `main` 元素的开始标签应该在 `h1` 元素的下方。 你把顺序写错了。

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling?.tagName, "H1");
```

你的 `main` 元素的开始标签应该在 `h2` 元素上方。 你把顺序写错了。

```js
const collection = [...document.querySelectorAll('main,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H2'));
```

你的 `main` 元素的开始标签应该在 `p` 元素的下方。 你把顺序写错了。

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/click here to view more cat photos/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
--fcc-editable-region--
  </body>
</html>
```

