---
id: 5ddbd81294d8ddc1510a8e56
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

链接的文本必须放置在锚点元素（`a`）的起始和闭合标签之间。 例如，`<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` 就是一个包含着文本 `click here to go to freeCodeCamp.org` 的链接。

将锚点文本 `link to cat pictures` 添加到锚点元素中。 这将成为链接的文本。

# --hints--

你的锚点元素（`a`）应该有一个起始标签。 起始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('a'));
```

你的锚点元素（`a`）应该有一个闭合标签。 闭合标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/a\>/));
```

你的锚点元素（`a`）中的文本应该是 `link to cat pictures`。 确保将链接的文本放置在锚点元素（`a`）的起始和闭合标签之间。

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

