---
id: 5dfa2407b521be39a3de7be1
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

要在新标签页中打开链接，你可以在锚元素（`a`）中使用 `target` 属性。

`target` 属性指定了在哪里打开链接文档。 `target="_blank"` 意思是在新标签页或窗口中打开链接的文档。

这是一个带有 `target` 属性的 `a` 元素的基本语法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

向锚元素（`a`）的开始标签中添加值为 `_blank` 的 `target` 属性，用以在新的标签页中打开链接。

# --hints--

你的 `p` 元素应该嵌套一个文本为 `cat photos` 的锚元素（`a`）。 你可能删除了它或有拼写错误。

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

你的锚元素（`a`）缺少一个 `target` 属性。 请检查在起始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert(document.querySelector('a').hasAttribute('target'));
```

`target` 属性的值应为 `_blank`。 你可能省略了这个值或者有拼写错误。 别忘了应该用引号包裹属性的值。

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

