---
id: 5dfa22d1b521be39a3de7be0
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

你现在可以在预览中看到图像旁边带下划线的字样 `cat photos`。 那是你的链接，点击它一下。 在你的 `p` 元素的文本中，将单词 `cat photos` 转换为指向 `https://freecatphotoapp.com` 的链接。 完成后，删除段落下方旧的锚标签和文本。

# --hints--

你的代码应该只包含一个锚点元素（`a`）。 删除多余的锚点元素。

```js
assert(document.querySelectorAll('a').length === 1);
```

你的锚点元素（`a`）应该嵌套在 `p` 元素之内。

```js
assert($('p > a').length);
```

链接的 href 值应为 `https://freecatphotoapp.com`。 你可能忽略了 href 值或者有拼写错误。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

链接的文本应该是 `cat photos`。 你可能忽略了文本或有拼写错误。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

在嵌套了锚点元素（`a`）之后，浏览器中显示的 `p` 元素应该仅为 `Click here to view more cat photos.`。再次检查 `p` 和嵌套其中的锚点元素的文本、间距或标点。

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

