---
id: 5dfa22d1b521be39a3de7be0
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

你可以将任何文本转化为一个链接，例如 `p` 元素内部的文本。

``` html
<p>I think <a href="https://www.freecodecamp.org">freeCodeCamp</a> is great.</p>
```

在 `p` 元素的文本中，在 `cat photos` 周围添加开始和结束锚标签（`a`），将这些字样转化为链接。 然后将 `href` 属性设置为 `https://freecatphotoapp.com`。


# --hints--

你应该在 `p` 元素中嵌套一个新的锚元素（`a`）。

```js
const nestedAnchor = document.querySelector(`p > a`);
assert.isNotNull(nestedAnchor)
```

链接的 `href` 值应为 `https://freecatphotoapp.com`。 你可能忽略了 `href` 值或者有拼写有误。

```js
assert(
  document.querySelector('p > a').href === 'https://freecatphotoapp.com/'
);
```

链接的文本应该是 `cat photos`。 你可能忽略了文本或有拼写错误。

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isTrue(innerContent.trim() === 'cat photos');
```

锚元素中的文本有额外的前导或尾随空格。 锚文本中唯一的空格应位于 `cat` 和 `photos` 之间。

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

在嵌套了锚元素（`a`）之后，浏览器中显示的 `p` 元素应该仅为 `See more cat photos in our gallery.`。再次检查 `p` 和嵌套在其中的锚元素的文本、间距或标点。

```js
assert.match(code, /<p>see more <a[^>]*>cat photos<\/a> in our gallery\.?<\/p>/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

