---
id: 5dfb6250eacea3f48c6300b2
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

在无序列表之后，添加一个新图像，它具有 `src` 属性，值为：

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

而它的 `alt` 属性值为：

`A slice of lasagna on a plate.`

# --hints--

在 `</ul>` 结束标签之后应该有一个 `img` 元素。

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

新图像没有 `alt` 属性。 检查开始标签名称后是否有空格和/或所有属性名称前是否有空格。

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

新图像的 `alt` 值应为 `A slice of lasagna on a plate.`。确保将 `alt` 属性的值放在引号中。

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

新图像没有 `src` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

新图像的 `src` 值应为 `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`。 确保将 `src` 属性的值放在引号中。

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

尽管你已将新图像的 `src` 设置为正确的 URL，但建议始终将属性值放在引号中。

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

