---
id: 5dfb6a35eacea3f48c6300b4
title: 步骤 24
challengeType: 0
dashedName: step-24
---

# --description--

图像标题（`figcaption`）元素用于添加标题以描述 `figure` 元素中包含的图像。 例如，`<figcaption>A cute cat</figcaption>` 添加标题 `A cute cat`。

在嵌套在 `figure` 元素中的图像之后，添加一个 `figcaption` 元素，其文本为：

`Cats love lasagna.`

# --hints--

你的 `figcaption` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('figcaption'));
```

你的 `figcaption` 元素应该有一个结束标签。 结束标签在 `<` 字符后面要有一个 `/`。

```js
assert(code.match(/<\/figcaption\>/));
```

`figcaption` 元素应该嵌套在 `figure` 元素中。

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

Lasagna `img` 元素应该嵌套在 `figure` 元素中。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

嵌套在 `figure` 元素中的 `figcaption` 元素应位于 `img` 元素下方。 你的顺序是错误的。

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

你的 `figcaption` 元素的文本应该是 `Cats love lasagna.`。 你可能忽略了文本或有拼写错误。

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

