---
id: 5ef9b03c81a63668521804e2
title: 步骤 54
challengeType: 0
dashedName: step-54
---

# --description--

对于可能有多个答案的问题，表单通常使用复选框。 例如，这是一个带有 `tacos` 选项的复选框：`<input type="checkbox"> tacos`。

在你刚刚添加的 `legend` 元素下，添加一个 `input`，将其 `type` 属性设置为 `checkbox`，并为其指定选项：

`Loving`

# --hints--

复选框的 `input` 元素应该有一个开始标签，而不是一个结束标签。

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

你应该只为你的复选框添加一个 input 元素。 删除多余的。

```js
assert($('fieldset > input').length === 1);
```

你的新 `input` 元素应位于 `legend` 元素下方，并带有文本 `What's your cat's personality?`。 你把顺序写错了。

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

你的新 `input` 元素没有 `type` 属性。 检查确认开始标签名称后应该有一个空格。

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

你的新 `input` 元素应该有一个 `type` 属性，其值为 `checkbox`。 你可能忽略了这个值或者有拼写错误。 请记住，属性值应该放在引号中。

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

尽管你已将新的 `input` 元素的 `type` 属性设置为 `checkbox`，但建议始终将属性的值放在引号中。

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

文本 `Loving` 应直接位于复选框的右侧。 请确保元素和文本之间有一个空格。 你可能忽略了文本或有拼写错误。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

