---
id: 5ef9b03c81a63668521804d0
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

通过将 `figcaption` 元素中的单词 `love` 包裹在强调 `em` 元素中来强调它。

# --hints--

你的强调 `em` 元素应该有一个开始标签。 开始标签的语法为：`<elementName>`。

```js
assert(document.querySelector('em'));
```

你的强调 `em` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert(code.match(/<\/em\>/));
```

你要么删除了 `figcaption` 元素，要么它缺少开始或结束标签。

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

你的强调 `em` 元素应包裹文本 `love`。 你可能忽略了文本或有拼写错误。

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

`figcaption` 的文本应该是 `Cats love lasagna`。 检查是否有拼写错误，以及 `em` 元素的开始和结束标签周围是否存在必要的空格。

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

