---
id: 5ef9b03c81a63668521804d6
title: 步骤 35
challengeType: 0
dashedName: step-35
---

# --description--

现在你将添加一个 Web 表单来收集来自用户的信息。

在 `Cat Form` 标题之后，添加一个 `form` 元素。

# --hints--

你的 `form` 元素应该有一个开始标签和一个结束标签。 你可能缺少一个或两个必需的标签，或者它们的顺序错误。

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

你的 `form` 元素标签的顺序不正确。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

嵌套在最后一个 `section` 元素中的 `form` 元素应位于 `h2` 元素之下。 你的 `h2` 元素和 `form` 元素的顺序错误。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

`form` 元素应该没有内容。 删除 `form` 元素标签之间的任何 HTML 元素或文本。

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

