---
id: 5ef9b03c81a63668521804d8
title: 步骤 37
challengeType: 0
dashedName: step-37
---

# --description--

`input` 元素允许你通过多种方式从 Web 表单中收集数据。 与 `img` 元素一样，`input` 元素是<dfn>自闭合的</dfn>，并且不需要结束标签。

在 `form` 元素中嵌套一个 `input` 元素。

# --hints--

你的 `form` 元素应具有正确顺序的开始标签和结束标签。 你可能缺少一个或两个必需的标签，或者它们的顺序错误。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

你的 `form` 元素的开始标签应该只有一个 `action` 属性。 删除你可能输入的任何其他内容。

```js
assert([...document.querySelector('form').attributes].length < 2);
```

你应该创建一个 `input` 元素。 请检查语法。

```js
assert(document.querySelector('input'));
```

你的 `input` 元素应该有一个开始标签，而不是一个结束标签。

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

你的 `input` 元素应该嵌套在 `form` 元素中。

```js
assert(document.querySelector('form > input'));
```

你的 `form` 应该只包含 `input` 元素。 删除 `form` 元素标签之间的任何 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">

        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

