---
id: 5ef9b03c81a63668521804d9
title: 步骤 40
challengeType: 0
dashedName: step-40
---

# --description--

占位符文本用于提示人们在输入框中输入什么样的信息。 例如，`<input type="text" placeholder="Email address">`。

将占位符文本 `cat photo URL` 添加到你的 `input` 元素中。

# --hints--

你要么删除了 `input` 元素，要么它的语法无效。 所有属性的值都应该放在引号中。

```js
assert($('input').length);
```

你的 `form` 应该只包含 `input` 元素。 删除 `form` 元素中的任何其他 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

你的 `input` 元素没有 `placeholder` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert($('input')[0].hasAttribute('placeholder'));
```

你的 `input` 元素应该有一个 `placeholder` 属性，其值为 `cat photo URL`。  你可能省略了这个值或者有拼写错误。 请记住，属性值应该放在引号中。

```js
assert(
  $('input')[0]
    .getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

尽管你已将 `input` 元素的 `placeholder` 属性设置为 `cat photo URL`，但建议始终将属性值放在引号中。

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

