---
id: 5ef9b03c81a63668521804dc
title: 步骤 44
challengeType: 0
dashedName: step-44
---

# --description--

对于需要从多个选项中获得一个答案的问题，你可以使用单选按钮。

下面是一个带有 `cat` 选项的单选按钮示例：`<input type="radio"> cat`。 请记住，`input` 元素是<dfn>自闭合的</dfn>。

在文本输入之前，添加一个具有如下选项的单选按钮：

`Indoor`

# --hints--

你应该为单选按钮创建一个 `input` 元素。 请检查语法。

```js
assert($('form > input').length >= 2);
```

你的 `input` 元素应该有一个开始标签，而不是一个结束标签。

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

你应该只为单选按钮添加一个 `input` 元素。 删除多余的。

```js
assert($('form > input').length === 2);
```

你的新 `input` 元素应位于现有 `input` 之上，其中 `type` 属性设置为 `text`。 你把顺序写错了。

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

你的新 `input` 元素没有 `type` 属性。 检查确认开始标签名称后应该有一个空格。

```js
assert($('input')[0].hasAttribute('type'));
```

Your new `input` element should have only one `type` attribute. Remove any extras.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Your new `input` element should have a `type` attribute with the value `radio`. You have either omitted the value or have a typo. Remember that attribute values should be surrounded with quotation marks.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Although you have set the new `input` element's `type` attribute to `radio`, it is recommended to always surround the value of an attribute with quotation marks.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

The `radio` button's `Indoor` text should be located after it instead of before it.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

The text `Indoor` should be located directly to the right of your `radio` button. You have either omitted the text or have a typo.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

