---
id: 5ef9b03c81a63668521804e9
title: 步骤 64
challengeType: 0
dashedName: step-64
---

# --description--

通过将文本 `freeCodeCamp.org` 包含在锚点（`a`）元素中，使其成为链接。 应把 `href` 属性设置为 `https://www.freecodecamp.org`。

# --hints--

你的锚（`a`）元素应该嵌套在 `footer` 元素中。 确保为锚（`a`）元素添加开始标签和结束标签。

```js
assert($('footer > p > a').length);
```

你的锚（`a`）元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert(aElemClosingTags && aElemClosingTags.length === 3);
```

你的锚点（`a`）元素应该有一个 `href` 属性，其值为 `https://www.freecodecamp.org`。 你可能忽略了属性/值，或者有拼写错误。

```js
const nestedAnchor = $('footer > p > a')[0];
assert(nestedAnchor.getAttribute('href').toLowerCase() === 'https://www.freecodecamp.org');
```

链接的文本应该是 `freeCodeCamp.org`。 你可能忽略了文本或有拼写错误。

```js
const nestedAnchor = $('footer > p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'freecodecamp.org'
);
```

嵌套锚（`a`）元素后，浏览器中唯一可见的 `p` 元素内容应该是 `No Copyright - freeCodeCamp.org`。 仔细检查 `p` 和嵌套锚元素的文本、间距或标点符号。

```js
const pText = $('footer > p')[0].innerText.toLowerCase().replace(/\s+/g, ' ');
assert(pText.match(/^no copyright - freecodecamp.org$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```

