---
id: 5ef9b03c81a63668521804d7
title: 步骤 36
challengeType: 0
dashedName: step-36
---

# --description--

`action` 属性指示应该将表单数据发送到哪里。 例如，`<form action="/submit-url"></form>` 告诉浏览器应该将表单数据发送到路径`/submit-url`。

将值为 `https://freecatphotoapp.com/submit-cat-photo` 的 `action` 属性添加到 `form` 元素。

# --hints--

你的 `form` 元素应具有正确顺序的开始标签和结束标签。 你可能缺少一个或两个必需的标签，或者它们的顺序错误。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

嵌套在最后一个 `section` 元素中的 `form` 元素应位于 `h2` 元素下方。 你的 `h2` 元素和 `form` 元素的顺序错误。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

你的 `form` 元素应该没有内容。 删除 `form` 元素标签之间的任何 HTML 元素或文本。

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

你的 `form` 元素没有 `action` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

你的 `form` 元素应该有一个 `action` 属性，其值为 `https://freecatphotoapp.com/submit-cat-photo`。

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

尽管你已将 `action` 属性设置为正确的 URL，但建议始终将属性的值放在引号中。

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

