---
id: 5efb2c990dc218d6c85f89b2
title: 第 43 步
challengeType: 0
dashedName: step-43
---

# --description--

即使你在文本输入下方添加了按钮，它们也会在页面上彼此相邻。 这是因为 `input` 和 `button` 元素都是<dfn>内联元素</dfn>，它们不会出现在新的行上。

你添加的按钮将默认提交表单。 然而，依赖默认行为可能造成混乱。 将值为 `submit` 的 `type` 属性添加到 `button` 以明确它是一个提交按钮。

# --hints--

你的 `button` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('button'));
```

你的 `button` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert(code.match(/<\/button\>/));
```

你的 `button` 元素没有 `type` 属性。 检查确认开始标签名称后应该有一个空格。

```js
assert($('button')[0].hasAttribute('type'));
```

你的 `button` 元素应该有一个 `type` 属性，其值为 `submit`。 你可能忽略了这个值或者有拼写错误。 请记住，属性值应该放在引号中。

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

尽管你已将 `button` 元素的 `type` 属性设置为 `submit`，但建议始终将属性的值放在引号中。

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

