---
id: 5efb23e70dc218d6c85f89b1
title: 步骤 38
challengeType: 0
dashedName: step-38
---

# --description--

你可以使用 `type` 属性创建多种输入。 你可以轻松创建密码字段、重置按钮或控件，让用户从他们的计算机中选择文件。

通过将值为 `text` 的 `type` 属性添加到 `input` 元素，创建一个文本字段以从用户那里获取文本输入。

# --hints--

你要么删除了 `input` 元素，要么它的语法无效。 如果你添加了属性，请确保将它们的值放在引号中。

```js
assert($('input').length);
```

你的 `form` 元素应该只包含 `input` 元素。 删除 `form` 元素标签之间的任何额外 HTML 元素或文本。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

你的 `input` 元素没有值为 `text` 的 `type` 属性。 请检查在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
assert($('input')[0].hasAttribute('type'));
```

你的 `input` 元素应该有一个 `type` 属性，其值为 `text`。 你可能忽略了这个值或者有拼写错误。 请记住，属性值应该放在引号中。

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

尽管你已将 `input` 元素的 `type` 属性设置为 `text`，但建议始终将属性的值放在引号中。

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

