---
id: 5efc4f528d6a74d05e68af74
title: 步骤 56
challengeType: 0
dashedName: step-56
---

# --description--

还有另一种方法可以将 `input` 元素的文本与元素本身相关联。 你可以将文本嵌套在 `label` 元素中，并添加与 `input` 元素的 `id` 具有相同值的 `for` 属性。

通过仅给 `label` 元素添加文本 `Loving`，并给它添加适当的 `for` 属性值，将文本 `Loving` 与复选框相关联。

# --hints--

你应该确保该复选框仍然存在。

```js
assert($('input[type="checkbox"]')[0]);
```

你的复选框应该仍然有一个 `id` 属性，其值为 `loving`。  应该是 `--fcc-expected--`，而不是 `--fcc-actual--`。

```js
assert.equal($('input[type="checkbox"]')[0].id, 'loving');
```

文本 `Loving` 不应再直接位于复选框的右侧。 它应该被包裹在 `label` 元素中。 确保两个元素之间有空格。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  !checkboxInputElem.nextSibling.nodeValue
    .replace(/\s+/g, ' ')
    .match(/ Loving/i)
);
```

你将需要添加一个新的 `label` 元素，在其中嵌套文本 `Loving`。 确保它有一个开始标签和一个结束标签。

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

新的 `label` 元素应直接位于复选框的右侧。 确保两个元素之间有空格。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

新的 `label` 元素没有 `for` 属性。 检查确认开始标签名称后应该有一个空格。

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

新的 `label` 元素应该有一个 `for` 属性，其值为 `loving`。  应该是 `--fcc-expected--`，而不是 `--fcc-actual--`。

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert.equal(labelElem.getAttribute('for'), 'loving');
```

文本 `Loving` 应该嵌套在新的 `label` 元素中。 你可能忽略了文本或有拼写错误。

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

