---
id: 5f05a1d8e233dff4a68508d8
title: 步骤 47
challengeType: 0
dashedName: step-47
---

# --description--

在第一个单选按钮下方创建另一个单选按钮。 将其放在 `label` 元素内部，`label` 内的文字为 `Outdoor`。 单选按钮的 `id` 属性为 `outdoor`。

# --hints--

应该只有原来的 `Indoor` 单选按钮 `input` 的 `id` 为 `indoor`。 如果需要，可以重新启动步骤来恢复原始的 HTML。

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

不应该修改 `Indoor` 单选按钮。 如果需要，可以重新启动步骤来恢复原始的 HTML。

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

应该在新创建的 `label` 元素内添加一个新的 `input` 元素。 请确保新的 `label` 同时有一个开始和结束标签。

```js
assert(document.querySelectorAll('label input').length === 2);
```

`label` 开始标签不应该有任何属性。

```js
assert(code.match(/<label\s*>/g).length === 2);
```

你的 `input` 应该有一个 `id` 属性，其值为 `outdoor`。

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

你的新 `input` 元素应该有一个 `type` 属性，其值为 `radio`。

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

除了 `input` 嵌套在 `label` 之外，不应该添加任何新的元素。

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

新的 `input` 的左边不应该有文本。

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

新的单选按钮的标签文本必须是 `Outdoor`。 首字母要大写。

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

新的单选按钮和标签应该在 `Indoor` 单选按钮和标签下面/后面。 两者之间不应有其他标签。

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

