---
id: 5f05a1d8e233dff4a68508d8
title: 步骤 46
challengeType: 0
dashedName: step-46
---

# --description--

在新的 `label` 元素中嵌套另一个带有选项 `Outdoor` 的单选按钮。 新的单选按钮应放在第一个单选按钮之后。 此外，将其 `id` 属性值设置为：

`outdoor`

# --hints--

你将需要添加一个新的 `label` 元素来嵌套新的单选按钮。 确保它有一个开始标签和一个结束标签。

```js
assert(
  document.querySelectorAll('label').length === 2 &&
    code.match(/<\/label\>/g).length === 2
);
```

`Outdoor` 文本应直接位于新 `radio` 按钮的右侧。 请确保元素和文本之间有一个空格。 你可能忽略了文本或有拼写错误。

```js
const radioButtons = [...$('input')];
assert(
  radioButtons.filter((btn) =>
    btn.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ *Outdoor/i)
  ).length
);
```

你的新单选按钮和相关标签应位于第一个的下方。 你把顺序写错了。

```js
const collection = [
  ...document.querySelectorAll('input[type="radio"]')
].map((node) => node.nextSibling.nodeValue.replace(/\s+/g, ''));
assert(collection.indexOf('Indoor') < collection.indexOf('Outdoor'));
```

你的新单选按钮应该有一个 `id` 属性。 请检查在开始标签的名称后面要有一个空格，且/或所有的属性名称前面也要有一个空格。

```js
assert($('input')[1].hasAttribute('id'));
```

你的新单选按钮应该有一个 `id` 属性，其值为 `outdoor`。 你可能省略了这个值或者有拼写错误。 请记住，属性值应该放在引号中。

```js
assert($('input')[1].id.match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

