---
id: 5f07be6ef7412fbad0c5626b
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

在添加任何新内容之前，你应该使用 `section` 元素将猫咪照片内容与未来的内容分开。

`section` 元素用于在文档中定义各部分，如章节、标题、页脚或文档的任何其他部分。 它是一个对 SEO 和无障碍特性有帮助的语义元素。

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

将你的 `h2` 元素、注释、`p` 元素和锚元素（`a`）嵌套在一个 `section` 元素中。

# --hints--

你的 `section` 元素应该有一个开始标签。 开始标签的书写语法为：`<elementName>`。

```js
assert(document.querySelector('section'));
```

你的 `section` 元素应该有一个结束标签。 结束标签在 `<` 字符之后有一个 `/`。

```js
assert(code.match(/<\/section\s*>/i));
```

整个 `section` 元素应该在 `main` 元素的开始标签和结束标签之间。

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

现有的 `h2` 元素、注释、`p` 元素和锚（`a`）元素应位于 `section` 元素的开始标签和结束标签之间。

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

`h1` 元素不应该被嵌套在 `section` 元素中。

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

