---
id: 5f1a80975fc4bcae0edb3497
title: 步骤 49
challengeType: 0
dashedName: step-49
---

# --description--

如果你选择 `Indoor` 单选按钮并提交表单，则该按钮的表单数据基于其 `name` 和 `value` 属性。 由于你的单选按钮没有 `value` 属性，因此表单数据将包含 `indoor-outdoor=on`，这在你有多个按钮时没有用处。

为两个单选按钮添加 `value` 属性。 为方便起见，将按钮的 `value` 属性设置为与其 `id` 属性相同的值。

# --hints--

两个单选按钮仍应位于 `label` 元素开始和结束标签之间。

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

两个单选按钮都应该有一个 `value` 属性。 请检查在开始标签的名称后面要有一个空格，且/或所有的属性名称前面也要有一个空格。

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('value')));
```

两个单选按钮都应该有 `type` 属性。 请确保在开始标签的名称后面要有一个空格，且所有的属性名称前面也要有一个空格。

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

`Indoor` 单选按钮的 `value` 属性应设置为 `indoor`。 你可能忽略了这个值或者有拼写错误。 请记住，属性值应该放在引号中。

```js
const indoorRadioButton = document.querySelector('#indoor');
assert(indoorRadioButton.getAttribute('value').match(/^indoor$/));
```

`Outdoor` 单选按钮的 `value` 属性应设置为 `outdoor`。 你可能落下了这个值或者有拼写错误。 所有属性的值都应该放在引号中。

```js
const outdoorRadioButton = document.querySelector('#outdoor');
assert(outdoorRadioButton.getAttribute('value').match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

