---
id: 62dabe2ef403a12d5d295273
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

现在将 `p` 元素中的文本 `cat photos` 变成了链接， 已经不需要 `p` 元素下面的第二个链接了。 删除 `p` 元素下面的锚元素。

# --hints--

你的代码应该只包含一个段落元素（`p`）。

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

你的代码应该只包含一个锚元素（`a`）。

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

不要对你在上一步创建的元素做任何更改。

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

`p` 元素下面不应该有 `cat photos` 文本。

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
