---
id: 655a1a6d8c44db154b00c909
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

Escreva outra declaração `if` que verifique se a chave de destino é menor que a chave do nó atual.

Dentro do bloco `if`, retorne o resultado da chamada do método `_search` com o filho à esquerda do nó atual e `key` como argumentos.

# --hints--

Você deve escrever outra instrução `if` para verificar se `key` < `node.key`.

```js
const after_split = code.split('def search(self, key):')[1];
assert.match(after_split, /if\s+key\s+<\s+node\.key/);
```

Você deve retornar `self._search(node.left, key)` do novo bloco `if`.

```js
({ test: () => assert.match(code, /return\s+self\._search\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node                
--fcc-editable-region--
```
