---
id: 60f027c87bc98f050395c139
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

接下来，在 `html` 元素中添加 `head` 和 `body` 的开始和结束标签。

# --hints--

应该有一个 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是同级元素。

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
