---
id: 60f027099a15b00485563dd2
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `DOCTYPE`下， 添加一个 `html` 元素，元素内有一个设置为 `en` 的 `lang` 属性，这样就可以编码了。

# --hints--

`DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 元素应该有一个开始标签，`lang` 属性的值应该为 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

`html` 标签顺序应该正确。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

应该只有一个 `html` 元素。

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
