---
id: 60f02e7361b68405e27b62a5
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `body` 中， 添加一个 `h1` 文本 `Registration Form`，来给内容添加一个标题。

# --hints--

应该在 `body` 内添加一个 `h1`。

```js
assert.exists(document.querySelector('body > h1'));
```

应该给 `h1` 元素添加文本 `Registration Form`。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
