---
id: 60f0286404aefb0562a4fdf9
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `head` 元素中添加一个 `title` 和 `meta` 元素。 设置项目的标题为 `Registration Form`，给 `meta` 元素添加值为 `utf-8` 的 `charset` 属性。

# --hints--

代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 元素应该在 `head` 元素内。

```js
assert.exists(document.querySelector('head > title'));
```

你的项目应该有标题 `Registration Form`。 记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

你应该在 `head` 元素中创建一个 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你应该给 `meta` 元素一个值为 `utf-8` 的 `charset` 属性。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

你的 `meta` 元素应该是一个空元素，它不应该有结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
