---
id: 60f027099a15b00485563dd2
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `DOCTYPE` 下方添加开始和结束 `html` 标签，以便添加其它代码。

# --hints--

`DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 元素应该有一个开始标签。

```js
assert(code.match(/<html\s*>/gi));
```

`html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

`html` 标签顺序应该正确。

```js
assert(code.match(/<html\s*>\s*<\/html\s*>/));
```

应该只有一个 `html` 元素。

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
