---
id: 60f0286404aefb0562a4fdf9
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

将 `title` 元素添加到 `head` 内，并将项目标题设置为 `Registration Form`。 此外，在 `head` 元素中嵌套一个自闭合 `link` 元素。 设置其 `rel` 属性值为 `stylesheet`， `href` 属性值为 `styles.css`。

# --hints--

代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 元素应该在 `head` 元素内。

```js
assert.exists(document.querySelector('head > title'));
```

项目的标题应该为 `Registration Form`。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

代码应该有一个 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

不应更改现有的 `head` 标签。 确保没有删除结束标签。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

应该有一个自闭合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 元素应该在 `head` 元素内。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  --fcc-editable-region--
  <head>

  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
