---
id: 60f0286404aefb0562a4fdf9
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

给 `head` 添加一个 `title` 和一个 `meta` 元素。 设置项目的标题为 `Registration Form`，添加一个 `meta` 元素设置其`charset` 属性的值为 `UTF-8`。

# --hints--

代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 元素应该在 `head` 元素内。

```js
assert.exists(document.querySelector('head > title'));
```

项目的标题应该为 `Registration Form`。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
