---
id: 60f1922fcbd2410527b3bd89
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--
`vh` 单位代表视口高度（viewport height），相当于视口 `height` 高度的 1%。

现在是时候用 CSS 来装饰项目了。 首先设置 `body` 的 `width` 为 `100%`，`height` 为 `100vh`。

# --hints--

你应该使用 `body` 元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你应该把 `body` 元素的 `width` 属性值设置为 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你应该把 `body` 的 `height` 设置为 `100vh`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
