---
id: 60f1a9cbd23023082e149fee
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

看起来好多了。 现在，把 `body` 的 `background-color` 设置为 `#1b1b32`，使背景更护眼。 然后相应的把 `color` 设置为 `#f5f6f7`，让文字显示出来。

# --hints--

应该在 `body` 元素选择器中添加 `background-color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

应该设置 `background-color` 的值为 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

应该在 `body` 元素选择器中添加 `color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

`color` 属性应该设置为 `#f5f6f7`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
