---
id: 60f1a5e2d2c23707a4f9a660
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

现在，通过设置 `body` 的默认 `margin` 为 `0` 来重置一些浏览器的默认设置，从而去掉水平滚动条。

# --hints--

应该在 `body` 元素选择器中添加 `margin`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

应该将 `margin` 属性设置为 `0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
