---
id: 60f5d2776c854e069560fbe6
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

第一个 `fieldset` 将包含名字（name）、邮箱（email）和密码（password）字段。 首先，在第一个 `fieldset` 内部添加四个 `label` 元素。

# --hints--

总共应该有三个 `fieldset` 元素。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

所有的 `fieldset` 元素应为 `form` 元素的直接子元素。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

四个 `label` 元素都应该是第一个 `fieldset` 元素的直接子元素。 请确保 `label` 有结束标签。

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

最后两个 `fieldset` 元素应该为空。

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

请确保您的label标签闭合。

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
