---
id: 60f803d5241e6a0433a523a1
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--
`rem` 单位代表根 `em`，与 `html` 元素的字体大小有关。

由于 `label` 元素默认是行内元素，它们出现在其标签文本的同一行，使得文本难以阅读。 给 `label` 元素添加 `display: block`，并设置其 `margin` 为 `0.5rem 0`，来使其显示在不同的行，并且行之间有一定的距离。

# --hints--

你应该使用 `label` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

你应该将 `display` 属性设置为 `block`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

你应该添加 `margin` 属性，值为 `0.5rem 0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
