---
id: 60f83e7bfc09900959f41e20
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

一些 `type` 属性值含有内置表单验证。 例如， `type="email"` 要求该值是一个有效的电子邮件地址。

给密码 `input` 元素添加自定义验证，新增 `minlength` 属性设置其值为 `8`。 这样密码少于 8 个字符的时候会阻止提交。

# --hints--

密码 `input` 元素应该添加 `minlength` 属性。

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

应该将 `minlength` 属性设置为 `8`。

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
