---
id: 60f80e0081e0f2052ae5b505
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

指定表单元素的 `type` 属性对浏览器预期表单的数据类型至关重要。 如果未指定 `type`，浏览器会使用默认值 `text`。

给前两个 `input` 元素设置其 `type` 属性值为 `text`，第三个 `type` 属性值为 `email`，第四个 `type` 属性值为 `password`。

`email` 类型只允许包含 `@` 以及域名中包含 `.` 的电子邮件。 `password` 类型会屏蔽输入，如果网站没有启用 HTTPS 会警告。

# --hints--

应该给第一个 `input` 添加 `type` 其值为 `text`。

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

应该给第二个 `input` 添加 `type` 其值为 `text`。

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

应该给第三个 `input` 添加 `type` 其值为 `email`。

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

应该给第四个 `input` 添加 `type` 其值为 `password`。

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
