---
id: 60f81167d0d4910809f88945
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

距离父级 `form` 元素最近的第一个 `type` 为 `submit` 的 `input` 元素会被当作提交按钮。

为了处理表单提交，在最后 `fieldset` 元素后添加一个 `input` 元素，其 `type` 属性设置为 `submit`，`value` 属性设置为 `Submit`。

# --hints--

应该将 `input` 元素添加到最后一个 `fieldset` 元素后。

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

`input` 的 `type` 属性值应该为 `submit`。

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

`input` 元素的 `value` 属性值应该为 `Submit`。

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input type="text" /></label>
        <label>Enter Your Last Name: <input type="text" /></label>
        <label>Enter Your Email: <input type="email" /></label>
        <label>Create a New Password: <input type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>

--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}

label {
    display: block;
    margin: 0.5rem 0;
}

```
