---
id: 60fab4a123ce4b04526b082b
title: 步骤 28
challengeType: 0
dashedName: step-28
---

# --description--

目前，用户可以不选中单选按钮而提交表格。 虽然你先前能够使用 `required` 属性来表明一个输入框是必须填写的，但是在这里不行，因为给两个输入都添加 `required` 会向用户传递错误的信息。

为了解决这个问题，你可以在第二个 `fieldset` 中的 `label` 元素之前添加一个有文本 `Account type (required)` 的 `legend` 元素，从而提供所需的上下文。 然后在 `Personal` 输入中添加 `checked` 属性，以确保提交表单时包含所需数据。

# --hints--

你的 `input` 元素应该保持 `type` 为 `radio`。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]')?.length, 2);
```

你应该在第二个 `fieldset` 元素中添加一个 `legend` 元素。

```js
assert.equal(document.querySelectorAll('fieldset:nth-of-type(2) > legend')?.length, 1);
```

你应该给 `legend` 元素添加 `Account type (required)` 文本。

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

你应该给 `Personal` 输入添加 `checked` 属性。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

        <label><input type="radio" name="account-type" /> Personal</label>
        <label><input type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
