---
id: 60fabf0dd4959805dbae09e6
title: 步骤 32
challengeType: 0
dashedName: step-32
---

# --description--

在第一个 `label` 在添加另一个，其文本为 `Input your age (years):`。 在它后面添加一个 `input` 并设置其 `type` 为 `number`。

由于我们不希望 13 岁以下的用户注册，给 `input` 添加一个 `min` 属性并设置其值为 `13`。 此外，我们假设年龄超过 120 岁的用户不会来注册，添加一个 `max` 属性，并设置其值为 `120`。

现在，如果提交的表单里面的年龄超过了这个范围，会出现一个警告，并阻止提交。 试一下吧。

# --hints--

应该在第三个 `fieldset` 内已存在的 `label` 后面在添加一个 `label`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

`label` 的文字应该为 `Input your age (years):`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

`label` 内应该有一个 `type` 为 `number` 的 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

`input` 应该有一个 `min` 属性且其值为 `13`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

`input` 应该有一个 `max` 属性且其值为 `120`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
