---
id: 60fab8367d35de04e5cb7929
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

在新添加的标签中，紧接着 `input` 元素之后添加文本 `I accept the terms and conditions`。 然后将文本 `terms and conditions` 链接到以下位置：

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

你应在第三个 fieldset 之后的标签中添加 `I accept the terms and conditions` 文本。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

你应该使用一个 `a` 元素来给条款和条件添加链接。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

你应该将新文本紧接在 `label` 中的 `input` 元素之后。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

你应该给 `a` 元素添加一个 `href` 属性并设置其值为 `https://www.freecodecamp.org/news/terms-of-service/`。

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

你应该仅把文本 `terms and conditions` 放在 `a` 元素中。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

锚元素内的文本有额外的前导或尾部空白。 锚文本中唯一的空格应位于 `terms`、`and` 和 `conditions` 之间。

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
