---
id: 62cc5b1779e4d313466f73c5
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `head` 元素中嵌套一个自闭合 `link` 元素。 设置其 `rel` 属性值为 `stylesheet`， `href` 属性值为 `styles.css`。

# --hints--

代码应该有一个 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

不应更改现有的 `head` 标签。 确保没有删除结束标签。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

应该有一个自闭合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 元素应该在 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
