---
id: 62ff8e998d3e7eae14d6ae3b
title: 第 30 步
challengeType: 0
dashedName: step-30
---

# --description--

你需要确认用户已阅读条款和条件。

添加一个 `label` 元素。 在新创建的 `label` 元素中添加一个 `input` 元素，并将 `type` 属性设置为 `checkbox`。 将这个 `input` 元素设置为 `required`，这样用户在不同意条款和条件的情况下无法注册。

为元素添加 `id` 和 `for` 属性，值为 `terms-and-conditions`，以增加可访问性。

# --hints--

你应该在第三个 `fieldset` 元素之后添加一个 `label`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

你应该给 `label` 元素添加一个 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

你应该给 `input` 元素添加值为 `checkbox` 的 `type` 属性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

你应该给 `input` 元素添加一个 `required` 属性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` 元素应该有一个值为 `terms-and-conditions` 的 `id`。

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` 元素应该有一个值为 `terms-and-conditions` 的 `for` 属性。

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
