---
id: 63541ef4f96cd82e8e6c788a
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

`method` 属性指定了如何将表单数据发送到 `action` 属性中指定的 URL。 表单数据可以通过 `GET` 请求作为 URL 参数发送（`method="get"`）或通过 `POST` 请求作为请求正文中的数据发送（`method="post"`）。

设置 `method` 属性通过 `POST` 请求发送你的表单数据。

# --hints--

你不应该添加一个新的 `form` 元素。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

你的 `form` 元素应有一个 `method` 属性。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

你的 `method` 属性应设置为 `post`。

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
