---
id: 671faf4ac1e19245df4db7f5
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Inside the `displayStory` function, first check if the `storyObj` object has a property that matches the `genre` parameter.

If the `storyObj` object has a property that matches the `genre` parameter, set the `result` paragraph's text content to the story of the selected genre.

Also, set the `storyContainer`'s border color to the `borderColor` property of the selected genre.

# --hints--

`displayStory(scary)` should result in the `result` paragraph displaying the scary story `"In the dark woods, a group of friends stumbled upon an old, abandoned cabin. They enter the cabin and awaken something malevolent that had been dormant for centuries"` with a border color of `#EE4B2B`.

```js
const ans = displayStory(scary);
ans.
assert.equal(storyContainer.style.borderColor, "#EE4B2B");

```

`displayStory(funny)` should result in the `result` paragraph displaying the funny story `"During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire. However, his attempt caused him to burn the dinner as well as his eyebrows off."` with a border color of `#f1be32`.

```js

```

`displayStory(adventure)` should result in the `result` paragraph displaying the adventure story `"Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple. They braved deadly traps and encountered strange wildlife, all while deciphering cryptic clues left behind by a mysterious civilization." ` with a border color of `#acd157`.

```js

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Storyteller app</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>

<h1>Want to hear a short story?</h1>

<main class="story-container">
        <div class="btn-container">
            <button class="btn" id="scary-btn">Scary Story</button>
            <button class="btn" id="funny-btn">Funny Story</button>
            <button class="btn" id="adventure-btn">Adventure Story</button>
        </div>
    		<p id="result"></p>

</main>
<script src="./script.js"></script>

</body>

</html>
```

```js
const storyContainer = document.querySelector(".story-container");

const scaryStoryBtn = document.getElementById("scary-btn");
const funnyStoryBtn = document.getElementById("funny-btn");
const adventureStoryBtn = document.getElementById("adventure-btn");

const result = document.getElementById("result");

const storyObj = {
  scary: {
    story: `In the dark woods, a group of friends stumbled upon an old, abandoned cabin.
        They enter the cabin and awaken something malevolent that had been dormant for centuries.`,
    borderColor: "#EE4B2B",
  },
  funny: {
    story: `During a camping trip, Mark decided to show off his culinary skills by cooking dinner over an open fire.
        However, his attempt caused him to burn the dinner as well as his eyebrows off.`,
    borderColor: "#f1be32",
  },
  adventure: {
    story: `Lost in the heart of the Amazon rain forest, Sarah and Jake stumbled upon an ancient temple.
        They braved deadly traps and encountered strange wildlife,
        all while deciphering cryptic clues left behind by a mysterious civilization.`,
    borderColor: "#acd157",
  },
};

--fcc-editable-region--
function displayStory(genre) {

}
--fcc-editable-region--

scaryStoryBtn.addEventListener("click", displayStory);


```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --dark-grey: #1b1b32;
    --black: #000;
    --white: #fff;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
}

body {
    background-color: var(--dark-grey);
    color: var(--white);
}

h1,
#result {
    text-align: center;
}

h1 {
    margin: 10px 0 15px;
}

.story-container {
    margin: auto;
    padding: 10px;
    width: 80%;
    border-style: double;
    border-width: 14px;
    border-color: var(--white);
}

.btn-container {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
}

@media (min-width: 760px) {
    .btn-container {
        flex-direction: row;
    }
}

#result {
    margin-top: 15px;
    font-size: 1.2rem;
    line-height: 30px;
}

.btn {
    cursor: pointer;
    width: 200px;
    margin: 10px 0 10px 0.5rem;
    color: var(--black);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}
```
