---
id: 6557709b0aee699a6a00528c
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

For example, a graph can be used to represent two points in the space, `A` and `B`, connected by a path. A graph like this will be made of two nodes connected by an edge.

Replace the existent `'species'` key with the strings `'A'`. Then, replace the correspondent value with the string `'B'` to represent the connection between the `'A'` and `'B'` nodes.

# --hints--

Your dictionary should have an `'A'` key.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    "A" in graph
  `))
})
```

Your `'A'` key should have `'B'` as the value.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    graph["A"] == "B"
  `))
})
```

Your dictionary should have two keys.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    len(graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
