---
id: 6579dd49fa8a8e1fd06b85a9
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Graphs are data structures representing relations between pairs of elements.
These elements, called *nodes*, can be real-life objects, entities, points in space or others. The connections between the nodes are called the *edges*.

Here's a visual representation of a graph:

<img class="img-responsive center-block" alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px;">


Rename the `copper` dictionary into `my_graph`. This will represent the graph to test your algorithm.

# --hints--

You should rename your `copper` dictionary into `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Your `my_graph` variable should be a dictionary.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
