---
id: 615f357957e370510f21ea16
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Створіть селектор `body` та надайте йому `font-family` зі значенням `Open Sans` з резервним шрифтом `sans-serif`.

Пам’ятайте, що для CSS шрифти з пробілами в назві повинні бути в лапках.

# --hints--

Ви повинні мати селектор `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Ваш селектор `body` повинен мати властивість `font-family` зі значенням `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
