---
id: 615f51e4e5b24a6e80eccce1
title: Крок 30
challengeType: 0
dashedName: step-30
---

# --description--

Створіть елемент `div` всередині свого елемента `.calories-info`. Надайте цьому елементу `div` атрибут `class` зі значенням `left-container`. В межах недавно створеного елемента `div` створіть елемент `h2` з текстом `Amount per serving`. Надайте елементу `h2` атрибут `class` зі значенням `bold small-text`.

# --hints--

Ви повинні створити новий елемент `div` в елементі `.calories-info`.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.localName === 'div');
```

Ваш новий елемент `div` повинен мати атрибут `class` зі значенням `left-container`.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.classList.contains('left-container'));
```

Ви повинні створити новий елемент `h2` в елементі `.left-container`.

```js
assert(document.querySelector('.left-container')?.children?.[0]?.localName === 'h2');
```

Ваш новий елемент `h2` повинен мати атрибут `class` зі значенням `bold small-text`.

```js
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('bold'));
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('small-text'));
```

Ваш новий елемент `h2` повинен мати текст `Amount per serving`.

```js
assert(document.querySelector('.left-container')?.children?.[0]?.innerText === 'Amount per serving');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
--fcc-editable-region--
    <div class="calories-info">

    </div>
--fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}
```
