---
id: 615f2abbe7d18d49a1e0e1c8
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Ми надали вам шаблонний код HTML.

Створіть елемент `h1` в елементі `body` та надайте йому текст `Nutrition Facts`.

# --hints--

Ви повинні додати новий елемент `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Ваш елемент `h1` повинен бути в межах елемента `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Ваш елемент `h1` повинен мати текст `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
