---
id: 615f3cafd794015aa9547a6d
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Пам’ятайте, що використання `h1`, `h2` та подібних теґів визначає семантичну структуру вашого HTML. Однак ви можете налаштувати CSS цих елементів, щоб контролювати візуальний потік та ієрархію.

Створіть правило `h1` та встановіть властивість `font-weight` на `800`. Це зробить ваш текст `h1` жирним.

# --hints--

Ви повинні створити селектор `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ваш селектор `h1` повинен мати властивість `font-weight` зі значенням `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
