---
id: 615f34ecc1091b4fd5a8a484
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

В своєму елементі `head` додайте елемент `link` з атрибутом `rel` зі значенням `stylesheet` та атрибут `href` зі значенням `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Це імпортує сімейство шрифтів `Open Sans` зі значеннями товщини шрифту `400`, `700` та `800`.

Також додайте елемент `link`, щоб прив’язати свій файл `styles.css`.

# --hints--

Ваш код повинен мати два самозакривні елементи `link`.

```js
assert(document.querySelectorAll('link').length === 2);
```

Обидва ваших елементи `link` повинні мати атрибут `rel` зі значенням `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Один з ваших елементів `link` повинен мати атрибут `href` зі значенням `./styles.css`.

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Один з ваших елементів `link` повинен мати атрибут `href` зі значенням `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
