---
id: 615f34948891834dd77655a6
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Додайте другий елемент `p` із текстом `Serving size 2/3 cup (55g)`.

# --hints--

Ви повинні мати другий елемент `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Ваш другий елемент `p` повинен бути в межах елемента `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Ваш другий елемент `p` повинен йти після наявного елемента `p`.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Ваш другий елемент `p` повинен мати текст `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
