---
id: 615f50473cc0196c6dd3892a
title: Крок 28
challengeType: 0
dashedName: step-28
---

# --description--

Ви можете помітити, що внизу елемента `.lg` досі є невеликий кордон. Щоб скинути його, надайте своєму селектору `.lg, .md` властивість `border` зі значенням `0`.

Примітка: (середній) клас `md` буде утилізовано в кроці 37 для тонших смужок на етикетці.

# --hints--

Ваш селектор `.lg, .md` повинен мати властивість `border` зі значенням `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.lg, .md')?.borderWidth === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

--fcc-editable-region--
.lg, .md {
  background-color: black;
}
--fcc-editable-region--
```
