---
id: 68a23004e96a105df3a21fc4
title: Build a Card Counting Assistant
challengeType: 26
dashedName: lab-counting-cards
---

# --description--

In the casino game Blackjack, a player can determine whether they have an advantage on the next hand over the house by keeping track of the relative number of high and low cards remaining in the deck. This is called Card Counting.

Having more high cards remaining in the deck favors the player. When the count is positive, the player should bet high. When the count is zero or negative, the player should bet low.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You will create a function called `cc`.
1. Your function will receive a `card` parameter which can either be a number or string. 
1. Your function will modify a global `count` variable based on certain criteria. 
1. The global `count` variable should be increased by `1` for the cards `2`, `3`,`4`, `5`, or `6`
1. The global `count` variable should remain unchanged for the cards `7`, `8`, `9`.
1. The global `count` variable should be decreased for the cards `10`, `'J'`, `'Q'`, `'K'`, `'A'`
1. Your function will return a string with current count and the string `Bet` if the count is positive.
1. Your function will return a string with current count and the string `Hold` if the count is less than or equal to `0`. 
1. The current count and the player's decision (`Bet` or `Hold`) should be separated by a space `-3 Hold`. 

# --hints--

Your function should return a value for count and the text (`Bet` or `Hold`) with one space character between them.

```js
count = 0;
let out = cc(10);
const pattern = /-?\d+ (Bet|Hold)/;
assert.match(out, pattern);
```

Cards Sequence `2`, `3`, `4`, `5`, `6` should return the string `5 Bet`

```js
count = 0;
cc(2);
cc(3);
cc(4);
cc(5);
const out = cc(6);
assert.strictEqual(out, '5 Bet')
```

Cards Sequence `7`, `8`, `9` should return the string `0 Hold`

```js
count = 0;
cc(7);
cc(8);
const out = cc(9);
assert.strictEqual(out, '0 Hold');
```

Cards Sequence `10`, `'J'`, `'Q'`, `'K'`, `'A'` should return the string `-5 Hold`

```js
count = 0;
cc(10);
cc('J');
cc('Q');
cc('K');
const out = cc('A');
assert.strictEqual(out, '-5 Hold');
```

Cards Sequence `3`, `7`, `'Q'`, `8`, `'A'` should return the string `-1 Hold`

```js
count = 0;
cc(3);
cc(7);
cc('Q');
cc(8);
const out = cc('A');
assert.strictEqual(out, '-1 Hold');
```

Cards Sequence `2`, `'J'`, `9`, `2`, `7` should return the string `1 Bet`

```js
count = 0;
cc(2);
cc('J');
cc(9);
cc(2);
const out = cc(7);
assert.strictEqual(out, '1 Bet');
```

Cards Sequence `2`, `2`, `10` should return the string `1 Bet`

```js
count = 0;
cc(2);
cc(2);
const out = cc(10);
assert.strictEqual(out,'1 Bet')
```

Cards Sequence `3`, `2`, `'A'`, `10`, `'K'` should return the string `-1 Hold`

```js
count = 0;
cc(3);
cc(2);
cc('A');
cc(10);
const out = cc('K');
assert.strictEqual(out,'-1 Hold');
```

# --seed--

## --seed-contents--


```js
let count = 0;


```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
