---
id: 673b02b03134b04637bf7055
title: Build a Reusable Footer using React
challengeType: 14
dashedName: build-a-reusable-footer-using-react
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `Footer`component.
2. Your `Footer` component should return a `footer` element.
3. The `footer` should have:

- Three unordered lists each with two list items.
- A paragraph element with a copyright `©` symbol.
- Three links with the `href` value set to `#` and the link content set to an icon or text of your choice.

# --hints--

You should have a `Footer` component.

```js
// passes

(getUserInput) =>
  assert(
    __helpers
      .removeWhiteSpace(getUserInput('index'))
      .includes('Footer(')
  );
```

The `Footer` component should return a `footer` element.

```js
assert(Enzyme.shallow(React.createElement(Footer)).type() === 'footer');
```

The `footer` should have three unordered lists each with two list items.

```js
assert(
    /<footer>.*<ul>.*<li>.*<\/li>.*<li>.*<\/li>.*<\/ul>.*<ul>.*<li>.*<\/li>.*<li>.*<\/li>.*<\/ul>.*<ul>.*<li>.*<\/li>.*<li>.*<\/li>.*<\/ul>.*<\/footer>/.test(
        Enzyme.shallow(React.createElement(Footer)).html()
    )
)
```

The `footer` should have a paragraph element with a copyright `©` symbol.

```js
assert(
    /<footer>.*<p>.*©.*<\/p>.*<\/footer>/.test(
        Enzyme.shallow(React.createElement(Footer)).html()
    )
)
```

The `footer` should have three links with the `href` value set to `#` and the link content set to an icon or text of your choice.

```js
assert(
    /<footer>.*<a href="#">.*<\/a>.*<a href="#">.*<\/a>.*<a href="#">.*<\/a>.*<\/footer>/.test(
        Enzyme.shallow(React.createElement(Footer)).html()
    )
)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>

    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>

    <script type="text/babel">
        const container = document.getElementById("root");
        const root = ReactDOM.createRoot(container);
        root.render(<Footer />);
    </script>
</body>

</html>
```

```css
.footer {
    background-color: #e3e1eb;
    padding: 20px;
    text-align: center;
    border-top: 1px solid #e5e7eb;
    font-family: Arial, sans-serif;
}

.footer-section {
    display: grid;
    grid-template-columns: repeat(3, 1fr);
    /* Three equal-width columns */
    gap: 20px;
    justify-content: center;
    margin-bottom: 20px;
}

.footer ul {
    list-style: none;
    padding: 0;
    margin: 0;
}

.footer li {
    margin-bottom: 10px;
}

.footer a {
    text-decoration: none;
    color: #1f2937;
}

.footer a:hover {
    color: #6366f1;
}

.footer p {
    color: #6b7280;
    margin-top: 20px;
}

.footer-icons {
    margin-top: 20px;
    display: flex;
    justify-content: center;
}

.footer-icons a {
    margin: 0 10px;
    font-size: 24px;
    /* Make the Unicode icons larger */
    color: #6b7280;
    text-decoration: none;
}

.footer-icons a:hover {
    color: #6366f1;
}
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>

    <!-- Don't use this in production: -->
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>

    <!-- Link to the external CSS file -->
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>

    <script type="text/babel">
        function Footer() {
            return (
                <footer className="footer">
                    <div className="footer-section">
                        <ul>
                            <li><a href="#">Fitness Dashboard</a></li>
                            <li><a href="#">Services</a></li>
                        </ul>
                        <ul>
                            <li><a href="#">Watch Videos</a></li>
                            <li><a href="#">Discord</a></li>
                        </ul>
                        <ul>
                            <li><a href="#">Privacy Policy</a></li>
                            <li><a href="#">Terms & Conditions</a></li>
                        </ul>
                    </div>
                    <p>© 2024 Fitness Dashboard. All Rights Reserved.</p>
                    <div className="footer-icons">
                        <a href="#">&#x1F3AE;</a>
                        <a href="#">&#x1F426;</a>
                        <a href="#">&#x1F4BB;</a>
                        <a href="#">&#x1F3C0;</a>
                    </div>
                </footer>
            );
        }

        const container = document.getElementById("root");
        const root = ReactDOM.createRoot(container);
        root.render(<Footer />);
    </script>
</body>

</html>
```

```css
.footer {
    background-color: #e3e1eb;
    padding: 20px;
    text-align: center;
    border-top: 1px solid #e5e7eb;
    font-family: Arial, sans-serif;
}

.footer-section {
    display: grid;
    grid-template-columns: repeat(3, 1fr);
    gap: 20px;
    justify-content: center;
    margin-bottom: 20px;
}

.footer ul {
    list-style: none;
    padding: 0;
    margin: 0;
}

.footer li {
    margin-bottom: 10px;
}

.footer a {
    text-decoration: none;
    color: #1f2937;
}

.footer a:hover {
    color: #6366f1;
}

.footer p {
    color: #6b7280;
    margin-top: 20px;
}

.footer-icons {
    margin-top: 20px;
    display: flex;
    justify-content: center;
}

.footer-icons a {
    margin: 0 10px;
    font-size: 24px;
    color: #6b7280;
    text-decoration: none;
}

.footer-icons a:hover {
    color: #6366f1;
}
```
