---
id: 66a93bbe65a26169dbf3bc39
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Your hotel feedback form should also collect an email address from the user. 

Start by adding a new `label` element with the text `Email address (required):` to the form. Your `label` element should have a `for` attribute set to the value of `"email"`.

# --hints--

You should have a `label` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Your `label` element should have the text `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

Your `label` element should have a `for` attribute set to the value of `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
