---
id: 66a94ea5df66236ebed260e8
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

The next section in the form will be responsible for asking users if they have stayed at the hotel before. 

Start by adding a `fieldset` element.

Inside the `fieldset` element, add a `legend` element with the text of `Was this your first time at our hotel?`. 

# --hints--

You should have a `fieldset` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 2);
```

You should have a `legend` element inside of your `fieldset` element.

```js
assert.lengthOf(document.querySelectorAll('fieldset legend'), 2);
```

Your `legend` element should have the text of `Was this your first time at our hotel?`.

```js
assert.strictEqual(document.querySelectorAll('fieldset legend')[1]?.textContent.trim(), 'Was this your first time at our hotel?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
