---
id: 66ad24c7eb8c121000c603a6
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

To make a checkbox input checked by default, you can add the `checked` attribute.

Here is an example of using the `checked` attribute:

```html
<input checked type="checkbox" id="checked" name="checked">
```

Add the `checked` attribute to the checkbox input with the `id` of `"reputation"` to make it checked by default.

# --hints--

Your checkbox input with the `id` of `"reputation"` should have the `checked` attribute.

```js
assert.exists(document.querySelector("input#reputation[checked]"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          --fcc-editable-region--
          <input type="checkbox" id="reputation" name="reputation" value="reputation" />
          --fcc-editable-region--
          <label for="reputation">Reputation</label>
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
