---
id: 67a51d3a8a6fe123b77b0c6e
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

`input` elements can have a `size` attribute. This attribute defines the number of characters that should be visible as the user types into the input. The value of `size` should be a non-negative integer greater than zero. If `size` is not specified, or is specified with an invalid value, the input will have the default width set by the browser.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Give the name and email inputs a `size` attribute with a value of `"20"`.

# --hints--

You should give the name `input` a `size` attribute with a value of `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

You should give the email `input` a `size` attribute with a value of `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
