---
id: 61fd77f7ad2aeb5ae34d07d6
title: 步驟 14
challengeType: 0
dashedName: step-14
---

# --description--

在第二個 `tr` 元素中，添加一個帶有文本 `Checking Our primary transactional account.` 的 `th` 元素。 將該文本（`Checking` 除外）包裝在 `span` 元素中，並將 `class` 屬性設置爲 `description`。

之後，添加三個 `td` 元素和以下文本（按順序）：`$54`、`$56`、`$53`。 給第三個 `td` 元素一個 `class` 屬性設置爲 `current`。

# --hints--

第二個 `tr` 應該有一個 `th` 元素。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th'));
```

`th` 元素應包含文本 `Checking Our primary transactional account.`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th')?.innerText === 'Checking Our primary transactional account.');
```

應該將文本 `Our primary transactional account.` 包裝在 `span` 元素中。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.textContent === 'Our primary transactional account.');
```

`span` 元素應該將 `class` 屬性設置爲 `description`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.classList?.contains('description'));
```

應該有三個 `td` 元素。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td').length === 3);
```

第一個 `td` 元素應該包含文本 `$54`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[0]?.textContent === '$54');
```

第二個 `td` 元素應該包含文本 `$56`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[1]?.textContent === '$56');
```

第三個 `td` 元素應該包含文本 `$53`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.textContent === '$53');
```

第三個 `td` 元素應該將 `class` 設置爲 `current`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
