---
id: 61fd70336ebb3e4f62ee81ba
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML 表格使用 `caption` 元素來描述表格的內容。 `caption` 元素應該始終是 `table` 的第一個子元素，但可以使用 `caption-side` CSS 屬性定位。

將 `caption` 元素添加到第一個 `table`，併爲其指定文本 `Assets`。

# --hints--

第一個 `table` 元素應該有一個 `caption` 元素。

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

`caption` 元素應該包含文本 `Assets`。

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
