---
id: 61fd6ab779390f49148773bb
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `h1` 元素下方，創建一個 `div` 元素。 給它一個 `id` 屬性設置爲 `years`。 你希望對屏幕閱讀器隱藏此特定元素，因此將 `aria-hidden` 屬性設置爲 `true`。

# --hints--

應該在 `h1` 元素之後創建一個新的 `div` 元素。

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

新的 `div` 元素應該將 `id` 屬性設置爲 `years`。

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

新的 `div` 元素應該將 `aria-hidden` 屬性設置爲 `true`。

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
