---
id: 61fd719788899952e67692b9
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

`thead` 和 `tbody` 元素用於指示表格的哪一部分是標題，哪一部分包含主要數據或內容。

將 `thead` 和 `tbody` 添加到第一個 `table`，在 `caption` 元素下方。

# --hints--

第一個 `table` 元素應該有一個 `thead` 元素。

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

第一個 `table` 元素應該有一個 `tbody` 元素。

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

`thead` 元素應該在 `caption` 元素的正下方。

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

`tbody` 元素應該在 `thead` 元素的正下方。

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
