---
id: 61fd67a656743144844941cb
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

屏幕閱讀器根據文檔流閱讀 HTML 元素。 我們最終希望資產負債表有一個 “Balance Sheet” 的標題和一個 “AcmeWidgetCorp” 的子標題。 但是，如果屏幕閱讀器宣佈此命令，這沒有任何意義。

將現有 `span` 的 `class` 屬性設置爲 `flex`，並在其中添加兩個 `span` 元素。 給第一個文本 `AcmeWidgetCorp`。 給第二個文本 `Balance Sheet`。 你將使用 CSS 來反轉頁面上文本的順序，但 HTML 順序對於屏幕閱讀器來說更有意義。

# --hints--

已有的 `span` 元素應該將 `class` 屬性設置爲 `flex`。

```js
assert(document.querySelector('h1')?.children?.[0]?.classList?.contains('flex'));
```

現有的 `span` 元素中應該有兩個新的 `span` 元素。

```js
assert(document.querySelector('.flex')?.children?.[0]?.localName === 'span');
assert(document.querySelector('.flex')?.children?.[1]?.localName === 'span');
```

新的 `span` 元素不應具有 `class` 屬性。

```js
assert(!document.querySelector('.flex')?.children?.[0]?.classList?.length);
assert(!document.querySelector('.flex')?.children?.[1]?.classList?.length);
```

第一個新的 `span` 元素應該包含文本 `AcmeWidgetCorp`。

```js
assert(document.querySelector('.flex')?.children?.[0]?.textContent === 'AcmeWidgetCorp');
```

第二個新的 `span` 元素應包含文本 `Balance Sheet`。

```js
assert(document.querySelector('.flex')?.children?.[1]?.textContent === 'Balance Sheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
