---
id: 61fd6b7c83dbf54a08cf0498
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `div` 元素中，添加三個 `span` 元素。 給它們每個一個 `class` 屬性設置爲 `year`，並添加以下文本（按順序）：`2019`、`2020` 和 `2021`。

# --hints--

`div` 元素應該有三個 `span` 元素。

```js
assert(document.querySelector('div')?.children?.length === 3);
```

每個 `span` 元素都應該有一個 `class` 屬性設置爲 `year`。

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

第一個 `span` 應該有文本 `2019`。

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

第二個 `span` 應該有文本 `2020`。

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

第三個 `span` 應該有文本 `2021`。

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
