---
id: 61fd66c687e610436494c6f1
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `section` 元素中，添加一個帶有嵌套 `span` 元素的 `h1` 元素。

# --hints--

`section` 元素應該有一個 `h1` 元素。

```js
assert(document.querySelector('section')?.children?.[0]?.localName === 'h1');
```

`h1` 元素應該有一個 `span` 元素。

```js
assert(document.querySelector('h1')?.children?.[0]?.localName === 'span');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
