---
id: 61fd7a160ed17960e971f28b
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

在第四個 `tr` 元素中，添加一個帶有文本 `Total Assets` 的 `th` 元素。 將文本 `Assets` 包裹在 `span` 元素中，並將 `class` 屬性設置爲 `sr-only`。

之後，添加三個 `td` 元素以及以下文本（按順序）：`$579`、`$736`、`$809`。 給第三個 `td` 元素一個 `class` 屬性設置爲 `current`。

# --hints--

第四個 `tr` 應該有一個 `th` 元素。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th'));
```

`th` 元素應包含文本 `Total Assets`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th')?.innerText === 'Total Assets');
```

應該將文本 `Assets` 包裝在 `span` 元素中。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th > span')?.textContent === 'Assets');
```

`span` 元素應該將 `class` 屬性設置爲 `sr-only`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th > span')?.classList?.contains('sr-only'));
```

應該有三個 `td` 元素。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td').length === 3);
```

第一個 `td` 元素應該包含文本 `$579`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[0]?.textContent === '$579');
```

第二個 `td` 元素應該包含文本 `$736`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[1]?.textContent === '$736');
```

第三個 `td` 元素應該包含文本 `$809`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[2]?.textContent === '$809');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
