---
id: 61968e9243a4090cc805531c
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

プロパティを 1 つだけ使用して、水平方向、垂直方向両方のスクロールバーを取り除いてください。

# --hints--

`body` の `overflow` を `--fcc-expected--` に設定する必要があります。 現状、`--fcc-actual--` が設定されています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;

}
--fcc-editable-region--
```
