---
id: 6196ce0415498d2463989e84
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

山脈のように見せるために、山をもう 1 つ追加しましょう。`.left-mountain` の直後に新しい `div` を作成し、その新しい `div` の `class` を `back-mountain` に設定してください。

# --hints--

`body` 内に `div` を新たに 1 つ追加する必要があります。

```js
assert.equal(document.querySelectorAll('body > div')?.length, 4);
```

新しい `div` の `class` を `back-mountain` に設定する必要があります。

```js
assert.include(document.querySelector('div:not(.left-mountain, .ground, .penguin)')?.className, 'back-mountain');
```

`.back-mountain` は `.left-mountain` の後に置く必要があります。

```js
assert.strictEqual(document.querySelector('.left-mountain')?.nextElementSibling, document.querySelector('.back-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>
    <div class="left-mountain"></div>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
