---
id: 619691693bc14b0e528f5a20
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`.ground` 要素を選択し、その `width` がビューポートの幅全体に広がるように設定してください。 次に、`height` を `400px` に設定してください。

# --hints--

`.ground` セレクターを使用する必要があります。

```js
assert.match(code, /\.ground\s*\{/);
```

`.ground` の要素の `width` を `100vw` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

`.ground` の要素の `height` を `400px` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
