---
id: 619665c9abd72906f3ad30f9
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

You will be building a happy Flappy Penguin, and further exploring CSS transforms and animations in the process.

Begin with your basic HTML boilerplate. Include the `DOCTYPE` declaration, `html` element with a language set to English, the appropriate `meta` tags, a `head`, `body`, and `title` element. Also, link your stylesheet to the page.

# --hints--

Your code should have a `<!DOCTYPE html>` declaration.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Your code should have an `html` element.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/gi));
```

Your code should have a `head` element within the `html` element.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Your code should have a `body` element within the `html` element.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Your `head` element should come before your `body` element.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

`meta` 要素のひとつは `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もうひとつの `meta` 要素は `charset` 属性を `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your `title` should have some text.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

コードには `link` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('link'));
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
