---
id: 619665c9abd72906f3ad30f9
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

ぱたぱたと翼を動かす楽しそうなペンギンを作成しながら、CSS の座標変換とアニメーションについてより深く学びましょう。

基本的な HTML ボイラープレートから始めましょう。 `DOCTYPE` 宣言、言語を英語に設定した `html` 要素、適切な `meta` タグ一式、`head` 要素、`body` 要素、`title` 要素を追加してください。 また、スタイルシートをページにリンクしてください。

# --hints--

コードには `<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

コードには `html` 要素が 1 つ必要です。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/gi));
```

コードの `html` 要素内に `head` 要素が必要です。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

コードの `html` 要素内に `body` 要素が必要です。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

`head` 要素は `body` 要素の前に置く必要があります。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

`meta` 要素が 2 つ必要です。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

`meta` 要素のひとつは `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もうひとつの `meta` 要素は `charset` 属性を `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` には何らかのテキストが必要です。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

コードには `link` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('link'));
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
