---
id: 61969e7451455614217e901b
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

もっと山らしく見えるようにするには、`skew` 変換関数が使用できます。これは 2 つの引数を取ります。 1 つ目は X 軸に沿ってゆがめる角度で、2 つ目は Y 軸に沿ってゆがめる角度です。

`transform` プロパティを使用して、山を X 軸方向に `0deg`、Y 軸方向に `44deg` ゆがめてください。

# --hints--

`.left-mountain` には `transform` プロパティが必要です。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform);
```

`transform` に `skew` 関数を使用する必要があります。

```js
assert.include(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform, 'skew');
```

`.left-mountain` の `transform` を `skew(0deg, 44deg)` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.getPropVal('transform', true), 'skew(0deg,44deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
