---
id: 6196cee94c6da1253809dff9
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

`.back-mountain` 要素を選択し、その `width` と `height` を `300px` に設定してください。 次に、`background` を線形グラデーションに設定してください。`rgb(203, 241, 228)` で始まり `rgb(47, 170, 255)` で終わるようにしてください。

# --hints--

`.back-mountain` セレクターを使用する必要があります。

```js
assert.match(code, /\.back-mountain\s*\{/);
```

`.back-mountain` の `width` を `300px` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.width, '300px');
```

`.back-mountain` の `height` を `300px` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.height, '300px');
```

`.back-mountain` の `background` を `linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255))` に設定する必要があります。

```js
assert.include(['linear-gradient(rgb(203,241,228),rgb(47,170,255))', 'rgba(0,0,0,0)linear-gradient(rgb(203,241,228),rgb(47,170,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.back-mountain')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

--fcc-editable-region--

--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
