---
id: 6196d32d1340d829f0f6f57d
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

太陽の `position` プロパティを、ページレイアウトのスペースを占有しないように設定してください。また、太陽の形が円になるように `border-radius` を設定してください。

# --hints--

`.sun` の `position` を `absolute` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sun')?.position, 'absolute');
```

`.sun` の `border-radius` を `50%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sun')?.borderRadius, '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

--fcc-editable-region--
.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
