---
id: 61968f8877c6720d6d61aaf5
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`body` 内に `div` を追加して、その `class` を `ground` に設定してください。

# --hints--

`div` を新たに 1 つ追加する必要があります。

```js
assert.exists(document.querySelector('div'));
```

`div` の `class` を `ground` に設定する必要があります。

```js
assert.include(document.querySelector('div')?.className, 'ground');
```

`div` は `body` の中に置く必要があります。

```js
assert.exists(document.querySelector('body > div.ground'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}
```
