---
id: 619692ff79f5770fc6d8c0b4
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`.ground` 要素の上に `class` が `penguin` の `div` を 1 つ追加してください。 この `div` の中にペンギンが入ります。

# --hints--

`body` 内に `div` を新たに 1 つ追加する必要があります。

```js
assert.equal(document.querySelectorAll('body > div')?.length, 2);
```

`div` の `class` を `penguin` に設定する必要があります。

```js
assert.include(document.querySelector('body > div:not(.ground)')?.className, 'penguin');
```

`div.penguin` を `div.ground` の前に置く必要があります。

```js
assert.strictEqual(document.querySelector('.ground')?.previousElementSibling, document.querySelector('.penguin'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
