---
id: 61993dbb35adf30b10d49e38
title: ステップ 32
challengeType: 0
dashedName: step-32
---

# --description--

`.penguin-body` 要素を選択し、`width` を `53%` に、`height` を `45%` に設定してください。 次に、`background` に線形グラデーションを設定してください。`45deg` の角度で、色は `0%` の位置で `rgb(134, 133, 133)`、`25%` からは `rgb(234, 231, 231)`、`67%` からは `white` になるよう設定してください。

# --hints--

`.penguin-body` セレクターを使用する必要があります。

```js
assert.match(code, /\.penguin-body\s*\{/);
```

`.penguin-body` の `width` を `--fcc-expected--` に設定する必要がありますが、`--fcc-actual--` となっています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin-body')?.width, '53%');
```

`.penguin-body` の `height` を `--fcc-expected--` に設定する必要がありますが、`--fcc-actual--` となっています。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin-body')?.height, '45%');
```

`.penguin-body` の `background` を `linear-gradient(45deg, rgb(134, 133, 133) 0%, rgb(234, 231, 231) 25%, white 67%)` に設定する必要があります。

```js
assert.include(['linear-gradient(45deg,rgb(134,133,133)0%,rgb(234,231,231)25%,white67%)', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(134,133,133)0%,rgb(234,231,231)25%,white67%)repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.penguin-body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin">
      <div class="penguin-head"></div>
      <div class="penguin-body"></div>
    </div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
  z-index: 4;
  position: relative;
}

.penguin-head {
  width: 50%;
  height: 45%;
  background: linear-gradient(
    45deg,
    gray,
    rgb(239, 240, 228)
  );
  border-radius: 70% 70% 65% 65%;
}

--fcc-editable-region--

--fcc-editable-region--

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
