---
id: 5d8a4cfbe6b6180ed9a1c9de
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Welcome to the dashboard project! You will be using the JavaScript data visualization library, D3, to build a visualization of your social media followers. It will consist of a line graph, a pie chart, and a legend.

First, you need to create the HTML file. Start by adding the `<!DOCTYPE html>` declaration at the top of the file to tell the browser what type of document it's reading.

# --hints--

test-text

```js
assert(/<!DOCTYPE\s+html\s*>/gi.test(code));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
