---
id: 63b6075a62883218d282504c
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

L'elemento `input` ha bisogno di qualche altro attributo. Dagli un `type` impostato su `number` per consentire solo input numerici, un attributo `min` impostato su `0` per consentire solo numeri positivi, e un `placeholder` impostato su `Daily calorie budget`.

Infine, contrassegna l'elemento `input` come `required`.

# --hints--

L'elemento `input` dovrebbe avere un attributo `type` impostato su `number`.

```js
assert.equal(document.querySelector('form input').type, 'number');
```

L'elemento `input` dovrebbe avere un attributo `min` impostato su `0`.

```js
assert.equal(document.querySelector('form input').min, '0');
```

L'elemento `input` dovrebbe avere un attributo `placeholder` impostato su `Daily calorie budget`.

```js
assert.equal(document.querySelector('form input').placeholder, 'Daily calorie budget');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input id="budget" />
        </form>
--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
