---
id: 63b60af1a0b9f7238a9dd294
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Il menu di selezione ha bisogno di opzioni per ciascuno degli elementi `fieldset` di cibi ed esercizi che hai creato nei passaggi precedenti. Usa l'elemento `option` per creare una nuova opzione per ogni `fieldset`. L'attributo `value` di ogni opzione dovrebbe essere l'`id` del `fieldset` e il testo di ogni opzione dovrebbe essere il testo di `legend`.

Imposta l'opzione `Breakfast` come opzione `selected`.

# --hints--

Dovresti creare cinque elementi `option` all'interno dell'elemento `select`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.length, 5);
```

Il primo `option` dovrebbe avere il testo `Breakfast`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.textContent, 'Breakfast');
```

Il primo `option` dovrebbe avere l'attributo `value` impostato su `breakfast`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.value, 'breakfast');
```

Il secondo `option` dovrebbe avere il testo `Lunch`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.textContent, 'Lunch');
```

Il secondo `option` dovrebbe avere l'attributo `value` impostato su `lunch`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.value, 'lunch');
```

Il terzo `option` dovrebbe avere il testo `Dinner`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.textContent, 'Dinner');
```

Il terzo `option` dovrebbe avere l'attributo `value` impostato su `dinner`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.value, 'dinner');
```

Il quarto `option` dovrebbe avere il testo `Snacks`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.textContent, 'Snacks');
```

Il quarto `option` dovrebbe avere l'attributo `value` impostato su `snacks`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.value, 'snacks');
```

Il quinto `option` dovrebbe avere il testo `Exercise`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.textContent, 'Exercise');
```

Il quinto `option` dovrebbe avere l'attributo `value` impostato su `exercise`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.value, 'exercise');
```

Il primo `option` dovrebbe essere selezionato di default.

```js
assert.isTrue(document.querySelectorAll('.controls select option')?.[0]?.getAttributeNames()?.includes('selected'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">

              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
