---
id: 63b60aaaa65f8922bfce6b7e
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

All'interno dell'elemento `span` crea un elemento `label` per un `entry-dropdown` e dagli il testo `Add food or exercise:`. Quindi crea un elemento `select` con l'`id` impostato su `entry-dropdown` e un `name` impostato su `options`. Sotto questo, aggiungi un elemento `button` con l'`id` impostato su `add-entry` e il testo `Add Entry`.

Dai all'elemento `button` un attributo `type` impostato su `button` per evitare l'invio automatico del modulo.

# --hints--

Dovresti aggiungere un elemento `label` all'elemento `span`.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Il nuovo elemento `label` dovrebbe avere un attributo `for` impostato su `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Il nuovo elemento `label` dovrebbe avere il testo `Add food or exercise:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Dovresti aggiungere un elemento `select` all'elemento `span`.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

L'elemento `select` dovrebbe trovarsi dopo l'elemento `label`.

```js
assert(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName === 'LABEL');
```

Il nuovo elemento `select` dovrebbe avere un attributo `id` impostato su `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Il nuovo elemento `select` dovrebbe avere un attributo `name` impostato su `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Dovresti aggiungere un elemento `button` all'elemento `span`.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

L'elemento `button` dovrebbe trovarsi dopo l'elemento `select`.

```js
assert(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName === 'SELECT');
```

Il nuovo elemento `button` dovrebbe avere un attributo `id` impostato su `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Il nuovo elemento `button` dovrebbe avere il testo `Add Entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
