---
id: 63b606f09a14cc1781aea1fb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Crea un elemento `label`, dagli un attributo `for` impostato su `budget` e il testo `Budget`, poi crea un elemento `input` con l'`id` impostato su `budget`.

# --hints--

Dovresti creare un elemento `label` nel `form`.

```js
assert.exists(document.querySelector('form label'));
```

L'elemento `label` dovrebbe avere un attributo `for` impostato su `budget`.

```js
assert.equal(document.querySelector('form label')?.getAttribute('for'), 'budget');
```

L'elemento `label` dovrebbe avere il testo `Budget`.

```js
assert.equal(document.querySelector('form label')?.innerText, 'Budget');
```

Dovresti creare un elemento `input` all'interno del `form`.

```js
assert.exists(document.querySelector('form input'));
```

L'elemento `input` dovrebbe trovarsi dopo l'elemento `label`.

```js
assert.equal(document.querySelector('form input')?.previousElementSibling?.tagName, "LABEL");
```

L'elemento `input` dovrebbe avere un `id` impostato su `budget`.

```js
assert.equal(document.querySelector('form input')?.id, "budget");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--
        <form id="calorie-counter">

        </form>
--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
