---
id: 63b607af6fcdb119aae9b16a
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Crea un elemento `fieldset` con l'`id` impostato su `breakfast`.

Al suo interno, crea un elemento `legend` con il testo `Breakfast` e un `div` vuoto con l'attributo `class` impostato su `input-container`.

# --hints--

Dovresti creare un elemento `fieldset` all'interno del `form`.

```js
assert.exists(document.querySelector('form fieldset'));
```

L'elemento `fieldset` dovrebbe trovarsi dopo l'elemento `input`.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

L'elemento `fieldset` dovrebbe avere un `id` impostato su `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Dovresti creare un elemento `legend` all'interno del `fieldset`.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

L'elemento `legend` dovrebbe avere il testo `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Dovresti creare un elemento `div` all'interno del `fieldset`.

```js
assert.exists(document.querySelector('form fieldset div'));
```

L'elemento `div` dovrebbe trovarsi dopo l'elemento `legend`.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

L'elemento `div` dovrebbe avere l'attributo `class` impostato su `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
