---
id: 63b61490e633a22b4593e62f
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Seguendo lo stesso modello, assegna l'elemento `#add-entry` a `addEntryButton`, l'elemento `#clear` a `clearButton`, e l'elemento `#output` a `output`.

# --hints--

Dovresti dichiarare una variabile chiamata `addEntryButton`.

```js
assert.isDefined(addEntryButton);
```

Dovresti usare `document.getElementById()` per selezionare l'elemento `#add-entry`.

```js
assert.match(code, /document\.getElementById\(\s*('|")add-entry\1\s*\)/g);
```

Dovresti assegnare l'elemento `#add-entry` a `addEntryButton`.

```js
assert.equal(addEntryButton, document.getElementById('add-entry'));
```

Dovresti dichiarare una variabile chiamata `clearButton`.

```js
assert.isDefined(clearButton);
```

Dovresti usare `document.getElementById()` per selezionare l'elemento `#clear`.

```js
assert.match(code, /document\.getElementById\(\s*('|")clear\1\s*\)/g);
```

Dovresti assegnare l'elemento `#clear` a `clearButton`.

```js
assert.equal(clearButton, document.getElementById('clear'));
```

Dovresti dichiarare una variabile chiamata `output`.

```js
assert.isDefined(output);
```

Dovresti usare `document.getElementById()` per selezionare l'elemento `#output`.

```js
assert.match(code, /document\.getElementById\(\s*('|")output\1\s*\)/g);
```

Dovresti assegnare l'elemento `#output` a `output`.

```js
assert.equal(output, document.getElementById("output"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">
                <option value="breakfast" selected>Breakfast</option>
                <option value="lunch">Lunch</option>
                <option value="dinner">Dinner</option>
                <option value="snacks">Snacks</option>
                <option value="exercise">Exercise</option>
              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
          <div>
            <button type="submit" id="calculate-calories">
              Calculate Remaining Calories
            </button>
            <button type="button" id="clear">Clear</button>
          </div>
        </form>
        <div id="output" class="output hide"></div>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js
--fcc-editable-region--
const calorieCounter = document.getElementById('calorie-counter');
const budgetNumberInput = document.getElementById('budget');
const entryDropdown = document.getElementById('entry-dropdown');

--fcc-editable-region--
```
