---
id: 63b6088bb56e2d1cac364043
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Continuando il modello, crea un `fieldset` per `dinner` con gli stessi elementi annidati.

# --hints--

Dovresti creare un terzo elemento `fieldset` all'interno del `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[2]);
```

Il terzo elemento `fieldset` dovrebbe trovarsi dopo il secondo elemento `fieldset`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.previousElementSibling?.tagName, "FIELDSET");
```

Il terzo elemento `fieldset` dovrebbe avere un `id` impostato su `dinner`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.id, "dinner");
```

Il terzo elemento `fieldset` dovrebbe contenere un elemento `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[2]?.querySelector('legend'));
```

Il nuovo elemento `legend` dovrebbe avere il testo `Dinner`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.querySelector('legend')?.innerText, "Dinner");
```

Il terzo elemento `fieldset` dovrebbe contenere un elemento `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[2]?.querySelector('div'));
```

Il nuovo elemento `div` dovrebbe avere l'attributo `class` impostato su `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[2]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
