---
id: 6989f10a72ad424d4ad6d514
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Next, you need to validate that the number of pairs is at least one, since you can't generate parentheses combinations with zero or negative pairs.

Add another `if` statement to check if `pairs` is less than `1`. If this condition is true, return the string `The number of pairs should be at least 1`.

# --hints--

Your second `if` statement should check if `pairs` is less than `1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

You should return the string `The number of pairs should be at least 1` when `pairs` is less than `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
    if (!Number.isInteger(pairs)) {
        return 'The number of pairs should be an integer';
    }
--fcc-editable-region--

--fcc-editable-region--
    return [];
}
```
