---
id: 6989eb65324d97775c4965fe
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you'll implement a function that generates all valid combinations of parentheses using a breadth-first search (BFS) approach. For example, the valid combinations of two pairs of parentheses are `(())` and `()()`.

Start by creating a function named `genParentheses` with a single parameter `pairs`. For now, return an empty array from the function.

# --hints--

You should define a function named `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;

assert.exists(genParentheses);
```

Your `genParentheses` function should have a single parameter named `pairs`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const genParenthesesParams = genParentheses.parameters;

assert.equal(genParenthesesParams.length, 1);
assert.equal(genParenthesesParams[0].toString(), 'pairs');
```

Your function should return an empty array.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
