---
id: 6989ef885db3a3488a13b426
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Before implementing the core algorithm, you need to validate the input. The `pairs` parameter should be an integer, as it represents the number of parentheses pairs to generate.

Add an `if` statement at the beginning of your function to check if `pairs` is not an integer. Use the `Number.isInteger()` method for that.

If the condition is true, return the string `The number of pairs should be an integer`.

# --hints--

Your `if` statement should check if `pairs` is not an integer using `!Number.isInteger(pairs)`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

You should return the string `The number of pairs should be an integer` when the input is not an integer.

```js
let a=genParentheses("i")
assert.equal(a, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--

--fcc-editable-region--
    return [];
}
```

