---
id: 6989f207306332c06b0add39
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

For the breadth-first search approach, you'll use a queue to track different states as you build the parentheses combinations. Each state will be represented as an array containing three elements:

- The current string being built
- The number of opening parentheses used so far
- The number of closing parentheses used so far

Create a variable named `queue` and initialize it with an array containing one array: `['', 0, 0]`. This represents the starting state with an empty string and zero parentheses used.

# --hints--

You should declare a variable named `queue` in your `genParentheses` function.

```js
assert.match(genParentheses.toString(), /queue/);
```

You should initialize `queue` with an array containing the array `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
    if (!Number.isInteger(pairs)) {
        return 'The number of pairs should be an integer';
    }
    if (pairs < 1) {
        return 'The number of pairs should be at least 1';
    }
--fcc-editable-region--
    
--fcc-editable-region--
    let result = [];
    return result;
}
```
