---
id: 6989eb65324d97775c4965fe
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you'll implement a function that generates all valid combinations of parentheses using a breadth-first search (BFS) approach. For example, the valid combinations of two pairs of parentheses are `(())` and `()()`.

Start by creating a function named `genParentheses` with a single parameter `pairs`. For now, return an empty list from the function.

# --hints--

You should define a function named `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Your `genParentheses` function should have a single parameter named `pairs`.

```js
assert.match(genParentheses.toString(), /\(\s*pairs\s*\)/);
```

Your function should return an empty list.

```js
assert.typeOf(genParentheses(1), 'object');
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
