---
id: 6989f2e35e6941e13e445156
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now you'll implement the main BFS loop. Create a `while` loop that continues as long as the `queue` is not empty (check if `queue.length` is greater than `0`).

Inside the loop, log `queue` to the console.

# --hints--

Your `while` loop should have the condition `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

You should log `queue` to the console inside your `while` loop.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
