---
id: 6989f1ba358acba651d5ee6b
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now you'll set up the data structure to store your results. Create a variable named `result` and initialize it to an empty array. This array will store all the valid parentheses combinations you generate.

Update your return statement to return `result` instead of an empty array.

# --hints--

You should declare a variable named `result` in your `genParentheses` function.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses.variables;

assert.exists(result);
```

You should initialize `result` to an empty array.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses.variables;
const resultValue = eval(result.value.toString());

assert.isArray(resultValue);
assert.isEmpty(resultValue);
```

You should return `result` at the end of your function.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;

assert.isTrue(genParentheses.hasReturn('result'));
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
