---
id: bad87fee1348bd9aedc08826
title: Mirar elementos por classe usando o jQuery
challengeType: 6
forumTopicId: 18316
required:
  - 
    link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-class-using-jquery
---

# --description--

Você vê como fizemos todos os seus elementos `button` quicar? Nós os selecionamos com `$("button")`, em seguida adicionamos algumas classes CSS a eles com `.addClass("animated bounce");`.

Você acabou de usar a função `.addClass()` do jQuery, a qual o permite adicionar classes a elementos.

Primeiro, vamos ter como alvo seus elementos `div` com a classe `well` usando o seletor `$(".well")`.

Observe que, assim como com declarações CSS, você digita um `.` antes do nome da classe.

Em seguida, use a função `.addClass()` do jQuery para adicionar as classes `animated` e `shake`.

Por exemplo, você poderia fazer todos os elementos com a classe `text-primary` sacudir adicionando o seguinte à sua `document ready function`:

```js
$(".text-primary").addClass("animated shake");
```

# --hints--

Você deve usar a função `addClass()` do jQuery para dar as classes `animated` e `shake` para todos os seus elementos com a classe `well`.

```js
assert($('.well').hasClass('animated') && $('.well').hasClass('shake'));
```

Você deve usar apenas jQuery para adicionar essas classes ao elemento.

```js
assert(!code.match(/class\.\*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
