---
id: 6859a4c514603bc6d95a80da
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Now it is time associate the tab panels with their respective tabs. 

For the first `tabpanel` element, add the `id` attribute with the value `"panel-earth"` and an `aria-labelledby` attribute with the value `"tab-earth"`.

# --hints--

Your first `tabpanel` element should have an `id` attribute.

```js
assert.isTrue(document.querySelector('[role="tabpanel"]').hasAttribute("id"));
```

Your `id` attribute should have the value `"panel-earth"`.

```js
assert.equal(document.querySelector('[role="tabpanel"]').getAttribute("id"), "panel-earth");
```

Your first `tabpanel` element should have an `aria-labelledby` attribute.

```js
assert.isTrue(document.querySelector('[role="tabpanel"]').hasAttribute("aria-labelledby"));
```

Your `aria-labelledby` attribute should have the value `"tab-earth"`.

```js
assert.equal(document.querySelector('[role="tabpanel"]').getAttribute("aria-labelledby"), "tab-earth");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--
      <div role="tabpanel">
      --fcc-editable-region--
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It’s the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div role="tabpanel">
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It’s a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div role="tabpanel">
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It’s a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
