---
id: 6859ae3ef795c4c9d4badf9c
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Now it is time to start adding the functionality to the tablist. The first step is to access all of the tab and panel elements in the tablist. 

Create a variable called `tabs` and assign it the result of querying for all elements with the `role` attribute set to `tab`.

Then, create a variable called `panels` and assign it the result of querying for all elements with the `role` attribute set to `tabpanel`.

Remember that you can use the `querySelectorAll` method to select multiple elements at once.

# --hints--

You should have a variable called `tabs`.

```js
assert.isDefined(tabs);
```

You should assign `document.querySelectorAll('[role="tab"]')` to the `tabs` variable.

```js
assert.deepEqual(
  Array.from(tabs),
  Array.from(document.querySelectorAll('[role="tab"]'))
);
```

You should have a variable called `panels`.

```js
assert.isDefined(panels);
```

You should assign `document.querySelectorAll('[role="tabpanel"]')` to the `panels` variable.

```js
assert.deepEqual(
  Array.from(panels),
  Array.from(document.querySelectorAll('[role="tabpanel"]'))
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn" hidden>
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars" hidden>
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
