---
id: 6859ab9ba7a915c84fa69851
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

For the second `tabpanel` element, add the `id` attribute with the value `"panel-saturn"` and an `aria-labelledby` attribute with the value `"tab-saturn"`.

For the third `tabpanel` element, add the `id` attribute with the value `"panel-mars"` and an `aria-labelledby` attribute with the value `"tab-mars"`.

# --hints--

Your second `tabpanel` element should have an `id` attribute.

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[1].hasAttribute("id"));
```

Your second `id` attribute should have the value `"panel-saturn"`.

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[1].getAttribute("id"), "panel-saturn");
```

Your second `tabpanel` element should have an `aria-labelledby` attribute.

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[1].hasAttribute("aria-labelledby"));
```

Your second `aria-labelledby` attribute should have the value `"tab-saturn"`.

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[1].getAttribute("aria-labelledby"), "tab-saturn");
``` 

Your third `tabpanel` element should have an `id` attribute.

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[2].hasAttribute("id"));
```

Your third `id` attribute should have the value `"panel-mars"`.

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[2].getAttribute("id"), "panel-mars");
```

Your third `tabpanel` element should have an `aria-labelledby` attribute.

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[2].hasAttribute("aria-labelledby"));
```

Your third `aria-labelledby` attribute should have the value `"tab-mars"`.

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[2].getAttribute("aria-labelledby"), "tab-mars");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      --fcc-editable-region--
      <div role="tabpanel">
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div role="tabpanel">
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
      --fcc-editable-region--
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
