---
id: 6859acc690736fc911a70967
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Right now, all three facts are visible at the same time. But only one fact should be visible at a time.

To achieve this, you need to set the `hidden` attribute on the second and third `tabpanel` elements.

# --hints--

Your second `tabpanel` element should have the `hidden` attribute.

```js
assert.isTrue(document.querySelectorAll('div[role="tabpanel"]')[1].hasAttribute("hidden"));
```

Your third `tabpanel` element should have the `hidden` attribute.

```js
assert.isTrue(document.querySelectorAll('div[role="tabpanel"]')[2].hasAttribute("hidden"));
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It’s the only planet in the solar system known to harbor life.
        </p>
      </div>
      --fcc-editable-region--
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn">
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It’s a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars">
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It’s a key focus for exploration in the search for past or present life.
        </p>
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
