---
id: 61fd67a656743144844941cb
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

屏幕阅读器根据文档流阅读 HTML 元素。 我们最终希望资产负债表有一个 “Balance Sheet” 的标题和一个 “AcmeWidgetCorp” 的子标题。 但是，如果屏幕阅读器宣布此命令，这没有任何意义。

将现有 `span` 的 `class` 属性设置为 `flex`，并在其中添加两个 `span` 元素。 给第一个文本 `AcmeWidgetCorp`。 给第二个文本 `Balance Sheet`。 你将使用 CSS 来反转页面上文本的顺序，但 HTML 顺序对于屏幕阅读器来说更有意义。

# --hints--

已有的 `span` 元素应该将 `class` 属性设置为 `flex`。

```js
assert(document.querySelector('h1')?.children?.[0]?.classList?.contains('flex'));
```

现有的 `span` 元素中应该有两个新的 `span` 元素。

```js
assert(document.querySelector('.flex')?.children?.[0]?.localName === 'span');
assert(document.querySelector('.flex')?.children?.[1]?.localName === 'span');
```

新的 `span` 元素不应具有 `class` 属性。

```js
assert(!document.querySelector('.flex')?.children?.[0]?.classList?.length);
assert(!document.querySelector('.flex')?.children?.[1]?.classList?.length);
```

第一个新的 `span` 元素应该包含文本 `AcmeWidgetCorp`。

```js
assert(document.querySelector('.flex')?.children?.[0]?.textContent === 'AcmeWidgetCorp');
```

第二个新的 `span` 元素应包含文本 `Balance Sheet`。

```js
assert(document.querySelector('.flex')?.children?.[1]?.textContent === 'Balance Sheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
