---
id: 61fd6343fb42013d99bcd7f3
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `body` 元素中，将 `section` 元素嵌套在 `main` 元素中。

# --hints--

`body` 元素内的第一个元素应该是 `main` 元素。

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

`main` 元素中应该有一个 `section` 元素。

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
