---
id: 61fd71d596e8f253b9408b39
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

`tr` 元素用于指示表格行。 在 `thead` 元素中添加 `tr` 元素。 在新的 `tr` 元素中，添加一个 `td` 元素，后跟三个 `th` 元素。

`td` 元素表示一个数据单元，而 `th` 元素表示一个标题单元。

# --hints--

`thead` 元素应该有一个 `tr` 元素。

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

`tr` 元素应该有一个 `td` 元素作为第一个子元素。

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

`tr` 元素应该具有三个 `th` 元素，在 `td` 元素之后。

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
