---
id: 61fd778081276b59d59abad6
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

在第一个 `tr` 中，添加一个 `th` 元素，其中包含文本 `Cash This is the cash we currently have on hand.`。 将除 `Cash` 之外的所有文本包装在 `span` 元素中，并将 `class` 设置为 `description`。

之后，添加三个 `td` 元素和以下文本（按顺序）：`$25`、`$30`、`$28`。 给第三个 `td` 元素一个 `class` 属性设置为 `current`。

# --hints--

第一个 `tr` 应该有一个 `th` 元素。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th'));
```

`th` 元素应该有文本 `Cash This is the cash we currently have on hand.`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th')?.innerText === 'Cash This is the cash we currently have on hand.');
```

应该将文本 `This is the cash we currently have on hand.` 包装在 `span` 元素中。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.textContent === 'This is the cash we currently have on hand.');
```

`span` 元素应该将 `class` 属性设置为 `description`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.classList?.contains('description'));
```

应该有三个 `td` 元素。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td').length === 3);
```

第一个 `td` 元素应该包含文本 `$25`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[0]?.textContent === '$25');
```

第二个 `td` 元素应该包含文本 `$30`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[1]?.textContent === '$30');
```

第三个 `td` 元素应该包含文本 `$28`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.textContent === '$28');
```

第三个 `td` 元素应该将 `class` 设置为 `current`。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
