---
id: 61fd5a93fd62bb35968adeab
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

使用 `DOCTYPE`、表明此文件是英文的 `html`、`head` 和 `body` 元素设置你的 HTML。

为 `head` 元素提供合适的 `meta` 元素，包含 `charset` 和 `viewport`，一个具有适当标题的 `title` 元素，以及样式表 `link` 元素。

# --hints--

代码应该有一个 `<!DOCTYPE html>` 声明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

代码应该有一个 `html` 元素。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

你的 `html` 元素应该有一个带有值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的代码应该在 `html` 元素中包含一个 `head` 元素。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

你的代码应该在 `html` 元素中包含一个 `body` 元素。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

你的 `head` 元素应该在你的 `body` 元素之前。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

你应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

一个 `meta` 元素应将 `name` 设置为 `viewport`，并将 `content` 设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一个 `meta` 元素应将 `charset` 属性设置为 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 应该有一些文本。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

你的代码应该有一个 `link` 元素。

```js
assert(/<link/.test(code))
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

你的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
