---
id: 612e9d142affc44a453655db
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

是时候设计按键本身的样式了。 为 `.key` 元素创建一个 `class` 选择器。 将 `background-color` 设置为值 `#ffffff`，`position` 属性设置为 `relative`，`width` 属性为 `41px`，`height` 属性为 `175px`。

# --hints--

应该有一个 `.key` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

`.key` 选择器应该有一个 `background-color` 属性设置为 `#ffffff`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

`.key` 选择器应该将 `position` 属性设置为 `relative`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

`.key` 选择器应该有一个 `width` 属性设置为 `41px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

`.key` 选择器应该将 `height` 属性设置为 `175px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
