---
id: 612e6afc009b450a437940a1
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

从基本的 HTML 结构开始。 添加 `DOCTYPE` 声明，以及 `html`、`head`、`body` 和 `title` 元素。 将 `title` 设置为 `Piano`。

# --hints--

代码应包含 `DOCTYPE` 参考。

```js
assert(code.match(/<!DOCTYPE/gi));
```

应该在 `DOCTYPE` 引用之后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

应该将文档类型定义为 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

应该在类型之后使用 `>` 关闭 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个开始标签，带有值为 `en` 的 `lang` 属性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素应该有一个结束标记。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

应该有一个打开的 `head` 标记。

```js
assert(code.match(/<head\s*>/i));
```

应该有一个结束 `head` 标记。

```js
assert(code.match(/<\/head\s*>/i));
```

应该有一个打开的 `body` 标记。

```js
assert(code.match(/<body\s*>/i));
```

应该有一个结束 `body` 标记。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

项目应该有一个 `Piano` 的标题。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

请记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
