---
id: 612e89562043183c86df287c
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

浏览器可以将默认边距和填充值应用于特定元素。 为确保你的钢琴看起来正确，你需要重置盒子模型。

将 `html` 规则选择器添加到 CSS 文件，并将 `box-sizing` 属性设置为 `border-box`。

# --hints--

应该有一个 `html` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

`html` 选择器应该将 `box-sizing` 属性设置为 `border-box`。

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
