---
id: 612e813b3ba67633222cbe54
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

请记住，`class` 属性可以有多个值。 要将白键与黑键分开，需要添加第二个 `class` 值 `black--key`。 将此添加到第二个、第三个、第五个、第六个和第七个 `.key` 元素。

# --hints--

第二个 `.key` 元素也应该有一个 `black--key` 的 `class`。

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

第三个 `.key` 应该在 `class` 中有 `black--key`。

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

第五个 `.key` 应该在 `class` 中有 `black--key`。

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

第六个 `.key` 应该在 `class` 中有 `black--key`。

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

第七个 `.key` 应该在 `class` 中有 `black--key`。

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

应该有五个 `.black--key` 元素。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

应该有七个 `.key` 元素。

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
