---
id: 612e96fc87fe8e44f69f7ec5
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

是时候给按键设计样式了。 在 `#piano` 规则下方，使用 `class` 选择器定位 `.keys` 元素。 为新规则赋予 `#040404` 的 `background-color` 属性、`949px` 的 `width` 属性和 `180px` 的 `height` 属性。

# --hints--

应该有一个 `.keys` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

`.keys` 选择器应该将 `background-color` 属性设置为 `#040404`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

`.keys` 选择器应该将 `width` 属性设置为 `949px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

`.keys` 选择器应该将 `height` 属性设置为 `180px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
