---
id: 612ea4c4993aba52ab4aa32e
title: 步骤 18
challengeType: 0
dashedName: step-18
---

# --description--

现在是时候使用你之前准备的伪选择器了。 要创建黑键，请添加一个新的 `.key.black--key::after` 选择器。 这将针对具有类 `key black--key` 的元素，并在 HTML 中选择这些元素之后的伪元素。

在新的选择器中，将 `background-color` 设置为 `#1d1e22`。 还将 `content` 属性设置为 `""`。 这将使伪元素为空。

`content` 属性用于设置或覆盖元素的内容。 默认情况下，`::before` 和 `::after` 伪选择器创建的伪元素是空的，这些元素不会被渲染到页面上。 将 `content` 属性设置为空字符串 `""` 将确保元素在仍然为空的情况下呈现到页面。

如果你想进行实验，请尝试删除 `background-color` 属性并为 `content` 属性设置不同的值，例如 `"♥"`。 完成后请记住撤消这些更改，以便测试通过。

# --hints--

应该有一个 `.key.black--key::after` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

`.key.black--key::after` 选择器应该有一个 `background-color` 属性设置为 `#1d1e22`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

`.key.black--key::after` 选择器应该有一个 `content` 属性设置为 `""`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
