---
id: 612e8279827a28352ce83a72
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

现在复制七个 `.key` 元素的集合，并粘贴两个新的集合到 `.keys` 中。

# --hints--

应该有 21 个 `.key` 元素。

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

应该有 15 个 `.black--key` 元素。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

所有 `.key` 元素都应该在 `.keys` 元素中。

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

所有 `.black--key` 元素都应该在 `.keys` 元素中。

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
