---
id: 68e6698b29e562d21b1d260c
title: Major Browsers List 
challengeType: 0
dashedName: major-browsers-list
# demoType can either be 'onClick' or 'onLoad'. If the project or lab doesn't have a preview, delete the property
demoType: onLoad
---

# --description--
<!-- step 1 instructions -->
In this challenge we are going to get you familiar with some major browsers and you will also learn how to use semantic tags like `<dl>`, `<dt>`, and `<dd>` visually.

Through out this workshop you will learn about;

The basic HTML document structure (`<!DOCTYPE html>`, `<html>`, `<head>`, `<body>`).
How to use definition list correctly. These are semantic tags like `<dl>`, `<dt>`, and `<dd>`.
The concept of definition lists (this is usually good for terms and descriptions). 
You will also get to practice writing valid HTML using modern browser standards.

In this workshop, you'll learn how to create a definition list of major web browsers using HTML semantic elements.

# --instructions--

Create a list using the `<dl>` tag.
Inside it, add `<dt>` elements for each browser name and `<dd>` elements for their descriptions.


```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>List of Browsers and Descriptions</title>
</head>
<body>
    <h1>List of Major Web Browsers</h1>
    <dl>
        <dt>Google Chrome</dt>
        <dd>This is a free web browser developed by Google and first released in 2008.</dd>

        <dt>Firefox</dt>
        <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

        <dt>Safari</dt>
        <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

        <dt>Brave</dt>
        <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

        <dt>Arc</dt>
        <dd>This is a free Chromium based web browser first released in 2023 by The Browser Company.</dd>
    </dl>
</body>
</html>
```

# --hints--

Test 1

```js
const dt = document.querySelector('dt');
assert.exists(dt);
```

Test 2

```js
const dd = document.querySelector('dd');
assert.exists(dd);
```

Test 3

```js
const h1 = document.querySelector('h1');
assert.exists(dl);
assert.equal(h1.htmlContent, "List of Major Web Browsers", "H1 tag contains text List of Major Web Browsers")
```

```html
--fcc-editable-region--

--fcc-editable-region--
```
