---
id: 68e6698b29e562d21b1d260c
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop you are going to practice using semantic tags like `dl`, `dt`, and `dd`.

Start by adding a heading to your page that reads `List of Major Web Browsers` using a `h1` element inside the `body` element.

# --hints--

Your page should have a `h1` element.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Your `h1` element should contain the text `List of Major Web Browsers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent, 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
        
    </body>
</html>
--fcc-editable-region--
```
