---
id: 68e6698b29e562d21b1d260c
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--
<!-- step 1 instructions -->
In this lesson you are going to get familiar with some major browsers and you will also learn how to use semantic tags like `dl`, `dt`, and `dd` visually.

Through out this workshop you will learn about;

The basic HTML document structure (`!DOCTYPE html`, `html`, `head`, `body` elemnts).
How to use definition list correctly. These are semantic tags like `dl`, `dt`, and `dd`.
The concept of definition lists (this is usually good for terms and descriptions). 
You will also get to practice writing valid HTML using modern browser standards.

In this workshop, you'll learn how to create a definition list of major web browsers using HTML semantic elements.

By the end of this project you will be able to 

- Use `dl` element to create a definition list.
- Use `dt` element to define a term.
- Use `dd` element to describe the term
- Create a well structured and semantic HTML content.

Add a heading to your page that reads "List of Major Web Browsers" using a `h1` element element inside the `body` tag.

# --hints--

Your page should have a `h1` element element

```js
const h1 = document.querySelector('h1');
assert.exists(h1)
```

Your `h1` element should contain the text "List of Major Web Browsers".

```js
const h1 = document.querySelector('h1');
assert.include(h1.textContent, 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
        
        </body>
</html>
--fcc-editable-region--

--fcc-editable-region--
```
