---
id: 68ecfa1003d88990bb96a33e
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Add four more `dt` and `dd` element pairs for the following browsers:

| Browser           | Description                                                                                        |
| ----------------- | -------------------------------------------------------------------------------------------------- |
| `Google Chrome`   | `This is a free web browser developed by Google and first released in 2008.`                       |
| `Firefox`         | `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`       |
| `Safari`          | `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.` |
| `Brave`           | `This is a free web browser first released in 2016 that is based on the Chromium web browser.`     |
| `Arc`             | `This is a free Chromium based web browser first released in 2023 by The Browser Company.`         |

# --hints--

Your `dl` element should contain five `dt` elements.

```js
const terms = document.querySelectorAll('dl dt');
assert.lengthOf(terms, 5);
```

Your `dl` element should contain five `dd` elements.

```js
const descs = document.querySelectorAll('dl dd');
assert.lengthOf(descs, 5);
```

Your `h1` element should be followed by a `dl` element.

```js
const h1 = document.querySelector('h1');
assert(h1?.nextElementSibling.tagName === 'DL'); 
```

Your first `dt` element should contain the text `Google Chrome`.

```js
const dts = document.querySelectorAll('dl dt');
assert.strictEqual(dts[0]?.textContent, 'Google Chrome');
```

Your first `dd` element should describe Google Chrome.

```js
const dds = document.querySelectorAll(`dl dd`);
assert.strictEqual(dds[0]?.textContent, "This is a free web browser developed by Google and first released in 2008."); 
```

Your second `dt` element should contain the text `Firefox`.

```js
const dts = document.querySelectorAll('dl dt');
assert.strictEqual(dts[1]?.textContent, 'Firefox');
```

Your second `dd` element should describe Firefox.

```js
const dds = document.querySelectorAll(`dl dd`);
assert.strictEqual(dds[1]?.textContent, "This is a free web browser developed by the Mozilla Corporation and first created in 2004."); 
```

Your third `dt` element should contain the text `Safari`.

```js
const dts = document.querySelectorAll('dl dt');
assert.strictEqual(dts[2]?.textContent, 'Safari');
```

Your third `dd` element should describe Safari.

```js
const dds = document.querySelectorAll(`dl dd`);
assert.strictEqual(dds[2]?.textContent, "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices."); 
```

Your fourth `dt` element should contain the text `Brave`.

```js
const dts = document.querySelectorAll('dl dt');
assert.strictEqual(dts[3]?.textContent, 'Brave');
```

Your fourth `dd` element should describe Brave.

```js
const dds = document.querySelectorAll(`dl dd`);
assert.strictEqual(dds[3]?.textContent, "This is a free web browser first released in 2016 that is based on the Chromium web browser."); 
```

Your fifth `dt` element should contain the text `Arc`.

```js
const dts = document.querySelectorAll('dl dt');
assert.strictEqual(dts[4]?.textContent, 'Arc');
```

Your fifth `dd` element should describe Arc.

```js
const dds = document.querySelectorAll(`dl dd`);
assert.strictEqual(dds[4]?.textContent, "This is a free Chromium based web browser first released in 2023 by The Browser Company."); 
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>
        </dl>
    </body> 
</html>
--fcc-editable-region--
```

