---
id: 68ecfa1003d88990bb96a33d
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Inside your `dl`, add a `dt` element containing the text `Google Chrome`, followed by a `dd` element dscribing it, containing the following text:

`This is a free web browser developed by Google and first released in 2008.`

# --hints--

Your page should contain a `dl` element.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Your `dl` element should contain at least one `dt` element.

```js
const dl = document.querySelector('dl');
const dt = dl?.querySelector('dt');
assert.exists(dt);
```

Your `dl` element should contain at least one `dd` element.

```js
const dl = document.querySelector('dl');
const dd = dl?.querySelector('dd');
assert.exists(dd);
```

Your first `dt` element should mention `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.innerHTML, "Google Chrome");
```

The `dd` element should contain the text `This is a free web browser developed by Google and first released in 2008.`.

```js
assert.strictEqual(document.querySelector('dd')?.innerText, `This is a free web browser developed by Google and first released in 2008.`);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            
        <dl>
    </body> 
</html>
--fcc-editable-region--
```

