---
id: 68ecfa0f03d88990bb96a33c
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--
Create a definition list. 

Below the heading create a `<dl>` element to hold the list of browsers

# --hints--

Test 1

```js
const h1 = document.querySelector('h1');
assert.exists(h1, "Your page should have a <h1> element")
assert.include(h1.textContent, 'List of Major Web Browsers', `Your <h1> shouykd contain the text "List of Major Web Browsers".`);
```

Test 2

```js
const dl = document.querySelector('dl')
assert.exists(dl, "Your page should contain a <dl> element")
```


```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 

    </body> 
</html>

--fcc-editable-region--

--fcc-editable-region--
```

