---
id: 68ecfa0f03d88990bb96a33c
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Create a definition list. 

Below the heading create a `dl` element to hold the list of browsers

# --hints--

Your page should have a `h1` element.

```js
const h1 = document.querySelector('h1');
assert.exists(h1)
```

Your `h1` should contain the text "List of Major Web Browsers".

```js
const h1 = document.querySelector('h1');
assert.include(h1.textContent, 'List of Major Web Browsers');
```

Your page should contain a `dl` element

```js
const dl = document.querySelector('dl')
assert.exists(dl)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 

    </body> 
</html>

--fcc-editable-region--

--fcc-editable-region--
```

