---
id: 68ecfa1003d88990bb96a33f
title: Step 5
challengeType: 4
dashedName: step-5
---

# --description--
Verify semantic structure.

Make sure your document is correctly structured and valid. 
`<h1>` before `<dl>`
All `<dt>` and `<dd>` are inside `<dl>`
Properly closed HTML tags.

# --hints--

Test 1

```js
const h1 = document.querySelector('h1');
assert.exists(h1, "Your page should have a <h1> element")
assert.include(h1.textContent, 'List of Major Web Browsers', `Your <h1> shouykd contain the text "List of Major Web Browsers".`);
```

Test 2

```js
const dl = document.querySelector('dl')
assert.exists(dl, "Your page should contain a <dl> element")
```

Test 3

```js
const dl = document.queryselector('dl');
const dt = dl.querySelector('dt')
const ds = dl.querySelector('dd')

assert.exists(dt, "Your <dl> should contain at least one <dt> element")
assert.exists(dd, "Your <dl> should contain at least one <dd> element.")
assert.match(dl.innerHTML, /Google Chrome/i, `Your first <dt> should mention "Google Chrome".`)
```

Test 4

```js
const terms = document.querySelectorAll('dl dt');
assert.equal(terms.length, 5, 'Your <dl> should contain five <dt> elements.');
const descs = document.querySelectorAll('dl dd');
assert.equal(descs.length, 5, 'Your <dl> should contain five <dd> elements.');
```

Test 5

```js
const h1 = document.querySelector('h1');
assert(h1.nextElementSibling.tagName === 'DL', 'Your <h1> should be followed by a <dl> element.');
```


```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            <dt>Brave</dt>
            <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

            <dt>Arc</dt>
            <dd>This is a free Chromium based web browser first released in 2023 by The Browser Company.</dd>
        </dl>  
    </body> 
</html>

--fcc-editable-region--

--fcc-editable-region--

```

