---
id: 68ecfa0f03d88990bb96a33c
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Below the heading create a `dl` element. This will hold the list of browsers.

The `dl` element is called definition list.

# --hints--

Your page should have a `h1` element.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Your `h1` should contain the text `List of Major Web Browsers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent, 'List of Major Web Browsers');
```

Your page should contain a `dl` element.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

The `dl` element should be below the `h1` element.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 

    </body> 
</html>
--fcc-editable-region--
```

