---
id: 68ecfa1003d88990bb96a33d
title: Step 3
challengeType: 2
dashedName: step-3
---

# --description--
Add your first term and description.
Inside your `<dl>`, add a `<dt>` element for **Google Chrome**, followed by a `<dd>` element dscribing it.

> "This is a free web browser developed by Google and first released in 2008."

# --hints--

Test 1

```js
const h1 = document.querySelector('h1');
assert.exists(h1, "Your page should have a <h1> element")
assert.include(h1.textContent, 'List of Major Web Browsers', `Your <h1> shouykd contain the text "List of Major Web Browsers".`);
```

Test 2

```js
const dl = document.querySelector('dl')
assert.exists(dl, "Your page should contain a <dl> element")
```


Test 3

```js
const dl = document.queryselector('dl');
const dt = dl.querySelector('dt')
const ds = dl.querySelector('dd')

assert.exists(dt, "Your <dl> should contain at least one <dt> element")
assert.exists(dd, "Your <dl> should contain at least one <dd> element.")
assert.match(dl.innerHTML, /Google Chrome/i, `Your first <dt> should mention "Google Chrome".`)
```

# --seed--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            
        <dl>
            </body> 
</html>

--fcc-editable-region--

--fcc-editable-region--
```

