---
id: 68ecfa1003d88990bb96a33e
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Add four more `dt` and `dd` element pairs for the following browsers

| Browser     | Description                                                                                      |
| ----------- | ------------------------------------------------------------------------------------------------ |
| `Firefox` | This is a free web browser developed by the Mozilla Corporation and first created in 2004.       |
| `Safari`  | This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices. |
| `Brave`   | This is a free web browser first released in 2016 that is based on the Chromium web browser.     |
| `Arc`     | This is a free Chromium based web browser first released in 2023 by The Browser Company.         |

# --hints--

Your page should have a `h1` element

```js
const h1 = document.querySelector('h1');
assert.exists(h1)
```

Your `h1` element should contain the text "List of Major Web Browsers".

```js
const h1 = document.querySelector('h1');
assert.include(h1.textContent, 'List of Major Web Browsers');
```

Your page should contain a `dl` element

```js
const dl = document.querySelector('dl')
assert.exists(dl)
```

Your `dl` element should contain at least one `dt` element

```js
const dl = document.querySelector('dl');
const dt = dl.querySelector('dt')

assert.exists(dt)
```

Your `dl` element should contain at least one `dd` element.

```js
const dl = document.querySelector('dl');
const dd = dl.querySelector('dd')
assert.exists(dd)
```

Your first `dt` element should mention "Google Chrome".

```js
const dl = document.querySelector('dl');
assert.match(dl.innerHTML, /Google Chrome/i)
```

Your `dl` element should contain five `dt` elements.

```js
const terms = document.querySelectorAll('dl dt');
assert.equal(terms.length, 5);
```

Your `dl` element should contain five `dd` elements.

```js
const descs = document.querySelectorAll('dl dd');
assert.equal(descs.length, 5);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>
        </dl>
    </body> 
</html>


--fcc-editable-region--

```

