---
id: 68ecfa1003d88990bb96a340
title: Add Browser Entries
challengeType: 5
dashedName: add-browser-entries
---

# --description--

Add each `<dt>` with its `<dd>` description.



# --instructions--

Create a list using the `<dl>` tag.
Inside it, add `<dt>` elements for each browser name and `<dd>` elements for their descriptions.

# --seed--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>List of Browsers and Descriptions</title>
</head>
<body>
    <h1>List of Major Web Browsers</h1>
    <dl>
        
    </dl>
</body>
</html>
```

# --hints--

Test 1

```js
const dt = document.querySelector('dt');
assert.exists(dt);
```

Test 2

```js
const dd = document.querySelector('dd');
assert.exists(dd);
```

Test 3

```js
const h1 = document.querySelector('h1');
assert.exists(dl);
assert.equal(h1.htmlContent, "List of Major Web Browsers", "H1 tag contains text List of Major Web Browsers")
```

# --seed--
