---
id: 68ecfa1003d88990bb96a33f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
Verify semantic structure.

Make sure your document is correctly structured and valid. 

- `h1` element before the `dl` element.
- All `dt` and `dd` elements are inside the `dl` element.
- Properly closed HTML tags.

# --hints--

Your page should have a `h1` element.

```js
const h1 = document.querySelector('h1');
assert.exists(h1)
```

Your `h1` element should contain the text "List of Major Web Browsers".

```js
const h1 = document.querySelector('h1');
assert.include(h1.textContent);
```

Your page should contain a `dl` element.

```js
const dl = document.querySelector('dl')
assert.exists(dl)
```

Your `dl` element should contain at least one `dt` element

```js
const dl = document.querySelector('dl');
const dt = dl.querySelector('dt')

assert.exists(dt)
```

Your `dl` element should contain at least one `dd` element

```js
const dl = document.querySelector('dl');
const dd = dl.querySelector('dd')

assert.exists(dd)
```

Your first `dt` element should mention "Google Chrome".

```js
const dl = document.querySelector('dl');
assert.match(dl.innerHTML, /Google Chrome/i)
```

Your `dl` element should contain five `dt` elements.

```js
const terms = document.querySelectorAll('dl dt');
assert.equal(terms.length, 5);
```

Your `dl` element should contain five `dd` elements.

```js
const descs = document.querySelectorAll('dl dd');
assert.equal(descs.length, 5);
```

Your `h1` element should be followed by a `dl` element.

```js
const h1 = document.querySelector('h1');
assert(h1.nextElementSibling.tagName === 'DL');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
  
    </body> 
</html>

--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```html

<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            <dt>Brave</dt>
            <dd>This is a free web browser first released in 2016 that is based on the Chromium web browser.</dd>

            <dt>Arc</dt>
            <dd>This is a free Chromium based web browser first released in 2023 by The Browser Company.</dd>
        </dl>  
    </body> 
</html>

```

