---
id: 69038f5efad6f722def4a214
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now it is time to add your fourth browser to the list.

Add another `dt` element containing the text `Brave`.

Below your `dt` element, add a `dd` element with the following text:

```md
This is a free web browser first released in 2016 that is based on the Chromium web browser.
```

# --hints--

You should have four `dt` elements in your code.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 4);
```

Your fourth `dt` element should have the text `Brave`.

```js
const dt = document.querySelectorAll('dt')[3];
assert.strictEqual(dt?.textContent.trim(), "Brave");
```

You should have four `dd` elements in your code.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 4);
```

Your fourth `dd` element should contain the text `This is a free web browser first released in 2016 that is based on the Chromium web browser.`.

```js
const ddElement = document.querySelectorAll('dd')[3];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser first released in 2016 that is based on the Chromium web browser.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            <dt>Firefox</dt>
            <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

            <dt>Safari</dt>
            <dd>This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
