---
id: 5d822fd413a79914d39e98e9
title: Schritt 33
challengeType: 0
dashedName: step-33
---

# --description--

Du solltest deinen Code optimieren. Verschiebe die `position` und `top`-Eigenschaften und Werte aus `.foreground-buildings` nach `.background-buildings`. Wähle dort dann beide `.background-buildings` und `.foreground-buildings`, wobei du diese Styles effektiv auf beide Elemente anwendest. Du kannst ein Komma (`,`) verwenden, um die Selektoren wie folgt zu trennen: `selector1, selector2`.

# --hints--

Du solltest nicht die `.foreground-buildings`-Deklaration entfernen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.foreground-buildings'));
```

Du solltest die `position`-Eigenschaft aus `.foreground-buildings` entfernen.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.position);
```

Du solltest die `top`-Eigenschaft aus `.foreground-buildings` entfernen.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.top);
```

Du solltest die `position`-Eigenschaft von `absolute` zu `.background-buildings, .foreground-buildings` hinzufügen.

```js
function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle('.background-buildings, .foreground-buildings') ?? a.getStyle('.foreground-buildings, .background-buildings');
}
assert.equal(eitherOr()?.position, 'absolute');
```

Du solltest die `top`-Eigenschaft von `0` zu `.background-buildings, .foreground-buildings` hinzufügen.

```js
function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle('.background-buildings, .foreground-buildings') ?? a.getStyle('.foreground-buildings, .background-buildings');
}
assert.equal(eitherOr()?.top, '0px');
```

Du solltest ein Komma benutzen, um beide `.foreground-buildings` und `.background-buildings`-Selektoren in der gleichen Stil-Deklaration zu verwenden.

```js
function eitherOr() {
  const a = new __helpers.CSSHelp(document)
  return a.getStyle('.background-buildings, .foreground-buildings') ?? a.getStyle('.foreground-buildings, .background-buildings');
}
assert.exists(eitherOr());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--
.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}
--fcc-editable-region--
.fb1 {
  width: 10%;
  height: 60%;
}

.fb2 {
  width: 10%;
  height: 40%;
}

.fb3 {
  width: 10%;
  height: 35%;
}

.fb4 {
  width: 8%;
  height: 45%;
}

.fb5 {
  width: 10%;
  height: 33%;
}

.fb6 {
  width: 9%;
  height: 38%;
}

```

