---
id: 5d822fd413a79914d39e98d6
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Jetzt hast du etwas, das einem Bauelement ähnelt. Du bist bereit dafür, deine erste Variable zu erstellen. Variable-Deklarationen beginnen mit zwei Bindestrichen (`-`) und erhalten einen Namen und einen Wert wie diese: `--variable-name: value;`. Erstelle in der Regel deiner `bb1`-Klasse eine Variable namens `--building-color1` und gib ihr den Wert `#999`.

# --hints--

Du solltest eine neue Variable namens `--building-color1` erstellen.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

Du solltest die `--building-color1`-Variable innerhalb `.bb1` definieren.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

Du solltest `--building-color1` einen Wert von `#999` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
