---
id: 5d822fd413a79914d39e98d0
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Erstelle ein `div`-Element in dem `body` mit einer Klasse von `background-buildings`. Dies wird ein Container für eine Gruppe von Bauelementen sein.

# --hints--

Du solltest ein `div`-Element erstellen.

```js
assert.exists(document.querySelector('div'));
```

Dein `div`-Element sollte sich innerhalb des `body` befinden.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Dein `div`-Element sollte eine Klasse von `background-buildings` haben

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

