---
id: 5d822fd413a79914d39e98e7
title: Schritt 31
challengeType: 0
dashedName: step-31
---

# --description--

Gib den sechs neuen Elementen folgende `width` und `height`-Werte: `10%` und `60%` an `.fb1`, `10%` und `40%` an `.fb2`, `10%` und `35%` an `.fb3`, `8%` und `45%` an `.fb4`, `10%` und `33%` an `.fb5`, und `9%` und `38%` an `.fb6`.

# --hints--

Du solltest einen `.fb1`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

Du solltest dem `.fb1`-Selektor eine `width` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

Du solltest dem `.fb1`-Selektor eine `height` von `60%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

Du solltest einen `.fb2`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

Du solltest dem `.fb2`-Selektor eine `width` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

Du solltest dem `.fb2`-Selektor eine `height` von `40%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

Du solltest einen `.fb3`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

Du solltest dem `.fb3`-Selektor eine `width` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

Du solltest dem `.fb3`-Selektor eine `height` von `35%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

Du solltest einen `.fb4`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

Du solltest dem `.fb4`-Selektor eine `width` von `8%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

Du solltest dem `.fb4`-Selektor eine `height` von `45%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

Du solltest einen `.fb5`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

Du solltest dem `.fb5`-Selektor eine `width` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

Du solltest dem `.fb5`-Selektor eine `height` von `33%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

Du solltest einen `.fb6`-Selektor erstellen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

Du solltest dem `.fb6`-Selektor eine `width` von `9%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

Du solltest dem `.fb6`-Selektor eine `height` von `38%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--

```

