---
id: 5d822fd413a79914d39e98cb
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Füge als nächstes öffnende und abschließende `head` und `body`-Tags innerhalb des `html`-Elements hinzu.

# --hints--

Du solltest einen öffnenden `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest einen abschließenden `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest einen öffnenden `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest einen abschließenden `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Das `head` und `body`-Element sollten Geschwister sein.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">

</html>
--fcc-editable-region--

```

