---
id: 5d822fd413a79914d39e98cf
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Du kannst den `body` (die innerste Box deiner Seite) sehen; die Box um ihn herum ist das `html`-Element. Lasse deinen `body` das gesamte Anzeigefenster füllen, indem du ihm eine `height` von `100vh` gibst. Entferne die voreingestellte `margin` aus dem `body`, indem du `margin` auf `0` setzt. Setzte schließlich die `overflow`-Eigenschaft auf `hidden`, um jegliche Scroll-Leisten zu verstecken, die erscheinen, wenn etwas über das Anzeigefenster hinausgeht.

# --hints--

Du solltest den `body`-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Dein `body` sollte eine `height` von `100vh` haben.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Dein `body` sollte eine `margin` von `0` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Dein `body` sollte eine `overflow`-Eigenschaft auf `hidden` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

