---
id: 5d822fd413a79914d39e98cc
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Niste innerhalb des `head` ein `meta`-Element mit einem `charset` von `UTF-8` ein, ein `title`-Element mit dem Titel `City Skyline`, und ein `link`-Element, das deine `styles.css` Datei verlinkt.

# --hints--

Du solltest ein `meta`-Element innerhalb des `head`-Elements erstellen.

```js
assert.exists(document.querySelector('head > meta'));
```

Du solltest dem `meta`-Tag ein `charset` von `UTF-8` geben.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Dein Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Das `title`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > title'));
```

Dein Projekt sollte den Titel `City Skyline` tragen.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Dein Code sollte ein `link`-Element haben.

```js
assert.match(code, /<link/)
```

Du solltest die vorhandenen `head`-Tags nicht ändern. Stelle sicher, dass du den abschließenden Tag nicht gelöscht hast.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Du solltest ein selbstschießendes `link`-Element haben.

```js
assert(document.querySelectorAll('link').length === 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
