---
id: 5d822fd413a79914d39e98df
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Das hat nicht funktioniert. Du solltest einen Fallback-Wert zu einer Variable hinzuf端gen, indem du ihn als zweiten Wert angibst, wobei du die Variable wie folgt verwenden w端rdest: `var(--variable-name, fallback-value)`. Die Eigenschaft wird den Fallback-Wert verwenden, wenn es ein Problem mit der Variable gibt. F端ge zur `background-color` von `.bb2` einen Fallback-Wert von `green` hinzu.

# --hints--

Du solltest einen Fallback-Wert von `green` zur `background-color`-Eigenschaft hinzuf端gen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.getPropVal('background-color', true), 'var(--building-color2,green)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}
--fcc-editable-region--
.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}
--fcc-editable-region--
.bb3 {
  width: 10%;
  height: 55%;
}

.bb4 {
  width: 11%;
  height: 58%;
}

```
