---
id: 5d822fd413a79914d39e98cd
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

In CSS kannst du alles mit einem Sternchen anvisieren. Füge einen Rand zu allem hinzu, indem du den `*`-Selekor verwendest, und gib ihm einen `border` von `1px solid black`. Dies ist ein Trick, der hilft den Standort und die Größe der Elemente zu visualisieren. Du wirst dies später entfernen.

# --hints--

Du solltest den `*`-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Du solltest die `border`-Eigenschaft verwenden, um alle Elemente zu gestalten.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Alle Elemente sollten einen `1px solid black`-Rand haben.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
