---
id: 5d822fd413a79914d39e98e1
title: Schritt 25
challengeType: 0
dashedName: step-25
---

# --description--

Das funktionierte nicht, weil die Variablen, die du in `.bb1` deklariert hast, nicht auf Geschwister-Elemente von `.bb2` und `.bb3` angewendet werden. So funktioniert CSS nun einmal. Deswegen werden Variablen oft im `:root`-Selektor deklariert. In CSS ist dies der Selektor mit dem höchsten Level; wenn du deine Variablen dort platzierst, sind sie überall verwendbar. Füge einen `:root`-Selektor oben auf deinem Stylesheet hinzu und verschiebe dort alle Deklarationen deiner Variablen.

# --hints--

Du solltest oben auf dem Stylesheet einen `:root`-Selektor deklarieren.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(':root'));
```

Du solltest `--building-color1` mit einem Wert von `#aa80ff` im `:root`-Selektor definieren.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color1')?.trim(), '#aa80ff');
```

Du solltest `--building-color2` mit einem Wert von `#66cc99` im `:root`-Selektor definieren.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color2')?.trim(), '#66cc99');
```

Du solltest `--building-color3` mit einem Wert von `#cc6699` im `:root`-Selektor definieren.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle(':root')?.getPropertyValue('--building-color3')?.trim(), '#cc6699');
```

Du solltest die benutzerdefinierte Eigenschaftsvariablen aus `.bb1` entfernen.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color2'));
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2, green);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3, pink);
}

.bb4 {
  width: 11%;
  height: 58%;
}

```

