---
id: 5d822fd413a79914d39e98d3
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Niste vier `div` Elemente in den `.bb1` Container. Gib ihnen die Klassen `bb1a`, `bb1b`, `bb1c` und `bb1d` in dieser Reihenfolge. Dieses Bauelement wird aus vier Teilen bestehen.

# --hints--

Du solltest vier neue `div`-Elemente erstellen.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Du solltest einem der neuen `div`-Elemente eine Klasse von `bb1a` geben.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Du solltest einem der neuen `div`-Elemente eine Klasse von `bb1b` geben.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Du solltest einem der neuen `div`-Elemente eine Klasse von `bb1c` geben.

```js
assert.exists(document.querySelector('div.bb1c'));
```

You should give one of the new `div` elements a class of `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Du solltest die neuen `div`-Elemente in der richtigen Reihenfolge platzieren.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Du solltest die neuen `div`-Elemente innerhalb des `.bb1`-Elements platzieren.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

