---
id: 5d822fd413a79914d39e98d1
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Gib deinem `.background-buildings`-Element eine `width` und `height` von `100%`, um die volle Breite und Höhe ihres übergeordneten Elements, dem `body`, einzunehmen.

# --hints--

Du solltest die `background-buildings`-Klasse zur Auswahl des richtigen Elements verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

Dein `.background-buildings`-Element sollte eine `width` von `100%` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

Dein `.background-buildings`-Element sollte eine `height` von `100%` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

