---
id: 5d822fd413a79914d39e98d4
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Gib den Teilen deiner Bauelemente `width` und `height`-Eigenschaften mit diesen Werten: `70%` und `10%` von `.bb1a`, `80%` und `10%` von `.bb1b`, `90%` und `10%` von `.bb1c`, und `100%` und `70%` von `.bb1d`. Denke daran, dass diese Prozentsätze zum übergeordneten Element verwandt sind und notiere, dass die Höhen bis zu 100% ausmachen - vertikal den Container füllen.

# --hints--

Du solltest einen Klassen-Selektor verwenden, um `.bb1a` zu gestalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Du solltest `.bb1a` eine `width` von `70%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Du solltest `.bb1a` eine `height` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Du solltest einen Klassen-Selektor verwenden, um `.bb1b` zu gestalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Du solltest `.bb1b` eine `width` von `80%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Du solltest `.bb1b` eine `height` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Du solltst einen Klassen-Selektor verwenden, um `.bb1c` zu gestalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Du solltest `.bb1c` eine `width` von `90%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Du solltest `.bb1c` eine `height` von `10%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Du solltest einen Klassen-Selektor verwenden, um `.bb1d` zu gestalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Du solltest `.bb1d` eine `width` von `100%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Du solltest `.bb1d` eine `height` von `70%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

