---
id: 6140c7e645d8e905819f1dd4
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Comece com o boilerplate padrão. Adicione a declaração `DOCTYPE` e os elementos `html` (com o idioma definido para o inglês), `head` e `body`.

Adicione seu elemento `meta` com o `charset` correto, o elemento `title` e um elemento `link` que vincule o arquivo `./styles.css`.

Defina o título `title` para `Ferris Wheel`.

# --hints--

O código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento para `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

O elemento `html` deve ter uma tag de abertura com um atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

Sua declaração `DOCTYPE` deve ser no início do seu HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Você deve acrescentar uma tag de encerramento ao seu `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Você deve acrescentar uma tag de abertura para o elemento `body`.

```js
assert(code.match(/<body\s*>/i));
```

Você deve acrescentar uma tag de fechamento para o elemento `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

O elemento `head` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

O código deve ter dois elementos `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

O elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O projeto deve ter o título `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

Você deve ter um elemento `link` de autofechamento.

```js
assert(document.querySelectorAll('link').length === 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
