---
id: 6140c7e645d8e905819f1dd4
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Comece com o boilerplate padrão. Adicione a declaração `DOCTYPE` e os elementos `html` (com o idioma definido para o inglês), `head` e `body`.

Adicione seu elemento `meta` com o `charset` correto, o elemento `title` e um elemento `link` que vincule o arquivo `./styles.css`.

Defina o título `title` para `Ferris Wheel`.

# --hints--

O código deve conter a declaração de `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Você deve incluir um espaço após a declaração de `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Você deve definir o tipo de documento para `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

O elemento `html` deve ter uma tag de abertura com um atributo `lang` de `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert.match(code, /<\/html\s*>/);
```

Sua declaração `DOCTYPE` deve ser no início do seu HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert.match(code, /<head\s*>/i);
```

Você deve acrescentar uma tag de encerramento ao seu `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Você deve acrescentar uma tag de abertura para o elemento `body`.

```js
assert.match(code, /<body\s*>/i);
```

Você deve acrescentar uma tag de fechamento para o elemento `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

O elemento `body` deve ser o primeiro elemento depois do elemento `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

O elemento `head` deve ser o primeiro elemento dentro do elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

O elemento `head` deve estar dentro do elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

O código deve ter dois elementos `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

O elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

O elemento `meta` deve estar dentro do elemento `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O projeto deve ter o título `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/);
```

Você deve ter um elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
