---
id: 6140cbeec34e970dfe75e710
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Crie um seletor para o elemento `.wheel`. Comece definindo a `border` para `2px solid black`, o `border-radius` para `50%` e a `margin-left` para `50px`.

# --hints--

Você deve ter um seletor `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

O seletor `.wheel` deve ter a propriedade `border` definida para `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

O seletor `.wheel` deve ter a propriedade `border-radius` definida para `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

O seletor `.wheel` deve ter a propriedade `margin-left` definida para `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
