---
id: 6140d36b8b747718b50d4b7a
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Defina o `.cabin` para ter uma propriedade `transform-origin` de `50% 0%`. Isso definirá que o ponto de origem seja deslocado `50%` da esquerda e `0%` a partir de cima, posicionando-o no meio da borda superior do elemento.

# --hints--

O seletor `.cabin` deve ter a propriedade `transform-origin` definida para `50% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.cabin')?.transformOrigin;
assert(transformOrigin === '50% 0%' || transformOrigin === '50% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

--fcc-editable-region--
.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
}
--fcc-editable-region--
```
