---
id: 6140c9d35015ae0ba0c250e8
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione uma `div` dentro do elemento `body` e dê a ela uma `class` com o valor `wheel`.

Dentro do novo `div`, adicione seis elementos `span` com uma `class` definida como `line` e seis elementos `div` com uma `class` definida para `cabin`.

# --hints--

Você deve criar um elemento `div` dentro de `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

Adicione um `div` no elemento `body` e dê a ele uma `class` com o valor `wheel`.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

O elemento `.wheel` deve ter seis elementos `span` dentro dele.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

Os elementos `span` dentro do elemento `.wheel` devem ter uma `class` definida como `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

O elemento `.wheel` deve ter seis elementos `div` dentro dele.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

Seu seis elementos `div` dentro do elemento `.wheel` deve ter uma `class` definida como `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
