---
id: 6140d0069049f5139d78da40
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

A propriedade `transform-origin` é usada para definir o ponto em torno do qual uma transformação CSS é aplicada. Por exemplo, ao executar um `rotate` (o que você fará mais tarde neste projeto), o `transform-origin` determina em torno de qual ponto o elemento é girado.

Defina o seletor `.line` para ter uma `transform-origin` de `0% 0%`. Isso definirá que o ponto de origem seja deslocado `0%` da esquerda e `0%` a partir de cima, posicionando-o no meio da borda superior do elemento.

# --hints--

O seletor `.line` deve ter a propriedade `transform-origin` definida para `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
