---
id: 6140d1a351e88f159ed54fca
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Usando o mesmo padrão, crie um seletor separado para a terceira `.line`, quarta `.line`, quinta `.line` e sexta `.line`.

Defina a propriedade `transform` para a terceira `.line` para `rotate(120deg)`, para a quarta para `rotate(180deg)`, para a quinta para `rotate(240deg)` e para a sexta para `rotate(300deg)`.

# --hints--

Você deve criar um seletor `.line:nth-of-type(3)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

O seletor `.line:nth-of-type(3)` deve ter a propriedade `transform` definida para `rotate(120deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform === 'rotate(120deg)');
```

Você deve criar um seletor `.line:nth-of-type(4)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

O seletor `.line:nth-of-type(4)` deve ter a propriedade `transform` definida para `rotate(180deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform === 'rotate(180deg)');
```

Você deve criar um seletor `.line:nth-of-type(5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

O seletor `.line:nth-of-type(5)` deve ter a propriedade `transform` definida para `rotate(240deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform === 'rotate(240deg)');
```

Você deve criar um seletor `.line:nth-of-type(6)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

O seletor `.line:nth-of-type(6)` deve ter a propriedade `transform` definida para `rotate(300deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform === 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
