---
id: 6140d4bc9db3c81c51a09ab7
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Continuando o padrão, selecione os elementos `.cabin` a seguir e aplique as regras específicas a eles:

- O segundo `.cabin` deve ter a propriedade `right` definida como `17%` e a propriedade `top` definida como `93.5%`.
- O terceiro `.cabin` deve ter a propriedade `right` definida como `67%` e a propriedade `top` definida como `93.5%`.
- O quarto `.cabin` deve ter a propriedade `left` definida como `-8.5%` e a propriedade `top` definida como `50%`.
- O quinto `.cabin` deve ter a propriedade `left` definida como `17%` e a propriedade `top` definida como `7%`.
- O sexto `.cabin` deve ter a propriedade `right` definida como `17%` e a propriedade `top` definida como `7%`.

# --hints--

Você deve criar um seletor `.cabin:nth-of-type(2)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

O seletor `.cabin:nth-of-type(2)` deve ter a propriedade `right` definida para `17%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right === '17%');
```

O seletor `.cabin:nth-of-type(2)` deve ter a propriedade `top` definida para `93.5%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top === '93.5%');
```

Você deve criar um seletor `.cabin:nth-of-type(3)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

O seletor `.cabin:nth-of-type(3)` deve ter a propriedade `right` definida para `67%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right === '67%');
```

O seletor `.cabin:nth-of-type(3)` deve ter a propriedade `top` definida para `93.5%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top === '93.5%');
```

Você deve criar um seletor `.cabin:nth-of-type(4)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

O seletor `.cabin:nth-of-type(4)` deve ter a propriedade `left` definida para `-8.5%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left === '-8.5%');
```

O seletor `.cabin:nth-of-type(4)` deve ter a propriedade `top` definida para `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top === '50%');
```

Você deve criar um seletor `.cabin:nth-of-type(5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

O seletor `.cabin:nth-of-type(5)` deve ter a propriedade `left` definida para `17%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left === '17%');
```

O seletor `.cabin:nth-of-type(5)` deve ter a propriedade `top` definida para `7%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top === '7%');
```

Você deve criar um seletor `.cabin:nth-of-type(6)`.

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

O seletor `.cabin:nth-of-type(6)` deve ter a propriedade `right` definida para `17%`.

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

O seletor `.cabin:nth-of-type(6)` deve ter a propriedade `top` definida para `7%`.

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
