---
id: 6140cdebd39d6a101e747529
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Crie um seletor para os elementos `.line`. Comece definindo a `background-color` para `black`, `width` para `50%` e `height` para `2px`.

# --hints--

Você deve ter um seletor `.line`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

O seletor `.line` deve ter a propriedade `background-color` definida como `black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

O seletor `.line` deve ter a propriedade `width` definida para `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

O seletor `.line` deve ter a propriedade `height` definida como `2px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
