---
id: 65687e069cf4e0c85b0a06b1
title: Step 30
challengeType: 20
dashedName: step-30
---

# --description--

Outside of the `if` statement, add `number` to `sum_of_even_digits`. Also, remove the `print` call.

# --hints--

You should have `sum_of_even_digits += number` within the `for` loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const for_loop = __helpers.python.getBlock("\n" + transformedCode, "for digit in even_digits");
        const {block_body} = for_loop;

        assert.match(block_body, /sum_of_even_digits\s*\+=\s*number/);
    }
})
```

You should not have `print(number)` within the `for` loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\s*\(\s*number\s*\)/);
    }
})
```

You should not have `sum_of_even_digits += number` within the `if` statement.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const if_statement = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const {block_body} = if_statement;

        assert.notMatch(block_body, /sum_of_even_digits *\+= *number/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
            print(number)
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
