---
id: 656879613a96aba21cbe80b9
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Within your `main` function, call the `verify_card_number` function and pass in the `translated_card_number` variable as an argument.

# --hints--

You should have `verify_card_number(translated_card_number)` within the `main` function.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +verify_card_number\(\s*translated_card_number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

--fcc-editable-region--
main()
```
