---
id: 60a3e3396c7b40068ad69972
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Ongeza sifa ya `height` yenye thamani `600px` kwa sheria yako ya `.canvas`.

# --hints--

Unapaswa kuweka sifa ya `height` kuwa `600px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

Kipengele chako cha `.canvas` kinapaswa kuwa na `height` ya `600px`.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
