---
id: 60a3e3396c7b40068ad69975
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Andika sheria mpya kwa kutumia kichaguzi cha darasa cha `.frame`.

Tumia tamko la mkato `border` ili kukipa kipengele cha `.frame` mpaka thabiti na mweusi wenye upana wa `50px`.

# --hints--

Msimbo wako unapaswa kuwa na kichaguzi cha `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Unapaswa kuweka sifa ya `border` kuwa `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Kipengele chako cha `.frame` kinapaswa kuwa na `border` ya `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
