---
id: 60a3e3396c7b40068ad69970
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Kabla ya kuanza kutengeneza `div` uliyoongeza, unahitaji kuunganisha CSS yako kwenye HTML yako.

Ongeza kipengele cha `link` ili kuunganisha faili yako ya `styles.css`. Weka `href` kuwa `styles.css`, na kumbuka kuweka sifa ya `rel` kuwa `stylesheet`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert(/<link/.test(code))
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya `rel` yenye thamani ya `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
