---
id: 60a3e3396c7b40068ad6998b
title: Hatua ya 33
challengeType: 0
dashedName: step-33
---

# --description--

Inasaidia kuweka kando ukingo wako kwa kuusukuma kwa mwelekeo mmoja.

Katika hali hii, ukingo wa chini wa kipengele cha `.one` husukuma `.two` chini pikseli 20.

Katika kichaguzi cha `.two`, tumia sifa ya mkato ya `margin` kuweka ukingo wa juu kuwa `0`, ukingo mlalo hadi `auto`, na ukingo wa chini hadi `20px`. Hii itaondoa ukingo wake wa juu, na kuiweka katikati mwa mlalo, na kuweka ukingo wake wa chini hadi pikseli 20.

# --hints--

Unapaswa kuweka sifa ya `margin` kuwa `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

Kipengele chako cha `.two` kinapaswa kuwa na thamani ya `margin` ya `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
