---
id: 60a3e3396c7b40068ad6997f
title: Hatua ya 22
challengeType: 0
dashedName: step-22
---

# --description--

Kuongeza pikseli 1 ya padding ya juu, chini, kushoto na kulia kwa turubai kulibadilisha vipimo vyake hadi pikseli 502 x 602.

Badilisha sifa ya `padding` na `overflow` iliyowekwa kuwa `hidden` - kubadilisha turubai hadi vipimo vyake asili.

# --hints--

Unapaswa kuondoa sifa ya `padding` kutoka kwa kichagusi cha `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Unapaswa kuweka sifa ya `overflow` kuwa `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

Kipengele chako cha `.canvas` kinapaswa kuwa na thamani ya `overflow` ya `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
