---
id: 60a3e3396c7b40068ad6997e
title: Hatua ya 21
challengeType: 0
dashedName: step-21
---

# --description--

Sasa `.one` imewekwa katikati kwa mlalo, lakini ukingo wake wa juu unasukuma turubai na kuingia kwenye mpaka wa fremu, ikihamisha turubai nzima chini ya pikseli 20.

Ongeza `padding` ya `1px` kwenye kipengele cha `.canvas` ili kukipa kipengele cha `.one` kitu thabiti cha kukisukuma.

# --hints--

Unapaswa kuweka sifa ya `padding` kuwa `1px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

Kipengele chako cha `.canvas` kinapaswa kuwa na thamani ya `padding` ya `1px`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
