---
id: 60a3e3396c7b40068ad6996d
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Maudhui yamezingirwa na nafasi inayoitwa padding, sawa na jinsi ufunikaji wa viputo unavyotenganisha kipengee kutoka kwa kisanduku kinachokizunguka.

Chukulia mpaka kama sanduku la kadibodi la bidhaa yako iliyosafirishwa ndani.

Badilisha sifa ya `src` iwe `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

Kipengele chako cha `img` kinafaa kuwa na `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
