---
id: 60a3e3396c7b40068ad69978
title: Hatua ya 15
challengeType: 0
dashedName: step-15
---

# --description--

Tumia pambizo (margins) kurekebisha nafasi nje ya kipengele.

Kwa kutumia kipengele cha `margin`, kipe kipengele cha `.frame` ukingo wima wa `20px`, na ukingo mlalo wa `auto`. Hii itasogeza fremu chini kwa pikseli 20 na kuiweka katikati kwa usawa kwenye ukurasa.

# --hints--

Unapaswa kuweka sifa ya `margin` kuwa `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Kipengele chako cha `.frame` kinapaswa kuwa na thamani ya `margin` ya `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
