---
id: 60a3e3396c7b40068ad6997d
title: Hatua ya 20
challengeType: 0
dashedName: step-20
---

# --description--

Tumia pambizo ili kuweka kipengele cha `.one` kwenye turubai (canvas).

Ongeza kipengele cha mkato cha `margin` chenye ukingo wima (vertical margin) wa `20px` na ukingo mlalo (horizontal margin) wa `auto`.

# --hints--

Unapaswa kuweka sifa ya `margin` kuwa `20px auto`.

```js
const marginFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.margin === '20px auto');
assert(marginFilter.length === 2);
```

Kipengele chako cha `.one` kinapaswa kuwa na thamani ya `margin` ya `20px auto`.

```js
const oneMargin = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('margin');
assert(oneMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
