---
id: 60a3e3396c7b40068ad6996f
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Ongeza kipengele cha `div` katika `body`.

Weka sifa ya `class` sawa na `canvas`. Kwa mfano, `<div class="my-div">`.

Hii itafanya kazi kama turubai(canvas) ya mchoro wako.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `div`.

```js
assert(document.querySelectorAll('div').length === 1)
```

Kipengele chako cha `div` kinafaa kuwa na sifa ya `class` yenye thamani ya `canvas`.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
