---
id: 60a3e3396c7b40068ad6996a
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Kufikia sasa, unapaswa kufahamu vipengele vya msingi ambavyo ukurasa wa HTML unapaswa kuwa navyo.

Weka msimbo wako ya kwa tamko la `DOCTYPE`, kipengele cha `html` chenye lugha iliyowekwa kwa Kiingereza, kipengele cha `head` na kipengele cha `body`.

# --hints--

Msimbo wako unapaswa kuwa na tamko la `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Unapaswa kuwa na tagi ya ufunguzi ya `<html>` yenye sifa ya `lang` ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Unapaswa kuwa na kipengele cha `head` ndani ya kipengele chako cha `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Msimbo wako unapaswa kuwa na kipengele cha `body` ndani ya kipengele cha `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

Kipengele chako cha `head` kinapaswa kuja kabla ya kipengele chako cha `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
