---
id: 60a3e3396c7b40068ad6996c
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Katika kisanduku cha CSS, kila kipengele cha HTML kinachukuliwa kama kisanduku chenye maeneo manne.

Waza unapokea kisanduku kutoka kwa muuzaji wako unayempenda mtandaoni -- maudhui ni bidhaa iliyo kwenye kisanduku, au kwa upande wetu, kichwa, aya, au kipengele cha picha.

Badilisha sifa ya `src` katika `<img>` kutoka `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` hadi `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Kipengele chako kipya cha `img` kinafaa kuwa na `src` ya `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

