---
id: 660f46b9c417a8341729a3ab
title: Step 88
challengeType: 1
dashedName: step-88
---

# --description--

When `done` has reached the value of `count`, we want the loop to stop executing.

Inside your `if` body, assign the boolean `false` to your `continueLoop` variable.

# --hints--

Your `if` body should assign `false` to your `continueLoop` variable.

```js
assert.match(code, /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{\s*continueLoop\s*=\s*false;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done === count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
