---
id: 6610b8017d1671f2814e8c77
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

When you declare a variable without initializing it, it is considered <dfn>uninitialized</dfn>. Currently, your `secondCharacter` variable is uninitialized.

Add a `console.log()` to see what the value of your `secondCharacter` variable is.

# --hints--

You should add a second `console.log`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

You should log your `secondCharacter` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
