---
id: 6610b9f7619764fad5fd516d
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

You can also assign the value of a variable to another variable. Por exemplo:

```js
let first = "One";
let second = "Two";
second = first;
```

The `second` variable would now have the value `"One"`.

To see this in action, change your `secondCharacter` assignment from `"Test"` to your `character` variable.

Then open the console to see what gets logged.

# --hints--

You should not assign the value `"Test"` to your `secondCharacter` variable.

```js
assert.notEqual(secondCharacter, "Test");
```

You should assign the value of the `character` variable to your `secondCharacter` variable. Don't forget your semi-colon.

```js
assert.match(code, /secondCharacter\s*=\s*character;/);
```

Your `secondCharacter` variable should now have the value `"World"`.


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
