---
id: 6738a5c704680355e29e077c
title: Step 30
challengeType: 1
dashedName: step-30
---

# --description--

As you already learned, the `delete` method is used to remove an entry from a map. Sets have a `delete` method the you can use to remove an element from them as well.

Sometimes you may need to remove all elements from a map or a set. The `clear` method does exactly that job.

Call the `clear` method on `plantsSet` to remove all the elements from it.

# --hints--

You should call the `clear` method on `plantsSet`.

```js
assert.match(__helpers.removeJSComments(code), /plantsSet\s*\.\s*clear\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}

const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });
catalog.set(willowVale, { small: 3, medium: 5, large: 0 });
catalog.set(hidcote, { small: 33, medium: 13, large: 18 });
catalog.set(imperialGem, { small: 19, medium: 35, large: 28 });
catalog.set(royalCrown, { small: 40, medium: 22, large: 9 });

const sellPlants = (plant, size, potsNo) => {
    if (!catalog.has(plant)) return "Item not found.";
    const name = `${plant.scientificName} '${plant.cultivar}'`
    const pots = catalog.get(plant);
    if (pots[size] - potsNo < 0) {
        return `Not enough ${size} size pots for ${name}. Only ${pots[size]} left.`
    }
    pots[size] -= potsNo;
    return `Catalog successfully updated.`
}

const removePlant = plant => catalog.delete(plant);

const displayCatalog = () => {
    let catalogString = "";
    catalog.forEach((val, key) => {
        catalogString += `${key.scientificName} '${key.cultivar}': ${val.small} S, ${val.medium} M, ${val.large} L
`
    })
    return catalogString
}

const displayPlantsSet = () => {
    const catalogSet = new Set();
    catalogSet.add(ballerina);
    catalogSet.add(ballerina);
    catalogSet.add(prettyPolly);
    return catalogSet
};
--fcc-editable-region--
const plantsSet = displayPlantsSet();

console.log(plantsSet.has(ballerina));
console.log(plantsSet);
--fcc-editable-region--
```
