---
id: 5d5a813321b9e3db6c106a46
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript – це потужна мова, яка дозволяє створювати інтерактивні вебсайти.

*Note*: For all remaining projects in this curriculum, you will need a basic level of knowledge in HTML and CSS. If you are new to HTML and CSS, please go through the [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Для початку створіть стандартний шаблон HTML з `DOCTYPE`, `html`, `head` та `body`, а потім додайте тег `meta` для `charset`. Додайте елемент `title` та використайте для нього текст `RPG - Dragon Repeller`. Не забудьте тег `link` для таблиці стилів, щоб додати файл `styles.css`.

Наприкінці створіть елемент `div`, що має `id` зі значенням `game`, в межах `body`.

# --hints--

Код повинен містити `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Додайте пробіл після `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Визначте тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Закрийте оголошення `DOCTYPE` за допомогою `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Елемент `html` повинен мати початковий тег. Не забудьте атрибут `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Елемент `html` повинен мати кінцевий тег.

```js
assert(code.match(/<\/html\s*>/));
```

Оголошення `DOCTYPE` має бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви повинні мати початковий тег `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий тег `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий тег `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий тег `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` та `body` мають бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Елемент `head` має бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` має бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Код повинен мати елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ви повинні мати кінцевий тег `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Код повинен мати елемент `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Ви повинні мати елемент `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Елемент `div` повинен мати атрибут `id` зі значенням `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Елемент `div` має бути в межах елемента `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
