---
id: 62a1166ed9a56d439c0770e7
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Створіть три елементи `span` в межах елемента `#stats`. Надайте кожному з них клас `stat`. Потім надайте першому текст `XP: 0`, другому – `Health: 100`, а третьому – `Gold: 50`.

# --hints--

Ви повинні мати три елементи `span` в межах елемента `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Надайте трьом новим елементам `span` клас зі значенням `stat`.

```js
assert.exists(document.querySelectorAll('#stats > .stat')?.length, 3);
```

Перший елемент `.stat` повинен мати наданий текст `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Другий елемент `.stat` повинен мати наданий текст `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Третій елемент `.stat` повинен мати наданий текст `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
