---
id: 62a115879a6d51422652cbfc
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Створіть чотири елементи `div` у своєму елементі `#game`. Надайте їм наступні значення `id`, у такому порядку: `stats`, `controls`, `monsterStats` та `text`.

# --hints--

Ви повинні створити чотири нові елементи `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Одному з нових елементів `div` надайте `id` зі значенням `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Одному з нових елементів `div` надайте `id` зі значенням `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Одному з нових елементів `div` надайте `id` зі значенням `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Одному з нових елементів `div` надайте `id` зі значенням `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Ви повинні розмістити нові елементи `div` в правильному порядку.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Ви повинні розмістити нові елементи `div` в межах елемента `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
--fcc-editable-region--
<body>
  <div id="game">
  </div>
</body>
--fcc-editable-region--
</html>
```
