---
id: 62a23d1c5f1c93161f3582ae
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Схоже до елемента `#stats`, вашому елементу `#monsterStats` потрібно два елементи `span`. Надайте їм клас `stat` та надайте першому елементу текст `Monster Name:`, а другому надайте текст `Health:`. Після тексту додайте елемент `strong` з пустим вкладеним елементом `span`.

# --hints--

Ваш елемент `#monsterStats` повинен мати два елементи `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

Ваші нові елементи `span` повинні мати `class` зі значенням `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Ваш перший елемент `span` повинен мати текст `Monster Name:`. Переконайтеся, що інтервал правильний.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].innerText, 'Monster Name: ');
```

Ваш другий елемент `span` повинен мати текст `Health:`. Переконайтеся, що інтервал правильний.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[1].textContent, 'Health: ');
```

Ваш перший елемент `span` повинен мати елемент `strong` із пустим вкладеним елементом `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Ваш другий елемент `span` повинен мати елемент `strong` із пустим вкладеним елементом `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Елементи `strong` та `span` повинні йти після тексту.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name: <strong>/);
assert.match(spans[1].innerHTML, /Health: <strong>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
--fcc-editable-region--
        <div id="monsterStats">

        </div>
--fcc-editable-region--
        <div id="text"></div>
    </div>
</body>
</html>
```
