---
id: 62a255dae245b52317da824a
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Одним із найпотужніших інструментів є консоль розробника. Її можна відкрити за допомогою `F12` або `Ctrl+Shift+I`, залежно від вашого браузера. На Mac ви можете натиснути `Option + ⌘ + C` та обрати «Консоль». Ви також можете натиснути кнопку «Консоль» над вікном попереднього перегляду, щоб побачити нашу вбудовану консоль.

Консоль розробника містить помилки, створені вашим кодом, але ви також можете використовувати її, щоб переглянути значення змінних у своєму коді, що корисно для налагодження.

Додайте рядок `console.log("Hello World");` між тегами `script`. Потім натисніть кнопку «Консоль», щоб відкрити консоль. You should see the text `"Hello World"`.


# --hints--

Додайте рядок `console.log("Hello World");` до коду.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

Рядок `console.log("Hello World");` має бути між тегами `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
