---
id: 62a2509ba163e020bb9d84ea
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Тепер ви можете почати писати свій JavaScript. Розпочніть зі створення елемента `script`. Цей елемент використовують, щоб завантажити JavaScript у файл HTML.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Ви повинні мати елемент `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

Елемент `script` повинен мати початковий тег.

```js
assert.match(code, /<script\s*>/i);
```

Елемент `script` повинен мати кінцевий тег.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
