---
id: 62a3b41c9494f937560640ab
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Оголосіть ще дві змінні `monsterHealth` та `inventory`, але не ініціалізовуйте їх.

# --hints--

Використайте `let`, щоб оголосити змінну `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Використайте верблюдячийРегістр, щоб назвати змінну `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` не повинна мати значення.

```js
assert.isUndefined(monsterHealth);
```

Не присвоюйте значення змінній `monsterHealth`. Не забудьте про крапку з комою.

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

Використайте `let`, щоб оголосити змінну `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` не повинна мати значення.

```js
assert.isUndefined(inventory);
```

Не присвоюйте значення змінній `inventory`. Не забудьте про крапку з комою.

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
