---
id: 62a3b3eab50e193608c19fc6
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Declare a variable called `fighting` but do not initialize it with a value. Не забудьте поставити крапку з комою в кінці рядка.

# --hints--

Використайте `let`, щоб оголосити змінну `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

Змінна `fighting` не повинна мати значення.

```js
assert.isUndefined(fighting);
```

Не присвоюйте значення змінній `fighting`. Не забудьте про крапку з комою в кінці рядка.

```js
assert.match(code, /let\s+fighting\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
--fcc-editable-region--
```
