---
id: 62a3bec30ea7f941412512dc
title: Крок 33
challengeType: 0
dashedName: step-33
---

# --description--

Так само як і з кнопками, створіть змінні для наступних `id` та використайте `querySelector()`, щоб надати їм елемент як значення:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats` та `monsterName`.

Не забудьте оголосити їх за допомогою ключового слова `const` та назвіть змінні відповідно до `id`.

# --hints--

Ви повинні оголосити змінну `text` за допомогою `const`.

```js
assert.match(code, /const text/);
```

Ваша змінна `text` повинна мати значення вашого елемента `#text`.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

Ви повинні оголосити змінну `xpText` за допомогою `const`.

```js
assert.match(code, /const xpText/);
```

Ваша змінна `xpText` повинна мати значення вашого елемента `#xpText`.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

Ви повинні оголосити змінну `healthText` за допомогою `const`.

```js
assert.match(code, /const healthText/);
```

Ваша змінна `healthText` повинна мати значення вашого елемента `#healthText`.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

Ви повинні оголосити змінну `goldText` за допомогою `const`.

```js
assert.match(code, /const goldText/);
```

Ваша змінна `goldText` повинна мати значення вашого елемента `#goldText`.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

Ви повинні оголосити змінну `monsterStats` за допомогою `const`.

```js
assert.match(code, /const monsterStats/);
```

Ваша змінна `monsterStats` повинна мати значення вашого елемента `#monsterStats`.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

Ви повинні оголосити змінну `monsterName` за допомогою `const`.

```js
assert.match(code, /const monsterName/);
```

Ваша змінна `monsterName` повинна мати значення вашого елемента `#monsterName`.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
