---
id: 62a3bec30ea7f941412512dc
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

Так само як і з кнопками, створіть змінні для наступних `id` та використайте `querySelector()`, щоб надати їм елемент як значення:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats` та `monsterName`.

Не забудьте оголосити їх за допомогою ключового слова `const` та назвіть змінні відповідно до `id`.

# --hints--

Оголосіть змінну `text` за допомогою `const`.

```js
assert.match(code, /const\s+text/);
```

Змінна `text` повинна мати значення елемента `#text`.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

Оголосіть змінну `xpText` за допомогою `const`.

```js
assert.match(code, /const\s+xpText/);
```

Змінна `xpText` повинна мати значення елемента `#xpText`.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

Оголосіть змінну `healthText` за допомогою `const`.

```js
assert.match(code, /const\s+healthText/);
```

Змінна `healthText` повинна мати значення елемента `#healthText`.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

Оголосіть змінну `goldText` за допомогою `const`.

```js
assert.match(code, /const\s+goldText/);
```

Змінна `goldText` повинна мати значення елемента `#goldText`.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

Оголосіть змінну `monsterStats` за допомогою `const`.

```js
assert.match(code, /const\s+monsterStats/);
```

Змінна `monsterStats` повинна мати значення елемента `#monsterStats`.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

Оголосіть змінну `monsterName` за допомогою `const`.

```js
assert.match(code, /const\s+monsterName/);
```

Змінна `monsterName` повинна мати значення елемента `#monsterName`.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
