---
id: 62a3b5843544ce3a77459c27
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Наразі гравець повинен починати тільки з `stick`. Змініть масив `inventory`, щоб він мав єдине значення `stick`.

# --hints--

Змінна `inventory` досі має бути масивом.

```js
assert.isArray(inventory);
```

Змінна `inventory` повинна мати лише одне значення.

```js
assert.lengthOf(inventory, 1);
```

Змінна `inventory` повинна містити рядок `stick`.

```js
assert.include(inventory, "stick");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick", "dagger", "sword"];
--fcc-editable-region--
```
