---
id: 62a3a488b24fb32b91155d56
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Змінним можна <dfn>присвоїти</dfn> значення. Якщо так зробити під час оголошення, це називається <dfn>ініціалізацією</dfn>. Наприклад:

```js
let age = 32;
```

This would initialize the `age` variable with a value of `32`, a number.

Ініціалізуйте змінну `xp`, щоб вона мала значення `0` (число).

# --hints--

`xp` повинна мати значення `0`.

```js
assert.equal(xp, 0);
```

Ініціалізуйте змінну `xp` значенням `0`. Не забудьте про крапку з комою в кінці рядка.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp;
--fcc-editable-region--
```
