---
id: 62a3b506dbaead396f58a701
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

В інвентарі гравця вашої гри буде декілька предметів. Вам потрібно використати той тип даних, який може так зробити. Для утримання декількох значень можна використати <dfn>масив</dfn>. Наприклад:

```js
let order = ["first", "second", "third"];
```

Це масив, який містить три значення. Зверніть увагу, що значення розділені комами.

Change your `inventory` variable to be an array with the strings `"stick"`, `"dagger"`, and `"sword"`.

# --hints--

Змінна `inventory` має бути масивом.

```js
assert.isArray(inventory);
```

Змінна `inventory` повинна мати три значення.

```js
assert.lengthOf(inventory, 3);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

Your `inventory` variable should include the string `"dagger"`.

```js
assert.include(inventory, "dagger");
```

Your `inventory` variable should include the string `"sword"`.

```js
assert.include(inventory, "sword");
```

Змінна `inventory` повинна мати значення в правильному порядку.

```js
assert.deepEqual(inventory, ["stick", "dagger", "sword"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = "stick";
--fcc-editable-region--
```
