---
id: 62a3a75d8466a12e009eff76
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Ініціалізуйте іншу змінну `health` значенням `100` та змінну `gold` значенням `50`.

# --hints--

You should use `let` to declare a variable called `health`.

```js
assert.match(code, /let\s+health/);
```

Ініціалізуйте змінну під назвою `health` значенням `100`.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

You should use `let` to declare a variable called `gold`.

```js
assert.match(code, /let\s+gold/);
```

Ініціалізуйте змінну під назвою `gold` значенням `50`.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` повинна мати значення `100`.

```js
assert.equal(health, 100);
```

`gold` повинна мати значення `50`.

```js
assert.equal(gold, 50);
```

`xp` досі повинна мати значення `0`.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
