---
id: 62a3cfc8328d3351b95d4f61
title: Крок 42
challengeType: 0
dashedName: step-42
---

# --description--

Властивість `innerText` контролює текст, який з'являється в елементі HTML. Наприклад:

```js
const info = document.querySelector("#info");
info.innerText = "Hello World";
```

Цей код змінить елемент, який присвоєний змінній `info`, щоб той мав текст `Hello World`.

Якщо гравець натискає на кнопку `Go to store`, потрібно, щоб змінювались кнопки та текст. Видаліть код всередині функції `goStore` та додайте рядок, який оновлює текст для `button1` на `Buy 10 health (10 gold)`.

# --hints--

Вам не потрібен рядок `console.log("Going to store.");` у своєму коді.

```js
assert.notMatch(code, /console.log\(('|")Going to store\.\1\)/);
```

Ви повинні використати крапкову нотацію, щоб отримати доступ до властивості `innerText` для `button1`.

```js
assert.match(code, /button1\.innerText/);
```

Ви не повинні використовувати `let` чи `const` для доступу до властивості `innerText` для `button1`.

```js
assert.notMatch(code, /(let|const)\s+button1.innerText/);
```

Ви повинні оновити властивість `innerText` для `button1` на `Buy 10 health (10 gold)`.

```js
assert.match(code, /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

Ви повинні оновити властивість `innerText` у своїй функції `goStore`.

```js
assert.match(goStore.toString(), /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    border: 1px solid black;
    padding: 5px;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
