---
id: 62a3b79d520a7f3d0e25afd6
title: Крок 29
challengeType: 0
dashedName: step-29
---

# --description--

JavaScript взаємодіє із HTML завдяки <dfn>об’єктній моделі документа</dfn>, або DOM. DOM – це дерево об’єктів, які представляють HTML. Ви можете отримати доступ до HTML завдяки об’єкту `document`, який являє собою весь документ HTML.

Один зі способів пошуку конкретних елементів у HTML – використати метод `querySelector()`. Метод `querySelector()` приймає CSS-селектор як аргумент та повертає перший елемент, який відповідає селектору. Наприклад, щоб знайти елемент `<h1>` у своєму HTML, ви б написали:

```js
let h1 = document.querySelector("h1");
```

Зауважте, що `h1` є рядком і відповідає CSS-селектору, який ви б використали.

Створіть змінну `button1` та використайте `querySelector()`, щоб присвоїти елемент з `id` зі значенням `button1` до неї. Пам'ятайте, що `id` CSS-селекторів розпочинається з `#`.

# --hints--

Ви повинні використати `let`, щоб оголосити змінну `button1`.

```js
assert.match(code, /let button1/);
```

Ви повинні використати `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Ви повинні використати селектор `#button1`.

```js
assert.match(code, /querySelector\(('|")#button1\1\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
