---
id: 62a3a0a3c0a4b32915d26a6e
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

Вилучіть рядок `console.log("Hello World");`, щоб почати писати код свого проєкту.

В JavaScript <dfn>змінну</dfn> використовують для утримання значення. Щоб використати змінну, потрібно <dfn>оголосити</dfn> її. Наприклад, щоб оголосити змінну `camperbot`, ви б написали:

```js
var camperbot;
```

Ключове слово `var` повідомляє JavaScript, що ви оголошуєте змінну. Оголосіть змінну під назвою `xp`.

# --hints--

Вам не потрібен рядок `console.log("Hello World");` у своєму коді.

```js
assert.notMatch(code, /console\.log\("Hello World"\);/);
```

Ви повинні використати ключове слово `var`, щоб оголосити змінну.

```js
assert.match(code, /var/);
```

Ви повинні оголосити змінну під назвою `xp`.

```js
assert.match(code, /xp/);
```

Ви не повинні присвоювати значення своїй змінній.

```js
assert.notMatch(code, /var xp =/);
```

Не забудьте про крапку з комою в кінці рядка.

```js
assert.match(code, /var xp;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
