---
id: 62a3a0a3c0a4b32915d26a6e
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Вилучіть рядок `console.log("Hello World");`, щоб почати писати код свого проєкту.

В JavaScript <dfn>змінну</dfn> використовують для того, щоб зберігати значення. Щоб використати змінну, потрібно <dfn>оголосити</dfn> її. Наприклад, щоб оголосити змінну `camperbot`, ви б написали:

```js
let camperbot;
```

The `let` keyword tells JavaScript you are declaring a variable. Оголосіть змінну під назвою `xp`.

# --hints--

Вам не потрібен рядок `console.log("Hello World");` у коді.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

You should use the `let` keyword to declare your variable.

```js
assert.match(code, /let/);
```

Оголосіть змінну під назвою `xp`.

```js
assert.match(code, /let\s+xp/);
```

Не присвоюйте значення змінній.

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

Не забудьте про крапку з комою в кінці рядка.

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
