---
id: 62a3b49686792938718b90d3
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Всі присвоєні змінні мають значення, що є числами. JavaScript має декілька різних типів даних. Наступний, який ви використаєте, — це <dfn>рядок</dfn>. Рядки використовують, щоб зберегти слова та текст. Рядки оточені подвійними, одинарними або зворотними лапками. Ось приклад оголошення змінної рядком:

```js
let developer = "Naomi";
```

Присвойте змінній `inventory` значення `stick`.

# --hints--

Встановіть `inventory` на рядок `stick`.

```js
assert.equal(inventory, "stick");
```

Ініціалізуйте змінну `inventory` рядком `stick`.

```js
assert.match(code, /let\s+inventory\s*=\s*('|"|`)stick\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
