---
id: 62a3a7e4f1060e2fc5ffb34b
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Create another variable called `currentWeapon` and set it to `0`.

# --hints--

You should use `let` to declare a variable called `currentWeapon`.

```js
assert.match(code, /let\s+currentWeapon/i);
```

Використайте верблюдячийРегістр, щоб назвати змінну.

```js
assert.match(code, /currentWeapon/);
```

Ваша змінна `currentWeapon` повинна бути встановлена на `0`.

```js
assert.equal(currentWeapon, 0);
```

Ініціалізуйте змінну значенням `0`.

```js
assert.match(code, /let\s+currentWeapon\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
