---
id: 62a7c071219da921758a35bb
title: Крок 48
challengeType: 0
dashedName: step-48
---

# --description--

У своїй функції `goTown` змініть властивості `innerText` елементів `button` на `Go to store`, `Go to cave` та `Fight dragon`. Оновіть значення `onclick` на `goStore`, `goCave` та `fightDragon`.

Зрештою, оновіть властивість тексту `innerText` на `You are in the town square. You see a sign that says Store.`.

# --hints--

Ви повинні встановити властивість `button1.innerText` на `Go to store` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /button1\.innerText\s*=\s*('|")Go to store\1/);
```

Ви повинні встановити властивість `button2.innerText` на `Go to cave` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /button2\.innerText\s*=\s*('|")Go to cave\1/);
```

Ви повинні встановити властивість `button3.innerText` на `Fight dragon` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /button3\.innerText\s*=\s*('|")Fight dragon\1/);
```

Ви повинні встановити властивість `button1.onclick` на `goStore` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /button1\.onclick\s*=\s*goStore/);
```

Ви повинні встановити властивість `button2.onclick` на `goCave` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /button2\.onclick\s*=\s*goCave/);
```

Ви повинні встановити властивість `button3.onclick` на `fightDragon` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /button3\.onclick\s*=\s*fightDragon/);
```

Ви повинні встановити властивість `text.innerText` на `You are in the town square. You see a sign that says Store.` у своїй функції `goTown`.

```js
assert.match(goTown.toString(), /text\.innerText\s*=\s*('|")You are in the town square. You see a sign that says Store\.\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goTown() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
  button1.onclick = buyHealth;
  button2.onclick = buyWeapon;
  button3.onclick = goTown;
  text.innerText = "You enter the store.";
}
--fcc-editable-region--

function goStore() {
  button1.innerText = "Buy 10 health (10 gold)";
  button2.innerText = "Buy weapon (30 gold)";
  button3.innerText = "Go to town square";
  button1.onclick = buyHealth;
  button2.onclick = buyWeapon;
  button3.onclick = goTown;
  text.innerText = "You enter the store.";
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}

function buyHealth() {

}

function buyWeapon() {

}
```
