---
id: 62a8d81f539f004776dd9b1e
title: Крок 108
challengeType: 0
dashedName: step-108
---

# --description--

Тепер ви можете розпочати код для боротьби з монстрами. Щоб код був організованим, ваша функція `fightDragon` була переміщена, щоб бути ближче до інших функцій `fight`.

Під своїм масивом `weapons` визначте змінну `monsters` та призначте їй масив. Масив повинен мати три об'єкти, кожен із властивостями `name`, `level` та `health`. Значеннями першого об'єкту повинні бути `slime`, `2` та `15`, за порядком. Значення другого: `fanged beast`, `8` та `60`. Значення третього: `dragon`, `20` та `300`.

# --hints--

Ви повинні використати `const`, щоб оголосити змінну `monsters`.

```js
assert.match(code, /const\s*monsters/);
```

Ваша змінна `monsters` повинна бути масивом.

```js
assert.isArray(monsters);
```

Ваша змінна `monsters` повинна мати 3 значення.

```js
assert.lengthOf(monsters, 3);
```

Ваш масив `monsters` повинен мати 3 об’єкти.

```js
assert(monsters.every(val => typeof val === "object"));
```

Перше значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `name` зі значенням `slime`.

```js
assert.equal(monsters[0].name, "slime");
```

Перше значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `level` зі значенням `2`.

```js
assert.equal(monsters[0].level, 2);
```

Перше значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `health` зі значенням `15`.

```js
assert.equal(monsters[0].health, 15);
```

Друге значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `name` зі значенням `fanged beast`.

```js
assert.equal(monsters[1].name, "fanged beast");
```

Друге значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `level` зі значенням `8`.

```js
assert.equal(monsters[1].level, 8);
```

Друге значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `health` зі значенням `60`.

```js
assert.equal(monsters[1].health, 60);
```

Третє значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `name` зі значенням `dragon`.

```js
assert.equal(monsters[2].name, "dragon");
```

Третє значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `level` зі значенням `20`.

```js
assert.equal(monsters[2].level, 20);
```

Третє значення у вашому масиві `monsters` повинне бути об'єктом з властивістю `health` зі значенням `300`.

```js
assert.equal(monsters[2].health, 300);
```

Не додавайте додаткові властивості до своїх об'єктів.

```js
assert.deepEqual(monsters, [
  {name: "slime", level: 2, health: 15},
  {name: "fanged beast", level: 8, health: 60},
  {name: "dragon", level: 20, health: 300}
])
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText =document.querySelector("#monsterHealth");
const weapons = [
  { name: 'stick', power: 5 },
  { name: 'dagger', power: 30 },
  { name: 'claw hammer', power: 50 },
  { name: 'sword', power: 100 }
];
--fcc-editable-region--

--fcc-editable-region--
const locations = [
    {
        name: "town square",
        "button text": ["Go to store", "Go to cave", "Fight dragon"],
        "button functions": [goStore, goCave, fightDragon],
        text: "You are in the town square. You see a sign that says \"Store\"."
    },
    {
        name: "store",
        "button text": ["Buy 10 health (10 gold)", "Buy weapon (30 gold)", "Go to town square"],
        "button functions": [buyHealth, buyWeapon, goTown],
        text: "You enter the store."
    },
    {
        name: "cave",
        "button text": ["Fight slime", "Fight fanged beast", "Go to town square"],
        "button functions": [fightSlime, fightBeast, goTown],
        text: "You enter the cave. You see some monsters."
    }
];

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

function update(location) {
  button1.innerText = location["button text"][0];
  button2.innerText = location["button text"][1];
  button3.innerText = location["button text"][2];
  button1.onclick = location["button functions"][0];
  button2.onclick = location["button functions"][1];
  button3.onclick = location["button functions"][2];
  text.innerText = location.text;
}

function goTown() {
  update(locations[0]);
}

function goStore() {
  update(locations[1]);
}

function goCave() {
  update(locations[2]);
}

function buyHealth() {
  if (gold >= 10) {
    gold -= 10;
    health += 10;
    goldText.innerText = gold;
    healthText.innerText = health;
  } else {
    text.innerText = "You do not have enough gold to buy health.";
  }
}

function buyWeapon() {
  if (currentWeapon < weapons.length - 1) {
    if (gold >= 30) {
      gold -= 30;
      currentWeapon++;
      goldText.innerText = gold;
      let newWeapon = weapons[currentWeapon].name;
      text.innerText = "You now have a " + newWeapon + ".";
      inventory.push(newWeapon);
      text.innerText += " In your inventory you have: " + inventory;
    } else {
      text.innerText = "You do not have enough gold to buy a weapon.";
    }
  } else {
    text.innerText = "You already have the most powerful weapon!";
    button2.innerText = "Sell weapon for 15 gold";
    button2.onclick = sellWeapon;
  }
}

function sellWeapon() {
  if (inventory.length > 1) {
    gold += 15;
    goldText.innerText = gold;
    let currentWeapon = inventory.shift();
    text.innerText = "You sold a " + currentWeapon + ".";
    text.innerText += " In your inventory you have: " + inventory;
  } else {
    text.innerText = "Don't sell your only weapon!";
  }
}

function fightSlime() {

}

function fightBeast() {

}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
