---
id: 62fc20387ef88d1d1998aac5
title: Крок 32
challengeType: 0
dashedName: step-32
---

# --description--

Використайте `querySelector()`, щоб отримати два інші елементи `button`, використовуючи їхні `id` (`button2` та `button3`). Розмістіть їх у змінних під назвою `button2` та `button3`. Не забудьте використати `const`.

# --hints--

Ви повинні оголосити змінну `button2` за допомогою `const`.

```js
assert.match(code, /const button2/);
```

Ваша змінна `button2` повинна мати значення вашого елемента `#button2`.

```js
assert.deepEqual(button2, document.querySelector('#button2'));
```

Ви повинні оголосити змінну `button3` за допомогою `const`.

```js
assert.match(code, /const button3/);
```

Ваша змінна `button3` повинна мати значення вашого елемента `#button3`.

```js
assert.deepEqual(button3, document.querySelector('#button3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
