---
id: 66ece0b518f5dcf5930612ab
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

For the last portion of the workshop, you will review how to get the minimum and maximum values of a set of numbers.

Start by adding a `console.log()` that logs the message `"The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range."` to the console.

# --hints--

You should log the message `"The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range."` to the console.

```js
assert.match(code, /console\.log\((['"`])The\s+Math\.max\(\)\s+and\s+Math\.min\(\)\s+methods\s+are\s+used\s+to\s+get\s+the\s+maximum\s+and\s+minimum\s+number\s+from\s+a\s+range\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);
--fcc-editable-region--

--fcc-editable-region--
```
