---
id: 5d822fd413a79914d39e98d0
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `body` 中创建一个带有 `background-buildings` 类的 `div` 元素。 这将是一组建筑物的容器。

# --hints--

你应该创建一个 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该在 `body` 内。

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

你的 `div` 元素应该有一个 `background-buildings` 类。

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

