---
id: 5d822fd413a79914d39e98d7
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

要使用变量，请将变量名放在括号中，并在它们前面加上 `var`，例如：`var(--variable-name)`。 添加变量作为 `.bb1a` 类的 `background-color` 属性的值。 你给变量的任何值都将应用于你使用它的任何属性。 在这个例子中，你的变量的值为 `#999`。 因此 `#999` 将用作 `background-color` 属性的值。

# --hints--

`bb1a` 元素的 `background-color` 应该被设置。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

你应该使用 `var(--building-color1)` 来设置 `.bb1a` 元素的 `background-color`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
