---
id: 5d822fd413a79914d39e98cd
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 CSS 中，你可以使用星号定位所有内容。 使用 `*` 选择器为所有内容添加边框，并为其提供 `1px solid black` 的 `border`。 这是一个有助于可视化元素的位置及其大小的技巧。 你稍后会删除这个。

# --hints--

你应该使用 `*` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你应该使用 `border` 属性来设置所有元素的样式。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

所有元素都应该有一个 `1px solid black` 的边框。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
