---
id: 5d822fd413a79914d39e98d2
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在背景建筑物容器中嵌套一个具有 `bb1` 类的 `div`。 打开你的 `styles.css` 文件，将 `.bb1` 的 `width` 设置为 `10%`，`height` 设置为 `70%`。 “bb” 代表“背景建筑”，这将是你的第一栋建筑。

# --hints--

你应该创建一个新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

你应该给新的 `div` 一个 `bb1` 类。

```js
assert.exists(document.querySelector('div.bb1'));
```

你应该使用 `.bb1` 类选择器来设置元素的样式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

你应该将 `.bb1` 元素的 `width` 设置为 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

你应该将 `.bb1` 元素的 `height` 设置为 `70%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

