---
id: 5d822fd413a79914d39e98ca
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `DOCTYPE` 下方添加开始和结束 `html` 标签，以便你可以开始放置一些代码。 确保将语言设置为英语。

# --hints--

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你的 `html` 元素应该有一个开始标签，并且 `lang` 属性的值为 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `html` 标签应该有正确的顺序。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

你应该只有一个 `html` 元素。

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
