---
id: 5d822fd413a79914d39e98e7
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

将这六个新元素的 `width` 和 `height` 值设置为：`.fb1` 的是 `10%` 和 `60%`、`.fb2` 的是 `10%` 和 `40%`、`.fb3` 的是 `10%` 和 `35%`、`.fb4` 的是 `8%` 和 `45%`、`.fb5` 的是 `10%` 和 `33%`、`.fb6` 的是 `9%` 和 `38%`。

# --hints--

你应该创建一个 `.fb1` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

你应该将 `.fb1` 选择器的 `width` 设置为 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

你应该将 `.fb1` 选择器的 `height` 设置为 `60%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

你应该创建一个 `.fb2` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

你应该将 `.fb2` 选择器的 `width` 设置为 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

你应该将 `.fb2` 选择器的 `height` 设置为 `40%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

你应该创建一个 `.fb3` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

你应该将 `.fb3` 选择器的 `width` 设置为 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

你应该将 `.fb3` 选择器的 `height` 设置为 `35%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

你应该创建一个 `.fb4` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

你应该为 `.fb4` 选择器提供 `8%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

您应该为 `.fb4` 选择器提供 `45%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

你应该创建一个 `.fb5` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

你应该为 `.fb5` 选择器提供 `10%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

你应该为 `.fb5` 选择器提供 `33%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

你应该创建一个 `.fb6` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

你应该为 `.fb6` 选择器提供 `9%` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

你应该为 `.fb6` 选择器提供 `38%` 的 `height`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--

```

