---
id: 5d822fd413a79914d39e98c9
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

欢迎构建 CSS 变量轮廓项目！ 首先，在文档的顶部添加 `!DOCTYPE html` 声明，这样浏览器就知道它在阅读什么类型的文档。

# --hints--

你的代码应包含 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应该在 `DOCTYPE` 声明之后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应该将文档类型定义为 `html`。

```js
assert(code.match(/html/gi));
```

你应该在类型之后使用 `>` 关闭 `DOCTYPE` 声明。

```js
assert(code.match(/html\s*>/gi));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
