---
id: 5d822fd413a79914d39e98d6
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

现在你有一些类似于建筑物的东西。 可以准备创建第一个变量了。 变量声明以两个破折号 (`-`) 开头，并被赋予名称和值，例如：`--variable-name: value;`。 在 `bb1` 类中，创建一个名为 `--building-color1` 的变量，并为其赋值 `#999`。

# --hints--

您应该创建一个名为 `--building-color1` 的新变量。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

你应该在 `.bb1` 中定义 `--building-color1` 变量。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

你应该给 `--building-color1` 一个值为 `#999`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
