---
id: 5f332b23c2045fb843337579
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

لتمكين الزوار من معرفة أن المقهى قد تأسس في عام 2020، أضف عنصر `p` أسفل عنصر `h1` مع النص `Est. 2020`.

# --hints--

يجب أن يكون لديك العلامة الأفتاح (opening tag) الآتية `<p>`.

```js
assert(code.match(/<p>/i));
```

يجب أن يكون لديك العلامة الإغلاق (closing tag) الآتي `</p>`.

```js
assert(code.match(/<\/p>/i));
```

لا يجب عليك تغيير عنصر `h1` الحالي. تيقن من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
assert($('h1').length === 1);
```

يجب أن يكون عنصرك `p` تحت عنصرك `h1`.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

يجب أن يحتوي عنصرك `p` على النص `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

