---
id: 5f33294a6af5e9188dbdb8f3
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

حان الوقت لإضافة محتوى للقائمة. أضف عنصر `main` داخل عنصر `body` الحالي. سيحتوي في النهاية على معلومات تسعير حول القهوة والحلويات التي يقدمها المقهى.

# --hints--

يجب أن يكون لديك العلامة الأفتاح (opening tag) الآتية `<main>`.

```js
assert(code.match(/<main>/i));
```

يجب أن يكون لديك العلامة الإغلاق (closing tag) الآتي `</main>`.

```js
assert(code.match(/<\/main>/i));
```

لا يجب عليك تغيير عنصر `body`. تيقن من عدم حذف علامتك إغلاق (closing tag) عن طريق الخطأ.

```js
assert($('body').length === 1);
```

يجب أن يأتي وسمك `main` داخل وسمك `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
