---
id: 5f33071498eb2472b87ddee4
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

كما تعلمت في الخطوات القليلة الأخيرة من Cat App Photo، هناك بنية أساسية مطلوبة لبدء بناء صفحة الويب الخاصة بك.

أضف العلامة (tag) الآتي `<!DOCTYPE html>`، وعنصر `html` مع سمة (attribute) تسمى `lang` بقيمة `en`.

# --hints--

يجب أن يكون لديك تعريف `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يكون لديك علامة `<html>` فتح مع سمة `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يكون لديك علامة إغلاق الآتي `<html>`. تذكر أن علامات الإغلاق (closing tags) تحتوي على `/` بعد فتح `<` القوس.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
