---
id: 5f331e55dfab7a896e53c3a1
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

يكون `title` واحد من عدة عناصر توفر معلومات إضافية غير مرئية على صفحة الويب، لكنه مفيد لمحركات البحث أو كيف يتم عرض الصفحة.

داخل عنصر `head`، قم بدمج عنصر `meta` مع سمة تسمى `charset` و تعيين قيمتها إلي `utf-8` لإخبار المتصفح كيفية ترميز الأحرف للصفحة. لاحظ أن عناصر `meta` مغلقة ذاتيا (self closing).

# --hints--

يجب أن يكون لديك وسم `meta`.

```js
assert(code.match(/<meta.*?\/?>/is));
```

يجب أن يحتوي وسم `meta` الخاص بك علي سمة `charset`.

```js
assert(code.match(/<meta charset=/i));
```

يجب أن تحتوي سمة `charset` الخاصة بك على القيمة `utf-8`.

```js
assert(code.match(/charset=('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
