---
id: 5f3313e74582ad9d063e3a38
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف عنصر `head` ضمن عنصر `html`، حتى تتمكن من إضافة عنصر `title`. نص عنصر `title` يجب أن يكون `Cafe Menu`.

# --hints--

يجب أن يكون لديك العلامة المفتوحة (opening tag) الآتية `<head>`.

```js
assert(code.match(/<head>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتية `</head>`.

```js
assert(code.match(/<head>/i));
```

يجب أن يكون لديك العلامة المفتوحة (opening tag) الآتية `<title>`.

```js
assert(code.match(/<title>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتي `</title>`.

```js
assert(code.match(/<\/title>/i));
```

عنصر `title` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

يجب أن يحتوي العنصر `title` على النص `Cafe Menu`. قد تحتاج إلى التحقق من الإملاء.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
