---
id: 5f331e55dfab7a896e53c3a1
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

يكون `title` واحد من عدة عناصر توفر معلومات إضافية غير مرئية على صفحة الويب، لكنه مفيد لمحركات البحث أو كيف يتم عرض الصفحة.

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8` to tell the browser how to encode characters for the page.

# --hints--

يجب أن يكون لديك وسم `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

يجب أن يحتوي وسم `meta` الخاص بك علي سمة `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

يجب أن تحتوي سمة `charset` الخاصة بك على القيمة `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
