---
id: 5f3326b143638ee1a09ff1e3
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

للتحضير لإنشاء بعض المحتوى الفعلي، أضف عنصر `body` أسفل عنصر `head`.

# --hints--

يجب أن يكون لديك العلامة المفتوحة (opening tag) الآتية `<body>`.

```js
assert(code.match(/<body>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتي `</body>`.

```js
assert(code.match(/<\/body>/i));
```

لا يجب عليك تغيير عنصر `head`. تأكد من أنك لم تقم بحذف علامة الإغلاق الخاصة بك.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

عنصر `body` الخاص بك يجب أن يأتي بعد عنصر `head` الخاص بك.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

