---
id: 5f3477aefa51bfc29327200b
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

لقد قمت بتصميم ثلاثة عناصر عن طريق كتابة CSS داخل (tags) علامات `style`. هذا يعمل، ولكن بما أنه سيكون هناك الكثير من الأنماط الأخرى، من الأفضل وضع جميع الأنماط في ملف منفصل وربطها به.

لقد قمنا بإنشاء ملف `styles.css` منفصل لك وقمنا بتغيير طريقة عرض المحرر إلى ذلك الملف. يمكنك التغيير بين الملفات من علامات التبويب في الجزء العلوي من المحرر.

ابدأ بإعادة كتابة الأنماط التي أنشأتها في ملف `styles.css`. تأكد من استبعاد علامات الفتح و الإغلاق للـ `style`.

# --hints--

ملف `styles.css` الخاص بك يجب أن يحتوي على `h1, h2, p` كـ type selectors.

```js
(getUserInput) => {
  assert(getUserInput('editableContents').replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
}
```

يجب أن يقوم المنتقي الخاص بك بتعيين خاصية `text-align` إلى `center`.

```js
(getUserInput) => {
  assert(getUserInput('editableContents').match(/text-align:\s*center;?/));
}
```

يجب ألا يحتوي الكود الخاص بك على وسوم `<style>`.

```js
  assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
