---
id: 5f3477cbcb6ba47918c1da92
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

لكي يبدو تصميم الصفحة مشابه على الهاتف المحمول كما هو الحال على حاسوب مكتبي أو حاسوب محمول، تحتاج إلى إضافة عنصر `meta` مع سمة `content` خاصة.

يضاف ما يلي ضمن عنصر `head`:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصرين `meta`.

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

يجب أن يحتوي عنصر `meta` الخاص بك على سمة `name` بقيمة `viewport`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

يجب أن يحتوي عنصر `meta` الخاص بك على `content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
