---
id: 5f3477ae9675db8bb7655b30
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

في الخطوة السابقة، استخدمت <dfn>type selector</dfn> لتصميم عنصر `h1`. أوضع عناصر `h2` و `p` في المنتصف, عن طريق إضافة منتقي type جديد لكل عنصر في عنصر `style`.

# --hints--

لا يجب عليك تغيير عنصر `h1` الحالي.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

يجب عليك إلا تضيف علامة `style` جديدة. أضف قواعد CSS الجديدة إلى علامة `style`.

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

يجب عليك إضافة منتقي `h2` جديد.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

يجب عليك إضافة منتقي `p` جديد.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

يجب أن يحتوي عنصرك `h1` على خاصية `text-align` بقيمة `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

يجب أن يحتوي عنصرك `h2` على خاصية `text-align` بقيمة `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

يجب أن يحتوي عنصرك `p` على خاصية `text-align` بقيمة `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

