---
id: 5f344fbc22624a2976425065
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

قم بإنشاء عنصر `h2` في عنصر `section` وقم بإعطائه النص `Coffee`.

# --hints--

يجب أن يكون لديك العلامة المفتوحة (opening tag) الآتية `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتي `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

لا يجب عليك تغيير عنصر `section` الحالي. تأكد من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
assert($('section').length === 1);
```

عنصر `h2` الخاص بك يجب أن يكون داخل عنصر `section` الخاص بك.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

يجب أن يحتوي عنصر `h2` الخاص بك على النص `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
