---
id: 5f3477cb2e27333b1ab2b955
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

الآن تحتاج إلى ربط ملف `styles.css` حتى يتم تطبيق الأنماط مرة أخرى. أدخل عنصر `link` ذاتي الأغلاق في عنصر `head`. اعطيها سمة `rel` بقيمة `stylesheet` و `href` بقيمة `styles.css`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
const link = document.querySelector('link');
assert(link);
```

لا يجب عليك تغيير عنصر `head` الحالي. تأكد من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
assert($('head').length === 1);
```

الكود الخاص بك يجب أن يحتوي على عنصر `link` مغلق ذاتيا.

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

يجب أن يكون عنصر `link` الخاص بك داخل عنصر `head` الخاص بك.

```js
const link = document.querySelector('head > link');
assert(link);
```

يجب أن يكون عنصر `link` الخاص بك على سمة `rel` بقيمة `stylesheet`.

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
