---
id: 5f3477ae34c1239cafe128be
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

لديك الآن ثلاثة type selectors بنفس النمط بالضبط. يمكنك إضافة نفس المجموعة من الأنماط إلى العديد من العناصر عن طريق إنشاء قائمة من الـ selectors. يتم فصل كل selector بفواصل مثل هذا:

```css
selector1, selector2 {
  property: value;
}
```

احذف الـ type selectors الثلاثة الموجودة واستبدلها بـ selector list واحدة تقوم بتركيز النصوص لعناصر `h1` و `h2` و `p`.

# --hints--

يجب عليك استخدام type selector واحد لكل العناصر الثلاثة، `h1, h2, p`. تأكد من استخدام هذا الترتيب.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

يجب أن يكون لديك منتقي (selector) واحد فقط في عنصر `style` الخاص بك.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

