---
id: 5f344fc1520b6719f2e35605
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

سيكون هناك قسمان (sections) في القائمة، قسم للقهوة (coffees) وآخر للحلويات (desserts). أضف عنصر `section` ضمن عنصر `main` حتى يكون لديك مكان لوضع جميع القهوة المتاحة.

# --hints--

يجب أن يكون لديك العلامة المفتوحة (opening tag) الآتية `<section>`.

```js
assert(code.match(/<section\s*>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتي `</section>`.

```js
assert(code.match(/<\/section\s*>/i));
```

لا يجب عليك تغيير عنصر `main` الحالي. تأكد من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
assert($('main').length === 1);
```

عنصر `section` الخاص بك يجب أن يكون ضمن عنصر `main` الخاص بك.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
