---
id: 5f3477cb2e27333b1ab2b955
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

Now you need to link the `styles.css` file, so the styles will be applied again. Inside the `head` element, add a `link` element. Give it a `rel` attribute with the value of `"stylesheet"` and a `href` attribute with the value of `"styles.css"`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

The `link` element is a void element, it should not have an end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

لا يجب عليك تغيير عنصر `head` الحالي. تأكد من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

You should have one `link` element.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

يجب أن يكون عنصر `link` الخاص بك داخل عنصر `head` الخاص بك.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

يجب أن يكون عنصر `link` الخاص بك على سمة `rel` بقيمة `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
