---
id: 5f344fad8bf01691e71a30eb
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

حتى الآن، كنت مقيدًا فيما يتعلق بالعرض التقديمي وشكل المحتوى الذي تقوم بإنشائه. لبدء التحكم، أضف عنصر `style` داخل العنصر `head`.

# --hints--

يجب أن يكون لديك الـ opening tag الآتي `<style>`.

```js
assert(code.match(/<style\s*>/i));
```

يجب أن يكون لديك الـ closing tag الآتي `</style>`.

```js
assert(code.match(/<\/style\s*>/));
```

عنصر `style` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
