---
id: 5f3477ae8466a9a3d2cc953c
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

الآن بعد أن أصبح لديك CSS في ملف `styles.css`، قم بإزالة عنصر `style` وجميع محتوياته. بمجرد إزالته، سينتقل النص الذي تم توسيطه مرة أخرى إلى اليسار.

# --hints--

يجب ألا يكون لديك أي tag (علامة) `style` في الكود الخاص بك.

```js
assert(!code.match(/style/i));
```

يجب ألا يكون لديك أي محددات CSS في مِلَفّ HTML الخاص بك.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

