---
id: 5f3477cb303c5cb61b43aa9b
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

يتم توسيط النص مرة أخرى بحيث يعمل الرابط إلى مِلَفّ CSS. أضف تصميم آخر إلى المِلَفّ الذي يغير خاصية `background-color` إلى `brown` لعنصر `body`.

# --hints--

يجب عليك استخدام منتقي `body`.

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

يجب عليك تعيين خاصية `background-color` إلى `brown`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

يجب أن يحتوي عنصر `body` الخاص بك على خلفية `brown`.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

