---
id: 5f35e5c44359872a137bd98f
title: الخطوة 28
challengeType: 0
dashedName: step-28
---

# --description--

بما أن المنتج الرئيسي للكافيه للبيع هو القهوة، يمكنك استخدام صورة لحبوب البن على خلفية الصفحة.

احذف التعليق ومحتوياته داخل منتقي `body`. الآن أضف خاصية `background-image` وقم بتعيين قيمتها إلى `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

يجب عليك إزالة خاصية `background-color` التي تم تعليقها.

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

يجب ألا يكون في منتقي `body` الخاص بك أي تعليقات.

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

يجب عليك تعيين خاصية `background-image` إلى `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

يجب أن يحتوي عنصر `body` على صورة خلفية لحبوب القهوة.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

