---
id: 5f35e5c4321f818cdc4bed30
title: الخطوة 29
challengeType: 0
dashedName: step-29
---

# --description--

يبدو الأمر جيدا. حان الوقت لبدء إضافة بعض عناصر القائمة. قم بإضافة عنصر `article` فارغ تحت عنوان `Coffee`. سوف تحتوي على نكهة (flavor) وسعر (price) لكل قهوة (coffee) تعرضها حالياً.

# --hints--

يجب أن يكون لديك العلامة المفتوحة الآتية `<article>`.

```js
assert(code.match(/<article>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتي `</article>`.

```js
assert(code.match(/<\/article>/i));
```

لا يجب عليك تغيير عنصر `h2` الحالي.

```js
assert($('h2').length === 1);
```

عنصر `article` الخاص بك يجب أن يكون تحت عنصر `h2` الخاص بك.

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

