---
id: 5f3c866dbf362f99b9a0c6d0
title: الخطوة 37
challengeType: 0
dashedName: step-37
---

# --description--

العناصر `p` متداخلة في عنصر `article` مع سمة class بقيمة `item`. يمكنك تصميم جميع عناصر `p` المتداخلة في أي مكان في العناصر مع class يسمى `item` مثل هذا:

```css
.item p { }
```

باستخدام المنتقي (selector) أعلاه، أضف خاصية `display` بقيمة `inline-block` بحيث تتصرف عناصر `p` مثل العناصر الـ inline.

# --hints--

يجب عليك استخدام منتقي `.item p`.

```js
const hasItemP = new __helpers.CSSHelp(document).getStyle('.item p');
assert(hasItemP);
```

يجب عليك تعيين خاصية `display` إلى `inline-block`.

```js
const hasDisplay = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.display === 'inline-block');
assert(hasDisplay);
```

يجب أن يقوم منتقي `.item p` بتعيين خاصية `display` إلى `inline-block`.

```js
const itemPDisplay = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('display');
assert(itemPDisplay === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

--fcc-editable-region--

--fcc-editable-region--

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
