---
id: 5f3c866dd0d0275f01d4d847
title: الخطوة 39
challengeType: 0
dashedName: step-39
---

# --description--

حسنا لم ينجح ذلك. صمم عناصر `p` لتكون `inline-block` ووضعها على سطر منفصل في الكود البرمجي يخلق مساحة إضافية ليمين عنصر `p` الأول، مما تسبب في انتقال الثاني إلى السطر التالي. إحدى طرق إصلاح هذا هو جعل عرض كل عنصر `p` أقل قليلاً من `50%`.

قم بتغيير قيمة `width` إلى `49%` لكل class لمعرفة ما يحدث.

# --hints--

يجب عليك تعيين خاصية `width` إلى `49%` في منتقي `.flavor` الخاص بك.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert(flavorWidth === '49%');
```

يجب عليك تعيين خاصية `width` إلى `49%` في منتقي `.price` الخاص بك.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert(priceWidth === '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
--fcc-editable-region--
```

