---
id: 5f3c866daec9a49519871816
title: الخطوة 30
challengeType: 0
dashedName: step-30
---

# --description--

عناصر `article` تحتوي عادة على عناصر متعددة تحتوي على معلومات ذات صلة. في هذه الحالة، سوف تحتوي على coffee flavor و price لتلك الـ flavor. ادمج عنصرين `p` داخل عنصر `article` الخاص بك. النص الأول يجب أن يكون `French Vanilla`، والنص الثاني `3.00`.

# --hints--

لا يجب عليك تغيير عنصر `article` الحالي.

```js
assert($('article').length === 1);
```

عنصر `article` الخاص بك يجب أن يحتوي على عنصرين `p`.

```js
assert($('article').children('p').length === 2);
```

يجب أن يحتوي العنصر `p` الاول على النص `French Vanilla`.

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

يجب أن يحتوي العنصر `p` الثاني على النص `3.00`.

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
