---
id: 5f76967fad478126d6552b0d
title: الخطوة 34
challengeType: 0
dashedName: step-34
---

# --description--

بعد ذلك، تريد محاذاة السعر إلى اليمين. أضف class تسمى `price` إلى عنصر `p` الخاص بك الذي يحتوي على `3.00` كنص له.

# --hints--

يجب عليك إضافة class تسمى `price` إلى عنصر `p` الخاص بك.

```js
assert(code.match(/<p\s*class=('|")price\1\s*>/i));
```

يجب أن يكون لديك عنصر واحد فقط مع class تسمى `price`.

```js
assert($('.price').length === 1);
```

يجب أن يكون class مسمى `price` الخاص بك على عنصر `p` مع النص `3.00`.

```js
assert($('.price')[0].innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
--fcc-editable-region--
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}
```
