---
id: 612e89562043183c86df287c
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

瀏覽器可以將默認邊距和填充值應用於特定元素。 爲確保你的鋼琴看起來正確，你需要重置盒子模型。

將 `html` 規則選擇器添加到 CSS 文件，並將 `box-sizing` 屬性設置爲 `border-box`。

# --hints--

應該有一個 `html` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

`html` 選擇器應該將 `box-sizing` 屬性設置爲 `border-box`。

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
