---
id: 612e6afc009b450a437940a1
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

從基本的 HTML 結構開始。 添加 `DOCTYPE` 聲明，以及 `html`、`head`、`body` 和 `title` 元素。 將 `title` 設置爲 `Piano`。

# --hints--

代碼應包含 `DOCTYPE` 參考。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 引用之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

應該將文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

應該在類型之後使用 `>` 關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開始標籤，帶有值爲 `en` 的 `lang` 屬性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素應該有一個結束標記。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

應該有一個打開的 `head` 標記。

```js
assert(code.match(/<head\s*>/i));
```

應該有一個結束 `head` 標記。

```js
assert(code.match(/<\/head\s*>/i));
```

應該有一個打開的 `body` 標記。

```js
assert(code.match(/<body\s*>/i));
```

應該有一個結束 `body` 標記。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

項目應該有一個 `Piano` 的標題。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
