---
id: 612e96fc87fe8e44f69f7ec5
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

是時候給按鍵設計樣式了。 在 `#piano` 規則下方，使用 `class` 選擇器定位 `.keys` 元素。 爲新規則賦予 `#040404` 的 `background-color` 屬性、`949px` 的 `width` 屬性和 `180px` 的 `height` 屬性。

# --hints--

應該有一個 `.keys` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

`.keys` 選擇器應該將 `background-color` 屬性設置爲 `#040404`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

`.keys` 選擇器應該將 `width` 屬性設置爲 `949px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

`.keys` 選擇器應該將 `height` 屬性設置爲 `180px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
