---
id: 612ea4c4993aba52ab4aa32e
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

現在是時候使用你之前準備的僞選擇器了。 要創建黑鍵，請添加一個新的 `.key.black--key::after` 選擇器。 這將針對具有類 `key black--key` 的元素，並在 HTML 中選擇這些元素之後的僞元素。

在新的選擇器中，將 `background-color` 設置爲 `#1d1e22`。 還將 `content` 屬性設置爲 `""`。 這將使僞元素爲空。

`content` 屬性用於設置或覆蓋元素的內容。 默認情況下，`::before` 和 `::after` 僞選擇器創建的僞元素是空的，這些元素不會被渲染到頁面上。 將 `content` 屬性設置爲空字符串 `""` 將確保元素在仍然爲空的情況下呈現到頁面。

如果你想進行實驗，請嘗試刪除 `background-color` 屬性併爲 `content` 屬性設置不同的值，例如 `"♥"`。 完成後請記住撤消這些更改，以便測試通過。

# --hints--

應該有一個 `.key.black--key::after` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

`.key.black--key::after` 選擇器應該有一個 `background-color` 屬性設置爲 `#1d1e22`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

`.key.black--key::after` 選擇器應該有一個 `content` 屬性設置爲 `""`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
