---
id: 612e813b3ba67633222cbe54
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

請記住，`class` 屬性可以有多個值。 要將白鍵與黑鍵分開，需要添加第二個 `class` 值 `black--key`。 將此添加到第二個、第三個、第五個、第六個和第七個 `.key` 元素。

# --hints--

第二個 `.key` 元素也應該有一個 `black--key` 的 `class`。

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

第三個 `.key` 應該在 `class` 中有 `black--key`。

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

第五個 `.key` 應該在 `class` 中有 `black--key`。

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

第六個 `.key` 應該在 `class` 中有 `black--key`。

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

第七個 `.key` 應該在 `class` 中有 `black--key`。

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

應該有五個 `.black--key` 元素。

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

應該有七個 `.key` 元素。

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
