---
id: 612e6afc009b450a437940a1
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

從基本的 HTML 結構開始。 添加 `DOCTYPE` 聲明，以及 `html`、`head`、`body` 和 `title` 元素。

將此頁面的語言設置爲英文。 將 `title` 設置爲 `Piano`。

# --hints--

代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 聲明之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該把文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

應該在 `DOCTYPE` 聲明的類型後面加上一個 `>`。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個值爲 `en` 的 `lang` 屬性的開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 聲明應該位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

應該有一個 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/i));
```

應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同級元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

項目應該有一個 `Piano` 標題。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
