---
id: 612e9d142affc44a453655db
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

是時候設計按鍵本身的樣式了。 爲 `.key` 元素創建一個 `class` 選擇器。 將 `background-color` 設置爲值 `#ffffff`，`position` 屬性設置爲 `relative`，`width` 屬性爲 `41px`，`height` 屬性爲 `175px`。

# --hints--

應該有一個 `.key` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

`.key` 選擇器應該有一個 `background-color` 屬性設置爲 `#ffffff`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

`.key` 選擇器應該將 `position` 屬性設置爲 `relative`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

`.key` 選擇器應該有一個 `width` 屬性設置爲 `41px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

`.key` 選擇器應該將 `height` 屬性設置爲 `175px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
