---
id: 612e83ec2eca1e370f830511
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

在 `head` 元素中添加 `link` 元素。 對於該 `link` 元素，將 `rel` 屬性設置爲 `stylesheet`，並將 `href` 設置爲 `./styles.css`。

# --hints--

代碼應該有一個 `link` 元素。

```js
assert.match(code, /<link/)
```

不應更改現有的 `head` 標籤。 確保沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

你應該有一個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 元素應該在 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
