---
id: 646ce8bb4b35544d501c7184
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

`.cat-ears` の要素内に、`div` 要素を 2 つ作成してください。それぞれクラスを `cat-left-ear` と `cat-right-ear` にしてください。(それぞれ、向かって左の耳、向かって右の耳にあたります。)

# --hints--

既存の、`cat-ears` クラスを持つ `div` 要素は変更しないでください。

```js
assert(document.querySelectorAll('div.cat-ears').length === 1);
```

`.cat-ears` 要素内に`div` 要素を 2 つ作成してください。

```js
assert(document.querySelectorAll('.cat-ears div').length === 2);
```

1 番目の `div` 要素にクラス `cat-left-ear` を設定してください。

```js
assert(document.querySelectorAll('.cat-ears div')[0]?.classList.contains('cat-left-ear'));
```

2 番目の `div` 要素にクラス `cat-right-ear` を設定してください。

```js
assert(document.querySelectorAll('.cat-ears div')[1]?.classList.contains('cat-right-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        --fcc-editable-region--
        <div class="cat-ears">

        </div>
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
