---
id: 646c5ffef5598d449b52ec12
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

では猫の耳を作成しましょう。 左右の耳と、それぞれ内側にインナーイヤー (inner ear) の部分を作成します。

`.cat-head` の要素の中に、`div` 要素を作成して `cat-ears` というクラスを持たせてください。

# --hints--

既存の、`cat-head` クラスを持つ `div` 要素は変更しないでください。

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

`.cat-head` 要素内に`div` 要素を 1 つ作成してください。

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

追加した `div` 要素にクラス `cat-ears` を設定してください。

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">

      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
