---
id: 646ce9d790d2a44de5f99e04
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

`div` 要素を 2 つ作成しましょう。1 つは `.cat-left-ear` の要素内に作成し、`cat-left-inner-ear` というクラスを付与してください。もう 1 つは `.cat-right-ear` の要素内に作成し、`cat-right-inner-ear` というクラスを付与してください。

# --hints--

You should not change the existing `div` element with the class `cat-left-inner-ear`.

```js
assert(document.querySelectorAll('div.cat-left-ear').length === 1);
```

You should not change the existing `div` element with the class `cat-right-inner-ear`.

```js
assert(document.querySelectorAll('div.cat-right-ear').length === 1);
```

`.cat-left-ear` の要素内に`div` 要素を 1 つ作成してください。

```js
assert(document.querySelectorAll('.cat-left-ear div').length === 1);
```

`.cat-right-ear` の要素内に`div` 要素を 1 つ作成してください。

```js
assert(document.querySelectorAll('.cat-right-ear div').length === 1);
```

`.cat-left-ear` の要素内に作成した `div` 要素に、クラス `cat-left-inner-ear` を設定してください。

```js
assert(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

`.cat-right-ear` の要素内に作成した `div` 要素に、クラス `cat-right-inner-ear` を設定してください。

```js
assert(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">

          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
