---
id: 646c5ace05e4be4211407935
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`.cat-head` の要素が分かりやすいように、線形グラデーション (linear gradient) の背景色を追加しましょう。 `85%` 地点で `#5e5e5e`、`100%` 地点で `#45454f` にしてください。

色の違いがわかりにくいかもしれませんが、よく見ると違います。

# --hints--

`.cat-head` セレクターに `background` プロパティを追加してください。

```js
assert.match(code, /background:/)
```

`background` プロパティに `linear-gradient` 関数を使用してください。

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

`linear-gradient` 関数で、最初の色 `#5e5e5e` を `85%` の位置まで指定してください。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

`linear-gradient` 関数で、2 番目の色 `#45454f` を `100%` の位置まで指定してください。

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
