---
id: 646ceb843412c74edee27a79
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

ではここで、三角形を描く CSS のテクニックを学びましょう。

これを使って猫の耳を描きます。

クラスセレクターを使用して、`.cat-right-ear` の要素の `height` プロパティと `width` プロパティを `100px` に設定してください。 また、`background-color` は `white` に設定してください。 左右の境界線は `35px solid blue` に、 上下の境界線は `35px solid red` に設定してください。

# --hints--

`.cat-right-ear` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

`.cat-right-ear` セレクターの `height` プロパティを `100px` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

`.cat-right-ear` セレクターの `width` プロパティを `100px` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

`.cat-right-ear` セレクターの `background-color` プロパティを `white` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

`.cat-right-ear` セレクターの `border-left` プロパティを `35px solid blue` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

`.cat-right-ear` セレクターの `border-right` プロパティを `35px solid blue` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

`.cat-right-ear` セレクターの `border-top` プロパティを `35px solid red` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

`.cat-right-ear` セレクターの `border-bottom` プロパティを `35px solid red` に設定してください。 セミコロンを忘れないようにしましょう。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
