---
id: 646c47867800472a4ed5d2ea
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML の基本的な構造から始めましょう。 Add a `DOCTYPE` reference of `html` and an `html` element with its `lang` attribute set to `en`. Also, add a `head` and a `body` element within the `html` element.

# --hints--

You should have the `DOCTYPE` declaration of `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

`DOCTYPE` 宣言は HTML の先頭におく必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

You should have an opening `html` tag with a `lang` attribute of `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

You should have a closing `html` tag.

```js
assert(code.match(/<\/html>/i));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body>/i));
```

Your `body` element should come after the `head` element.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Your `head` and `body` elements should be inside the `html` element.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
