---
id: 646c47867800472a4ed5d2ea
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML の基本的な構造から始めましょう。 `html` を指定する `DOCTYPE` 参照と、`lang` 属性として `en` を設定した `html` 要素を追加してください。 また、その `html` 要素の中に、`head` 要素と `body` 要素を追加してください。

# --hints--

`html` の `DOCTYPE` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

`DOCTYPE` 宣言は HTML の先頭に置く必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`lang` 属性が `en` に設定された `html` の開始タグが必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` の終了タグが必要です。

```js
assert(code.match(/<\/html>/i));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body>/i));
```

`body` 要素を `head` 要素の後に置く必要があります。

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

`head` 要素と `body` 要素は `html` 要素の内側に置いてください。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
