---
id: 646c586be7180e40ddf74ff6
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`body` 要素の `background-color` を `#c9d2fc` にしてください。

# --hints--

`body` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` セレクターの `background-color` プロパティを `#c9d2fc` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(201, 210, 252)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```

