---
id: 646c48df8674cf2b91020ecb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Within your `head` element, add a `meta` tag with the `charset` attribute of `utf-8`. Also add a `title` element with the text `fCC Cat Painting`.

# --hints--

You should add a single `meta` element.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Your `meta` element should have a `charset` attribute set to `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

You should add a single `title` element.

```js
assert(document.querySelectorAll('title').length === 1);
```

Your `title` element should have the text `fCC Cat Painting`. Note that spelling and casing matter.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
