---
id: 646c48df8674cf2b91020ecb
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素の中に `meta` タグをネストして、`charset` 属性を `utf-8` にしてください。 また、`title` 要素を追加して、`fCC Cat Painting` というテキストを設定してください。

# --hints--

`meta` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 要素は `charset` 属性を `utf-8` に設定する必要があります。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

`title` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 要素は `fCC Cat Painting` というテキストを持つ必要があります。 スペルや大文字小文字の区別に気をつけましょう。

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
