---
id: 646c5d7057c45f432fcdd46c
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

クラスセレクターを使用して、`.cat-head` に該当する要素の幅を `205px`、高さを `180px` に設定してください。 また、境界線を `1px solid #000`、`border-radius` を `46%` に設定してください。

# --hints--

`.cat-head` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

`.cat-head` セレクターで `width` を `205px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

`.cat-head` セレクターで `height` を `180px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

`.cat-head` セレクターで `border` を `1px solid #000` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

`.cat-head` セレクターで `border-radius` を `46%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
