---
id: 6476fc5cf14b276e6d04e82a
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

クラスセレクターを使用して、`.box` に該当する要素の幅を `200px`、高さを `600px`、背景色を `#000` に設定してください。 また、`position` を `absolute` に、`top` を `800px` に、`left` を `650px` に設定してください。

# --hints--

`.box` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

`.box` セレクターの `width` プロパティを `200px` に設定してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

`.box` セレクターの `height` プロパティを `600px` に設定してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

`.box` セレクターの `background-color` プロパティを `#000` に設定してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

`.box` セレクターの `position` プロパティを `absolute` に設定してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

`.box` セレクターの `left` プロパティを `650px` に設定してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

`.box` セレクターの `top` プロパティを `800px` に設定してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

