---
id: 6476f5c17f99146071ee884c
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

何も起こらないことがわかります。 `top` と `left` をそれぞれ `100px` に設定したにも関わらず、`.cat-head` の要素の配置は変わりません。 しかし、位置指定を `relative` にすると状況が変わります。

`relative` という値を指定すると、要素がドキュメントの通常のフローに従って配置される点は変わりませんが、`top`、`left`、`bottom`、`right` の値が有効になります。

では、`.cat-head` の位置指定を `static` の代わりに `relative` に設定しましょう。`top` プロパティと `left` プロパティはそのままにしてください。

# --hints--

`.cat-head` セレクターの `position` プロパティを `relative` に設定する必要があります。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

`.cat-head` セレクターの `top` プロパティを `100px` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

`.cat-head` セレクターの `left` プロパティを `100px` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
