---
id: 6476f7a4827bcc61682f2347
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

次に学ぶ position プロパティは `absolute` です。 `position` プロパティに `absolute` の値を設定すると、要素が通常のドキュメントフローから取り除かれ、`top`、`right`、`bottom`、`left` のプロパティで配置が決まるようになります。

では、`.cat-head` の要素の position プロパティを `absolute` に設定し、`top` プロパティと `left` プロパティを任意のピクセル値に設定してください。

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

`.cat-head` セレクターの `position` プロパティを `absolute` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

`.cat-head` セレクターの `top` プロパティを、単位が `px` の数値に設定してください。 必ずセミコロンを追加してください。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

`.cat-head` セレクターの `left` プロパティを、単位が `px` の数値に設定してください。 必ずセミコロンを追加してください。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
