---
id: 6476fb7cbaafa36d65e9cf35
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`position` プロパティの値に `fixed` を使うと、ページがどこにスクロールされても要素の位置が変わらないように固定することができます。

`fixed` 位置指定の効果を確認するには、いくつかマークアップを追加する必要があります。 HTML に、`box` クラスを持つ `div` 要素を追加してください。

# --hints--

`div` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('div').length === 2);
```

`div` 要素には `box` というクラスが必要です。

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
