---
id: 6476fd4213318f6ee211028a
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

それでは、`.cat-head` に該当する要素の `position` プロパティの値を `fixed` に置き換えてください。 `top` と `left` はどちらも変更しないでください。

その後、ページを上下にスクロールして、`fixed` の効果を確認しましょう。

# --hints--

`.cat-head` セレクターの `position` プロパティを `fixed` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

`.cat-head` セレクターの `top` プロパティは `100px` に設定されている必要があります。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

`.cat-head` セレクターの `left` プロパティは `100px` に設定されている必要があります。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
