---
id: 6477062778c85972eb648030
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

では猫の頭を中央に表示させましょう。

`.cat-head` の要素の `position` プロパティを `absolute` に設定してください。 そして `right`、`left`、`top`、`bottom` それぞれのプロパティに `0` の値を設定し、さらに全方向の `margin` プロパティを `auto` に設定してください。 これは、要素を縦横中央に配置する手法の一例です。

# --hints--

`.cat-head` セレクターの `position` プロパティを `absolute` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

`.cat-head` セレクターの `top` プロパティを `0` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

`.cat-head` セレクターの `left` プロパティを `0` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

`.cat-head` セレクターの `right` プロパティを `0` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

`.cat-head` セレクターの `bottom` プロパティを `0` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

`.cat-head` セレクターの `margin` プロパティを `auto` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
