---
id: 64770351e8586671ec0911f0
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

最後の position プロパティの値は `sticky` です。 `sticky` 位置指定は、`relative` と `fixed` のハイブリッドのようなものです。 sticky を使うと、スクロール位置に応じて、要素を親要素またはビューポートの特定の位置に **stick** させる (くっ付ける) ことができます。

`.cat-head` の `position` プロパティの値を `sticky` に変更し、`top` を `0` に設定して、`left` プロパティとその値は削除してください。

**注**: `sticky` の効果を確認するには、 `.cat-head` に該当する `div` 要素の前後に、ある程度の長さのテキストを追加する必要があります。 その状態でページ下部へスクロールしていくと、`.cat-head` の要素が上端にくっ付いて残るのがわかります。

# --hints--

`.cat-head` セレクターの `position` プロパティを `sticky` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

`.cat-head` セレクターの `top` プロパティを `0` に設定してください。 必ずセミコロンを追加してください。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

コード内に `left` プロパティとその値がないようにしてください。

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
