---
id: 647d855ac12cd436059acd39
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

CSS で全称セレクターを使用して、`box-sizing: border-box;` を追加してください。 この指定により、要素の幅と高さにパディングと境界線が含まれるようになります。

# --hints--

全称セレクター (`*`) が必要です。

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

全称セレクターの `box-sizing` プロパティを `border-box` に設定する必要があります。

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
