---
id: 6676a8a8d9861319ee901a90
title: ステップ 24
challengeType: 0
dashedName: step-24
---

# --description--

では猫の耳を作成しましょう。

確認に使ったコードをクリアするため、`.cat-right-ear` セレクターとその中の全プロパティは削除してください。

クラスセレクターを使用して、`.cat-left-ear` の要素について、左側と右側の境界線をそれぞれ `35px solid transparent` に設定してください。 また、下側の境界線を `70px solid #5e5e5e` に設定してください。

# --hints--


`.cat-right-ear` セレクターが `height` プロパティを持たないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

`.cat-right-ear` セレクターが `width` プロパティを持たないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


`.cat-right-ear` セレクターが `border-left` プロパティを持たないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

`.cat-right-ear` セレクターが `border-right` プロパティを持たないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

`.cat-right-ear` セレクターが `border-top` プロパティを持たないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

`.cat-right-ear` セレクターが `border-bottom` プロパティを持たないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

`.cat-right-ear` セレクターは必要ありません。

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

`.cat-left-ear` セレクターが必要です。

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

`.cat-left-ear` セレクターの `border-left` プロパティを `35px solid transparent` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

`.cat-left-ear` セレクターの `border-right` プロパティを `35px solid transparent` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

`.cat-left-ear` セレクターの `border-bottom` プロパティを `70px solid #5e5e5e` に設定してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
