---
id: 641da7bfbc7f0f477438ad8a
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Ora è il momento di chiamare la funzione `displayAuthors`. Ma ancora, non vuoi popolare la pagina con tutti gli autori contemporaneamente. Invece, puoi estrarre una parte degli autori con le variabili `startingIndex` e `endingIndex`. Il modo migliore per farlo è il metodo per array `.slice()`.

Per prima cosa, rimuovi l'istruzione console log che mostra `authorDataArr`. Poi, chiama la funzione `displayAuthors` con l'array `authorDataArr` e `.slice()`. Usa la variabile `startingIndex` per il punto di partenza e la variabile `endingIndex` per il punto finale.

# --hints--

Dovresti assegnare `data` alla variabile `authorDataArr`

```js
assert.match(code, /authorDataArr\s*=\s*data\s*;?/)
```

Dovresti rimuovere l'istruzione console log che mostra `authorDataArr`

```js
assert.notMatch(code, /console\.log\(\s*("|'|`)Author\s+Data\s+Array:\s*\1\s*,\s*authorDataArr\s*\)\s*;?/)
```

Dovresti chiamare la funzione `displayAuthors`.

```js
assert.match(code, /displayAuthors\(/)
```

Dovresti chiamare la funzione `displayAuthors` e passare l'argomento `authorDataArr.slice()`.

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(/)
```

Il primo parametro del metodo `slice()` dovrebbe essere `startingIndex`.

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex/)
```

Il secondo parametro del metodo `slice()` dovrebbe essere `endingIndex`.

```js
assert.match(code, /displayAuthors\(\s*authorDataArr\.slice\(\s*startingIndex\s*,\s*endingIndex\s*\)\s*\)\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    --fcc-editable-region--
    authorDataArr = data;
    console.log("Author Data Array:", authorDataArr);

    --fcc-editable-region-- 
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
    </div>
  `;
  });
};
```
