---
id: 641da5462576784453146ec2
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Ora che hai i dati che vuoi, puoi usarli per popolare l'interfaccia utente. Ma i dati recuperati contengono un array di 26 autori e aggiungerli alla pagina tutti allo stesso tempo potrebbe determinare scarse prestazioni.

Invece, dovresti aggiungere 8 autori alla volta e avere un pulsante per aggiungerne altri 8 fino a quando non ci sono più dati da visualizzare.

Usa `let` per creare 2 variabili chiamate `startingIndex` e `endingIndex`, e assegna loro i valori numerici `0` e `8`, rispettivamente. Inoltre, crea una variabile `authorDataArr` con `let` e impostala su un array vuoto.

# --hints--

Dovresti usare `let` per dichiarare una variabile chiamata `startingIndex`.

```js
assert.match(code, /let\s+startingIndex/)
```

Dovresti impostare la variabile `startingIndex` a `0`.

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

Dovresti usare `let` per dichiarare una variabile chiamata `endingIndex`.

```js
assert.match(code, /let\s+endingIndex/)
```

Dovresti impostare la variabile `endingIndex` a `8`.

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

Dovresti usare `let` per dichiarare una variabile chiamata `authorDataArr`.

```js
assert.match(code, /let\s+authorDataArr/)
```

Dovresti impostare la variabile `authorDataArr` su un array vuoto (`[]`).

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
