---
id: 641da895fb7ec648a5bdf19c
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Ora hai tutto quello che occorre includere nell'interfaccia utente. Il prossimo passo è far sì che il pulsante `Load More Authors` recuperi altri autori ogni volta che viene cliccato. Puoi farlo aggiungendo un evento `click` al pulsante e incrementando attentamente le variabili `startingIndex` e `endingIndex`.

Crea una funzione `fetchMoreAuthors` con la sintassi delle funzioni freccia. Non aggiungere ancora nulla al suo interno. Assicurati di usare le parentesi graffe perché avrai più di un'espressione all'interno della funzione.

# --hints--

Dovresti usare `const` per creare una funzione `fetchMoreAuthors`.

```js
assert.match(code, /const fetchMoreAuthors\s*=\s*/)
```

`fetchMoreAuthors` dovrebbe essere una funzione.

```js
assert.isFunction(fetchMoreAuthors)
```

La funzione `fetchMoreAuthors` non dovrebbe prendere alcun parametro.

```js
assert.match(code, /const fetchMoreAuthors\s*=\s*\(\s*\)\s*/)
```

La funzione `fetchMoreAuthors` dovrebbe usare la sintassi freccia.

```js
assert.match(code, /const fetchMoreAuthors\s*=\s*\(\)\s*=>\s*/)
```

La funzione `fetchMoreAuthors` dovrebbe essere vuota.

```js
assert.match(code, /const fetchMoreAuthors\s*=\s*\(\)\s*=>\s*\{\s*\}/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    authorDataArr = data;
    displayAuthors(authorDataArr.slice(startingIndex, endingIndex));  
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

--fcc-editable-region--

--fcc-editable-region--  

const displayAuthors = (authors) => {
  authors.forEach(({ author, image, url, bio }, index) => {
    authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
      <img class="user-img" src="${image}" alt="${author} avatar" />
      <p class="bio">${bio}</p>
      <a class="author-link" href="${url}" target="_blank">${author}'s author page</a>
    </div>
  `;
  });
};
```
