---
id: 68c18b1d0fce081f85312d97
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

In a previous lesson, you learned that an attribute prefixed with a single underscore is meant for internal use by convention.

Modify both `name` and `level` attributes into `_name` and `_level`, since these are not supposed to be modified from outside their class.

Note that this does not prevent the attribute from being accessed or modified outside the class. Also, in Python there's always a way to access private attributes (prefixed with a double underscore) as well.

# --hints--

Your `__init__` method should set `self._name` to `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Your `__init__` method should set `self._level` to `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
