---
id: 68c15a2a1f0e70ca7154aba5
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you are going to build a salary tracking system for employees.

Start by creating a class named `Employee`. Inside it create an `__init__` method with `self`, `name`, and `level` parameters. Within the `__init__` method, assign `name` and `level` to the instance attributes with the same name.

# --hints--

You should have a class named `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Your `Employee` class should have an `__init__` method with three parameters `self`, `name`, and `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Your `__init__` method should set `self.name` to `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Your `__init__` method should set `self.level` to `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
