---
id: 6140cdebd39d6a101e747529
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Crea un selettore per gli elementi `.line`. Inizia impostando `background-color` su `black`, `width` al `50%` e `height` a `2px`.

# --hints--

Dovresti avere un selettore `.line`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

Il selettore `.line` dovrebbe avere una proprietà `background-color` impostata su `black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

Il selettore `.line` dovrebbe avere una proprietà `width` impostata al `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

Il selettore `.line` dovrebbe avere una proprietà `height` impostata a `2px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
