---
id: 6140c7e645d8e905819f1dd4
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Inizia con il testo HTML standard. Aggiungi la dichiarazione `DOCTYPE` e gli elementi `html` (specificando che la pagina è in inglese), `head` e `body`.

Aggiungi l'elemento `meta` per il `charset` corretto, l'elemento `title` e un elemento `link` per il file `./styles.css`.

Dai all'elemento `title` il testo `Ferris Wheel`.

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Dovresti definire il tipo di documento come `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un carattere `>`.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

L'elemento `html` dovrebbe avere una tag di apertura con un attributo `lang` di `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/html\s*>/);
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio dell'HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Dovresti avere un tag `head` di apertura.

```js
assert.match(code, /<head\s*>/i);
```

Dovresti avere un tag `head` di chiusura.

```js
assert.match(code, /<\/head\s*>/i);
```

Dovresti avere un tag `body` di apertura.

```js
assert.match(code, /<body\s*>/i);
```

Dovresti avere un tag `body` di chiusura.

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Il codice dovrebbe avere un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

L'elemento `meta` dovrebbe avere un attributo `charset` con il valore `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il progetto dovrebbe avere il titolo `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Fai attenzione all'uso delle maiuscole/minuscole e all'ortografia nel titolo.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Il codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
