---
id: 6140c9d35015ae0ba0c250e8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi un `div` all'interno del tuo elemento `body` e dargli una `class` di `wheel`.

All'interno del nuovo `div`, aggiungi sei elementi `span` con un attributo `class` di `line`, e sei elementi `div` con `class` con valore `cabin`.

# --hints--

Dovresti creare un nuovo `div` all'interno del tuo elemento `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

L'elemento `div` dentro l'elemento `body` dovrebebe avere un attributo `class` con valore di `wheel`.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

L'elemento `.wheel` dovrebbe avere sei elementi `span` all'interno.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

I sei elementi `span` dentro `.wheel` dovrebbero avere un attributo `class` di `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

Il tuo elemento `.wheel` dovrebbe avere sei elementi `div` all'interno.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

I tuoi sei elementi `div` all'interno dell'elemento `.wheel` dovrebbero avere un attributo `class` con valore di `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
