---
id: 6140cd32d018ed0f600eefce
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Imposta la proprietà `position` del selettore `.wheel` su `absolute`. Imposta entrambe le proprietà `height` e `width` a `55vw`.

# --hints--

Il selettore `.wheel` dovrebbe avere una proprietà `position` impostata su `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

Il selettore `.wheel` dovrebbe avere una proprietà `height` impostata su `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

Il selettore `.wheel` dovrebbe avere una proprietà `width` impostata su `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
