---
id: 67ecfbfe17c7d24465e0de88
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Inside the `form` element after the label, add an input with an `id` of `search-input` and a `type` of `search`.

# --hints--

You should have an `input` inside the `form` element with an `id` of `search-input`.

```js
const form = document.querySelector("form");
assert.exists(form);

const input = form.querySelector("input");
assert.exists(input);
assert.strictEqual(input.tagName, "INPUT");
assert.strictEqual(input.getAttribute("id"), "search-input");
```

Your `input` should have a `type` value of `search`.

```js
const form = document.querySelector("form");
assert.exists(form);

const input = form.querySelector("input");
assert.exists(input);
assert.strictEqual(input.tagName, "INPUT");
assert.strictEqual(input.getAttribute("type"), "search");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  
  return (
    <div id="search-container">
    --fcc-editable-region--
      <form>
        <label htmlFor="search-input">Search for fruits:</label>
        
      </form>
    --fcc-editable-region--
    </div>
  );
}
```
