---
id: 67f8ef57b9974052e5ef6b30
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Use the ternary operator inside the `results` `div` to conditionally display content based on the `results` array. When the array contains items, display each fruit name inside a `div` with a class of `result-item`. When the array is empty, display a message that reads exactly: `No results found`.

# --hints--

You should use the ternary operator to check if `results.length > 0`, and inside the true branch, use `results.map()` to return a `<p>` for each fruit. Each `<p>` must have the class `result-item` and display the fruit name.

```js
assert.match(
  code,
  /{results\.length\s*>\s*0\s*\?\s*\(\s*results\.map\(\s*\(?\w+,\s*\w+\)?\s*=>\s*\(\s*<p[^>]*className=["']result-item["'][^>]*>.*?<\/p>\s*\)\s*\)\s*\)/s
)
```

In the false branch (after the colon), return a single `<p>` element with the text `No results found`.

```js
assert.match(code, /:\s*\(\s*<p>\s*No results found\s*<\/p>\s*\)/s)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  const [query, setQuery] = useState('');
  const [results, setResults] = useState([]);

  function handleSubmit(e) {
    e.preventDefault();
  }

  return (
    <div id="search-container">
      <form onSubmit={handleSubmit}>
        <label htmlFor="search-input">Search for fruits:</label>
        <input
          id="search-input"
          type="search"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        />
      </form>
      --fcc-editable-region--
      <div id="results">
      </div>
      --fcc-editable-region--
    </div>
  );
}
```
