---
id: 6813fe47f24095890691642f
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

You should define a `handleSubmit` function that accepts an event object and calls `e.preventDefault()`. Next, pass the `handleSubmit` function to the `onSubmit` attribute of the `<form>` element so the form submission can be handled correctly.

# --hints--

You should define a `handleSubmit` function that calls `e.preventDefault()`.

```js
assert.match(
  code,
  /(function\s+handleSubmit\s*\(\s*\w+\s*\)|const\s+handleSubmit\s*=\s*\(?\s*\w+\s*\)?\s*=>)\s*{[^}]*\.\s*preventDefault\s*\(\s*\)/s
);
```

The `handleSubmit` function should be the value of the `onSubmit` attribute of the `<form>` element.

```js
assert.match(
  code,
  /<form[^>]*onSubmit\s*=\s*{handleSubmit}/
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Fruits Search</title>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
   <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { FruitsSearch } from './index.jsx';
    ReactDOM.createRoot(document.getElementById('root')).render(<FruitsSearch />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#search-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#search-input {
  padding: 10px;
  width: 80%;
  border: 1px solid #ccc;
  border-radius: 5px;
  margin-bottom: 10px;
}

#results {
  text-align: left;
  max-height: 150px;
  overflow-y: auto;
}
.result-item {
  padding: 5px;
  border-bottom: 1px solid #ddd;
}
```

```jsx
const { useState, useEffect } = React;

export function FruitsSearch() {
  const [query, setQuery] = useState('');
  const [results, setResults] = useState([]);
  --fcc-editable-region--

  return (
    <div id="search-container">
      <form>
        <label htmlFor="search-input">Search for fruits:</label>
        <input
          id="search-input"
          type="search"
          value={query}
          onChange={(e) => setQuery(e.target.value)}
        />
      </form>
  --fcc-editable-region--
    </div>
  );
}
```
