---
id: 66c644c08ebcd2ef90c750bd
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

In the previous lecture videos, you learned how to add items to the end of an array using the `push` method like this:

```js
fruits.push('pear');
```

Using the `push` method, add the string `"Apples"` to the `shoppingList` array.

# --hints--

Your `shoppingList` array should have one item in it.

```js
assert.lengthOf(shoppingList, 1);
```

Your `shoppingList` array should have the string `"Apples"` in it. Double check for spelling and capitalization errors.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
