---
id: 66c64095f61166ede6e39a84
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

For this shopping list, you will use an array to represent the items you need to buy.

In the previous lecture videos, you learned how to create arrays like this:

```js
const fruits = ['apple', 'banana', 'orange'];
```

In this step, create a variable called `shoppingList` and assign it an empty array.

# --hints--

You should have a variable called `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

You should assign an array to the `shoppingList` variable.

```js
assert.isArray(shoppingList);
```

Your array should be empty.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
