---
id: 66c73a7798f6f62b2ae58f22
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

Now it is time to add more items to the beginning of the grocery list.

Start by adding a `console.log` statement that logs the message `"It might be nice to get a dessert."`

Below that `console` statement, use the correct array method to add the string `"Chocolate Cake"` to the beginning of the `shoppingList` array.

Finally, add a `console.log` and call the `getShoppingListMsg` function inside of the `console.log` to see the updated list logged to the console.

# --hints--

You should have a `console.log` that logs the message `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Your `shoppingList` array should have a total of `6` items.

```js
assert.lengthOf(shoppingList, 6);
```

You should use the `unshift` method to add the string `"Chocolate Cake"` to the beginning of the `shoppingList` array.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

You should call `getShoppingListMsg(shoppingList)` inside of the `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
