---
id: 66cbe2319d3845545a293a0b
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now it is time to see the message logged to the console.

Add a `console.log` and call the `getShoppingListMsg` function inside of the `console.log` to see the message logged to the console.

# --hints--

You should call the `getShoppingListMsg(shoppingList)` function inside of the `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
