---
id: 661e131f068359c3ccf2f4d6
title: Basic Functions Exercise B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

Write a function, named `multiply`, that takes two arguments and returns their product.

# --hints--

You should have a function named `multiply` that takes two arguments.

```js
assert.strictEqual(multiply(2, 3), 6);
```

Your function should take in an integer as an argument.

```js
assert.strictEqual(multiply(5, 5), 25);
```

You should return the product of the two integers.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

## --solution--

```js 
function multiply(a, b) {
  return a * b;
}
```
