---
id: 61695ab9f6ffe951c16d03dd
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素中嵌套一個 `head` 元素。 在 `head` 元素之後，添加一個 `body` 元素。

# --hints--

你應該有一個打開的 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

你的 `head` 和 `body` 元素應該是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

你的 `head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

你的 `body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
