---
id: 61695197ac34f0407e339882
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

正如你在之前的項目中看到的，網頁應該以 `DOCTYPE html` 聲明開頭，然後是 `html` 元素。

在文檔頂部添加 `DOCTYPE html` 聲明，然後添加 `html` 元素。 給 `html` 元素一個值爲 `en` 的 `lang` 屬性。

# --hints--

你的代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 聲明後面加一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該將文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在類型之後使用 `>` 關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個帶有值爲 `en` 的 `lang` 屬性的開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的 `DOCTYPE` 聲明應該位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
