---
id: 616971b3cd990186b66c99a1
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

現在你的外部 CSS 文件已設置好，你可以開始設置頁面樣式了。

提醒一下，以下是定位段落元素並將其向右對齊的方法：

```css
p {
  text-align: right;
}
```

創建一個定位 `h1` 元素的新 CSS 規則，並將其 `text-align` 屬性設置爲 `center`。

# --hints--

你應該使用 `h1` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

你的 `h1` CSS 規則應該將 `text-align` 屬性設置爲 `center`。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
