---
id: 61696ef7ac756c829f9e4048
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

在這個項目中，你將使用外部 CSS 文件來設置頁面樣式。 我們已經爲你創建了一個 `styles.css` 文件。 但在你使用它之前，你需要將它鏈接到頁面上。

在 `head` 元素中嵌套一個 `link` 元素。 給它一個 `rel` 屬性，設置爲 `stylesheet`，和一個 `href` 屬性，設置爲 `styles.css`。

# --hints--

你應該有一個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

您的 `link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
