---
id: 61695c4aad56f95497c19583
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

請記住，`title` 元素爲搜索引擎提供了有關頁面的額外信息。 它還通過以下兩種方式顯示 `title` 元素的內容：

* 當頁面打開時，在標題欄中
* 當你把鼠標懸停在該頁面上時，在瀏覽器標籤中。 即使該標籤未被激活，一旦你將鼠標懸停在該標籤上，`title` 文本就會顯示出來。

在 `head` 元素中，嵌套一個帶有文本 `Colored Markers` 的 `title` 元素。

# --hints--

你應該有一個 `title` 開始標籤。

```js
assert(code.match(/<title\s*>/i));
```

你應該有一個 `title` 結束標籤。

```js
assert(code.match(/<\/title\s*>/i));
```

你的項目應該有標題 `Colored Markers`。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

請記住，標題文本的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
