---
id: 61695d1fbc003856628bf561
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

要告訴瀏覽器如何對頁面上的字符進行編碼，請將 `charset` 設置爲 `utf-8`。 `utf-8` 是一個通用字符集，幾乎包括所有人類語言中的每個字符。

在 `head` 元素內，嵌套一個 `meta` 元素，其中 `charset` 屬性值爲 `utf-8`。 請記住，`meta` 元素是自閉合的，不需要結束標籤。

# --hints--

你應該有一個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

你的 `meta` 元素應該是一個自閉合元素。

```js
assert(code.match(/<\/meta>/i) === null);
```

你的 `meta` 元素應該有一個 `charset` 屬性值爲 `utf-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
