---
id: 616d525007b8c5d8b3308b1c
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

如果你的筆在頁面上居中，看起來會更好。 一個簡單的方法是使用 `margin` <dfn>簡寫屬性</dfn>。

在上一個項目中，你使用 `margin-top` 和 `margin-left` 等屬性分別設置元素的邊距區域。 `margin` 簡寫屬性可以方便地同時設置多個邊距區域。

要將筆在頁面上居中，請將其 `margin` 屬性設置爲 `auto`。 這會將 `margin-top`、`margin-right`、`margin-bottom` 和 `margin-left` 全部設置爲 `auto`。

# --hints--

你的 `.marker` CSS 規則應該有一個 `margin` 屬性設置爲 `auto`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
