---
id: 616d595270d902f0e7086e18
title: 步驟 15
challengeType: 0
dashedName: step-15
---

# --description--

你有三個單獨的標記 `div` 元素，它們看起來像一個大矩形。 你應該在它們之間添加一些空間，以便更容易查看每個元素。

當簡寫 `margin` 屬性有兩個值時，它將 `margin-top` 和 `margin-bottom` 設置爲第一個值，並將 `margin-left` 和 `margin-right` 設置爲第二個值。

在你的 `.marker` CSS 規則中，將 `margin` 屬性設置爲 `10px auto`。

# --hints--

你的 `.marker` CSS 規則應該有一個 `margin` 屬性設置爲 `10px auto`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
