---
id: 616d50b93ba424d6282c99cf
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

The background color was applied, but since the marker `div` element is empty, it doesn't have any height by default.

In your `.marker` CSS rule, set the `height` property to `25px` and the `width` property to `200px`

# --hints--

你的 `.marker` CSS 規則應該將 `width` 屬性設置爲 `200px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

你的 `.marker` CSS 規則應該將 `height` 屬性設置爲 `25px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
