---
id: 616d4a84b756d9c4b8255093
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

是時候爲彩筆添加一些顏色了。 請記住，爲元素添加顏色的一種方法是使用 <dfn>color 關鍵字</dfn>，例如 `black`、`cyan` 或 `yellow`。

提醒一下，以下是定位類 `freecodecamp` 的方法：

```css
.freecodecamp {

}
```

創建一個以 `marker` 類爲目標的新 CSS 規則，並將其 `background-color` 屬性設置爲 `red`。

**注意：** 你將不會在添加 CSS 後看到任何更改。

# --hints--

你應該創建一個類選擇器來定位 `marker` 類。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

你的 `.marker` CSS 規則應該將 `background-color` 屬性設置爲 `red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
