---
id: 61764c602bee6974e7790f35
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

在學校裏，你可能已經知道紅色、黃色和藍色是原色，並學會了如何通過混合這些顏色來創造新的顏色。 然而，這是一個過時的模型。

目前，有兩種主要的顏色模型：用於電子設備的 <dfn>疊加</dfn> RGB（紅、綠、藍）模型，以及 <dfn>消減</dfn> CMYK（青色、品紅色、黃色，黑色）用於打印的模型。 在這個項目中，你將使用 RGB 模型。

首先，將 `class` `one` 添加到第一個標記 `div` 元素。

```html
<div class="exampleClass1 exampleClass2">
  ...content
</div>
```

# --hints--

你應該將 class​​ `one` 添加到第一個標記 `div` 元素。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

你的第一個標記 `div` 應該有 `marker` 和 `one` class。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
