---
id: 6176519636a76b810ab1219a
title: 步驟 20
challengeType: 0
dashedName: step-20
---

# --description--

創建一個定位 `two` class 的 CSS 規則，並將其 `background-color` 屬性設置爲 `green`。

此外，創建一個定位 `three` class 的單獨 CSS 規則，並將其 `background-color` 設置爲 `blue`。

# --hints--

你應該使用 class 選擇器來定位 class `two`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

你的 `.two` CSS 規則應該將 `background-color` 屬性設置爲 `green`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

你應該使用 class 選擇器來定位 class `three`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

你的 `.three` CSS 規則應該將 `background-color` 屬性設置爲 `blue`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
