---
id: 61764c602bee6974e7790f35
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

爲了讓彩筆有不同的顏色，需要給每支彩筆設置一個不同的 class。 多個 classes 可以以空格分隔放在 `class` 屬性內來一起設置給一個元素。 比如，如下代碼給 `div` 元素添加了 `animal` 和 `dog` 類。

```html
<div class="animal dog">
```

將 class `one` 設置給代表第一支彩筆的 `div` 元素。


# --hints--

應該給代表第一支彩筆的 `div` 元素添加 `one` class。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

代表第一支彩筆的 `div` 的 class 應該爲 `marker` 和 `one`。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
