---
id: 617bbb6b97a83f6d1f7d6e38
title: 步驟 31
challengeType: 0
dashedName: step-31
---

# --description--

現在你已經熟悉了間色，你將學習如何創建<dfn>複色</dfn>。 複色是通過組合一種原色和一種鄰近的間色來創建的。

要創建複色橙色，更新 `.one` CSS 規則中的 `rgb` 函數，以使紅色是最大值，綠色是 `127`。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(255, 127, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
