---
id: 617bc3386dc7d07d6469bf20
title: 步驟 32
challengeType: 0
dashedName: step-32
---

# --description--

注意，要創建橙色，你需要增加紅色的強度，並降低綠色 `rgb` 值的強度。 這是因爲橙色是紅色和黃色的結合，再色輪上處於兩種顏色之間。

要創建複色亮綠色，需要組合青色和綠色。 更新 `.two` CSS 規則中的 `rgb` 函數，使綠色是最大值，藍色的值是 `127`。

# --hints--

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(0, 255, 127)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
