---
id: 617bd6ec666b1da2587e4e37
title: 步驟 34
challengeType: 0
dashedName: step-34
---

# --description--

還有三種複色：黃綠色（綠色+黃色）、天藍色（藍色+青色）、玫瑰色（紅色+品紅色）。

要創建黃綠色，更新 `.one` CSS 規則中的 `rgb` 函數，使紅色的值是 `127`，綠色是最大值。

要創建天藍色，更新 `.two` CSS 規則中的 `rgb` 函數，使綠色的值是 `127`，藍色是最大值。

對於有時被稱爲亮粉紅色的玫瑰色，更新 `.three` CSS 規則中的 `rgb` 函數，使藍色的值爲 `127`，紅色爲最大值。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(127, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(0, 127, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

你的 `.three` CSS 規則應該將 `background-color` 屬性設置爲 `rgb(255, 0, 127)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
