---
id: 617b65579ce424bf5f02ca73
title: 步驟 21
challengeType: 0
dashedName: step-21
---

# --description--

你已經學習過 RGB 顏色模型是疊加的。 這意味着顏色從黑開始，隨着紅色、綠色和藍色數值水平的不同而變化。

一個簡單的方法是使用 CSS `rgb` 函數。

創建一個定位 `container` 類的新 CSS 規則，並使用 `rgb(0, 0, 0)` 將其 `background-color` 設置爲黑色。

# --hints--

你應該使用類選擇器來定位類名爲 `container`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

你的 `.container` CSS 容器應當包含設一個屬性 `background-color` ，它的值被設置爲 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
