---
id: 617b8e0d93a8d10d9a90e720
title: 步驟 23
challengeType: 0
dashedName: step-23
---

# --description--

請注意，筆的 `background-color` 仍然是紅色的。 這是因爲你將 `rgb` 函數的紅色值設置爲最大值 `255`，即 100% 紅色，並將綠色和藍色值都設置爲 `0`。

現在使用 `rgb` 函數來設置其他顏色。

在 `.two` CSS 規則中，使用 `rgb` 函數將 `background-color` 設置爲綠色是最大值，而其他顏色的值是 `0`。 在 `.three` CSS 規則中，使用 `rgb` 函數將 `background-color` 設置爲藍色是最大值，而其他顏色的值是 `0`。

# --hints--

你的 `.two` CSS 規則不應使用 `green` 顏色關鍵字來設置 `background-color` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(0, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

你的 `.three` CSS 規則不應使用 `blue` 顏色關鍵字來設置 `background-color` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

你的 `.three` CSS 規則應該將 `background-color` 屬性設置爲 `rgb(0, 0, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
