---
id: 617bb9fdef27bc6aa0470ac2
title: 步驟 30
challengeType: 0
dashedName: step-30
---

# --description--

要創建最後一個間色品紅色，請更新 `.three` CSS 規則中的 `rgb` 函數以組合純藍色和純紅色。

# --hints--

你的 `.three` CSS 規則應該將 `background-color` 屬性設置爲 `rgb(255, 0, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 255, 0);
}

.two {
  background-color: rgb(0, 255, 255);
}

--fcc-editable-region--
.three {
  background-color: rgb(0, 0, 255);
}
--fcc-editable-region--

```
