---
id: 617bb5624a75e86463b7e638
title: 步驟 28
challengeType: 0
dashedName: step-28
---

# --description--

<dfn>間色</dfn>是你組合主色得到的顏色。 在上一步更改紅色、綠色和藍色值時，你可能已經注意到一些間色。

要創建第一個間色，黃色，請在 `.one` CSS 規則中更新 `rgb` 函數，以組合純紅色和純綠色。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(255, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
