---
id: 618a00ed1ca871a2b3aca0cb
title: 步驟 35
challengeType: 0
dashedName: step-35
---

# --description--

現在你已經學習了色輪上的所有原色、間色和複色，更容易理解其他顏色理論概念和它們如何影響設計。

首先，在 CSS 規則 `.one`、`.two` 和 `.three` 中，調整 `rgb` 函數中的值，將每個元素的 `background-color` 設置爲純黑色。 請記住 `rgb` 函數使用加成色模型。顏色起始爲黑色，隨紅色、綠色和藍色的值增加而變化。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性，設置爲 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

你的 `.three` CSS 規則應該將 `background-color` 屬性設置爲 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
