---
id: 618a0927005553b74bfa05ae
title: 步驟 36
challengeType: 0
dashedName: step-36
---

# --description--

色輪是一個圓圈，其中相似的顏色或<dfn>色調</dfn>彼此靠近，而不同的顏色相距較遠。 例如，純紅色介於玫瑰色和橙色之間。

色輪上相互對立的兩種顏色稱爲<dfn>補色</dfn>。 如果將兩種互補色組合在一起，它們會產生灰色。 但當它們並排放置時，這些顏色會產生強烈的視覺對比，顯得更亮。

在 `.one` CSS 規則的 `rgb` 函數中，將紅色值設置爲 `255` 的最大值以生成純紅色。 在 `.two` CSS 規則的 `rgb` 函數中，將綠色和藍色的值設置爲 `255` 的最大值以生成青色。

# --hints--

你的 `.one` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(255, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

你的 `.two` CSS 規則應該有一個 `background-color` 屬性設置爲 `rgb(0, 255, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
