---
id: 619b761916dac02643940022
title: 步驟 46
challengeType: 0
dashedName: step-46
---

# --description--

使用 CSS 將顏色應用於元素的一種非常常見的方法是使用<dfn>十六進制</dfn>或 hex 值。 雖然十六進制值聽起來很複雜，但它們實際上只是 RGB 值的另一種形式。

十六進制顏色值以 `#` 字符開頭，從 0-9 和 A-F 取六個字符。 第一對字符代表紅色，第二對代表綠色，第三對代表藍色。 例如，`#4B5320`。

在 `.green` 類選擇器中，將 `background-color` 屬性設置爲十六進制顏色代碼，其值 `00` 表示紅色，`FF` 表示綠色，`00` 表示藍色。

# --hints--

你的 `.green` CSS 規則應該將 `background-color` 屬性設置爲 `#00FF00`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
