---
id: 619b7c3c83de403126b69c1e
title: 步驟 47
challengeType: 0
dashedName: step-47
---

# --description--

你可能已經熟悉從 0 到 9 的十進制或以 10 爲底的值。 十六進制或以 16 爲底的值，從 0 到 9，然後是 A - F：

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

對於十六進制顏色，`00` 是該顏色的 0%，`FF` 是 100%。 所以 `#00FF00` 轉換爲 0% 紅色、100% 綠色和 0% 藍色，與 `rgb(0, 255, 0)` 相同。

通過將十六進制顏色的綠色值設置爲 `7F` 來降低綠色強度。

# --hints--

你的 `.green` CSS 規則應該將 `background-color` 屬性設置爲 `#007F00`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
