---
id: 61a489b8579e87364b2d2cdb
title: 步驟 49
challengeType: 0
dashedName: step-49
---

# --description--

你已經瞭解了幾種在 CSS 中設置純色的方法，但你也可以在元素上使用顏色過渡或<dfn>漸變色</dfn>。

漸變是一種顏色過渡到另一種顏色。 CSS `linear-gradient` 函數可讓你控制沿線的過渡方向以及使用的顏色。

要記住的一件事是 `linear-gradient` 函數實際上創建了一個 `image` 元素，並且通常與可以接受圖像作爲一個值的 `background` 屬性配對。

在 `.red` CSS 規則中，將 `background-color` 屬性更改爲 `background`。

# --hints--

你的 `.red` CSS 規則應該有一個 `background` 屬性，其值爲 `rgb(255, 0, 0)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
