---
id: 61a493ead935c148d2b76312
title: 步驟 50
challengeType: 0
dashedName: step-50
---

# --description--

`linear-gradient` 函數非常靈活——這是你將在本教程中使用的基本語法：

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` 是表明過渡的線的方向。 `color1` 和 `color2` 是顏色參數，它們是將在過渡本身中使用的顏色。 這些可以是任何類型的顏色，包括顏色關鍵字、十六進制、`rgb` 或 `hsl`。

現在，你將沿 90 度線將紅到綠漸變應用到第一支筆。

首先，在 `.red` CSS 規則中，將 `background` 屬性設置爲 `linear-gradient()`，並將它的值設置爲 `90deg`，作爲 `gradientDirection`。

# --hints--

你的 `.red` CSS 規則應該有一個 `background` 屬性，其值爲 `linear-gradient(90deg)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
