---
id: 61a49d15bdbb5e57cc6fd280
title: 步驟 54
challengeType: 0
dashedName: step-54
---

# --description--

顏色中間點（color-stop）允許你微調顏色沿漸變線的位置。 它們是像 `px` 這樣的長度單位，或者是在 `linear-gradient` 函數中定位顏色的百分比。

例如，在這個紅-黑漸變中，從紅色到黑色的過渡發生在漸變線的 90% 處，因此紅色佔據了大部分可用空間：

```css
linear-gradient(90deg, red 90%, black);
```

在 `linear-gradient` 函數中，在第一個紅色參數之後添加 `75%` 顏色中間點。 不要將顏色中間點添加到其他顏色參數。

# --hints--

你的 `.red` CSS 規則應該有一個 `background` 屬性設置爲 `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`。

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
