---
id: 61a5be9833e7dc178de2af10
title: 步驟 63
challengeType: 0
dashedName: step-63
---

# --description--

這看起來更好，但綠色筆的底部邊緣需要更暗以添加更多尺寸。

在同一個 `linear-gradient` 函數中，添加一個十六進制顏色代碼，其值 `11` 表示紅色，`6C` 表示綠色，`31` 將藍色作爲第三個顏色參數。

# --hints--

你的 `.green` CSS 規則應將 `background` 屬性設置爲 `linear-gradient(180deg, #55680D, #71F53E, #116C31)`。

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background: linear-gradient(180deg, #55680D, #71F53E);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
