---
id: 61a5bfe091060f1d6a629dd0
title: 步驟 64
challengeType: 0
dashedName: step-64
---

# --description--

即使沒有顏色中間點，你可能已經注意到綠色筆的顏色在與紅色筆相同的點位過渡。 第一種顏色位於漸變線的開頭 (0%)，第二種顏色位於中間 (50%)，最後一種顏色位於漸變線的結尾 (100%)。

`linear-gradient` 函數會自動爲你計算這些值，並在默認情況下沿漸變線均勻分佈顏色。

在 `.red` CSS 規則中，從 `linear-gradient` 函數中移除三個色標，以稍微清理你的代碼。

# --hints--

你的 `.red` CSS 規則應該有一個 `background` 屬性設置爲 `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`。

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}
--fcc-editable-region--

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
