---
id: 61addaf7e83988b59a7d18ff
title: 步驟 82
challengeType: 0
dashedName: step-82
---

# --description--

`border-left` 簡寫屬性允許你同時設置左邊框的寬度、樣式和顏色。

語法如下：

```css
border-left: width style color;
```

在 `.sleeve` CSS 規則中，用 `border-left` 簡寫屬性替換 `border-left-width`、`border-left-style` 和 `border-left-color` 屬性。 左邊框的寬度、樣式和顏色的值應該相同。

# --hints--

你的 `.sleeve` CSS 規則不應具有 `border-left-width` 屬性和值。

```js
assert(!/border-left-width/g.test(code));
```

你的 `.sleeve` CSS 規則不應具有 `border-left-style` 屬性和值。

```js
assert(!/border-left-style/g.test(code));
```

你的 `.sleeve` CSS 規則不應具有 `border-left-color` 屬性和值。

```js
assert(!/border-left-color/g.test(code));
```

你的 `.sleeve` CSS 規則應該有一個 `border-left` 簡寫屬性，並且其值爲 `10px solid black`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left-width: 10px;
  border-left-style: solid;
  border-left-color: black;
}
--fcc-editable-region--

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
