---
id: 61b093329e7fc020b43b1bb5
title: 步驟 75
challengeType: 0
dashedName: step-75
---

# --description--

你已經熟悉使用 `rgb` 函數來設置顏色。 要將 alpha 通道添加到 `rgb` 顏色，請改用 `rgba` 函數。

`rgba` 函數的工作方式與 `rgb` 函數類似，但在 `0` 到 `1.0` 之間多取一個數字作爲 alpha 通道：

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

你也可以使用帶有 `hsl` 和 `hex` 顏色的 alpha 通道。 你將很快看到如何實現這一點。

在 `.sleeve` 規則中，使用 `rgba` 函數將 `background-color` 屬性設置爲具有 50% 不透明度的純白色。

# --hints--

你的 `.sleeve` CSS 規則應將 `background-color` 屬性設置爲 `rgba(255, 255, 255, 0.5)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor === 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
