---
id: 61b0934c9e7fc020b43b1bb7
title: 步驟 77
challengeType: 0
dashedName: step-77
---

# --description--

創建一個新的 CSS 規則來定位類 `cap`。 在新規則中，將 `width` 屬性設置爲 `60px`，將 `height` 設置爲 `25px`。

# --hints--

你應該使用類選擇器來定位類 `cap`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap'));
```

你的 `.cap` CSS 規則應該將 `width` 屬性設置爲 `60px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap')?.width === '60px');
```

你的 `.cap` CSS 規則應該將 `height` 屬性設置爲`25px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
