---
id: 61b0936d9e7fc020b43b1bb8
title: 步驟 78
challengeType: 0
dashedName: step-78
---

# --description--

看起來你的筆筒不見了，但別擔心——它還在那裏。 發生的情況是，你的新筆蓋 `div` 佔據了筆的整個寬度，並將筆筒推到了下一行。

這是因爲 `div` 元素的 `display` 屬性默認是 `block`。 因此，當兩個 `block` 元素彼此相鄰時，它們會像實際的塊一樣堆疊。 例如，你的彩筆元素都是堆疊在一起的。

要將兩個 `div` 元素放置在同一行上，需要將它們的 `display` 屬性設置爲 `inline-block`。

創建一個新規則以同時定位 `cap` 和 `sleeve` 類，並將 `display` 設置爲 `inline-block`。

# --hints--

你應該使用類選擇器來定位 `cap` 和 `sleeve` 類。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

你的 CSS 規則應該將 `display` 屬性設置爲 `inline-block`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
