---
id: 61b0a44a6b865738ba49b9fb
title: 步驟 86
challengeType: 0
dashedName: step-86
---

# --description--

你要做的最後一件事是爲每支筆添加輕微的陰影，使它們看起來更加真實。

`box-shadow` 屬性允許你在元素周圍應用一個或多個陰影。 這是基本語法：

```css
box-shadow: offsetX offsetY color;
```

下面是 `offsetX` 和 `offsetY` 值的作用：

* `offsetX` 和 `offsetY` 都接受數字值，單位是 `px` 和其他 CSS 單位
* 正 `offsetX` 值將陰影向右移動，負值將它向左移動
* 正 `offsetY` 值將陰影向下移動，負值將它向上移動
* 如果你想要 `offsetX` 和 `offsetY` 其中一個值爲零，或兩個都爲零（`0`），那麼你不需要添加單位。 每種瀏覽器都知道零意味着沒有變化。

陰影的高度和寬度由應用陰影的元素的高度和寬度決定。 你也可以使用可選的 `spreadRadius` 值來擴大陰影的範圍。 之後會介紹更多。

首先爲紅色筆添加一個簡單的陰影。

在 `.red` CSS 規則中，添加 `box-shadow` 屬性，其中 `offsetX` 的值爲 `5px`，`offsetY` 爲 `5px`，`color` 爲 `red`。

# --hints--

你的 `.red` CSS 規則應該有一個 `box-shadow` 簡寫屬性，其值爲 `5px 5px red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
