---
id: 5d5a813321b9e3db6c106a46
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript es un poderoso lenguaje que le permite construir sitios web que sean interactivos. Para comenzar, cree su plantilla HTML estándar con un `DOCTYPE`, `html`, `head`, and `body`. Incluye un elemento `title` y un `link` para tu hoja de estilo, y tu etiqueta `meta` para el `charset`. Luego, crea un elemento `div` con el id `juego` dentro de tu `cuerpo`. Usa el texto `RPG - Dragon Repeller` para tu elemento `title`.

# --hints--

Tu código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debe incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debe cerrar la declaración `DOCTYPE` con `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de apertura. No olvides el atributo `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debe tener una etiqueta de apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debe tener una etiqueta de cierre `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Debe tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debe tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

El elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Tu código debe tener un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Tu elemento `meta` debe tener un atributo `charset` con el valor `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu código debe tener un elemento `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Debe tener un elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Tu elemento `div` debe tener un atributo `id` con el valor `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Tu elemento `div` debe estar dentro del elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
