---
id: 62a115879a6d51422652cbfc
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Cree cuatro elementos `div` dentro del elemento `#game`. Dales los siguientes valores respectivos de `id`, en orden: `stats`, `controls`, `monsterStats` y `text`.

# --hints--

Debe crear cuatro nuevos elementos `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Debe asignar a uno de los nuevos elementos `div` un `id` de `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Debe asignar a uno de los nuevos elementos `div` un `id` de `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Debe asignar a uno de los nuevos elementos `div` un `id` de `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Debe asignar a uno de los nuevos elementos `div` un `id` de `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Debe colocar los nuevos elementos `div` en el orden correcto.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Debe colocar los nuevos elementos `div` dentro del elemento `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
--fcc-editable-region--
<body>
  <div id="game">
  </div>
</body>
--fcc-editable-region--
</html>
```
