---
id: 62a3a0a3c0a4b32915d26a6e
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Elimina la línea `console.log("Hello World");` para comenzar a escribir el código del proyecto.

En JavaScript, se almacena un valor mediante una <dfn>variable</dfn>. Para usar una variable, primero hay que <dfn>declararla</dfn>. Por ejemplo, para declarar una variable llamada `camperbot`, habría que escribir:

```js
var camperbot;
```

La palabra clave `var` le dice a JavaScript que se está declarando una variable. Declara una variable llamada `xp`.

# --hints--

No debes tener la línea `console.log("Hello World");` en tu código.

```js
assert.notMatch(code, /console\.log\("Hello World"\);/);
```

Debes usar la palabra clave `var` para declarar una variable.

```js
assert.match(code, /var/);
```

Debes declarar una variable llamada `xp`.

```js
assert.match(code, /xp/);
```

No le debes asignar ningún valor a la variable.

```js
assert.notMatch(code, /var xp =/);
```

No olvides el punto y coma al final de la línea.

```js
assert.match(code, /var xp;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
