---
id: 62a3bec30ea7f941412512dc
title: Paso 33
challengeType: 0
dashedName: step-33
---

# --description--

Just like you did with the buttons, create variables for the following `id`s and use `querySelector()` to give them the element as a value:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats`, and `monsterName`.

Remember to declare these with the `const` keyword, and name the variables to match the `id`s.

# --hints--

You should declare a `text` variable with `const`.

```js
assert.match(code, /const text/);
```

Your `text` variable should have the value of your `#text` element.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

You should declare a `xpText` variable with `const`.

```js
assert.match(code, /const xpText/);
```

Your `xpText` variable should have the value of your `#xpText` element.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

You should declare a `healthText` variable with `const`.

```js
assert.match(code, /const healthText/);
```

Your `healthText` variable should have the value of your `#healthText` element.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

You should declare a `goldText` variable with `const`.

```js
assert.match(code, /const goldText/);
```

Your `goldText` variable should have the value of your `#goldText` element.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

You should declare a `monsterStats` variable with `const`.

```js
assert.match(code, /const monsterStats/);
```

Your `monsterStats` variable should have the value of your `#monsterStats` element.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

You should declare a `monsterName` variable with `const`.

```js
assert.match(code, /const monsterName/);
```

Your `monsterName` variable should have the value of your `#monsterName` element.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
    <script src="./script.js"></script>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
