---
id: 62a3a0a3c0a4b32915d26a6e
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Elimina la línea `console.log("Hello World");` para comenzar a escribir el código del proyecto.

En JavaScript, se almacena un valor mediante una <dfn>variable</dfn>. Para usar una variable, primero hay que <dfn>declararla</dfn>. Por ejemplo, para declarar una variable llamada `camperbot`, habría que escribir:

```js
let camperbot;
```

The `let` keyword tells JavaScript you are declaring a variable. Declara una variable llamada `xp`.

# --hints--

No debes tener la línea `console.log("Hello World");` en tu código.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

You should use the `let` keyword to declare your variable.

```js
assert.match(code, /let/);
```

Debes declarar una variable llamada `xp`.

```js
assert.match(code, /let\s+xp/);
```

No le debes asignar ningún valor a la variable.

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

No olvides el punto y coma al final de la línea.

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
