---
id: 60b69a66b6ddb80858c51579
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do elemento `head`, adicione uma `meta` com o atributo `charset` definido como `utf-8`. Adicione também um elemento `title` com o texto `Picasso Painting`.

# --hints--

Você deve adicionar exatamente um elemento `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

O elemento `meta` deve ter o atributo `charset`.

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

O atributo `charset` deve ser definido como `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Você deve adicionar exatamente um elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

O elemento `title` deve ter o texto `Picasso Painting`. Atente-se a ortografia e letras maiúsculas.

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
