---
id: 60b69a66b6ddb80858c5157b
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Para começar com a pintura, dê ao elemento `body` uma `background-color` de `rgb(184, 132, 46)`.

# --hints--

Você deve usar o seletor `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

O elemento `body` deve ter a propriedade `background-color` definida como `rgb (184, 132, 46)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(184, 132, 46)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  </body>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

```
