---
id: 60b69a66b6ddb80858c51578
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Comece configurando a estrutura do HTML. Adicione uma declaração `<!DOCTYPE>` e um elemento `html` com um atributo `lang` definido como `en`. Dentro do elemento `html`, adicione um elemento `head` e um elemento `body`.

# --hints--

O código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento para `html`.

```js
assert(code.match(/html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/html\s*>/gi));
```

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

O elemento `html` deve ter uma tag de abertura com um atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

Você deve ter uma tag de abertura para `head`.

```js
assert(code.match(/<head\s*>/i));
```

Você deve ter uma tag de fechamento para `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Você deve ter uma tag de abertura para `body`.

```js
assert(code.match(/<body\s*>/i));
```

Você deve ter uma tag de fechamento para `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

O elemento `head` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

--fcc-editable-region--

```

```css

```
