---
id: 60b69a66b6ddb80858c5157f
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Normalmente, o HTML é renderizado de cima para baixo. Os elementos na parte superior do código são posicionados na parte superior da página. No entanto, muitas vezes você pode querer mover os elementos para posições diferentes. Você pode fazer isso com a propriedade `position`.

Defina a propriedade `position` do elemento `#back-wall` como `absolute`. Uma posição `absolute` retira o elemento daquele fluxo do documento de cima para baixo e permite que você o ajuste relativo ao contêiner.

Quando um elemento é posicionado manualmente, você pode mudar seu layout com `top`, `left`, `right`e `bottom`. Defina o elemento `#back-wall` de modo que tenha um valor de `top` de `0` e um valor de `left` de `0`.


# --hints--

O seletor `#back-wall` deve ter uma propriedade `position` definida como `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.position === 'absolute');
```

O seletor `#back-wall` deve ter a propriedade `top` definida como `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.top === '0px');
```

O seletor `#back-wall` deve ter a propriedade `left` definida como `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.left === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
