---
id: 60b80da8676fb3227967a731
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Vincule o arquivo CSS agora, mesmo que você ainda não tenha escrito o CSS.

Adicione um elemento `link` com um `rel` de `stylesheet` e um `href` de `styles.css`.

# --hints--


O código deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

Você deve ter um elemento `link` de autofechamento.

```js
assert(document.querySelectorAll('link').length === 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
