---
id: 60b80da8676fb3227967a731
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Vincule o arquivo CSS agora, mesmo que você ainda não tenha escrito o CSS.

Adicione um elemento `link` com um `rel` de `stylesheet` e um `href` de `styles.css`.

# --hints--


O código deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

O elemento `link` deve ser de fechamento automático.

```js
assert(code.match(/<link[\w\W\s]+\/>/i));
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
