---
id: 615f3b091162165948e1cb33
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

如果使用浏览器的开发人员工具检查 `.label` 元素，你可能会注意到它实际上是 288 像素宽而不是 270 像素。 这是因为默认情况下，浏览器在确定元素大小时会包含边框和内边距。

要解决这个问题，创建一个 `*` 选择器并给它一个 `border-box` 的 `box-sizing` 属性来重置盒子模型。

# --hints--

应该创建一个 `*` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` 选择器应该有一个 `box-sizing` 属性设置为 `border-box`。

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
