---
id: 615f522dea4f776f64dc3e91
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

`rem` 单位代表 `root em`，与 `html` 元素的字体大小有关。

创建一个 `.small-text` 选择器，并将 `font-size` 设置为 `0.85rem`，计算结果大致为 `13.6px`（记得你之前将 `html` 设置为具有 `16px` 的 `font-size`）。

# --hints--

你应该有一个 `.small-text` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text'));
```

你的 `.small-text` 选择器应该有一个 `font-size` 属性设置为 `0.85rem`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text')?.fontSize === '0.85rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
      </div>
    </div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
