---
id: 615f7c71eab8218f846e4503
title: 步骤 51
challengeType: 0
dashedName: step-51
---

# --description--

在最后一个 `.divider` 之后，创建另一个 `p` 元素，其中包含文本 `Trans Fat 0g`。 通过将单词 `Trans` 包装在 `i` 元素中，将其变为斜体。 给新的 `p` 元素的 `class` 属性设置为 `indent no-divider`。 将 `Trans Fat 0g` 放在一个 `span` 元素中，以对齐它。

# --hints--

你应该在 `.daily-value` 元素的末尾创建一个新的 `p` 元素。

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.localName === 'p');
```

新 `p` 元素应该包含文本 `Trans Fat 0g`。

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.innerText === 'Trans Fat 0g');
```

新的 `p` 元素应该将 `class` 属性设置为 `indent no-divider`。

```js
const p = document.querySelector('.daily-value.small-text')?.lastElementChild;
assert(p?.classList?.contains('indent'));
assert(p?.classList?.contains('no-divider'));
```

新的 `p` 元素应该有一个 `i` 元素。

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelector('i'));
```

`i` 元素应该包含文本 `Trans`。

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelector('i')?.innerText === 'Trans');
```

你的新 `p` 元素应该有一个 `span` 元素。

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelector('span'));
```

你的 `span` 元素应该包含文本 `Trans Fat 0g`。

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelector('span')?.innerText?.match(/Trans[\s|\n]+Fat 0g/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
        <p>Calories</p>
      </div>
      <span>230</span>
    </div>
    <div class="divider medium"></div>
    --fcc-editable-region--
    <div class="daily-value small-text">
      <p class="bold right no-divider">% Daily Value *</p>
      <div class="divider"></div>
      <p><span><span class="bold">Total Fat</span> 8g</span> <span class="bold">10%</span></p>
      <p class="indent no-divider">Saturated Fat 1g <span class="bold">5%</span></p>
      <div class="divider"></div>
    </div>    
    --fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

.medium {
  height: 5px;
}

.small-text {
  font-size: 0.85rem;
}


.calories-info {
  display: flex;
  justify-content: space-between;
  align-items: flex-end;
}

.calories-info h2 {
  margin: 0;
}

.left-container p {
  margin: -5px -2px;
  font-size: 2em;
  font-weight: 700;
}

.calories-info span {
  margin: -7px -2px;
  font-size: 2.4em;
  font-weight: 700;
}

.right {
  justify-content: flex-end;
}

.indent {
  margin-left: 1em;
}

.daily-value p:not(.no-divider) {
  border-bottom: 1px solid #888989;
}
```
