---
id: 615f3ed16592445e57941aec
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

线条可以帮助分隔和分组重要内容，尤其是在空间有限的情况下。

在 `h1` 元素下方创建一个 `div` 元素，并为其提供一个 `class` 属性，设置为 `divider`。

# --hints--

应该创建一个新的 `div`。

```js
assert(document.querySelectorAll('div')?.length === 2);
```

新的 `div` 应该将 `class` 属性设置为 `divider`。

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

`.divider` 元素应该在 `.label` 元素中。

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

`.divider` 元素应该在 `h1` 元素之后。

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
