---
id: 615f34948891834dd77655a6
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

添加第二个 `p` 元素，其文本为 `Serving size 2/3 cup (55g)`。

# --hints--

应该有第二个 `p` 元素。

```js
assert(document.querySelectorAll('p')?.length === 2);
```

第二个 `p` 元素应该在 `body` 元素中。

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

第二个 `p` 元素应该在现有的 `p` 元素之后。

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

第二个 `p` 元素应包含文本 `Serving size 2/3 cup (55g)`。

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
