---
id: 615f2d4150fe0d4cbd0f2628
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `h1` 元素下方，添加一个 `p` 元素，其文本为 `8 servings per container`。

# --hints--

你应该添加一个新的 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的 `p` 元素应该在你的 `body` 元素中。

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

`p` 元素应该在 `h1` 元素之后。

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

`p` 元素应该有文本 `8 servings per container`。

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
