---
id: 615f2abbe7d18d49a1e0e1c8
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

我们已经为你提供了一个基本的 HTML 模板。

在你的 `body` 元素中创建一个 `h1` 元素，并为其指定文本 `Nutrition Facts`。

# --hints--

你应该添加一个新的 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应该在你的 `body` 元素中。

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

你的 `h1` 元素应包含文本 `Nutrition Facts`。

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
