---
id: 615f4bfb9de4a16703b56eb6
title: 步骤 23
challengeType: 0
dashedName: step-23
---

# --description--

`float` 属性用于将元素放置在其容器的左侧或右侧，允许其他内容（例如文本）环绕它。

创建一个新的 `.right` 选择器，并将 `float` 属性设置为 `right`。

# --hints--

应该创建一个新的 `.right` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.right'));
```

`.right` 选择器应该有一个 `float` 属性设置为 `right`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.right')?.float === 'right');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

--fcc-editable-region--

--fcc-editable-region--
```
