---
id: 615f4ec58334106a4170c2a8
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

在你的 `header` 元素下创建一个新的 `div` 元素，并将其 `class` 属性设置为 `divider large`。

# --hints--

应该创建一个新的 `div` 元素。

```js
assert(document.querySelectorAll('div')?.length === 3);
```

新的 `div` 元素应位于 `header` 元素之后。

```js
assert(document.querySelector('.label')?.lastElementChild?.localName === 'div');
```

你的新 `div` 元素应该将 `class` 属性设置为 `divider large`。

```js
const div = document.querySelector('.label')?.lastElementChild;
assert(div?.classList?.contains('divider'));
assert(div?.classList?.contains('large'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
--fcc-editable-region--

--fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}
```
