---
id: 615f357957e370510f21ea16
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

创建一个 `body` 选择器，并将 `font-family` 设置为 `Open Sans`，以 `sans-serif` 为备用字体。

请记住，名称中带有空格的字体必须用 CSS 引号括起来。

# --hints--

应该有一个 `body` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 选择器应该有一个 `font-family` 属性设置为 `"Open Sans", sans-serif`。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
