---
id: 615f74a71f1e498619e38ee8
title: 步骤 47
challengeType: 0
dashedName: step-47
---

# --description--

`% Daily Value *` 和 `Saturated Fat 1g 5%` 元素下的底部边框不会扩展标签的整个宽度。 将 `no-divider` 添加到这两个元素的 `class` 中。

# --hints--

带有文本 `% Daily Value *` 的 `p` 元素应该将 `no-divider` 添加到 `class` 属性中。 不要删除现有的类。

```js
const p = document.querySelector('.daily-value.sm-text')?.firstElementChild;
assert(p?.classList?.contains('no-divider'));
assert(p?.classList?.contains('bold'));
assert(p?.classList?.contains('right'));
```

带有文本 `Saturated Fat 1g 5%` 的 `p` 元素应该将 `no-divider` 添加到 `class` 属性中。 不要删除现有的类。

```js
const p = document.querySelector('.daily-value.sm-text')?.lastElementChild;
assert(p?.classList?.contains('no-divider'));
assert(p?.classList?.contains('indent'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
      <div class="calories-info">
        <p class="bold sm-text">Amount per serving</p>
        <h1>Calories <span class="right">230</span></h1>
      </div>
      <div class="divider md"></div>
--fcc-editable-region--
      <div class="daily-value sm-text">
        <p class="right bold">% Daily Value *</p>
        <div class="divider"></div>
        <p><span class="bold">Total Fat</span> 8g<span class="bold right">10%</span></p>
        <p class="indent">Saturated Fat 1g <span class="bold right">5%</span></p>
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
  clear: right;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}

.md {
  height: 5px;
}

.sm-text {
  font-size: 0.85rem;
}

.calories-info h1 {
  margin: -5px -2px;
  overflow: hidden;
}

.calories-info span {
  font-size: 1.2em;
  margin-top: -7px;
}

.indent {
  margin-left: 1em;
}

.daily-value p {
  border-bottom: 1px solid #888989;
}
```
