---
id: 66faa30000000000000000cc
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--
Inside `update`, move the paddle when the `left` or `right` actions are active. Read **only** through `isDown` — the gameplay code shouldn't know which physical keys map to which action.

```js
if (this.isDown("left")) this.paddle.x -= 4;
else if (this.isDown("right")) this.paddle.x += 4;
```

Reload and try arrows, `A`, and `D`. All three should move the paddle.

# --hints--
The `update` body should call `this.isDown("left")` and decrement `this.paddle.x` by `4`.

```js
assert.match(code, /if\s*\(\s*this\.isDown\(\s*["']left["']\s*\)\s*\)\s*this\.paddle\.x\s*-=\s*4/);
```

The `update` body should call `this.isDown("right")` and increment `this.paddle.x` by `4`.

```js
assert.match(code, /else\s+if\s*\(\s*this\.isDown\(\s*["']right["']\s*\)\s*\)\s*this\.paddle\.x\s*\+=\s*4/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
  }

  update() {
--fcc-editable-region--

--fcc-editable-region--
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

