---
id: 66faa30000000000000000d7
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--
Make the ball collide with the paddle. `physics.add.collider(a, b, callback, processCallback, scope)` registers a hard physical collision between two physics-enabled objects, optionally calling a function on contact.

```js
this.physics.add.collider(
  this.ball,
  this.paddle,
  this.onPaddleBounce,
  null,
  this
);
```

The fifth argument (`this`) sets the scope `onPaddleBounce` will run with, so inside the method `this.tweens` etc. resolve correctly.

# --hints--
You should register a collider between `this.ball` and `this.paddle` with `this.onPaddleBounce` and scope `this`.

```js
assert.match(code, /this\.physics\.add\.collider\(\s*this\.ball\s*,\s*this\.paddle\s*,\s*this\.onPaddleBounce\s*,\s*null\s*,\s*this\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.ball = this.add.circle(240, 540, 8, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setBounce(1, 1);
    this.ball.body.setCollideWorldBounds(true);
    this.physics.world.checkCollision.down = false;
    this.ball.body.setVelocity(120, -260);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
    if (this.ball.y > 660) {
      this.ball.setPosition(240, 540);
      this.ball.body.setVelocity(120, -260);
    }
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

