---
id: 66faa30000000000000000e6
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--
Wire game-over. Replace `loseLife` so it calls a new `gameOver()` method when lives hit zero, and add `gameOver` and `persistHighScore` methods.

Replace the existing `loseLife` body with:

```js
loseLife() {
  this.lives -= 1;
  this.livesText.setText("LIVES " + this.lives);
  if (this.lives <= 0) {
    this.gameOver();
    return;
  }
  this.ball.setPosition(240, 540);
  this.ball.body.setVelocity(120, -260);
}
```

Then add the two new methods immediately after:

```js
gameOver() {
  this.persistHighScore();
  this.scene.start("GameOverScene", { score: this.score, win: false });
}

persistHighScore() {
  const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
  if (this.score > high) localStorage.setItem(SAVE_KEY, String(this.score));
}
```

`persistHighScore` reads the previous high, compares, and writes if the run beat it. `gameOver` persists the score and hands off to `GameOverScene`.

# --hints--
The `loseLife` method should branch to `this.gameOver()` when `this.lives <= 0`.

```js
assert.match(code, /loseLife\s*\([\s\S]*?if\s*\(\s*this\.lives\s*<=\s*0\s*\)[\s\S]*?this\.gameOver\(\s*\)/);
```

You should declare a `gameOver` method that calls `persistHighScore` then `scene.start("GameOverScene", ...)`.

```js
assert.match(code, /gameOver\s*\(\s*\)\s*\{[\s\S]*?this\.persistHighScore\(\s*\)[\s\S]*?this\.scene\.start\(\s*["']GameOverScene["'][\s\S]*?score:\s*this\.score[\s\S]*?win:\s*false/);
```

You should declare a `persistHighScore` method using `localStorage` and `SAVE_KEY`.

```js
assert.match(code, /persistHighScore\s*\(\s*\)\s*\{[\s\S]*?localStorage\.getItem\(\s*SAVE_KEY[\s\S]*?localStorage\.setItem\(\s*SAVE_KEY/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

const SAVE_KEY = "phaser-bb-high-score";

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.ball = this.add.circle(240, 540, 8, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setBounce(1, 1);
    this.ball.body.setCollideWorldBounds(true);
    this.physics.world.checkCollision.down = false;
    this.ball.body.setVelocity(120, -260);
    this.physics.add.collider(
      this.ball,
      this.paddle,
      this.onPaddleBounce,
      null,
      this
    );
    this.bricks = this.physics.add.staticGroup();
    this.spawnBricks();
    this.physics.add.collider(
      this.ball,
      this.bricks,
      this.onBrickHit,
      null,
      this
    );
    this.score = 0;
    this.scoreText = this.add.text(10, 10, "SCORE 0", {
      fontSize: "16px",
      color: "#ffffff"
    });
    this.lives = 3;
    this.livesText = this.add.text(10, 32, "LIVES 3", {
      fontSize: "16px",
      color: "#ffffff"
    });
  }

  update() {
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
    if (this.ball.y > 660) this.loseLife();
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }

  onPaddleBounce(ball, paddle) {
    const offset = (ball.x - paddle.x) / (paddle.width / 2);
    const speed = Math.sqrt(
      ball.body.velocity.x ** 2 + ball.body.velocity.y ** 2
    );
    const angle = -Math.PI / 2 + offset * (Math.PI / 3);
    ball.body.setVelocity(speed * Math.cos(angle), speed * Math.sin(angle));
    this.tweens.add({
      targets: paddle,
      scaleX: 1.2,
      scaleY: 0.8,
      duration: 80,
      yoyo: true
    });
  }

  spawnBricks() {
    for (let row = 0; row < 3; row++) {
      for (let col = 0; col < 10; col++) {
        const x = 40 + col * 40 + 20;
        const y = 80 + row * 22;
        const brick = this.add.rectangle(x, y, 36, 18, 0x44aaff);
        this.bricks.add(brick);
      }
    }
  }

  onBrickHit(ball, brick) {
    brick.destroy();
    this.score += 10;
    this.scoreText.setText("SCORE " + this.score);
  }

--fcc-editable-region--
  loseLife() {
    this.lives -= 1;
    this.livesText.setText("LIVES " + this.lives);
    this.ball.setPosition(240, 540);
    this.ball.body.setVelocity(120, -260);
  }
--fcc-editable-region--
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create(data) {
    this.add.rectangle(240, 320, 480, 640, 0x000000, 0.85);
    this.add
      .text(240, 240, data.win ? "YOU WIN" : "GAME OVER", {
        fontSize: "32px",
        color: data.win ? "#44ee88" : "#ff6644"
      })
      .setOrigin(0.5);
    this.add
      .text(240, 300, "Score: " + data.score, {
        fontSize: "20px",
        color: "#ffffff"
      })
      .setOrigin(0.5);
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    this.add
      .text(240, 340, "High: " + high, {
        fontSize: "16px",
        color: "#aaaaaa"
      })
      .setOrigin(0.5);
    this.add
      .text(240, 420, "Press R to restart", {
        fontSize: "14px",
        color: "#ffee44"
      })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-R", () => {
      this.scene.start("MainScene", { level: 1, score: 0, lives: 3 });
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene, GameOverScene]
};

const game = new Phaser.Game(config);
```

