---
id: 66faa30000000000000000c6
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the brick-breaker workshop. Over the next 50 steps you'll build a complete brick breaker that runs identically on keyboard, mouse, and touch — five levels of bricks, multi-hit and indestructible blocks, lives, score, a persistent high score, a virtual-controls overlay scene, and a game-over flow.

Start by writing the bare scene shell. Inside the editable region, declare a class `MainScene` that extends `Phaser.Scene`, with a constructor that calls `super("MainScene")` and empty `create` and `update` methods.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {}

  update() {}
}
```

The `config` object below already references `MainScene`, so once the class is defined the canvas mounts.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["']MainScene["']\s*\)/);
```

You should declare an empty `create` method.

```js
assert.match(code, /create\s*\(\s*\)\s*\{/);
```

You should declare an empty `update` method.

```js
assert.match(code, /update\s*\(\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

