---
id: 66faa30000000000000000c7
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
The canvas is `480 × 640`. Inside `create`, paint a dark background rectangle that fills the play area:

```js
this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
```

`add.rectangle(centerX, centerY, width, height, color)` creates a non-physics display rectangle. `0x111122` is a navy-leaning charcoal that lets the brighter game elements stand out without burning your eyes during a long playtest.

# --hints--
You should add a rectangle the size of the canvas at the canvas center, color `0x111122`.

```js
assert.match(code, /this\.bg\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*320\s*,\s*480\s*,\s*640\s*,\s*0x111122\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

