---
id: 66faa30000000000000000cb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Add a method on `MainScene` that resolves an action name to "is any bound key down right now?":

```js
isDown(action) {
  return ACTIONS[action].some(name => {
    const k = this.keys[name] || this.cursors[name.toLowerCase()];
    return k && k.isDown;
  });
}
```

The fallback `this.cursors[name.toLowerCase()]` lets the action map resolve `"LEFT"` against `this.cursors.left` (the arrow key the cursor-keys plugin returned), not just `this.keys.LEFT`. With one method the rest of the workshop can ask "is `left` down?" and not care which device or which keycode is supplying the input.

# --hints--
You should declare an `isDown(action)` method on `MainScene`.

```js
assert.match(code, /isDown\s*\(\s*action\s*\)\s*\{/);
```

The body should iterate `ACTIONS[action]` with `.some(...)`.

```js
assert.match(code, /ACTIONS\[\s*action\s*\]\.some\(/);
```

It should fall back to `this.cursors[name.toLowerCase()]` when `this.keys[name]` is missing.

```js
assert.match(code, /this\.keys\[\s*name\s*\]\s*\|\|\s*this\.cursors\[\s*name\.toLowerCase\(\s*\)\s*\]/);
```

It should return `k && k.isDown`.

```js
assert.match(code, /return\s+k\s*&&\s*k\.isDown/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
  }

--fcc-editable-region--

--fcc-editable-region--
  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

