---
id: 66faa30000000000000000cf
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Refactor `update` so both the keyboard and the pointer drive the same `targetX`, then have the paddle smoothly lerp toward it. Replace the contents of `update` with:

```js
if (this.isDown("left")) this.targetX -= 6;
else if (this.isDown("right")) this.targetX += 6;
this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
```

`Phaser.Math.Linear(a, b, t)` is the same as `a + (b - a) * t`. With `t = 0.3` the paddle covers about a third of the remaining distance per frame — fast enough that pointer drags feel responsive, slow enough that arrow taps don't snap.

# --hints--
The `update` body should accumulate `this.targetX` with arrows/WASD at speed `6`.

```js
assert.match(code, /if\s*\(\s*this\.isDown\(\s*["']left["']\s*\)\s*\)\s*this\.targetX\s*-=\s*6/);
assert.match(code, /else\s+if\s*\(\s*this\.isDown\(\s*["']right["']\s*\)\s*\)\s*this\.targetX\s*\+=\s*6/);
```

It should clamp `this.targetX` to `[40, 440]`.

```js
assert.match(code, /this\.targetX\s*=\s*Phaser\.Math\.Clamp\(\s*this\.targetX\s*,\s*40\s*,\s*440\s*\)/);
```

It should lerp `this.paddle.x` toward `this.targetX` with `Phaser.Math.Linear`.

```js
assert.match(code, /this\.paddle\.x\s*=\s*Phaser\.Math\.Linear\(\s*this\.paddle\.x\s*,\s*this\.targetX\s*,\s*0\.3\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
  }

  update() {
--fcc-editable-region--
    if (this.isDown("left")) this.paddle.x -= 4;
    else if (this.isDown("right")) this.paddle.x += 4;
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
--fcc-editable-region--
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

