---
id: 66faa30000000000000000df
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--
Update `onBrickHit` to award `10` points and refresh the HUD:

```js
onBrickHit(ball, brick) {
  brick.destroy();
  this.score += 10;
  this.scoreText.setText("SCORE " + this.score);
}
```

`Text.setText(value)` replaces the entire string. You'll keep this pattern for every HUD update — there's no two-way binding here, just direct mutations.

# --hints--
The `onBrickHit` method should increment `this.score` by `10`.

```js
assert.match(code, /this\.score\s*\+=\s*10/);
```

It should refresh the score HUD with `setText`.

```js
assert.match(code, /this\.scoreText\.setText\(\s*["']SCORE ["']\s*\+\s*this\.score\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.ball = this.add.circle(240, 540, 8, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setBounce(1, 1);
    this.ball.body.setCollideWorldBounds(true);
    this.physics.world.checkCollision.down = false;
    this.ball.body.setVelocity(120, -260);
    this.physics.add.collider(
      this.ball,
      this.paddle,
      this.onPaddleBounce,
      null,
      this
    );
    this.bricks = this.physics.add.staticGroup();
    this.spawnBricks();
    this.physics.add.collider(
      this.ball,
      this.bricks,
      this.onBrickHit,
      null,
      this
    );
    this.score = 0;
    this.scoreText = this.add.text(10, 10, "SCORE 0", {
      fontSize: "16px",
      color: "#ffffff"
    });
  }

  update() {
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
    if (this.ball.y > 660) {
      this.ball.setPosition(240, 540);
      this.ball.body.setVelocity(120, -260);
    }
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }

  onPaddleBounce(ball, paddle) {
    const offset = (ball.x - paddle.x) / (paddle.width / 2);
    const speed = Math.sqrt(
      ball.body.velocity.x ** 2 + ball.body.velocity.y ** 2
    );
    const angle = -Math.PI / 2 + offset * (Math.PI / 3);
    ball.body.setVelocity(speed * Math.cos(angle), speed * Math.sin(angle));
    this.tweens.add({
      targets: paddle,
      scaleX: 1.2,
      scaleY: 0.8,
      duration: 80,
      yoyo: true
    });
  }

  spawnBricks() {
    for (let row = 0; row < 3; row++) {
      for (let col = 0; col < 10; col++) {
        const x = 40 + col * 40 + 20;
        const y = 80 + row * 22;
        const brick = this.add.rectangle(x, y, 36, 18, 0x44aaff);
        this.bricks.add(brick);
      }
    }
  }

--fcc-editable-region--
  onBrickHit(ball, brick) {
    brick.destroy();
  }
--fcc-editable-region--
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

