---
id: 66faa30000000000000000c8
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Add the paddle. It's a rectangle near the bottom that the player will move side-to-side. Give it a physics body so the ball can collide with it later, and mark it `setImmovable(true)` so collisions don't push it around.

```js
this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
this.physics.add.existing(this.paddle);
this.paddle.body.setImmovable(true);
```

`physics.add.existing` retrofits an arcade body onto an existing display object — useful when you want the visual to be a primitive shape instead of a sprite. `setImmovable(true)` means the ball reflects off the paddle without dragging the paddle along.

# --hints--
You should add a paddle rectangle at `(240, 580)` sized `80 × 16` colored `0xeeeeee`.

```js
assert.match(code, /this\.paddle\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*580\s*,\s*80\s*,\s*16\s*,\s*0xeeeeee\s*\)/);
```

You should give the paddle an arcade physics body.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.paddle\s*\)/);
```

You should mark the paddle's body `setImmovable(true)`.

```js
assert.match(code, /this\.paddle\.body\.setImmovable\(\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

