---
id: 66faa30000000000000000ca
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
Declare an `ACTIONS` map at the top of the file (above `class MainScene`). Map each game action to the **physical keys** that drive it.

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};
```

This is the action-map abstraction from the lecture. The rest of the workshop will read from `ACTIONS` and never check key state directly.

# --hints--
You should declare a top-level `const ACTIONS` object before the class.

```js
assert.match(code, /const\s+ACTIONS\s*=\s*\{[\s\S]*\}/);
```

`ACTIONS.left` should bind to `["LEFT", "A"]`.

```js
assert.match(code, /left\s*:\s*\[\s*["']LEFT["']\s*,\s*["']A["']\s*\]/);
```

`ACTIONS.right` should bind to `["RIGHT", "D"]`.

```js
assert.match(code, /right\s*:\s*\[\s*["']RIGHT["']\s*,\s*["']D["']\s*\]/);
```

`ACTIONS.launch` should bind to `["SPACE"]`.

```js
assert.match(code, /launch\s*:\s*\[\s*["']SPACE["']\s*\]/);
```

`ACTIONS.pause` should bind to `["P"]`.

```js
assert.match(code, /pause\s*:\s*\[\s*["']P["']\s*\]/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

