---
id: 66faa30000000000000000d8
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--
Define `onPaddleBounce`. The classic brick-breaker trick is to vary the bounce angle based on **where** the ball hit the paddle: hit the left edge → ball goes left; hit the right edge → ball goes right; hit dead center → ball goes straight up.

Add the method on `MainScene`:

```js
onPaddleBounce(ball, paddle) {
  const offset = (ball.x - paddle.x) / (paddle.width / 2);
  const speed = Math.sqrt(
    ball.body.velocity.x ** 2 + ball.body.velocity.y ** 2
  );
  const angle = -Math.PI / 2 + offset * (Math.PI / 3);
  ball.body.setVelocity(speed * Math.cos(angle), speed * Math.sin(angle));
}
```

`offset` is a number from `-1` (far left) to `1` (far right). `-Math.PI / 2` is straight up; multiplying `offset` by `Math.PI / 3` (60°) caps how steep the angle can get on the edges. You preserve `speed` so the ball's energy is conserved through the bounce.

# --hints--
You should declare an `onPaddleBounce(ball, paddle)` method on `MainScene`.

```js
assert.match(code, /onPaddleBounce\s*\(\s*ball\s*,\s*paddle\s*\)\s*\{/);
```

The method should compute `offset` from `(ball.x - paddle.x) / (paddle.width / 2)`.

```js
assert.match(code, /const\s+offset\s*=\s*\(\s*ball\.x\s*-\s*paddle\.x\s*\)\s*\/\s*\(\s*paddle\.width\s*\/\s*2\s*\)/);
```

The method should compute the bounce angle as `-Math.PI / 2 + offset * (Math.PI / 3)`.

```js
assert.match(code, /-\s*Math\.PI\s*\/\s*2\s*\+\s*offset\s*\*\s*\(\s*Math\.PI\s*\/\s*3\s*\)/);
```

The method should set the ball's velocity from `speed * Math.cos(angle)` and `speed * Math.sin(angle)`.

```js
assert.match(code, /ball\.body\.setVelocity\(\s*speed\s*\*\s*Math\.cos\(\s*angle\s*\)\s*,\s*speed\s*\*\s*Math\.sin\(\s*angle\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.ball = this.add.circle(240, 540, 8, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setBounce(1, 1);
    this.ball.body.setCollideWorldBounds(true);
    this.physics.world.checkCollision.down = false;
    this.ball.body.setVelocity(120, -260);
    this.physics.add.collider(
      this.ball,
      this.paddle,
      this.onPaddleBounce,
      null,
      this
    );
  }

  update() {
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
    if (this.ball.y > 660) {
      this.ball.setPosition(240, 540);
      this.ball.body.setVelocity(120, -260);
    }
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }
--fcc-editable-region--

--fcc-editable-region--
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

