---
id: 66faa30000000000000000d0
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--
Add a `pointerdown` listener with the same thumb-zone gate so a tap snaps the paddle to that x without waiting for `pointermove`:

```js
this.input.on("pointerdown", pointer => {
  if (pointer.y > 540) this.targetX = pointer.x;
});
```

Tap-and-go feels much better than tap-then-wait on touch devices.

# --hints--
You should attach a `pointerdown` listener gated to `pointer.y > 540`.

```js
assert.match(code, /this\.input\.on\(\s*["']pointerdown["'][\s\S]*?if\s*\(\s*pointer\.y\s*>\s*540\s*\)\s*this\.targetX\s*=\s*pointer\.x/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

