---
id: 66faa30000000000000000cd
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--
The paddle currently slides off the canvas. Clamp `paddle.x` to `[40, 440]` so it can't escape:

```js
if (this.isDown("left")) this.paddle.x -= 4;
else if (this.isDown("right")) this.paddle.x += 4;
this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
```

`Phaser.Math.Clamp(value, min, max)` is the one you'll reach for any time gameplay logic produces a number that should never escape a range.

# --hints--
The `update` body should clamp `this.paddle.x` to `[40, 440]` after the movement.

```js
assert.match(code, /this\.paddle\.x\s*=\s*Phaser\.Math\.Clamp\(\s*this\.paddle\.x\s*,\s*40\s*,\s*440\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
  }

  update() {
--fcc-editable-region--
    if (this.isDown("left")) this.paddle.x -= 4;
    else if (this.isDown("right")) this.paddle.x += 4;
--fcc-editable-region--
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

