---
id: 66faa30000000000000000ef
title: Step 42
challengeType: 0
dashedName: step-42
---

# --description--
Add pause. At the top of the new `update`, check `JustDown(this.keys.P)` and early-return if paused. Initialize `this.paused` in `create`. Add a `togglePause()` method that pauses/resumes the physics world and shows/hides a "PAUSED" overlay.

Replace the start of `update`:

```js
const v = this.ball.body.velocity;
```

with:

```js
if (Phaser.Input.Keyboard.JustDown(this.keys.P)) this.togglePause();
if (this.paused) return;
const v = this.ball.body.velocity;
```

Add `this.paused = false;` to `create` (anywhere after the keys are wired). Then add the new method:

```js
togglePause() {
  this.paused = !this.paused;
  if (this.paused) {
    this.physics.world.pause();
    this.pauseText = this.add
      .text(240, 320, "PAUSED", {
        fontSize: "40px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(100);
  } else {
    this.physics.world.resume();
    if (this.pauseText) this.pauseText.destroy();
  }
}
```

`physics.world.pause()` halts all body integration; the rest of the scene (tweens, input) keeps running. `setDepth(100)` keeps the overlay above bricks and ball.

# --hints--
The `update` body should call `Phaser.Input.Keyboard.JustDown(this.keys.P)` and `this.togglePause()` on the first line.

```js
assert.match(code, /Phaser\.Input\.Keyboard\.JustDown\(\s*this\.keys\.P\s*\)/);
assert.match(code, /update\s*\(\s*\)\s*\{[\s\S]*?this\.togglePause\(\s*\)/);
```

It should early-return when `this.paused` is true.

```js
assert.match(code, /if\s*\(\s*this\.paused\s*\)\s*return/);
```

You should declare a `togglePause` method that pauses/resumes `physics.world` and toggles a "PAUSED" overlay.

```js
assert.match(code, /togglePause\s*\(\s*\)\s*\{[\s\S]*?this\.physics\.world\.pause\(\s*\)[\s\S]*?this\.physics\.world\.resume\(\s*\)/);
assert.match(code, /this\.add[\s\S]*?["']PAUSED["'][\s\S]*?fontSize:\s*["']40px["']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

const SAVE_KEY = "phaser-bb-high-score";

const LEVELS = [
  [
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
  ],
  [
    [2, 2, 2, 2, 2, 2, 2, 2, 2, 2],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [2, 2, 2, 2, 2, 2, 2, 2, 2, 2]
  ],
  [
    [9, 1, 1, 1, 1, 1, 1, 1, 1, 9],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [9, 1, 1, 1, 1, 1, 1, 1, 1, 9]
  ],
  [
    [3, 3, 2, 2, 1, 1, 2, 2, 3, 3],
    [9, 1, 1, 1, 1, 1, 1, 1, 1, 9],
    [1, 1, 1, 9, 9, 9, 9, 1, 1, 1]
  ],
  [
    [3, 3, 3, 3, 3, 3, 3, 3, 3, 3],
    [9, 2, 2, 2, 2, 2, 2, 2, 2, 9],
    [9, 1, 1, 9, 9, 9, 9, 1, 1, 9]
  ]
];

const COLORS = [0x666666, 0x44aaff, 0x44ee88, 0xffee44, 0xffaa44, 0xff6644];

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  init(data) {
    this.level = data.level || 1;
    this.score = data.score || 0;
    this.lives = data.lives || 3;
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.ball = this.add.circle(240, 540, 8, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setBounce(1, 1);
    this.ball.body.setCollideWorldBounds(true);
    this.physics.world.checkCollision.down = false;
    this.ball.body.setVelocity(
      (220 + (this.level - 1) * 12) * 0.5,
      -(220 + (this.level - 1) * 12)
    );
    this.physics.add.collider(
      this.ball,
      this.paddle,
      this.onPaddleBounce,
      null,
      this
    );
    this.bricks = this.physics.add.staticGroup();
    this.spawnBricks();
    this.physics.add.collider(
      this.ball,
      this.bricks,
      this.onBrickHit,
      null,
      this
    );
    this.scoreText = this.add.text(10, 10, "SCORE " + this.score, {
      fontSize: "16px",
      color: "#ffffff"
    });
    this.livesText = this.add.text(10, 32, "LIVES " + this.lives, {
      fontSize: "16px",
      color: "#ffffff"
    });
    this.highText = this.add
      .text(470, 10, "", {
        fontSize: "16px",
        color: "#aaaaaa"
      })
      .setOrigin(1, 0);
    this.updateHud();
  }

  update() {
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
--fcc-editable-region--
    const v = this.ball.body.velocity;
    if (Math.abs(v.y) < 60) v.y = v.y < 0 ? -120 : 120;
    const speed = Math.sqrt(v.x * v.x + v.y * v.y);
    const maxSpeed = 360 + (this.level - 1) * 20;
    if (speed > maxSpeed) {
      v.x *= maxSpeed / speed;
      v.y *= maxSpeed / speed;
    }
    if (this.ball.y > 660) this.loseLife();
--fcc-editable-region--
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }

  onPaddleBounce(ball, paddle) {
    const offset = (ball.x - paddle.x) / (paddle.width / 2);
    const speed = Math.sqrt(
      ball.body.velocity.x ** 2 + ball.body.velocity.y ** 2
    );
    const angle = -Math.PI / 2 + offset * (Math.PI / 3);
    ball.body.setVelocity(speed * Math.cos(angle), speed * Math.sin(angle));
    this.tweens.add({
      targets: paddle,
      scaleX: 1.2,
      scaleY: 0.8,
      duration: 80,
      yoyo: true
    });
  }

  spawnBricks() {
    const layout = LEVELS[(this.level - 1) % LEVELS.length];
    for (let row = 0; row < layout.length; row++) {
      for (let col = 0; col < layout[row].length; col++) {
        const cell = layout[row][col];
        if (cell === 0) continue;
        const x = 40 + col * 40 + 20;
        const y = 80 + row * 22;
        const isIndestructible = cell === 9;
        const color = isIndestructible ? COLORS[0] : COLORS[cell];
        const brick = this.add.rectangle(x, y, 36, 18, color);
        this.bricks.add(brick);
        brick.hp = isIndestructible ? Infinity : cell;
        brick.points = (3 - row + cell) * 10;
      }
    }
  }

  onBrickHit(ball, brick) {
    if (brick.hp === Infinity) return;
    brick.hp -= 1;
    if (brick.hp <= 0) {
      this.score += brick.points;
      brick.destroy();
      this.scoreText.setText("SCORE " + this.score);
      const remaining = this.bricks
        .getChildren()
        .filter(b => b.hp !== Infinity).length;
      if (remaining === 0) this.advanceLevel();
    } else {
      brick.setFillStyle(COLORS[Phaser.Math.Clamp(brick.hp, 0, 5)]);
    }
  }

  loseLife() {
    this.lives -= 1;
    this.updateHud();
    if (this.lives <= 0) {
      this.gameOver();
      return;
    }
    this.ball.setPosition(240, 540);
    const baseSpeed = 220 + (this.level - 1) * 12;
    this.ball.body.setVelocity(baseSpeed * 0.5, -baseSpeed);
  }

  gameOver() {
    this.persistHighScore();
    this.scene.start("GameOverScene", { score: this.score, win: false });
  }

  persistHighScore() {
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    if (this.score > high) localStorage.setItem(SAVE_KEY, String(this.score));
  }

  advanceLevel() {
    this.persistHighScore();
    this.level += 1;
    if (this.level > LEVELS.length) {
      this.scene.start("GameOverScene", { score: this.score, win: true });
      return;
    }
    this.scene.restart({
      level: this.level,
      score: this.score,
      lives: this.lives
    });
  }

  updateHud() {
    this.scoreText.setText("SCORE " + this.score);
    this.livesText.setText("LIVES " + this.lives);
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    this.highText.setText("HIGH " + Math.max(high, this.score));
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create(data) {
    this.add.rectangle(240, 320, 480, 640, 0x000000, 0.85);
    this.add
      .text(240, 240, data.win ? "YOU WIN" : "GAME OVER", {
        fontSize: "32px",
        color: data.win ? "#44ee88" : "#ff6644"
      })
      .setOrigin(0.5);
    this.add
      .text(240, 300, "Score: " + data.score, {
        fontSize: "20px",
        color: "#ffffff"
      })
      .setOrigin(0.5);
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    this.add
      .text(240, 340, "High: " + high, {
        fontSize: "16px",
        color: "#aaaaaa"
      })
      .setOrigin(0.5);
    this.add
      .text(240, 420, "Press R to restart", {
        fontSize: "14px",
        color: "#ffee44"
      })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-R", () => {
      this.scene.start("MainScene", { level: 1, score: 0, lives: 3 });
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene, GameOverScene]
};

const game = new Phaser.Game(config);
```

