---
id: 66faa30000000000000000c9
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
Wire two keyboard sources at once. `createCursorKeys()` gives you the four arrow keys. `addKeys("A,D,SPACE,P")` returns the rest of the bindings the workshop will need — WASD-style movement, `SPACE` for "launch ball" later, and `P` for pause much later.

```js
this.cursors = this.input.keyboard.createCursorKeys();
this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
```

The action map you build in the next step will OR these two sources together.

# --hints--
You should call `this.input.keyboard.createCursorKeys()` and assign the result to `this.cursors`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)/);
```

You should call `this.input.keyboard.addKeys("A,D,SPACE,P")` and assign the result to `this.keys`.

```js
assert.match(code, /this\.keys\s*=\s*this\.input\.keyboard\.addKeys\(\s*["']A,D,SPACE,P["']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Mobile-First Brick Breaker with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

