---
id: 66cc0a9e06e00b75d6782be9
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

In the previous lecture videos, you learned how to call (invoke) a function. Calling a function means to execute the code inside the function.

Here is a reminder of how to call a function:

```js
function myFunction() {
  return 'Hello World';
}

// function call
myFunction();
```

Add a `console.log` statement, and inside it, call the `addTwoAndSeven` function.

You should now see the sum of `2` and `7` logged to the console.

# --hints--

You should have a `console.log` in your code.

```js
assert.match(code, /console\.log(.*)/);
```

You should call the `addTwoAndSeven` function inside the `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
