---
id: 66cc095dfe1682753d2ab030
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

One of concepts you learned in the previous lecture videos was how to return values from a function.

Here is a reminder of how to return a value from a function:

```js
function myFunction() {
  return "Hello World";
}
```

Inside your `addTwoAndSeven` function, return the sum of `2` and `7`.

# --hints--

Your `addTwoAndSeven` function should return the sum of `2` and `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  
}
--fcc-editable-region--
```
