---
id: 66cc14ca14e65e7a9ebecb80
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Now that your calculator can add two numbers together, it is time to create a function that will subtract two numbers.

Declare a function called `calculateDifference` that takes two parameters, `num1` and `num2`.

Inside the function, return the difference between the two parameters.

# --hints--

You should have a function called `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Your `calculateDifference` function should have two parameters called `num1` and `num2`.

```js
assert.match(code, /calculateDifference\s*(?:=\s*)?\(\s*num1\s*,\s*num2\s*\)\s*/);
```

Your `calculateDifference` function should return the difference between the two parameters.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
