---
id: 66cc21d23238dc8240a8a182
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

It would be nice to have your calculator calculate the square of a number. The square of a number is the number multiplied by itself.

To calculate the square of a number in JavaScript, you can use the <dfn>Math.pow()</dfn> method. The `Math.pow()` method takes two arguments: the base number and the exponent. 

You can also use the exponentiation operator (`**`) to calculate the square of a number.

Here is an example:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Declare a function called `calculateSquare` that takes a `num` parameter and returns the square of `num`.

# --hints--

You should have a function called `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

Your `calculateSquare` function should take a `num` parameter.

```js
assert.match(code, /calculateSquare\s*(?:=\s*)?\(\s*num\s*\)/);
```

Your `calculateSquare` function should return the square of `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
