---
id: 66cc2739f687e484d50bb6f1
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

The last piece of functionality for your calculator will be the square root functionality.

A square is a number multiplied by itself. So, the square root would be the number that when multiplied by itself gives the original number.

For example, the square root of `9` is `3` because `3 * 3 = 9`.

To get the square root of a number in JavaScript, you can use the <dfn>Math.sqrt()</dfn> method.

Here is an example:

```js
// result: 3
Math.sqrt(9);
```

Declare a function called `calculateSquareRoot` that takes a `num` parameter and returns the square root of `num`.

# --hints--

You should have a function called `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Your `calculateSquareRoot` function should take a `num` parameter.

```js
assert.match(code, /calculateSquareRoot\s*(?:=\s*)?\(\s*num\s*\)/);
```

Your `calculateSquareRoot` function should return the square root of `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
