---
id: 66cc06012e46aa6bc9b8c001
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will review how to work with functions by building a calculator app. 

In the previous lecture videos, you learned how to declare functions like this:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Start by creating a function called `addTwoAndSeven`.

You can choose to use the regular function syntax or the arrow function syntax.

# --hints--

You should declare a function called `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
