---
id: 66cc0f1ae40802781b2ea972
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

In the previous lecture videos, you learned how to work with function parameters.

A function parameter is a variable that is defined in the function's declaration and acts as a placeholder.

Here is an example of a function that has a parameter:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Declare a function called `calculateSum` that takes two parameters, `num1` and `num2`.

# --hints--

You should declare a function called `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Your `calculateSum` function should take two parameters, `num1` and `num2`.

```js
assert.match(code, /calculateSum\s*(?:=\s*)?\(\s*num1\s*,\s*num2\s*\)\s*/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
