---
id: 60f0286404aefb0562a4fdf9
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

給 `head` 添加一個 `title` 和一個 `meta` 元素。 設置項目的標題爲 `Registration Form`，添加一個 `meta` 元素設置其`charset` 屬性的值爲 `UTF-8`。

# --hints--

代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 元素應該在 `head` 元素內。

```js
assert.exists(document.querySelector('head > title'));
```

項目的標題應該爲 `Registration Form`。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
