---
id: 60f0286404aefb0562a4fdf9
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

將 `title` 元素添加到 `head` 內，並將項目標題設置爲 `Registration Form`。 此外，在 `head` 元素中嵌套一個自閉合 `link` 元素。 設置其 `rel` 屬性值爲 `stylesheet`， `href` 屬性值爲 `styles.css`。

# --hints--

代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 元素應該在 `head` 元素內。

```js
assert.exists(document.querySelector('head > title'));
```

項目的標題應該爲 `Registration Form`。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

代碼應該有一個 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

不應更改現有的 `head` 標籤。 確保沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

應該有一個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 元素應該在 `head` 元素內。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  --fcc-editable-region--
  <head>

  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
