---
id: 60f0286404aefb0562a4fdf9
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `head` 元素中添加一個 `title` 和 `meta` 元素。 設置項目的標題爲 `Registration Form`，給 `meta` 元素添加值爲 `utf-8` 的 `charset` 屬性。

# --hints--

代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 元素應該在 `head` 元素內。

```js
assert.exists(document.querySelector('head > title'));
```

你的項目應該有標題 `Registration Form`。 記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

你應該在 `head` 元素中創建一個 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你應該給 `meta` 元素一個值爲 `utf-8` 的 `charset` 屬性。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

你的 `meta` 元素應該是一個空元素，它不應該有結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
