---
id: 60f027099a15b00485563dd2
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `DOCTYPE`下， 添加一個 `html` 元素，元素內有一個設置爲 `en` 的 `lang` 屬性，這樣就可以編碼了。

# --hints--

`DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 元素應該有一個開始標籤，`lang` 屬性的值應該爲 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

`html` 標籤順序應該正確。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

應該只有一個 `html` 元素。

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
