---
id: 60f02e7361b68405e27b62a5
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `body` 中， 添加一個 `h1` 文本 `Registration Form`，來給內容添加一個標題。

# --hints--

應該在 `body` 內添加一個 `h1`。

```js
assert.exists(document.querySelector('body > h1'));
```

應該給 `h1` 元素添加文本 `Registration Form`。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
