---
id: 60f1a5e2d2c23707a4f9a660
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

現在，通過設置 `body` 的默認 `margin` 爲 `0` 來重置一些瀏覽器的默認設置，從而去掉水平滾動條。

# --hints--

應該在 `body` 元素選擇器中添加 `margin`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

應該將 `margin` 屬性設置爲 `0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
