---
id: 60f1922fcbd2410527b3bd89
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--
`vh` 單位代表視口高度（viewport height），相當於視口 `height` 高度的 1%。

現在是時候用 CSS 來裝飾項目了。 首先設置 `body` 的 `width` 爲 `100%`，`height` 爲 `100vh`。

# --hints--

你應該使用 `body` 元素選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你應該把 `body` 元素的 `width` 屬性值設置爲 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你應該把 `body` 的 `height` 設置爲 `100vh`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
