---
id: 60f1a9cbd23023082e149fee
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

看起來好多了。 現在，把 `body` 的 `background-color` 設置爲 `#1b1b32`，使背景更護眼。 然後相應的把 `color` 設置爲 `#f5f6f7`，讓文字顯示出來。

# --hints--

應該在 `body` 元素選擇器中添加 `background-color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

應該設置 `background-color` 的值爲 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

應該在 `body` 元素選擇器中添加 `color`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

`color` 屬性應該設置爲 `#f5f6f7`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
