---
id: 60f5cb8875ab6a0610f05071
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

表單包含三個不同的部分，在 `form` 內部添加三個 `fieldset` 元素。

# --hints--

應該有三個 `fieldset` 元素。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

`fieldset` 元素應該在 `form` 內。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

三個 `fieldset` elements 應該是同級元素。

```js
assert.exists(document.querySelector('fieldset + fieldset + fieldset'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>

    </form>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
