---
id: 60f83e7bfc09900959f41e20
title: 步驟 22
challengeType: 0
dashedName: step-22
---

# --description--

一些 `type` 屬性值含有內置表單驗證。 例如， `type="email"` 要求該值是一個有效的電子郵件地址。

給密碼 `input` 元素添加自定義驗證，新增 `minlength` 屬性設置其值爲 `8`。 這樣密碼少於 8 個字符的時候會阻止提交。

# --hints--

密碼 `input` 元素應該添加 `minlength` 屬性。

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

應該將 `minlength` 屬性設置爲 `8`。

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
