---
id: 60f805f813eaf2049bc2ceea
title: 步驟 17
challengeType: 0
dashedName: step-17
---

# --description--

在每個 `label` 中嵌套一個 `input` 元素。 使 `input` 在 `label` 文字後面，冒號後添加一個空格。

# --hints--

應該 `fieldset` 元素內添加四個 `input` 元素。

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

應該在 `label` 元素內添加 `input` 元素。

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

應該將第一個 `input` 添加在 `label` 文字 `Enter Your First Name:` 後面，並在冒號後面加一個空格。

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

應該將第二個 `input` 添加在 `label` 文字 `Enter Your Last Name:` 後面，並在冒號後面加一個空格。

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

應該將第三個 `input` 添加在 `label` 文字 `Enter Your Email:` 後面，並在冒號後面加一個空格。

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

應該將第四個 `input` 添加在 `label` 文字 `Create a New Password:` 後面，並在冒號後面加一個空格。

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
