---
id: 60f8604682407e0d017bbf7f
title: 步驟 26
challengeType: 0
dashedName: step-26
---

# --description--

對於條款及條件，在第三個 `label` 元素內添加一個 `type` 爲 `checkbox` 的 `input` 元素。 另外，用戶需要在閱讀條款和條件之後才能註冊，所以將這個 `input` 框設置爲 `required`。

# --hints--

應該給第三個 `label` 元素添加一個 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(2) label:nth-child(3) input'));
```

應該給 `input` 元素添加一個 `type` 屬性並設置其值爲 `checkbox`。

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(3) input')?.type, 'checkbox');
```

應該給 `input` 元素添加一個 `required` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(3) input')?.required, true);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label><input type="radio" /></label>
        <label><input type="radio" /></label>
        <label></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
