---
id: 60f803d5241e6a0433a523a1
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--
`rem` 單位代表根 `em`，與 `html` 元素的字體大小有關。

由於 `label` 元素默認是行內元素，它們出現在其標籤文本的同一行，使得文本難以閱讀。 給 `label` 元素添加 `display: block`，並設置其 `margin` 爲 `0.5rem 0`，來使其顯示在不同的行，並且行之間有一定的距離。

# --hints--

你應該使用 `label` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

你應該將 `display` 屬性設置爲 `block`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

你應該添加 `margin` 屬性，值爲 `0.5rem 0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
