---
id: 60f81616cff80508badf9ad5
title: 步驟 21
challengeType: 0
dashedName: step-21
---

# --description--

到這裏，應該可以提交表單了。 實際上，會發現提交之後沒有任何效果。

爲了使表單更具交互性，在第一個 `fieldset` 字段中的 `input` 元素添加 `required` 屬性。

這樣，如果沒有填寫必需字段就提交表單，會展示一個錯誤信息。

# --hints--

應該給第一個 `input` 元素添加 `required` 屬性。

```js
assert.equal(document.querySelector('input')?.required, true);
```

應該給第二個 `input` 元素添加 `required` 屬性。

```js
assert.equal(document.querySelectorAll('input')?.[1]?.required, true);
```

應該給第三個 `input` 元素添加 `required` 屬性。

```js
assert.equal(document.querySelectorAll('input')?.[2]?.required, true);
```

應該給第四個 `input` 元素添加 `required` 屬性。

```js
assert.equal(document.querySelectorAll('input')?.[3]?.required, true);
```

應該給 `submit` `input` 添加 `required` 屬性。

```js
assert.equal(document.querySelector('input[type="submit"]')?.required, false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
