---
id: 60f84ec41116b209c280ba91
title: 步驟 23
challengeType: 0
dashedName: step-23
---

# --description--

在 `type="password"` 內可以使用 `pattern` 屬性來使用正則表達式來校驗密碼。

給密碼 `input` 元素添加 `pattern` 屬性，要求輸入匹配 `[a-z0-5]{8,}`。

上面是一個正則表達式，匹配 8 個以上的小寫字母或數字 `0` 到 `5`。 可以刪掉 `minlength` 屬性，測試一下。

# --hints--

密碼 `input` 元素應該添加 `pattern` 屬性。

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

你應該將 `pattern` 屬性的值設置爲 `[a-z0-5]{8,}`。

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

應該刪除密碼 `input` 元素的 `minlength` 屬性。

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
