---
id: 60fac4095512d3066053d73c
title: 步驟 29
challengeType: 0
dashedName: step-29
---

# --description--

使用 `select` 元素給表單添加一個下拉列表很方便。 `select` 元素是包含一組 `option` 元素的容器，`option` 元素就是下拉列表的內容。 每一個元素都需要一個結束標籤。

首先，在兩個 `label` 元素下面添加一個 `select` 元素。 然後，在 `select` 元素內添加 5 個 `option` 元素。

# --hints--

應該在第三個 `fieldset` 內添加一個 `select` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > select'));
```

應該在 `select` 元素內添加 5 個 `option` 元素。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > select > option')?.length, 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name: <input type="text" required /></label>
        <label>Enter Your Last Name: <input type="text" required /></label>
        <label>Enter Your Email: <input type="email" required /></label>
        <label>Create a New Password: <input type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label>
          <input type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
              </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" />
              </label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}

label {
    display: block;
    margin: 0.5rem 0;
}

```
