---
id: 60fab4a123ce4b04526b082b
title: 步驟 28
challengeType: 0
dashedName: step-28
---

# --description--

目前，用戶可以不選中單選按鈕而提交表格。 雖然你先前能夠使用 `required` 屬性來表明一個輸入框是必須填寫的，但是在這裏不行，因爲給兩個輸入都添加 `required` 會向用戶傳遞錯誤的信息。

爲了解決這個問題，你可以在第二個 `fieldset` 中的 `label` 元素之前添加一個有文本 `Account type (required)` 的 `legend` 元素，從而提供所需的上下文。 然後在 `Personal` 輸入中添加 `checked` 屬性，以確保提交表單時包含所需數據。

# --hints--

你的 `input` 元素應該保持 `type` 爲 `radio`。

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]')?.length, 2);
```

你應該在第二個 `fieldset` 元素中添加一個 `legend` 元素。

```js
assert.equal(document.querySelectorAll('fieldset:nth-of-type(2) > legend')?.length, 1);
```

你應該給 `legend` 元素添加 `Account type (required)` 文本。

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

你應該給 `Personal` 輸入添加 `checked` 屬性。

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

        <label><input type="radio" name="account-type" /> Personal</label>
        <label><input type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
