---
id: 60fab9f17fa294054b74228c
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

現在到了最後一個 `fieldset`。 如果你想允許用戶上傳個人資料圖片，應該怎麼做呢？

類型爲 `file` 的 `input` 可以實現。 添加一個 `label`，其文本爲 `Upload a profile picture:`，添加一個可以接收文件上傳的 `input`。

# --hints--

應該添加一個 `label` 並設置其文字爲 `Upload a profile picture:`。

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label')?.innerText, /Upload a profile picture:/i);
```

應該在 `label` 元素內添加一個 `input` 元素。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label > input'));
```

`input` 元素的 `type` 應該爲 `file`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label > input')?.type, 'file');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name: <input type="text" required /></label>
        <label>Enter Your Last Name: <input type="text" required /></label>
        <label>Enter Your Email: <input type="email" required /></label>
        <label>Create a New Password: <input type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label>
          <input type="checkbox" required /> I accept the
                    <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
              </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}

label {
    display: block;
    margin: 0.5rem 0;
}

```
