---
id: 62cc5b1779e4d313466f73c5
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `head` 元素中嵌套一個自閉合 `link` 元素。 設置其 `rel` 屬性值爲 `stylesheet`， `href` 屬性值爲 `styles.css`。

# --hints--

代碼應該有一個 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

不應更改現有的 `head` 標籤。 確保沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

應該有一個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 元素應該在 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
