---
id: 62ff8e998d3e7eae14d6ae3b
title: 第 30 步
challengeType: 0
dashedName: step-30
---

# --description--

你需要確認用戶已閱讀條款和條件。

添加一個 `label` 元素。 在新創建的 `label` 元素中添加一個 `input` 元素，並將 `type` 屬性設置爲 `checkbox`。 將這個 `input` 元素設置爲 `required`，這樣用戶在不同意條款和條件的情況下無法註冊。

爲元素添加 `id` 和 `for` 屬性，值爲 `terms-and-conditions`，以增加可訪問性。

# --hints--

你應該在第三個 `fieldset` 元素之後添加一個 `label`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

你應該給 `label` 元素添加一個 `input`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

你應該給 `input` 元素添加值爲 `checkbox` 的 `type` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

你應該給 `input` 元素添加一個 `required` 屬性。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` 元素應該有一個值爲 `terms-and-conditions` 的 `id`。

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` 元素應該有一個值爲 `terms-and-conditions` 的 `for` 屬性。

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
