---
id: 63541ef4f96cd82e8e6c788a
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

`method` 屬性指定了如何將表單數據發送到 `action` 屬性中指定的 URL。 表單數據可以通過 `GET` 請求作爲 URL 參數發送（`method="get"`）或通過 `POST` 請求作爲請求正文中的數據發送（`method="post"`）。

設置 `method` 屬性通過 `POST` 請求發送你的表單數據。

# --hints--

你不應該添加一個新的 `form` 元素。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

你的 `form` 元素應有一個 `method` 屬性。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

你的 `method` 屬性應設置爲 `post`。

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
