---
id: 5d822fd413a79914d39e98cc
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the CSS Variables Skyline project! Start by adding a `link` element that links your `styles.css` file within the `head` element.

# --hints--

Your code should have a `link` element.

```js
assert.exists(document.querySelector('link'), 'Your code should have a `link` element.');
```

You should not change your existing `head` tags. Make sure you did not delete the closing tag.

```js
const headElements = document.querySelectorAll('head');
assert.lengthOf(headElements, 1, 'You should not change your existing `head` tags. Make sure you did not delete the closing tag.');
```

You should have one void `link` element.

```js
assert.lengthOf(document.querySelectorAll('link'), 1, 'You should have one void `link` element.');
```

Your `link` element should be within your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);


const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const headElements = document.querySelectorAll('head');
assert.lengthOf(headElements, 1, 'You should not change your existing `head` tags. Make sure you did not delete the closing tag.');
const rel = linkElement?.getAttribute("rel");
assert.lengthOf(document.querySelectorAll('link'), 1, 'You should have one void `link` element.');
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute('rel');
const href = linkElement?.getAttribute('href');
assert.equal(rel, 'stylesheet', 'Your `link` element should have a `rel` attribute with the value `stylesheet`.');
assert.strictEqual(href, 'styles.css', 'Your `link` element should have an `href` attribute with the value `styles.css`.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
