---
id: 69f0a37c3bf194a7c9c62492
title: Build a Device Loan Ledger
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

You are a software developer tasked with creating a ledger application to manage IT hardware devices provisioned by your company's Service Desk. The ledger's structure is an object made up of `assetTag` properties, and values that are also objects:

```js
const ledger = {
  assetTag: {
    type: "type",
    status: "status",
    borrower: {
      name: "name",
      email: "email"
    },
    dueDate: "dueDate"
  }
};
```

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should model the ledger as `{ [assetTag]: { type, status, borrower: { name, email }, dueDate } }`.
1. You should implement a `checkoutDevice(ledger, assetTag, borrower)` function that clones the ledger, sets the borrower's `name` and `email`, updates the `status` to `"CheckedOut"`, and returns a confirmation string.
1. You should implement a `checkinDevice(ledger, assetTag)` function that clears the borrower's `name` and `email`, resets the `dueDate` to `""`, updates the `status` to `"CheckedIn"`, and returns a confirmation string.
1. You should implement a `listOverdueDevices(ledger, today)` function that returns an array of overdue entries sorted by due date ascending. Do not use the JavaScript `Date` object.
1. You should implement `serializeLedger(ledger)` and `loadLedger(json)` helper functions to convert between a JavaScript object and a JSON string.

# --hints--

You should create a function named `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

You should create a function named `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

You should create a function named `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

You should create a function named `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

You should create a function named `loadLedger`.

```js
assert.isFunction(loadLedger);
```

`serializeLedger` should return a JSON string.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

`loadLedger` should return a JavaScript object.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

`checkoutDevice` should not mutate the original ledger.

```js
const ledger = JSON.stringify({
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
});
const original = JSON.parse(ledger);
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(JSON.parse(ledger), original);
```

`checkoutDevice` should update the borrower's name and email and set status to `"CheckedOut"`.

```js
const ledger = JSON.stringify({
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
});
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isString(result);
assert.include(result, "Alice");
```

`checkoutDevice` should return a confirmation string that includes the asset tag and borrower name.

```js
const ledger = JSON.stringify({
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
});
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result);
assert.include(result, "2");
assert.include(result, "Bob");
```

`checkinDevice` should clear borrower info, reset `dueDate`, and set status to `"CheckedIn"`.

```js
const ledger = JSON.stringify({
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
});
const result = checkinDevice(ledger, "1");
assert.isString(result);
assert.include(result, "1");
```

`checkinDevice` should return a confirmation string that includes the asset tag.

```js
const ledger = JSON.stringify({
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
});
const result = checkinDevice(ledger, "3");
assert.isString(result);
assert.include(result, "3");
```

`listOverdueDevices` should return an array.

```js
const ledger = JSON.stringify({
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
});
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

`listOverdueDevices` should return only overdue `CheckedOut` devices.

```js
const ledger = JSON.stringify({
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
});
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

`listOverdueDevices` should return entries sorted by due date ascending.

```js
const ledger = JSON.stringify({
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
});
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```


# --seed--

## --seed-contents--

```js
const equipmentLedger = JSON.stringify({
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
});


```

# --solutions--

```js
const equipmentLedger = JSON.stringify({
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
});

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const loaded = loadLedger(ledger);
  loaded[assetTag].status = "CheckedOut";
  loaded[assetTag].borrower.name = borrower.name;
  loaded[assetTag].borrower.email = borrower.email;
  return `${loaded[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
}

function checkinDevice(ledger, assetTag) {
  const loaded = loadLedger(ledger);
  loaded[assetTag].status = "CheckedIn";
  loaded[assetTag].borrower.name = "";
  loaded[assetTag].borrower.email = "";
  loaded[assetTag].dueDate = "";
  return `${loaded[assetTag].type} with asset tag ${assetTag} has been checked in.`;
}

function reformatDate(date) {
  const [month, day, year] = date.split("/");
  return parseInt(`${year}${month.padStart(2, "0")}${day.padStart(2, "0")}`, 10);
}

function listOverdueDevices(ledger, today) {
  const loaded = loadLedger(ledger);
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in loaded) {
    const item = loaded[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
