---
id: 6895dffc64e97df0e62e9ef9
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Every class needs an initializer method, which is called when you create a new instance of the class. In Python, the initializer is named `__init__`.

Inside your `MusicalInstrument` class, remove `pass` and add an `__init__` method that takes three parameters: `self`, `name`, and `instrument_type`. Add `pass` inside the `__init__` method for now.

The `self` parameter refers to the instance of the class and is required as the first parameter for all methods in a class.

# --hints--

You should define an `__init__` method inside your `MusicalInstrument` class.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

Your `__init__` method should take three parameters: `self`, `name`, and `instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
