---
id: 61329d52e5010e08d9b9d66b
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Ein weiteres wichtiges `meta` Element für Accessibility und SEO ist die `description` Definition. Der Wert des Attributs `content` wird von Suchmaschinen verwendet, um eine Beschreibung deiner Seite bereitzustellen.

Füge ein `meta` Element mit einem `name` Attribut von `description` hinzu, and gebe ihm ein nützliches `content` Attribut.

# --hints--

Du solltest ein neues `meta` Element zu dem `head` hinzufügen.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Du solltest `meta` ein `name` Attribut mit `description` geben.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Du solltest `meta` ein `content` Attribut geben.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Der `content` Attributwert sollte nicht mehr als 165 Zeichen lang sein. _Das ist die maximale Beschreibungslänge von Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
    background: #f5f6f7;
    color: #1b1b32;
    font-family: Helvetica;
    margin: 0;
}
```
