---
id: 61329b210dac0b08047fd6ab
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Um mit den `meta`-Elementen fortzufahren: Durch eine `viewport`-Definition wird dem Browser mitgeteilt, wie er die Seite darzustellen hat. Indem du eine solche hinzufügst, verbesserst du die visuelle Barrierefreiheit deiner Seite auf mobilen Endgeräten und dein _SEO_ (Search Engine Optimization – Suchmaschinenoptimierung).

Füge eine `viewport`-Definition mit einem `content`-Attribut hinzu, das die `width` und den `initial-scale` der Seite beschreibt.

# --hints--

Du solltest ein weiteres `meta`-Element im `head`-Element erstellen.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Du solltest dem `meta`-Element ein `name`-Attribut von `viewport` geben.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Du solltest `meta` ein `content`-Attribut von `width=device-width, initial-scale=1` geben.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
