---
id: 613297a923965e0703b64796
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Vielleicht kennst du bereits das `meta`-Element; es wird dazu verwendet, bestimmte Informationen über die Seite, wie beispielsweise Titel, Beschreibung, Schlüsselwörter und Autor anzugeben.

Gib deiner Seite ein `meta`-Element mit einem entsprechenden `charset`-Wert.

Das `charset`-Attribut bestimmt über die Zeichenkodierung der Seite. Heutzutage ist `UTF-8` die einzige Kodierung, die von den meisten Browsern unterstützt wird.

# --hints--

Du solltest ein `meta`-Element innerhalb eines `head`-Elements erstellen.

```js
assert.exists(document.querySelector('head > meta'));
```

Du solltest dem `meta`-Tag ein `charset` von `UTF-8` geben.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
