---
id: 6133d11ef548f51f876149e3
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Die Seitennavigation ist ein wesentlicher Bestandteil der Barrierefreiheit, und Screenreader sind darauf angewiesen, dass du die Struktur deiner Seite angibst. Das geschieht mit semantischen HTML-Elementen.

F端ge die Elemente `header` und `main` deiner Seite hinzu.

Das `header`-Element wird verwendet, um die Seite vorzustellen und ein Navigationsmen端 bereitzustellen.

Das `main`-Element wird den Kerninhalt deiner Seite enthalten.

# --hints--

Du solltest ein `header`-Element dem `body` hinzuf端gen.

```js
assert.exists(document.querySelector('body > header'));
```

Du solltest ein `main`-Element dem `body` hinzuf端gen.

```js
assert.exists(document.querySelector('body > main'));
```

Das `header`-Element sollte vor `main` stehen.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
