---
id: 6133acc353338c0bba9cb553
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Schließlich ist das Element `title` im `head` für Bildschirmleseprogramme nützlich, um den Seiteninhalt zu verstehen. Außerdem ist es ein wichtiger Bestandteil des _SEO_.

Gib deiner Seite einen beschreibenden und prägnanten `title`.

# --hints--

Du sollest dem `head` ein `title`-Element hinzufügen.

```js
assert.exists(document.querySelector('head > title'));
```

Der `title` sollte nicht länger als 60 Zeichen sein.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Versuche, einen besser beschreibenden `title` zu finden. _Tipp: Mindestens 15 Zeichen_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
