---
id: 613e2546d0594208229ada50
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Verschachtele die Elemente `img`, `h1` und `nav` im `header`, um deiner Seite Kontext zu verleihen.

Hierbei sollte `img` auf `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` verweisen und eine auf `logo` gesetzte `id` haben.

Das Element `h1` sollte den Text `HTML/CSS Quiz` enthalten.

# --hints--

Du solltest das Element `img` dem `header` hinzufügen.

```js
assert.exists(document.querySelector('header > img'));
```

Du solltest das Element `h1` dem `header` hinzufügen.

```js
assert.exists(document.querySelector('header > h1'));
```

Du solltest das Element `nav` dem `header` hinzufügen.

```js
assert.exists(document.querySelector('header > nav'));
```

Du solltest die Elemente `img`, `h1` und `nav` in genau dieser Reihenfolge platzieren.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Du solltest dem Element `img` ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` zuweisen.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Du solltest dem Element `img` ein `id`-Attribut mit dem Wert `logo` zuweisen.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Du solltest dem Element `h1` den Text `HTML/CSS Quiz` geben.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
