---
id: 613e275749ebd008e74bb62e
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Eine nützliche Eigenschaft einer _SVG_ (scalable vector graphics, skalierbare Vektorgrafiken) ist, dass dieses ein `path`-Attribut enthält, mit welchem ein Bild, ohne die Auflösung des Resultats zu beeinflussen, skaliert werden kann.

Derzeit übernimmt das `img` die Ursprungsgröße, welche aber zu groß ist. CSS verfügt über eine `max`-Funktion, die den größte Wert einer durch Kommata getrennten Menge von Werten zurückgibt. Zum Beispiel:

```css
img {
  width: max(250px, 25vw);
}
```

Im obigen Beispiel wird die Breite des Bildes 250px sein, falls die Breite des Viewports kleiner als 1000 Pixel ist. Falls die Breite des Viewports größer als 1000 Pixel ist, wird die Breite des Bildes 25vw betragen. Das liegt daran, dass 25vw 25% der Breite des Viewports entspricht.

Skaliere das Bild mit seiner `id` als Selektor und setzte die `width` (Breite) auf das Maximum von `10rem` oder `18vw`.

# --hints--

Du solltest den `#logo`-Selektor verwenden, um das `img`-Element auszuwählen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Du solltest dem `img` eine `width` von `max(10rem, 18vw)` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
