---
id: 6140827cff96e906bd38fc2b
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Wie im <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp-Stil-Leitfaden</a> beschrieben, sollte das Logo ein Seitenverhältnis von `35 / 4` beibehalten und genug Abstand zum Text halten.

Ändere zuerst die `background-color` zu `#0a0a23`, damit du das Logo sehen kannst. Verwende anschließend die `aspect-ratio`-Eigenschaft, um das gewünschte Seitenverhältnis auf `35 / 4` einzustellen. Füge zum Schluss ein `padding`von `0.4rem` an allen Seiten hinzu.

# --hints--

Du solltest `#logo` eine `background-color` von `#0a0a23` zuweisen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

Du solltest die Eigenschaft `aspect-ratio` verwenden.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

Du solltest die Eigenschaft `height` nicht verwenden.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

Du solltest die Eigenschaft `width` nicht ändern.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

Du solltest dem `img` ein `aspect-ratio` von `35 / 4` zuweisen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

Du solltest dem `img` ein `padding` von `0.4rem` zuweisen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(100px, 18vw);

}
--fcc-editable-region--

```
