---
id: 6140883f74224508174794c0
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Lasse den `header` die volle Breite seines übergeordneten Containers annehmen, setze die `height` auf `50px` und setze die `background-color` auf `#1b1b32`. Setze `display` anschließend auf _Flexbox_.

# --hints--

Du solltest den `header`-Element-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

Du solltest dem `header` eine `width` von `100%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

Du solltest dem `header` eine `height` von `50px` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

Du solltest `header` eine `background-color` von `#1b1b32` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

Du solltest dem `header` ein `display` von `flex` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
