---
id: 6141fabd6f75610664e908fd
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Da es sich um ein Quiz handelt, brauchst du ein Formular, in das Nutzer ihre Antworten eintragen können. Du kannst den Inhalt innerhalb des Formulars mit `section`-Elementen semantisch trennen.

Erstelle innerhalb des `main`-Elements ein Formular, das drei verschachtelte `section`-Elemente enthält. Übermittle dann das Formular mit der richtigen Methode an `https://freecodecamp.org/practice-project/accessibility-quiz`.

# --hints--

Verschachtele das `form`-Element innerhalb deines `main`-Elements.

```js
assert.exists(document.querySelector('main > form'));
```

Du solltest drei `section`-Elemente in deinem `form`-Element verschachteln.

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

Du solltest dem `form`-Element ein `action`-Attribut geben.

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

Du solltest dem `action`-Attribut den Wert `https://freecodecamp.org/practice-project/accessibility-quiz` geben.

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

Du solltest dem `form`-Element ein `method`-Attribut geben.

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

Du solltest dem `form`-Element ein `method`-Attribut mit dem Wert `post` geben.

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
