---
id: 614396f7ae83f20ea6f9f4b3
title: Schritt 26
challengeType: 0
dashedName: step-26
---

# --description--

Füge innerhalb des zweiten `section`-Elements zwei `div`-Elemente der Klasse `question-block` ein.

Anschließend fügst du innerhalb eines jeden `div.question-block`-Elements ein `p`-Element ein, dessen Text aus jeweils aufsteigenden Zahlen besteht – angefangen mit `1`. Füge außerdem ein `fieldset`-Element der Klasse `question` ein.

# --hints--

Du solltest zwei `div`-Elemente innerhalb des zweiten `section`-Elements verschachteln.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.length, 2);
```

Du solltest dem ersten neuen `div`-Element die Klasse `question-block` zuweisen.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[0]?.className, 'question-block');
```

Du solltest dem zweiten neuen `div`-Element die Klasse `question-block` zuweisen.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div')?.[1]?.className, 'question-block');
```

Du solltest ein `p`-Element innerhalb eines jeden `div.question-block`-Elements verschachteln.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > p')?.length, 2);
```

Du solltest dem ersten `p`-Element den Text `1` zuweisen.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > p')?.[0]?.textContent, '1');
```

Du solltest dem zweiten `p` Element den Text `2` zuweisen.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > p')?.[1]?.textContent, '2');
```

Du solltest ein `fieldset`-Element innerhalb eines jeden `div.question-block`-Elements verschachteln.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.length, 2);
```

Du solltest das erste `fieldset`-Element nach dem ersten `p`-Element einfügen.

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block > p + fieldset'));
```

Du solltest das zweite `fieldset`-Element nach dem zweiten `p`-Element einfügen.

```js
assert.exists(document.querySelector('section:nth-of-type(2) > div.question-block:nth-of-type(2) > p + fieldset'));
```

Du solltest dem ersten `fieldset` die Klasse `question` zuweisen.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[0]?.className, 'question');
```

Du solltest dem zweiten `fieldset`-Element die Klasse `question` zuweisen.

```js
assert.equal(document.querySelectorAll('section:nth-of-type(2) > div.question-block > fieldset')?.[1]?.className, 'question');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input type="text" name="student-name" id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input type="email" name="student-email" id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">D.O.B.<span class="sr-only">(Date of Birth)</span></label>
            <input type="date" name="birth-date" id="birth-date" />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>

        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

.sr-only {
  position: absolute;
  width: 1px;
  height: 1px;
  padding: 0;
  margin: -1px;
  overflow: hidden;
  clip: rect(0, 0, 0, 0);
  white-space: nowrap;
  border: 0;
}

```
