---
id: 614ccc21ea91ef1736b9b578
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Willkommen zum ersten Teil des Barrierefreiheit-Quiz. Da du langsam ein erfahrener HTML- und CSS-Entwickler wirst, haben wir dich bereits mit dem grundlegenden Boilerplate-Code ausgestattet.

Beginne diese Einführung in die Barrierefreiheit, indem du deinem `html`-Element ein `lang`-Attribut hinzufügst. Das hilft Bildschirmleseprogrammen bei der Identifizierung der Seitensprache.

# --hints--

Du solltest dem Element `html` ein `lang`-Attribut geben. _Hinweis: Du kannst den Wert `en` für Englisch verwenden._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
