---
id: 6718d2d59337c822ecb697ff
title: Build a Bank Account Management Program
challengeType: 14
dashedName: build-a-bank-account-management-program
---


# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should define a class named `BankAccount` with a `constructor` that has a `balance` parameter. The default parameter value should be set to 0. The constructor should also initialize an empty `transactions` array to store transaction records.

2. You should define a method named `updateBalance`.

3. You should define a method named `deposit` with an `amount` parameter. This method should update `balance` accordingly.

4. You should define a method named `withdraw` with an `amount` parameter. This method should update `balance` accordingly.

5. You should define a method named `checkBalance` that returns the current balance.

6. You should define a method named `listAllDeposits` that returns all deposits.

7. You should define a method named `listAllWithdrawals` that returns all withdrawals.

# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor(balance = 0) {
      this.balance = balance;
      this.transactions = [];
    }
  
    deposit(amount) {
      if (amount > 0) {
        this.transactions.push({ type: 'deposit', amount });
        this.updateBalance();
        return `Successfully deposited $${amount}. New balance: $${this.balance}`;
      } else {
        return 'Deposit amount must be greater than zero.';
      }
    }
  
    withdraw(amount) {
      if (amount > 0 && amount <= this.balance) {
        this.transactions.push({ type: 'withdraw', amount });
        this.updateBalance();
        return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
      } else {
        return 'Insufficient balance or invalid amount.';
      }
    }
  
    checkBalance() {
      return `Current balance: $${this.balance}`;
    }
  
    updateBalance() {
      this.balance = this.transactions.reduce((acc, transaction) => {
        return transaction.type === 'deposit' ? acc + transaction.amount : acc - transaction.amount;
      }, 0);
    }
  
    listAllDeposits() {
      const deposits = this.transactions
        .filter(transaction => transaction.type === 'deposit')
        .map(transaction => transaction.amount);
      return `Deposits: ${deposits}`;

    }
  
    listAllWithdrawals() {
      const withdrawals = this.transactions
        .filter(transaction => transaction.type === 'withdraw')
        .map(transaction => transaction.amount);
      return `Withdrawals: ${withdrawals}`;
    }
}
  
const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
