---
id: 6718d2d59337c822ecb697ff
title: Build a Bank Account Management Program
challengeType: 14
dashedName: build-a-bank-account-management-program
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should define a class named `BankAccount` with a constructor that takes `balance` as parameter. The constructor should set the default `balance` to `0`and initialize an empty array named `transactions` to store transaction records as objects. 

2. Each transaction stored in the `transactions` array should be an object with two properties: `type` and `amount`. The `type` property should be either `deposit` or `withdraw`, and the `amount` property should be the amount deposited or withdrawn.

3. Inside the `BankAccount` class, you should define a method named `updateBalance` that updates the balance according to the type of transaction. If the transaction type is `deposit`, it should add the amount to the balance. If the transaction type is `withdraw`, it should subtract the amount from the balance.

4. You should define a method named `deposit` that takes the deposit amount as a parameter. This method should update the current balance according to deposits. When the deposit amount is greater than `0`, it should:

   - Push a new object to the transactions array with a type of deposit and the amount deposited.
   - Update the balance.   
   - Return `"Successfully deposited $[amount]. New balance: $[balance]"`.

5. In the `deposit` method, if the amount is less than or equal to `0`, it should return `"Deposit amount must be greater than zero."`.

6. You should define a method named `withdraw` that takes an amount as a parameter. This method should update the current balance according to withdrawals. When the amount to be withdrawn is greater than `0` and less than or equal to the current balance, it should:

   - Push a new object to the transactions array with a type of withdraw and the amount withdrawn.
   - Update the balance.
   - Return `"Successfully withdrew $[amount]. New balance: $[balance]"`. 

7. In the `withdraw` method, when the amount to be withdrawn is less than or equal to `0` or greater than the current balance, it should return `"Insufficient balance or invalid amount."`.


8. You should define a method named `checkBalance` that returns the current balance in the format `"Current balance: $[balance]"`.

9. You should define a method named `listAllDeposits` that iterates through the transactions array and returns all deposits in the format `"Deposits: deposit,deposit,..."`.

10. You should define a method named `listAllWithdrawals` that iterates through the transactions array and returns all withdrawals in the format `"Withdrawals: withdrawal,withdrawal,..."`.

# --hints--

You should define a class named `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

The `BankAccount` object should initially have a `balance` of `0` and an empty array `transactions` to store transaction records.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

You should have a `deposit` method that takes the deposit amount as a parameter.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

You should have a `withdraw` method that takes the withdrawal amount as a parameter.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

You should have a `checkBalance` method that checks the current balance.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```

You should have a `listAllDeposits` method that lists all deposits.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

You should have a `listAllWithdrawals` method that lists all withdrawals.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

You should have an `updateBalance` method that updates the balance according to the type of transaction.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.updateBalance);
```

`BankAccount.deposit(100)` should return `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")

// prevent hardcoding
assert.equal(myAccount.deposit(200), "Successfully deposited $200. New balance: $300")

```

`BankAccount.deposit(-50)` should return `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")

// prevent hardcoding
assert.equal(myAccount.deposit(-30), "Deposit amount must be greater than zero.")

```

`BankAccount.deposit(0)` should return `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(0), "Deposit amount must be greater than zero.")
```

When the account balance is `100`, `BankAccount.withdraw(150)` should return `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(250), "Insufficient balance or invalid amount.")

```

`BankAccount.withdraw(-50)` should return `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")

// prevent hardcoding
assert.equal(myAccount.withdraw(-30), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(0)` should return `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(0), "Insufficient balance or invalid amount.")
```

When the account balance is `200`, `BankAccount.withdraw(150)` should return `"Successfully withdrew $150. New balance: $50"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.withdraw(150), "Successfully withdrew $150. New balance: $50")
```

When the account balance is `200`, `BankAccount.checkBalance()` should return `"Current balance: $200"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(200);
assert.equal(myAccount.checkBalance(), "Current balance: $200")

// prevent hardcoding
myAccount.deposit(300);
assert.equal(myAccount.checkBalance(), "Current balance: $500")
```

When you deposit `10`, `35`, `90`, the `listAllDeposits` method should return `"Deposits: 10,35,90"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(10);
myAccount.deposit(35);
myAccount.deposit(90);
assert.equal(myAccount.listAllDeposits(), "Deposits: 10,35,90")
```

When you withdraw `20`, `50`, `100`, the `listAllWithdrawals` method should return `"Withdrawals: 20,50,100"`.

```js
const myAccount = new BankAccount();
myAccount.deposit(500);

myAccount.withdraw(20);
myAccount.withdraw(50);
myAccount.withdraw(100);
assert.equal(myAccount.listAllWithdrawals(), "Withdrawals: 20,50,100")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor(balance = 0) {
      this.balance = balance;
      this.transactions = [];
    }
  
    deposit(amount) {
      if (amount > 0) {
        this.transactions.push({ type: 'deposit', amount });
        this.updateBalance();
        return `Successfully deposited $${amount}. New balance: $${this.balance}`;
      } else {
        return 'Deposit amount must be greater than zero.';
      }
    }
  
    withdraw(amount) {
      if (amount > 0 && amount <= this.balance) {
        this.transactions.push({ type: 'withdraw', amount });
        this.updateBalance();
        return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
      } else {
        return 'Insufficient balance or invalid amount.';
      }
    }
    checkBalance() {
      return `Current balance: $${this.balance}`;
    }
  
    updateBalance() {
      this.balance = this.transactions.reduce((acc, transaction) => {
        return transaction.type === 'deposit' ? acc + transaction.amount : acc - transaction.amount;
      }, 0);
    }
  
    listAllDeposits() {
      const deposits = this.transactions
        .filter(transaction => transaction.type === 'deposit')
        .map(transaction => transaction.amount);
      return `Deposits: ${deposits}`;

    }
  
    listAllWithdrawals() {
      const withdrawals = this.transactions
        .filter(transaction => transaction.type === 'withdraw')
        .map(transaction => transaction.amount);
      return `Withdrawals: ${withdrawals}`;
    }
}
  
const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
