---
id: 6718d2d59337c822ecb697ff
title: Build a Bank Account Management Program
challengeType: 14
dashedName: build-a-bank-account-management-program
---


# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should define a class named `BankAccount` with a constructor that takes `balance` as parameter. The default `balance` value should be set to `0`. The constructor should also initialize an empty array `transactions` to store transaction records as objects. 

2. Inside the `BankAccount` class, You should define a method named `updateBalance` that updates the balance according to the type of transaction. If the transaction type is deposit, it should add the amount to the balance. If the transaction type is withdraw, it should subtract the amount from the balance.

3. You should define a method named `deposit` that takes the deposit amount as a parameter. This method should update the current balance according to deposits. When the deposit amount is greater than `0`, it should:

   - Push a new object to the transactions array with a type of deposit and the amount deposited.
   - Update the balance.   
   - Return `"Successfully deposited $[amount]. New balance: $[balance]"`.

4. In the `deposit` method, if the amount is less than or equal to `0`, it should return `"Deposit amount must be greater than zero."`.

5. You should define a method named `withdraw` that takes an amount as a parameter. This method should update the current balance according to withdrawals. When the amount to be withdrawn is greater than `0` and less than or equal to the current balance, it should:

   - Push a new object to the transactions array with a type of withdraw and the amount withdrawn.
   - Update the balance.
   - Return `"Successfully withdrew $[amount]. New balance: $[balance]"`. 

6. In the `withdraw` method, when the amount to be withdrawn is less than or equal to `0` or greater than the current balance, it should return `"Insufficient balance or invalid amount."`.


7. You should define a method named `checkBalance` that returns the current balance in the format `"Current balance: $[balance]"`.

8. You should define a method named `listAllDeposits` that iterates through the transactions array and returns all deposits in the format `"Deposits: deposit,deposit,..."`.

9. You should define a method named `listAllWithdrawals` that iterates through the transactions array and returns all withdrawals in the format `"Withdrawals: withdrawal,withdrawal,..."`.

# --hints--

You should define a class named `BankAccount`.

```js
assert.isFunction(BankAccount);
assert.isObject(new BankAccount());
```

The `BankAccount` object should initially have a `balance` of `0` and an empty array `transactions` to store transaction records.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.balance, 0);
assert.equal(myAccount.transactions.length, 0);
```

You should have a `deposit` method that takes the deposit amount as a parameter.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.deposit);
assert.lengthOf(myAccount.deposit, 1);
```

You should have a `withdraw` method that takes the withdrawal amount as a parameter.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.withdraw);
assert.lengthOf(myAccount.withdraw, 1);
```

You should have a `checkBalance` method that checks the current balance.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.checkBalance);
```


You should have a `listAllDeposits` method that lists all deposits.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllDeposits);
```

You should have a `listAllWithdrawals` method that lists all withdrawals.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.listAllWithdrawals);
```

You should have an `updateBalance` method that updates the balance according to the type of transaction.

```js
const myAccount = new BankAccount();
assert.isFunction(myAccount.updateBalance);
```

`BankAccount.deposit(100)` should return `"Successfully deposited $100. New balance: $100"`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(100), "Successfully deposited $100. New balance: $100")
```

`BankAccount.deposit(-50)` should return `"Deposit amount must be greater than zero."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.deposit(-50), "Deposit amount must be greater than zero.")
```

When the account balance is `100`, `BankAccount.withdraw(150)` should return `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
myAccount.deposit(100);
assert.equal(myAccount.withdraw(150), "Insufficient balance or invalid amount.")
```

`BankAccount.withdraw(-50)` should return `"Insufficient balance or invalid amount."`.

```js
const myAccount = new BankAccount();
assert.equal(myAccount.withdraw(-50), "Insufficient balance or invalid amount.")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class BankAccount {
    constructor(balance = 0) {
      this.balance = balance;
      this.transactions = [];
    }
  
    deposit(amount) {
      if (amount > 0) {
        this.transactions.push({ type: 'deposit', amount });
        this.updateBalance();
        return `Successfully deposited $${amount}. New balance: $${this.balance}`;
      } else {
        return 'Deposit amount must be greater than zero.';
      }
    }
  
    withdraw(amount) {
      if (amount > 0 && amount <= this.balance) {
        this.transactions.push({ type: 'withdraw', amount });
        this.updateBalance();
        return `Successfully withdrew $${amount}. New balance: $${this.balance}`;
      } else {
        return 'Insufficient balance or invalid amount.';
      }
    }
    checkBalance() {
      return `Current balance: $${this.balance}`;
    }
  
    updateBalance() {
      this.balance = this.transactions.reduce((acc, transaction) => {
        return transaction.type === 'deposit' ? acc + transaction.amount : acc - transaction.amount;
      }, 0);
    }
  
    listAllDeposits() {
      const deposits = this.transactions
        .filter(transaction => transaction.type === 'deposit')
        .map(transaction => transaction.amount);
      return `Deposits: ${deposits}`;

    }
  
    listAllWithdrawals() {
      const withdrawals = this.transactions
        .filter(transaction => transaction.type === 'withdraw')
        .map(transaction => transaction.amount);
      return `Withdrawals: ${withdrawals}`;
    }
}
  
const myAccount = new BankAccount();
  
console.log(myAccount.deposit(0)); 
console.log(myAccount.deposit(500));
console.log(myAccount.withdraw(30));
console.log(myAccount.withdraw(100));
console.log(myAccount.checkBalance());
console.log(myAccount.listAllDeposits());
console.log(myAccount.listAllWithdrawals());
```
