---
id: 66068fb0bfddba2b7977eb60
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you will learn about classes and objects by building a sudoku puzzle solver.

In Python, a class is a blueprint for creating objects. Objects created from a class are instances of that class. You can create a class using this syntax:

```py
class ClassName:
    pass
```

Where `class` is the keyword required to define the class and `ClassName` is the name of the class, written by convention in *PascalCase*.

Begin by creating a `Board` class.

# --hints--

You should create a class named `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
