---
id: 6143862a5e54455d262c212e
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Ndani ya kipengele chako cha `body`, unda kipengele cha `main`. Kisha katika kipengele hicho, unda `section` yenye `class` iliyowekwa kuwa `heading`.

# --hints--

Unapaswa kuwa na kipengele cha `main`.

```js
assert.exists(document.querySelector('main'));
```

Kipengele chako cha `main` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

Unapaswa kuwa na kipengele cha `section`.

```js
assert.exists(document.querySelector('section'));
```

Kipengele chako cha `section` kinapaswa kuwa ndani ya kipengele chako cha `main`.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

Kipengee chako kipya cha `section` kinapaswa kuwa na `class` iliyowekwa kuwa `heading`.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
