---
id: 614385513d91ae5c251c2052
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ongeza kipengele cha `title` chenye maandishi `Magazine`, kipengele cha `link` cha `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` laha ya fonti, `link` ya `https://use.fontawesome.com/releases/v5.8.2/css/all.css` Laha ya mtindo wa FontAwesome, na `link` ya laha yako ya `./styles.css`.

Laha yako ya fonti itapakia fonti tatu tofauti: `Anton`, `Baskervville` na `Raleway`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha kujifunga cha `link`.

```js
assert(document.querySelectorAll('link').length === 3);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert(document.querySelectorAll('head > link').length === 3);
```

Kipengele chako cha tatu cha `link` kinapaswa kuwa na sifa ya `rel` yenye thamani ya `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Moja ya vipengele vyako vya kiungo lazima kiwe na `href` iliyowekwa kuwa `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Moja ya vipengele vyako vya kiungo lazima kiwe na `href` iliyowekwa kuwa `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Moja ya vipengele vyako vya `link` lazima kiwe na sifa ya `href` yenye thamani `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/)
```

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Mradi wako unapaswa kuwa na kichwa cha `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

Kumbuka, ukubwa wa herufi na tahajia ni jambo la muhimu katika kichwa.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
