---
id: 6169cd8a558aa8434e0ad7f6
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

`Referer` ya kichwa cha HTTP kina maelezo kuhusu anwani au URL ya ukurasa ambao huenda mtumiaji anatembelea. Maelezo haya yanaweza kutumika katika uchanganuzi ili kufuatilia ni watumiaji wangapi kutoka kwa ukurasa wako wanaotembelea freecodecamp.org, kwa mfano. Kuweka sifa ya `rel` kuwa `noreferrer` huacha maelezo haya kutoka kwa ombi la HTTP. Kipe kipengele chako cha `a` sifa ya `rel` iliyowekwa kuwa `noreferrer`.

# --hints--

Kipengee chako kipya cha `a` kinapaswa kuwa na `rel` iliyowekwa kuwa `noreferrer`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel') === 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
