---
id: 67eaa957114d373deb3a9149
title: Build a Currency Converter
challengeType: 25
dashedName: build-a-currency-converter
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. Your CurrencyConverter component should render an input element to accept the amount to be converted.
1. Your input element should have a type attribute of "number" and should update the amount when changed.
1. Your CurrencyConverter component should render a select element to choose the currency to convert from.
1. Your select element should include options for "USD", "EUR", and "GBP".
1. Your CurrencyConverter component should use useMemo to optimize the calculation of the converted amount.
1. Your CurrencyConverter component should render a p element that displays the selected currency and appends "to JPY Conversion" to it. (e.g., "USD to JPY Conversion").
1. Your CurrencyConverter component should render a p element displaying the converted amount in JPY in the format Converted Amount: xx.00 JPY.
The converted amount should be rounded to two decimal places.

# --hints--

You should export a `CurrencyConverter` component.

```js
async () => {
  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const a = eval(script);

  assert.property(exports, "CurrencyConverter");
}
```

# --before-all--

```js
async function delay(time) {
  return new Promise((resolve) => setTimeout(resolve, time));
}

// Gets the text of the document excluding that of the matching selector.
function getInnerTextExcept(removingSelector) {
  const body = document.body.cloneNode(true);

  const els = body.querySelectorAll(`${removingSelector},script`);
  els.forEach(element => {
    element.parentNode.removeChild(element)
  });

  return body.innerText;
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Currency Converter</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { CurrencyConverter } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<CurrencyConverter />);
    </script>
</body>

</html>
```

```css

```

```jsx
export function CurrencyConverter() {

}
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Currency Converter</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { CurrencyConverter } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<CurrencyConverter />);
    </script>
</body>

</html>
```

```css
body {
    background-color: #edd5c9;
    font-family: Arial, sans-serif;
}

main {
    background-color: white;
    padding: 20px;
    border-radius: 8px;
    box-shadow: 0px 4px 6px rgba(0, 0, 0, 0.1);
    text-align: center;
    width: 400px;
    margin: 50px auto;
}

h2 {
    margin-bottom: 16px;
}

input,
select {
    width: 100%;
    display: block;
    box-sizing: border-box;
    padding: 8px;
    margin-bottom: 12px;
    border: 1px solid #ccc;
    border-radius: 4px;
}

.conversion-display {
    font-size: 16px;
    color: #555;
    margin-bottom: 10px;
}

p {
    font-size: 18px;
    font-weight: bold;
    color: green;
}
```

```jsx
export default function CurrencyConverter() {
    const [amount, setAmount] = React.useState(1);
    const [currency, setCurrency] = React.useState("USD");
  
    const exchangeRates = {
      USD: 150, // 1 USD = 150 JPY
      EUR: 155, // 1 EUR = 155 JPY
      GBP: 190, // 1 GBP = 190 JPY
    };

    const convertedAmount = React.useMemo(() => {
      return (amount * exchangeRates[currency]).toFixed(2);
    }, [amount, currency]);
  
    return (
      <main>
        <h1>Currency Converter</h1>
        <p className="conversion-display">{currency} to JPY Conversion</p>
        <input
          type="number"
          value={amount}
          onChange={(e) => setAmount(Number(e.target.value))}
        />
        <select value={currency} onChange={(e) => setCurrency(e.target.value)}>
          {Object.keys(exchangeRates).map((curr) => (
            <option key={curr} value={curr}>
              {curr}
            </option>
          ))}
        </select>
        <p>Converted Amount: <span>{convertedAmount}</span> JPY</p>
      </main>
    );
  }
```
