---
id: 68eaeaf8b317a4294f170f40
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

step 1 instructions

In this workshop, you will practice working with CSS flexbox by design a set of colored boxes. This is what you will be doing [](url)  

To begin the project, link the `style.css` into your `html` 

# --hints--

Your `link` element should have an `href` attribute with the value `styles.css`.

Test 1

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```


# --seed--

## --seed-contents--

--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
