---
id: 68eaeaf8b317a4294f170f40
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will practice working with CSS flexbox by designing a set of colored boxes.  

To begin the project, link the `style.css` into your `html` 

# --hints--

Create a `link` element inside `<head>`

```js
assert.exists('link')
```

Your `link` tag should have an `href` atribute

```js
const hrefAttribute = linkElement.getAttribute('href');
assert.exists(hrefAttribute, 'style.css');
```

Your `href` should link to `styles.css`.

```js
const linkElement = document.querySelector('head link');
const rel = linkElement.getAttribute('rel');
assert.equal(rel, 'stylesheet', 'Your link should have rel="stylesheet"');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```
