---
id: 68eaeaf8b317a4294f170f40
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will practice working with CSS flexbox by designing a set of colored boxes.  

To begin the project, link the `styles.css` into your `html` 

# --hints--

Create a `link` element inside `head`.

```js
assert.exists(document.querySelector('head>link'));
```


Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const link_element = document.querySelector('link');
assert.equal(link_element?.getAttribute('rel'), 'stylesheet');
```


Your `link` element should have an `href` attribute with the value `styles.css`.

```js
 const linkHrefValue = document.querySelector('link')?.dataset?.href; 
 assert.strictEqual(linkHrefValue, 'styles.css');
```


# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
--fcc-editable-region--  
  <body>
  </body>
</html>

```
