---
id: 68ed661e9f463c155bb677a1
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now, add the title `CSS flexbox colorful boxes` in a `h1` element inside the `header` element.

# --hints--

You should have a `header` element inside `body` element

```js
const bodyElement = document.querySelector('body');
const headerElement = bodyElement?.querySelector('header');
assert.exists(headerElement);

```

Your `header` element should have an `h1` element inside it.

```js
const headerElement = document.querySelector('header');
const h1Element = headerElement?.querySelector('h1');
assert.exists(h1Element);
```

Your `h1` element should have the text `CSS flexbox colorful boxes` inside it.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'CSS flexbox colorful boxes');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--
```
