---
id: 68ed661e9f463c155bb677a1
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop you will practice working with CSS flexbox by designing a set of colored boxes.

To begin, inside the `body` element, add a `header` element. Nest an `h1` element inside the `header` element with the text `Colored Boxes Layout`.

# --hints--

You should have a `header` element inside the `body` element.

```js
const headerEl = document.querySelector('body > header');
assert.exists(headerEl);
```

Your `header` element should have an `h1` element inside it.

```js
const h1El = document.querySelector('body > header > h1');
assert.exists(h1El);
```

Your `h1` element should have the text `Colored Boxes Layout` inside it.

```js
const h1El = document.querySelector('body > header > h1');
assert.equal(h1El?.textContent.trim(), 'Colored Boxes Layout');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
</head>
<body>
--fcc-editable-region--
	
--fcc-editable-region--
</body>
</html>
```
