---
id: 68ef216359687854fa2f3081
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Now we are going to organize the boxes inside our flexbox container. First, create a new class called `box1` in your `styles.css` file. Add a `background` property with the value `#f16e79`.

Go to your `html` file and add the class `box1` to the first box in your div.

Now add a property called `order` with the value `1` to the `box1` class. You'll see the first box move to the end. This happens because the `box` class already has an `order` value of `0`, and items with higher order values appear later.

Also add the `flex-grow` property with a value of `1`. This will make the box grow to fill extra space on its line.

# --hints--

Create a class called `box1`.

```js
assert.exist(new __helpers.CSSHelp(document).getStyle('.box1'));
```

Your class `box1` should have a property `background` with the value `#f16e79`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.background, '#f16e79');
```

Your class `box1` should have a property `order` with the value `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.order, '1');
```

Your class `box1` should have a property `flex-grow` with the value `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box1')?.flex-grow, '1');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>    
  </body>
</html>
--fcc-editable-region--

```

```css
--fcc-editable-region--
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}


--fcc-editable-region--
```
