---
id: 68ef9c7dbdab54899bdf1738
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Finally, let's create our last class called `box6`. Add a `background` property with the value `#3f51b5`.

Go to your `html` file and add the class `box6` to the last box in your div.

Now add the `order` property with a value of `5` to the `box6` class. You'll see this box move after the one with the `box5` class.

Add the `flex-grow` property with a value of `1`. This will make box 6 grow to fill any remaining space in the container.

# --hints--

Create a class called `box6`.

```js
assert.exist(new __helpers.CSSHelp(document).getStyle('.box6'));
```

Your class `box6` should have a property `background` with the value `#3f51b5`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.background, '#3f51b5');
```

Your class `box6` should have a property `order` with the value `5`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.order, '5');
```

Your class `box6` should have a property `flex-grow` with the value `1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box6')?.flex-grow, '1');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>
    <div class="flex-container">
      <div class="box box1">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box box2">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box box3">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box box4">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box box5">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
--fcc-editable-region--

```

```css
--fcc-editable-region--
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3; 
}

.box5 {
  background: #457b9d;
  order: 4;
}
--fcc-editable-region--

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Colored Boxes Layout</h1>

    <div class="flex-container">
      <div class="box box1">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box box2">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box box3">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box box4">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box box5">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box box6">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1; 
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2; 
  flex-shrink: 3; 
}

.box4 {
  background: #4caf50;
  order: 3;
}
.box5 {
  background: #457b9d;
  order: 4; 
}

.box6 {
  background: #3f51b5;
  order: 5; 
  flex-grow: 1; 
}
```
