---
id: 68ef0b3237fef991c92d9c19
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now add a new `div` element with the class `flex-container` to your `header`.

Inside that `flex-container`, create six div elements with the `box` class. Add an `h3` element inside each div with these labels in order: `Box 1`, `Box 2`, `Box 3`, `Box 4`, `Box 5`, and `Box 6`.

Also add a `p` element below each `h3` with these color names in order: `Red`, `Orange`, `Yellow`, `Green`, `Blue`, and `Indigo`. One color for each box.

We have already populated the properties for each class in the `styles.css` file.

# --hints--

Create a `div` element inside `header` with the class `flex-container`.

```js
const bodyElement = document.querySelector('div');
const headerElement = bodyElement?.querySelector('header');
const classElement = headerElement?.querySelector(".flex-container");
assert.exists(classElement);

```

Add six `div` elements with the class `box`.

```js
const matches = document.querySelectorAll("div.box");
assert.equal(matches.length, 6);
```

Add a `h3` element with the text `Box 1` inside the first `div` with the class `box`.

```js
const heading = document.querySelector("div.box > h3");
assert.equal(heading.textContent, "Box 1");
```

Add a `h3` element with the text `Box 2` inside the second `div` with the class `box`.

```js
const heading = document.querySelector("div.box > h3");
assert.equal(heading.textContent, "Box 2");
```

Add a `h3` element with the text `Box 3` inside the third `div` with the class `box`.

```js
const heading = document.querySelector("div.box > h3");
assert.equal(heading.textContent, "Box 3");
```

Add a `h3` element with the text `Box 4` inside the fourth `div` with the class `box`.

```js
const heading = document.querySelector("div.box > h3");
assert.equal(heading.textContent, "Box 4");
```

Add a `h3` element with the text `Box 5` inside the fifth `div` with the class `box`.

```js
const heading = document.querySelector("div.box > h3");
assert.equal(heading.textContent, "Box 5");
```

Add a `h3` element with the text `Box 6` inside the last `div` with the class `box`.

```js
const heading = document.querySelector("div.box > h3");
assert.equal(heading.textContent, "Box 6");
```

Add a `p` element with the text `Red` inside the first `div` with the `h3` element.

```js
const heading = document.querySelector("div.box > p");
assert.equal(heading.textContent, "Red");
```

Add a `p` element with the text `Orange` inside the second `div` with the `h3` element.

```js
const heading = document.querySelector("div.box > p");
assert.equal(heading.textContent, "Orange");
```

Add a `p` element with the text `Yellow` inside the third `div` with the `h3` element.

```js
const heading = document.querySelector("div.box > p");
assert.equal(heading.textContent, "Yellow");
```

Add a `p` element with the text `Green` inside the fourth `div` with the `h3` element.

```js
const heading = document.querySelector("div.box > p");
assert.equal(heading.textContent, "Green");
```

Add a `p` element with the text `Blue` inside the fifth `div` with the `h3` element.

```js
const heading = document.querySelector("div.box > p");
assert.equal(heading.textContent, "Blue");
```

Add a `p` element with the text `Indigo` inside the last `div` with the `h3` element.

```js
const heading = document.querySelector("div.box > p");
assert.equal(heading.textContent, "Indigo");
```



# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>

  </body>
</html>
--fcc-editable-region--
```

```css

h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}


```
