---
id: 68ef18bf52e0620eb8ae2694
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now we are going to check on the last alignment value. Change the `align-content` value from `space-around` to `space-evenly`.

With the `space-evenly` value, items are evenly distributed within of the alignment container. The space between all items is exactly the same, including the space at the start and the end.

# --hints--

Your class `flex-container` should have a property `align-content` with the value `space-evenly`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.align-content, 'space-evenly');
```


# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
--fcc-editable-region--
```

```css
--fcc-editable-region--
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-around;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}
--fcc-editable-region--

```
