---
id: 68ef96d3fc1b7b4db79e093a
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Now that you know how the `order` property works, let's create another class called `box2` in your `styles.css` file. Add a `background` property with the value `#f4a261`.

Now add the `order` property with a value of `0` to the `box2` class. You'll see this box move before the one with the `box1` class.

# --hints--

Create a class called `box2`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box2'));
```

Your class `box2` should have a property `background` with the value `#f4a261`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box2')?.background, 'rgb(244, 162, 97)');
```

Your class `box2` should have a property `order` with the value `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box2')?.order, '0');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box box1">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>    
  </body>
</html>

```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}
--fcc-editable-region--


--fcc-editable-region--
```
