---
id: 68ef0b3237fef991c92d9c19
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now add a new `div` element with the class `flex-container` to your `header`. Then create six `div` elements with the class `box` inside the `div` you just created.

We have already populated the `flex-container` and `box` class in the `styles.css` file.

# --hints--

Create a `div` element inside `header` with the class `flex-container`.

```js
assert.exists(document.querySelector('header > div.flex-container'));
```

Add six `div` elements with the class `box` inside the `div` with the `flex-container` class.

```js
const matches = document.querySelectorAll('div.box');
assert.strictEqual(matches.length, 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
--fcc-editable-region--  

--fcc-editable-region--
    </header>

  </body>
</html>
```

```css

h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}


```
