---
id: 68ef1b1c657f1b2396d98592
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

The `flex` property controls the size and behavior of the items inside a flexible container. It is compose by three properties: `flex-grow` (1), `flex-shrink` (1), and `flex-basis` (100px).

Now, getting back to the `.box` selector, change the first number in `flex` (corresponding to the `flex-grow` value) from `1` to `0`.

The `flex-grow` property controls how much extra space the flex item should take up if there is free space available in the container.

# --hints--

Your `.box` selector should have the property `flex` with a value of `0 1 100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '0 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h2>Box 1</h2>
        <p>Red</p>
      </div>
      <div class="box">
        <h2>Box 2</h2>
        <p>Orange</p>
      </div>
      <div class="box">
        <h2>Box 3</h2>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h2>Box 4</h2>
        <p>Green</p>
      </div>
      <div class="box">
        <h2>Box 5</h2>
        <p>Blue</p>
      </div>
      <div class="box">
        <h2>Box 6</h2>
        <p>Indigo</p>
      </div>
    </div>    
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}
--fcc-editable-region--
.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}
--fcc-editable-region--
```
