---
id: 68ef9855a1aa425c78b0bc4a
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Now let's continue adding properties to our boxes. Create a new class called `box3` and add a `background` property with the value `#ffd166`.

Go to your `html` file and add the class `box3` to the third box in your div.

Now add the `order` property with a value of `2` to the `box3` class. You'll see this box move after the ones with the `box1` and `box2` classes.

Also add the `flex-shrink` property with a value of `3`. This means that this box will shrink three times as much as the other boxes when there isn't enough space on its line.

# --hints--

Create a class called `box3`.

```js
assert.exist(new __helpers.CSSHelp(document).getStyle('.box3'));
```

Your class `box3` should have a property `background` with the value `#ffd166`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.background, '#ffd166');
```

Your class `box3` should have a property `order` with the value `2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.order, '2');
```

Your class `box3` should have a property `flex-shrink` with the value `3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.flex-shrink, '3');
```


# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>
    <div class="flex-container">
      <div class="box box1">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box box2">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>    
  </body>
</html>
--fcc-editable-region--

```

```css
--fcc-editable-region--
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}
--fcc-editable-region--

```
