---
id: 68f23761964a40289166e2fb
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now add an `h3` element inside each div with the class `box`.

Inside the six `h3` elements add the following texts, in order:

1. `Box 1`
2. `Box 2`
3. `Box 3`
4. `Box 4`
5. `Box 5`
6. `Box 6`

# --hints--

You should have an `h3` element with the text `Box 1` inside the first `div` with the class `box`.

```js
const firstH3El = document.querySelector("div.box:nth-of-type(1) > h3");
assert.equal(firstH3El.textContent, "Box 1");
```

You should have an `h3` element with the text `Box 2` inside the second `div` with the class `box`.

```js
const secondH3El = document.querySelector("div.box:nth-of-type(2) > h3");
assert.equal(secondH3El.textContent, "Box 2");
```

You should have an `h3` element with the text `Box 3` inside the third `div` with the class `box`.

```js
const thirdH3El = document.querySelector("div.box:nth-of-type(3) > h3");
assert.equal(thirdH3El.textContent, "Box 3");
```

You should have an `h3` element with the text `Box 4` inside the fourth `div` with the class `box`.

```js
const fourthH3El = document.querySelector("div.box:nth-of-type(4) > h3");
assert.equal(fourthH3El.textContent, "Box 4");
```

You should have an `h3` element with the text `Box 5` inside the fifth `div` with the class `box`.

```js
const fifthH3El = document.querySelector("div.box:nth-of-type(5) > h3");
assert.equal(fifthH3El.textContent, "Box 5");
```

You should have an `h3` element with the text `Box 6` inside the last `div` with the class `box`.

```js
const sixthH3El = document.querySelector("div.box:nth-of-type(6) > h3");
assert.equal(sixthH3El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
--fcc-editable-region--
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
}

.box {
  flex: 1 1 400px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0;
}
```
