---
id: 68f23761964a40289166e2fb
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Inside each `.box` element, add an `h2` element with the text `Box` followed by a space and a sequential number starting from `1`.

# --hints--

You should have an `h2` element with the text `Box 1` inside the first `div` with the class `box`.

```js
const firstH2El = document.querySelector("div.box:nth-of-type(1) > h2");
assert.equal(firstH2El.textContent, "Box 1");
```

You should have an `h2` element with the text `Box 2` inside the second `div` with the class `box`.

```js
const secondH2El = document.querySelector("div.box:nth-of-type(2) > h2");
assert.equal(secondH2El.textContent, "Box 2");
```

You should have an `h2` element with the text `Box 3` inside the third `div` with the class `box`.

```js
const thirdH2El = document.querySelector("div.box:nth-of-type(3) > h2");
assert.equal(thirdH2El.textContent, "Box 3");
```

You should have an `h2` element with the text `Box 4` inside the fourth `div` with the class `box`.

```js
const fourthH2El = document.querySelector("div.box:nth-of-type(4) > h2");
assert.equal(fourthH2El.textContent, "Box 4");
```

You should have an `h2` element with the text `Box 5` inside the fifth `div` with the class `box`.

```js
const fifthH2El = document.querySelector("div.box:nth-of-type(5) > h2");
assert.equal(fifthH2El.textContent, "Box 5");
```

You should have an `h2` element with the text `Box 6` inside the last `div` with the class `box`.

```js
const sixthH2El = document.querySelector("div.box:nth-of-type(6) > h2");
assert.equal(sixthH2El.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
--fcc-editable-region--
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
}
```
