---
id: 68f23744f67ee426f4fb6ae5
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now let's create the color name for the boxes. To do this, please add a `p` element below the `h3` element with these color names in order: `Red`, `Orange`, `Yellow`, `Green`, `Blue`, and `Indigo`. One color for each box.

# --hints--

Add a `p` element with the text `Red` inside the first `div` below `h3` element.

```js
const heading = document.querySelectorAll("div.box > p")[0];
assert.equal(heading.textContent, "Red");
```

Add a `p` element with the text `Orange` inside the second `div` below `h3` element.

```js
const heading = document.querySelectorAll("div.box > p")[1];
assert.equal(heading.textContent, "Orange");
```

Add a `p` element with the text `Yellow` inside the third `div` below `h3` element.

```js
const heading = document.querySelectorAll("div.box > p")[2];
assert.equal(heading.textContent, "Yellow");
```

Add a `p` element with the text `Green` inside the fourth `div` below `h3` element.

```js
const heading = document.querySelectorAll("div.box > p")[3];
assert.equal(heading.textContent, "Green");
```

Add a `p` element with the text `Blue` inside the fifth `div` below `h3` element.

```js
const heading = document.querySelectorAll("div.box > p")[4];
assert.equal(heading.textContent, "Blue");
```

Add a `p` element with the text `Indigo` inside the last `div` below `h3` element.

```js
const heading = document.querySelectorAll("div.box > p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
--fcc-editable-region--
    <div class="flex-container">
      <div class="box">
        <h3>Box 1</h3>
      </div>
      <div class="box">
        <h3>Box 2</h3>
      </div>
      <div class="box">
        <h3>Box 3</h3>
      </div>
      <div class="box">
        <h3>Box 4</h3>
      </div>
      <div class="box">
        <h3>Box 5</h3>
      </div>
      <div class="box">
        <h3>Box 6</h3>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
}

.box {
  flex: 1 1 400px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0;
}
```
