---
id: 68f23761964a40289166e2fb
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now add an `h3` element inside each div with the class `box`. Fill it with these text in order: `Box 1`, `Box 2`, `Box 3`, `Box 4`, `Box 5`, and `Box 6`.

# --hints--

Add a `h3` element with the text `Box 1` inside the first `div` with the class `box`.

```js
const heading = document.querySelectorAll("div.box > h3")[0];
assert.equal(heading.textContent, "Box 1");
```

Add a `h3` element with the text `Box 2` inside the second `div` with the class `box`.

```js
const heading = document.querySelectorAll("div.box > h3")[1];
assert.equal(heading.textContent, "Box 2");
```

Add a `h3` element with the text `Box 3` inside the third `div` with the class `box`.

```js
const heading = document.querySelectorAll("div.box > h3")[2];
assert.equal(heading.textContent, "Box 3");
```

Add a `h3` element with the text `Box 4` inside the fourth `div` with the class `box`.

```js
const heading = document.querySelectorAll("div.box > h3")[3];
assert.equal(heading.textContent, "Box 4");
```

Add a `h3` element with the text `Box 5` inside the fifth `div` with the class `box`.

```js
const heading = document.querySelectorAll("div.box > h3")[4];
assert.equal(heading.textContent, "Box 5");
```

Add a `h3` element with the text `Box 6` inside the last `div` with the class `box`.

```js
const heading = document.querySelectorAll("div.box > h3")[5];
assert.equal(heading.textContent, "Box 6");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>
--fcc-editable-region--    
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}
```
