---
id: 68f46e53fdfdc18a773282ba
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

Inside your `styles.css` file, create a `.box2` selector and add the `background` property with the value `#f4a261` to it.

Also, add the `order` property with a value of `0`.

# --hints--

You should have a `.box2` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box2'));
```

Your `.box2` selector should have the property `background` with the value `#f4a261`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box2')?.getPropVal('background'), 'rgb(244, 162, 97)');
```

Your `.box2` selector should have the property `order` with the value `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box2')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box box1">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box box2">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 0 0 150px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}
--fcc-editable-region--

--fcc-editable-region--
```
