---
id: 68f46fcc34da869b1005455d
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Now to continue adding properties to our boxes, create a new class called `box3` and add a `background` property with the value `#ffd166`.

Now add the `order` property with a value of `2` to the `box3` class. You'll see this box move after the ones with the `box1` and `box2` classes.

Also add the `flex-shrink` property with a value of `3`. This means that this box will shrink three times as much as the other boxes when there isn't enough space on its line.

# --hints--

Create a class called `box3`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box3'));
```

Your class `box3` should have a property `background` with the value `#ffd166`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.getPropVal('background'), 'rgb(255, 209, 102)');
```

Your class `box3` should have a property `order` with the value `2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.getPropVal('order'), '2');
```

Your class `box3` should have a property `flex-shrink` with the value `3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box3')?.getPropVal('flex-shrink'), '3');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box box1">
        <h2>Box 1</h2>
        <p>Red</p>
      </div>
      <div class="box box2">
        <h2>Box 2</h2>
        <p>Orange</p>
      </div>
      <div class="box box3">
        <h2>Box 3</h2>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h2>Box 4</h2>
        <p>Green</p>
      </div>
      <div class="box">
        <h2>Box 5</h2>
        <p>Blue</p>
      </div>
      <div class="box">
        <h2>Box 6</h2>
        <p>Indigo</p>
      </div>
    </div>    
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: 600px;
  padding: 10px;
  margin: 20px auto;
  align-content: space-evenly;
}

.box {
  flex: 0 0 150px;
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}
--fcc-editable-region--

--fcc-editable-region--
```
