---
id: 68f46fcc34da869b1005455d
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Now go to the `html` file and add the class `box3` to the third div with the class `box`.

# --hints--

Add the class `box3` to the third div box

```js
assert.exists(document.querySelector('div.box3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Boxes</title>
    
  </head>
  <body>
    <header>
      <h1>CSS flexbox colorful boxes</h1>
    </header>
    <div class="flex-container">
      <div class="box box1">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box box2">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
--fcc-editable-region--      
      <div class="box">
--fcc-editable-region--        
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>    
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}

.box {
  flex: 0 0 250px;
  max-height: 120px;
  color: #fff;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.box1 {
  background: #f16e79;
  order: 1; 
  flex-grow: 1;
}

.box2 {
  background: #f4a261;
  order: 0; 
}

.box3 {
  background: #ffd166;
  order: 2;
  flex-shrink: 3; 
}
```
