---
id: 68f93a8b58cbf330b4ba45c6
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

To finish styling the `flex-container` class, add a `padding` property with a value of `10px` and a `margin` property with a value of `20px auto`.

Remember: `padding` adds space inside the element (between the border and content), while `margin` adds space outside the element (between the border and other elements).

# --hints--

You should have a property `padding` with the value `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('padding'), '10px');
```

You should have a property `margin` with the value `20px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('margin'), '20px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--
.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
}
--fcc-editable-region--
```
