---
id: 68f9531aad56f929bc52506c
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Now you are going to style to the `div` elements with the class `box`. First create a `.box` class selector and add a `flex` property with the value `1 1 400px`.

The `flex` property controls the size and behavior of the items inside a flexible container. It is compose by three properties: `flex-grow` (1), `flex-shrink` (1), and `flex-basis` (400px). This means each box starts at `400px`, can grow to fill available space, and can shrink when needed.

# --hints--

You should create a `.box` class selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.box'));
```

Your `.box` selector should have a property `flex` with the value `1 1 400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '1 1 400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h3>Box 1</h3>
        <p>Red</p>
      </div>
      <div class="box">
        <h3>Box 2</h3>
        <p>Orange</p>
      </div>
      <div class="box">
        <h3>Box 3</h3>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h3>Box 4</h3>
        <p>Green</p>
      </div>
      <div class="box">
        <h3>Box 5</h3>
        <p>Blue</p>
      </div>
      <div class="box">
        <h3>Box 6</h3>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  width: 70%;
  height: auto;
  padding: 10px;
  margin: 20px auto;
  height: 400px;
  align-content: space-evenly;
}
--fcc-editable-region--

--fcc-editable-region--
```
