---
id: 68f932c803f4d3d9f511b436
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

To continue, add a `flex-wrap` property with the value `wrap` to the `.flex-container` class selector.

The `flex-wrap` property defines whether flex items are forced onto one line or can wrap onto multiple lines. With the `wrap` value, flex items wrap onto new lines when they don't fit in the container's width, creating a more flexible, responsive layout. 

By default, flex items have a value of `nowrap`, meaning they stay on a single line even if they overflow the container.

# --hints--

You should have a property `flex-wrap` with the value `wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('flex-wrap'), 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--
.flex-container {
  display: flex;
}
--fcc-editable-region--
```
