---
id: 68f92924265e9185f1bb391b
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Now go to the `styles.css` file and create a `flex-container` class. This workshop is focused on the CSS flex properties. To practice these, add a `display` property with the value `flex` to the class created. 

The `display` property controls how an element is rendered on the page. When you define it with `flex` value, this enables Flexbox layout mode, which gives you control to arrange and align the items inside the container.

# --hints--

Create a `h1` element

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.flex-container'));
```

You should have a property `display` with the value `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('display'), 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
