---
id: 68f9366fb3358401adf0ec79
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now you will set the width and height of the `flex-container` class. To do this, add a `width` property with a value `70%` and `height` property with a value `auto`.

Defining `width` and `height` properties will set the amount of width and height the `div` takes form its parent element, in this case the `body` element.

# --hints--

You should have a property `width` with the value `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('width'), '70%');
```

You should have a property `height` with the value `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('height'), 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--
.flex-container {
  display: flex;
  flex-wrap: wrap;
}
--fcc-editable-region--
```
