---
id: 68f92924265e9185f1bb391b
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Going back to the `styles.css` file, create a `.flex-container` selector and set the `display` property to the value `flex`. 

The `display` property controls how an element is rendered on the page. When you define it with the `flex` value, this enables Flexbox layout mode, which gives you control to arrange and align the items inside the container.

# --hints--

You should have a `.flex-container` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.flex-container'));
```

You should set the `display` property to the value `flex` within the `.flex-container` selector.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('display'), 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
      <div class="box">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
