---
id: 68f9366fb3358401adf0ec79
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Now you will set the width and height of the `.flex-container` element. To do this, add the `width` property with the value `70%` and the `height` property with the value `600px`.

Defining `width` and `height` properties will set the amount of width and height the `div` takes form its parent element, in this case the `body` element.

# --hints--

Your `.flex-container` selector should have the property `width` with the value `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('width'), '70%');
```

Your `.flex-container` selector should have the property `height` with the value `600px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.flex-container')?.getPropVal('height'), '600px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h2>Box 1</h2>
        <p>Red</p>
      </div>
      <div class="box">
        <h2>Box 2</h2>
        <p>Orange</p>
      </div>
      <div class="box">
        <h2>Box 3</h2>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h2>Box 4</h2>
        <p>Green</p>
      </div>
      <div class="box">
        <h2>Box 5</h2>
        <p>Blue</p>
      </div>
      <div class="box">
        <h2>Box 6</h2>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
  order: 0; 
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
--fcc-editable-region--

--fcc-editable-region--
}
```
