---
id: 692f2d89d3a3a86de7273e08
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Add a `justify-content` property with a value `center` to the `box` class.

The `justify-content` property aligns flex items along the main axis and distributes the extra space. Since the `box` uses `flex-direction: column`, the main axis is vertical, so `center` centers items vertically.

# --hints--

You should have a property `justify-content` with the value `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('justify-content'), 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h2>Box 1</h2>
        <p>Red</p>
      </div>
      <div class="box">
        <h2>Box 2</h2>
        <p>Orange</p>
      </div>
      <div class="box">
        <h2>Box 3</h2>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h2>Box 4</h2>
        <p>Green</p>
      </div>
      <div class="box">
        <h2>Box 5</h2>
        <p>Blue</p>
      </div>
      <div class="box">
        <h2>Box 6</h2>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

--fcc-editable-region--
.box {
  flex: 1 1 100px;
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex-direction: column;
}
--fcc-editable-region--
```
