---
id: 692f22933ba1cb86cd75a70c
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now it's time to create the color name for the boxes. To do this, add a `p` element below each `h2` element with these color names in order: `Red`, `Orange`, `Yellow`, `Green`, `Blue`, and `Indigo`. One color for each box.

# --hints--

You should add a `p` element with the text `Red` below the `h2` element inside the first `.box` element.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[0];
assert.equal(heading.textContent, "Red");
```

You should add a `p` element with the text `Orange` below the `h2` element inside the second `.box` element.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[1];
assert.equal(heading.textContent, "Orange");
```

You should add a `p` element with the text `Yellow` below the `h2` element inside the third `.box` element.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[2];
assert.equal(heading.textContent, "Yellow");
```

You should add a `p` element with the text `Green` below the `h2` element inside the fourth `.box` element.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[3];
assert.equal(heading.textContent, "Green");
```

You should add a `p` element with the text `Blue` below the `h2` element inside the fifth `.box` element.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[4];
assert.equal(heading.textContent, "Blue");
```

You should add a `p` element with the text `Indigo` below the `h2` element inside the sixth `.box` element.

```js
const heading = document.querySelectorAll("div.box > h2 + p")[5];
assert.equal(heading.textContent, "Indigo");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
--fcc-editable-region--
      <div class="box">
        <h2>Box 1</h2>
      </div>
      <div class="box">
        <h2>Box 2</h2>
      </div>
      <div class="box">
        <h2>Box 3</h2>
      </div>
      <div class="box">
        <h2>Box 4</h2>
      </div>
      <div class="box">
        <h2>Box 5</h2>
      </div>
      <div class="box">
        <h2>Box 6</h2>
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}
```
