---
id: 692f318859d5e4bf8fd77ba1
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

The `order` property specifies the order of a flex item relative to other flex items inside the same container. By default, all flex items have an `order` value of `0`, meaning they appear in the order they're written in the HTML. Items with lower `order` values appear first, while items with higher values appear last.

Add the property `order` with the value `0` to the `.box` selector.

# --hints--

Your `.box` selector should have the property `order` with the value `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('order'), '0');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Colored Boxes</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <h1>Colored Boxes Layout</h1>
    </header>
    <div class="flex-container">
      <div class="box">
        <h2>Box 1</h2>
        <p>Red</p>
      </div>
      <div class="box">
        <h2>Box 2</h2>
        <p>Orange</p>
      </div>
      <div class="box">
        <h2>Box 3</h2>
        <p>Yellow</p>
      </div>
      <div class="box">
        <h2>Box 4</h2>
        <p>Green</p>
      </div>
      <div class="box">
        <h2>Box 5</h2>
        <p>Blue</p>
      </div>
      <div class="box">
        <h2>Box 6</h2>
        <p>Indigo</p>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
  border: 1px solid #000;
  display: flex;
  flex: 1 1 100px;
  flex-direction: column;
  align-items: center;
  margin: 10px;
  font-weight: bold;
  font-size: 1.125rem;
  border-radius: 5px;
--fcc-editable-region--

--fcc-editable-region--
}
```
