---
id: 694b46a414bb2f5079868ed1
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now add the `flex` property with the value `1 1 100px` to the `.box` selector.

The `flex` property controls the size and behavior of the items inside a flexible container. It is compose by three properties: `flex-grow` (1), `flex-shrink` (1), and `flex-basis` (100px). This means each box starts at `100px`, can grow to fill available space, and can shrink when needed.

# --hints--

Your `.box` selector should have a property `flex` with the value `1 1 100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.box')?.getPropVal('flex'), '1 1 100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
	<meta charset="UTF-8" />
	<meta name="viewport" content="width=device-width, initial-scale=1.0" />
	<title>Colored Boxes</title>
	<link rel="stylesheet" href="styles.css" />
</head>
<body>
	<header>
		<h1>Colored Boxes Layout</h1>
	</header>
	<div class="flex-container">
		<div class="box">
			<h2>Box 1</h2>
			<p>Red</p>
		</div>
		<div class="box">
			<h2>Box 2</h2>
			<p>Orange</p>
		</div>
		<div class="box">
			<h2>Box 3</h2>
			<p>Yellow</p>
		</div>
		<div class="box">
			<h2>Box 4</h2>
			<p>Green</p>
		</div>
		<div class="box">
			<h2>Box 5</h2>
			<p>Blue</p>
		</div>
		<div class="box">
			<h2>Box 6</h2>
			<p>Indigo</p>
		</div>
	</div>
</body>
</html>
```

```css
h1 {
  text-align: center;
  margin-bottom: 10px;
}

.box {
  max-height: 120px;
  color: #000;
	border: 1px solid #000;
	display: flex;
--fcc-editable-region--
	
--fcc-editable-region--
}
```
