---
id: 66f686b8ebdb982fa8e14330
title: Build a Date Conversion Program
challengeType: 14
dashedName: lab-date-conversion
---

# --description--

In this lab, you will code a date conversion program that converts a given date to long and short date formats. The program will output the current date and time, the date in `MM/DD/YYYY` format, and the date in Month Day, Year format.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should create a variable named `currentDate` and assign it the current date and time using the `Date` object.
1. You should create a variable named `currentDateFormat` and assign it the string `Current Date and Time: [current date]`. Replace `[current date]` with the result of the `currentDate` variable.
1. You should create a function named `formatDateMMDDYYYY` that takes a parameter. You can name this parameter whatever you like.
1. Your `formatDateMMDDYYYY` function should return a string that formats the date in a `month/day/year` format. You can review the previous lecture videos to help you with which date methods to use in constructing your string.
1. You should create a variable named `shortDateFormat` and assign it the string `Formatted Date (MM/DD/YYYY): [formatted date]`. Replace `[formatted date]` with the result of calling your `formatDateMMDDYYYY` function with the `currentDate` variable.
1. You should create a function named `formatDateLong` that takes a parameter. You can name this parameter whatever you like.
1. Your `formatDateLong` function should return a string that formats the date in a `Month Day, Year` format. You can review the previous lecture videos on the different ways to format dates.
1. You should create a variable named `longDateFormat` and assign it the string `Formatted Date (Month Day, Year): [formatted date]`. Replace `[formatted date]` with the result of calling your `formatDateLong` function with the `currentDate` variable.
1. You should log the value of `currentDateFormat`, `shortDateFormat` and `longDateFormat` to the console.

# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `${month}/${day}/${year}`;
}
const shortDateFormat = `Formatted Date (MM/DD/YYYY): ${formatDateMMDDYYYY(currentDate)}`;
console.log(shortDateFormat);

function formatDateLong(date) {
  return date.toLocaleDateString(undefined, {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
}
const longDateFormat = `Formatted Date (Month Day, Year): ${formatDateLong(currentDate)}`;
console.log(longDateFormat);
```

