---
id: 6196aead7ac7bf1584b17a7f
title: Крок 18
challengeType: 0
dashedName: step-18
---

# --description--

Щоб краще перекривати гору та елементи `.ground`, надайте горі `margin-top` зі значенням `100px` та елементу `.ground` надайте `margin-top` зі значенням `-58px`.

# --hints--

Ви повинні надати `.left-mountain` властивість `margin-top` зі значенням `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.marginTop, '100px');
```

Ви повинні надати `.ground` властивість `margin-top` зі значенням `-58px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.marginTop, '-58px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;

}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;

}
--fcc-editable-region--
```
