---
id: 619691693bc14b0e528f5a20
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Націльте елемент `.ground` та встановіть його `width` так, щоб він займав всю ширину вюпорту. Потім встановіть `height` на `400px`.

# --hints--

Ви повинні використати селектор `.ground`.

```js
assert.match(code, /\.ground\s*\{/);
```

Ви повинні надати елементу `.ground` властивість `width` зі значенням `100vw`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

Ви повинні надати елементу `.ground` властивість `height` зі значенням `400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
