---
id: 619d237a107c10221ed743fa
title: 햌쮏 86
challengeType: 0
dashedName: step-86
---

# --description--

햕혰햨햟쒫쟳 혟햟햨혝: 혰햫햡쒬뒗쫧 햫햣 햪쮏웸혝혧 햩혰혝햟혝햦 햠햣향 햨햦햩.

 햪햣햤햟혠 `.penguin-body` 햣햣햢 햣햩햣햪햣햫혝햟햪햦 `.foot` 햢쮏얧썛좜햣 햢쒫 햣햩햣햪햣햫혝햦 `div`, 햨쮏웷왏 향 `class` 향혰 향햫햟혢햣햫햫혪햪 `arm`. 햣혣쮏쩀 `.arm` 햫햟햢햟햧혝햣 `class` 향혰 향햫햟혢햣햫햫혪햪 `left`, 햟 햢혞햡쮏쩀 `.arm` 햫햟햢햟햧혝햣 `class` 향혰 향햫햟혢햣햫햫혪햪 `right`.

# --hints--

뉋 쮏쒫쟳쫧쫨 햢쮏얧썜햦 햢쒫 햣햩햣햪햣햫혝햦 `div`  `.penguin-body`. 뤰혰햨혞쒫썛햦혜혧 햣햩햣햪햣햫혝햦 `div` `--fcc-expected--`, 향햫햟햧햢햣햫 `--fcc-actual--`.

```js
assert.equal(document.querySelectorAll('.penguin-body > div')?.length, 4);
```

햣혣쮏쩀 햫쮏쒫쮏쩀 `div` 쮐혰햠햫 햫햟햢햟혝햦 `class` 향혰 향햫햟혢햣햫햫혪햪 `arm`.

```js
assert.exists(document.querySelector('.penguin-body > div.arm'));
```

햆혞햡쮏쩀 햫쮏쒫쮏쩀 `div` 쮐혰햠햫 햫햟햢햟혝햦 `class` 향혰 향햫햟혢햣햫햫혪햪 `arm`.

```js
assert.equal(document.querySelectorAll('.penguin-body > div.arm')?.length, 2);
```

뤯얧쫧쮏쩀 `div` 쮐혰햠햫 햫햟햢햟혝햦 `class` 향혰 향햫햟혢햣햫햫혪햪 `left`.

```js
assert.exists(document.querySelector('.penguin-body > div.arm.left'));
```

핹햫혣쮏쩀 `div` 쮐혰햠햫 햫햟햢햟혝햦 `class` 향혰 향햫햟혢햣햫햫혪햪 `right`.

```js
assert.exists(document.querySelector('.penguin-body > div.arm.right'));
```

뉋 쮏쒫쟳쫧쫨 쮏햪혰혜혝햦혝햦 `.arm.right` 혰혜햩혪 `.arm.left`.

```js
assert.exists(document.querySelector('.arm.left + .arm.right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin">
      <div class="penguin-head">
        <div class="face left"></div>
        <div class="face right"></div>
        <div class="chin"></div>
        <div class="eye left">
          <div class="eye-lid"></div>
        </div>
        <div class="eye right">
          <div class="eye-lid"></div>
        </div>
        <div class="blush left"></div>
        <div class="blush right"></div>
        <div class="beak top"></div>
        <div class="beak bottom"></div>
      </div>
      <div class="shirt">
        <div>游눞</div>
        <p>I CSS</p>
      </div> 
--fcc-editable-region--
      <div class="penguin-body">

        <div class="foot left"></div>
        <div class="foot right"></div>
      </div>
--fcc-editable-region--
    </div>

    <div class="ground"></div>
  </body>
</html>
```

```css
:root {
  --penguin-face: white;
  --penguin-picorna: orange;
}

body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
  z-index: 4;
  position: relative;
}

.penguin * {
  position: absolute;
}

.penguin-head {
  width: 50%;
  height: 45%;
  background: linear-gradient(
    45deg,
    gray,
    rgb(239, 240, 228)
  );
  border-radius: 70% 70% 65% 65%;
  top: 10%;
  left: 25%;
  z-index: 1;
}

.face {
  width: 60%;
  height: 70%;
  background-color: var(--penguin-face);
  border-radius: 70% 70% 60% 60%;
  top: 15%;
}

.face.left {
  left: 5%;
}

.face.right {
  right: 5%;
}

.chin {
  width: 90%;
  height: 70%;
  background-color: var(--penguin-face);
  top: 25%;
  left: 5%;
  border-radius: 70% 70% 100% 100%;
}

.eye {
  width: 15%;
  height: 17%;
  background-color: black;
  top: 45%;
  border-radius: 50%;
}

.eye.left {
  left: 25%;
}

.eye.right {
  right: 25%;
}

.eye-lid {
  width: 150%;
  height: 100%;
  background-color: var(--penguin-face);
  top: 25%;
  left: -23%;
  border-radius: 50%;
}

.blush {
  width: 15%;
  height: 10%;
  background-color: pink;
  top: 65%;
  border-radius: 50%;
}

.blush.left {
  left: 15%;
}

.blush.right {
  right: 15%;
}

.beak {
  height: 10%;
  background-color: var(--penguin-picorna);
  border-radius: 50%;
}

.beak.top {
  width: 20%;
  top: 60%;
  left: 40%;
}

.beak.bottom {
  width: 16%;
  top: 65%;
  left: 42%;
}

.shirt {
  font: bold 25px Helvetica, sans-serif;
  top: 165px;
  left: 127.5px;
  z-index: 1;
  color: #6a6969;
}

.shirt div {
  font-weight:  initial;
  top: 22.5px;
  left: 12px;
}

.penguin-body {
  width: 53%;
  height: 45%;
  background: linear-gradient(
    45deg,
    rgb(134, 133, 133) 0%,
    rgb(234, 231, 231) 25%,
    white 67%
  );
  border-radius: 80% 80% 100% 100%;
  top: 40%;
  left: 23.5%;
}

.penguin-body::before {
  content: "";
  position: absolute;
  width: 50%;
  height: 45%;
  background-color: gray;
  top: 10%;
  left: 25%;
  border-radius: 0% 0% 100% 100%;
  opacity: 70%;
}

.foot {
  width:  15%;
  height: 30%;
  background-color: var(--penguin-picorna);
  top: 85%;
  border-radius: 50%;
  z-index: -1;
}

.foot.left {
  left: 25%;
  transform: rotate(80deg);
}

.foot.right {
  right: 25%;
  transform: rotate(-80deg);
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
