---
id: 61a8fe15a6a31306e60d1e89
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Нормалізуйте свою сторінку, встановивши `width` на `100%` та `height` на `100vh`.

# --hints--

Ви повинні надати `body` властивість `width` зі значенням `--fcc-expected--`, але знайдено `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Ви повинні надати `body` властивість `height` зі значенням `--fcc-expected--`, але знайдено `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;

}
--fcc-editable-region--
```
