---
id: 60a3e3396c7b40068ad6998c
title: 步驟 34
challengeType: 0
dashedName: step-34
---

# --description--

你的畫的顏色和形狀太銳利了，不能像 Rothko 一樣通過。

使用 `filter` 屬性在 `.canvas` 元素中使其 `blur` 爲 `2px`。

下面是添加一個 3px `blur` 規則的示例：

```css
p {
  filter: blur(3px);
}
```

# --hints--

應該將 `filter` 屬性設置爲 `blur(2px)`。

```js
const hasFilter = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.filter === 'blur(2px)');
assert(hasFilter);
```

`.canvas` 元素的 `filter` 值應爲 `blur(2px)`。

```js
const canvasFilter = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('filter');
assert(canvasFilter === 'blur(2px)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
