---
id: 60a3e3396c7b40068ad6996a
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

到目前爲止，你應該熟悉 HTML 頁面應具有的基本元素。

使用 `DOCTYPE` 聲明、`html` 元素、`head` 元素和 `body` 元素設置你的代碼。

# --hints--

代碼應該有一個 `<!DOCTYPE html>` 聲明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

代碼應該有一個 `html` 元素。

```js
assert(document.querySelectorAll('html').length === 1);
```

代碼應該在 `html` 元素中包含一個 `head` 元素。

```js
assert(document.querySelectorAll('head').length === 1);
```

代碼應該在 `html` 元素中包含一個 `body` 元素。

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 元素應該在你的 `body` 元素之前。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
