---
id: 60a3e3396c7b40068ad6996a
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

到目前爲止，你應該熟悉 HTML 頁面應具有的基本元素。

添加 `DOCTYPE` 聲明、語言設置爲英語的 `html` 元素、 `head` 和 `body` 元素。

# --hints--

代碼應該有一個 `<!DOCTYPE html>` 聲明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個值爲 `en` 的 `lang` 屬性的 `<html>` 開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

代碼應該在 `html` 元素中包含一個 `head` 元素。

```js
assert(document.querySelectorAll('head').length === 1);
```

代碼應該在 `html` 元素中包含一個 `body` 元素。

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 元素應該在你的 `body` 元素之前。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
