---
id: 60a3e3396c7b40068ad69978
title: 步驟 15
challengeType: 0
dashedName: step-15
---

# --description--

使用邊距調整元素外部的間距。

使用 `margin` 屬性，爲 `.frame` 元素提供 `20px` 的垂直邊距和 `auto` 的水平邊距。 這會將框架向下移動 20 像素並在頁面上水平居中。

# --hints--

應該將 `margin` 屬性設置爲 `20px auto`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

`.frame` 元素的 `margin` 值應爲 `20px auto`。

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
