---
id: 60a3e3396c7b40068ad69977
title: 步驟 14
challengeType: 0
dashedName: step-14
---

# --description--

使用填充來調整元素內的間距。

在 `.frame` 中，使用 `padding` 簡寫屬性來增加 `.frame` 和 `.canvas` 元素之間的空間到 `50px`。 簡寫方式將增加元素邊框和畫布的頂部、底部、左側和右側的空間。

# --hints--

應該將 `padding` 屬性設置爲 `50px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

`.frame` 元素的 `padding` 值應爲 `50px`。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
