---
id: 60a3e3396c7b40068ad69986
title: 步驟 28
challengeType: 0
dashedName: step-28
---

# --description--

在 `.two` 元素下方創建一個新的 `div`，其 `class` 值爲 `three`。 這將是你的第三個矩形。

# --hints--

不應更改現有的 `.one` 和 `.two` 元素。

```js
assert(document.querySelectorAll('.one').length === 1);
assert(document.querySelectorAll('.two').length === 1);
```

新的 `div` 應該嵌套在您的 `.canvas` 元素中。

```js
assert(document.querySelector('.canvas').children[2].tagName === 'DIV');
```

新的 `div` 應該在 `.two` 元素之後。

```js
assert(document.querySelector('.two').nextElementSibling.tagName === 'DIV');
```

新的 `div` 元素應該有一個 `class`，其值爲 `three`。

```js
assert(document.querySelector('.canvas').children[2].className.split(' ').includes('three'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
