---
id: 60a3e3396c7b40068ad6996e
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

邊距是框外的區域，可用於控制其他框或元素之間的間距。

此處底部元素具有較大的上邊距，將其推到頁面下方。

現在你瞭解了 CSS 盒子模型，讓我們開始 Rothko 繪畫。

移除 `<img>` 元素。

# --hints--

代碼中不應包含 `img` 元素。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
