---
id: 60a3e3396c7b40068ad69971
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

CSS 的時間到了。

即使你的 `<div>` 沒有文本，它仍然被視爲包含內容的框。 編寫一個使用 `.canvas` 類選擇器的 CSS 規則，並將其 `width` 設置爲 500 像素。 這是一個 CSS 規則，將類 `card` 的寬度設置爲 300 像素：

```css
.card {
  width: 300px;
}
```

# --hints--

代碼應該有一個 `.canvas` 選擇器。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

應該將 `width` 屬性設置爲 `500px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

`.canvas` 元素的 `width` 應爲 `500px`。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
