---
id: 60a3e3396c7b40068ad69973
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

將畫布的 `background-color` 更改爲 `#4d0f00`。

# --hints--

應該將 `background-color` 屬性設置爲 `#4d0f00`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

`.canvas` 元素應該有一個 `#4d0f00` 的 `background-color`。

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
