---
id: 68cabc534d2e741f33a3fe1d
title: Build a Game Character Stats Tracker
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

In this lab, you'll build a game character stats tracker. The character stats tracker will allow you to create a character with specific attributes, update those attributes, and retrieve the current stats of the character.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. Create a class named `GameCharacter` that represents a game character and manages character stats.
2. When instantiated, a new `GameCharacter` object should have the following private attributes:
    - `_name` set to the string given at the moment of the instantiation.
    - `_health` set to `100`.
    - `_mana` set to `50`.
    - `_level` set to `1`.
3. Create a `name` property for read-only access to the character's name.
4. For the `health` property:
     - Define a getter that returns the current health.
     - Define a setter that prevents health from being set below `0`, and caps the health at `100`.
5. For the `mana` property:
     - Define a getter that returns the current mana.
     - Define a setter that prevents mana from being set below `0`, and caps the mana at `50`.
6. Create a getter for `level` to return the character's current level.
7. Define a method named `level_up()` that:
    - Increases the character's level by 1.
    - Resets health to `100` and mana to `50` using their corresponding property setters.
    - Prints a message in the form of `<name> leveled up to <level>!` (where `<name>` and `<level>` should be replaced by the character's name and new level, respectively).
8. Define a `__str__()` method that returns a formatted string including:
    - The character's name.
    - The character's level.
    - The character's current health.
    - The character's current mana.
   For example, a character named `Kratos`, right after the instantiation, should be represented as the following:

   ```md
    Name: Kratos
    Level: 1
    Health: 100
    Mana: 50
   ```

## Usage example

```python
hero = GameCharacter("Kratos") # Creates a new character named "Kratos"
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

You should have a `GameCharacter` class.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

You should have an `__init__` method in the `GameCharacter` class.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

The `__init__` method should have two parameters, `self` and `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__init__").has_args("self, name")`) 
})
```

You should assign `name` to `self._name` inside the `__init__` method.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

You should assign `100` to `self._health` inside the `__init__` method.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

You should assign `50` to `self._mana` inside the `__init__` method.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

You should assign `1` to `self._level` inside the `__init__` method.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

You should have a `name` method in the `GameCharacter` class.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

The `name` method should have only one parameter, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

You should return `self._name` in the `name` method.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

The `name` method should have a `@property` decorator.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

You should have a `health` method in the `GameCharacter` class.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

The `health` method should have only one parameter, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

You should return `self._health` in the `health` method.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

The `health` method should have a `@property` decorator.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

You should create a `health` setter method.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

The `health` setter method should have a `@health.setter` decorator.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

The `health` setter method should have two parameters, `self` and `value`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_args("self, value")`) 
})
```

You should set `self._health` to `0` if `value` is less than `0`, `100` if `value` is greater than `100`, and to `value` otherwise in the `health` setter method.

```js
({ 
    test: () => runPython(`
    _cond = _Node(_code).find_class("GameCharacter").find_functions("health")[1].find_ifs()[0]
    assert _cond.is_equivalent("if value < 0:\\n self._health = 0\\nelif value > 100:\\n self._health = 100\\nelse:\\n self._health = value")
`)})
```

You should have a `mana` method in the `GameCharacter` class.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

The `mana` method should have only one parameter, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

You should return `self._mana` in the `mana` method.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

The `mana` method should have a `@property` decorator.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self.__name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
