---
id: 6870dd38187258b547f82d76
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Instead of manually listing each `Card` component, use the `map()` method to loop through your profiles array.

```js

function UserList() {
  const users = [
    { id: "user-001-employee", name: "Alice", email: "alice@example.com" },
    { id: "user-002-employee", name: "Bob", email: "bob@example.com" },
    { id: "user-003-employee", name: "John", email: "john@example.com" },
  ];
  return (
    <div>
      {users.map((user) => (
        <div key={user.id}>
          <h3>{user.name}</h3>
          <p>{user.email}</p>
        </div>
      ))}
    </div>
  );
}

```

Make sure to include a unique `key` prop. This helps React track list items efficiently.

Ensure each `Card` receives its `name`, `title`, and `bio` props from the current `profile` object.

With that, your reusable profile card component project is complete!

# --hints--

You should map through `profiles`.

```js
assert.match(code, /\{\s*profiles\s*\.\s*map\s*\(/);
```

You should pass in `profile` as the parameter of the map.

```js
assert.match(code, /\{\s*profiles\s*\.\s*map\s*\(\s*(profile|\(\s*profile\s*\))\s*=>/);
```

Your `map()` should render `Card` components with all required props (`key`, `name`, `title`, and `bio`).

```js 
assert.match(code, /<\s*Card\s+key\s*=\s*\{\s*profile\.id\s*\}\s+name\s*=\s*\{\s*profile\.name\s*\}\s+title\s*=\s*\{\s*profile\.title\s*\}\s+bio\s*=\s*\{\s*profile\.bio\s*\}\s*\/>/);
```

Each card should display the correct profile data.

```js 
const cardElements = document?.querySelectorAll(".card");
const profiles = [
  { id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
  { id: '2', name: 'Tiffany', title: "Engineering manager", bio: "I have worked in tech for 15 years and love to help people grow in this industry." },
  { id: '3', name: "Doug", title: "Backend developer", bio: "I have been a software developer for over 20 years and I love working with Go and Rust." }
];

profiles.forEach((profile, index) => {
  const card = cardElements[index];
  assert.strictEqual(card?.querySelector("h2")?.textContent, profile.name);
  assert.strictEqual(card?.querySelector(".card-title")?.textContent, profile.title);
  assert.strictEqual(card?.querySelector("p:not(.card-title)")?.textContent, profile.bio);
});
```

You should render the correct number of cards.

```js 
const profilesLength = 3;

const cardElements = document?.querySelectorAll(".card");
assert.lengthOf(cardElements, profilesLength); 
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
    { id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
    { id: '2', name: 'Tiffany', title: "Engineering manager", bio: "I have worked in tech for 15 years and love to help people grow in this industry." },
    { id: '3', name: "Doug", title: "Backend developer", bio: "I have been a software developer for over 20 years and I love working with Go and Rust." }
  ];

  return (
    <div className="flex-container">
      --fcc-editable-region--
      <Card
        name="Mark"
        title="Frontend developer"
        bio="I like to work with different frontend technologies and play video games."
      />

      <Card
        name="Tiffany"
        title="Engineering manager"
        bio="I have worked in tech for 15 years and love to help people grow in this industry."
      />

      <Card
        name="Doug"
        title="Backend developer"
        bio="I have been a software developer for over 20 years and I love working with Go and Rust."
      />
      --fcc-editable-region--
    </div>
  );
}
```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
    { id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
    { id: '2', name: 'Tiffany', title: "Engineering manager", bio: "I have worked in tech for 15 years and love to help people grow in this industry." },
    { id: '3', name: "Doug", title: "Backend developer", bio: "I have been a software developer for over 20 years and I love working with Go and Rust." }
  ];

  return (
    <div className="flex-container">
      {profiles.map((profile) => (
        <Card
          key={profile.id}
          name={profile.name}
          title={profile.title}
          bio={profile.bio}
        />
      ))}
    </div>
  );
}
```
