---
id: 6877bfc3d277a59b3900f5a7
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Now, let's prepare our data. Inside your `App` component, create a `profiles` array. This array should contain one profile object with the following properties: `id` set to `"1"`, `name` set to `"Mark"`, `title` set to `"Frontend developer"`, and `bio` set to `"I like to work with different frontend technologies and play video games."`.

# --hints--

You should create a `profiles` array.

```js
assert.match(code, /(const|let|var)\s+profiles\s*=\s*\[/i);
```

You should have an object within the `profiles` array.

```js
assert.match(code, /(const|let|var)\s+profiles\s*=\s*\[\s*\{[^}]*\}\s*\]/s);
```

The value of the `id` property should be `"1"`.

```js
assert.match(code, /(const|let|var)\s+profiles\s*=\s*\[\s*\{[^}]*id\s*:\s*(1|"1"|'1')[^}]*\}\s*\]/s);
```

The value of the `name` property should be `"Mark"`.

```js
assert.match(code, /(const|let|var)\s+profiles\s*=\s*\[\s*\{[^}]*name\s*:\s*["']Mark["'][^}]*\}\s*\]/s);
```

The value of the `title` property should be `"Frontend developer"`.

```js
assert.match(code, /(const|let|var)\s+profiles\s*=\s*\[\s*\{[^}]*title\s*:\s*["']Frontend developer["'][^}]*\}\s*\]/s);
```

The value of the `bio` property should be `"I like to work with different frontend technologies and play video games."`.

```js
assert.match(code, /(const|let|var)\s+profiles\s*=\s*\[\s*\{[^}]*bio\s*:\s*["']I like to work with different frontend technologies and play video games\.["'][^}]*\}\s*\]/s);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  --fcc-editable-region--
  
  --fcc-editable-region--
  return (
    <div className="flex-container"></div>
  );
}
```
