---
id: 687942c174eb0172ef4ded6a
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Add the other two profiles data to your `profiles` array in the `App` component. 

Use the following data:

| Props  | First Card Values | Second Card Values |
|--------|-------------------|--------------------|
|`name`  | `"Tiffany"` | `"Doug"`|
|`title` | `"Engineering manager"` | `"Backend developer"` |
|`bio`   | `"I have worked in tech for 15 years and love to help people grow in this industry."` | `"I have been a software developer for over 20 years and I love working with Go and Rust."`|


This will demonstrate how powerful the `map()` function is, as all your cards will appear dynamically without you having to manually add individual `Card` components.

With that, your reusable profile card component project is complete!

# --hints--

The `profiles` array should contain three profile objects.

```js 
async () => {
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const cards = testElem.querySelectorAll('.card');
  assert.lengthOf(cards, 3);
}
```

The first `Card` should display Mark's information.

```js
async () => {
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const firstCard = testElem.querySelectorAll('.card')[0];
  assert.equal(firstCard.querySelector('h2').textContent, 'Mark');
  assert.equal(firstCard.querySelector('.card-title').textContent, 'Frontend developer');
  assert.equal(firstCard.querySelector('p:last-child').textContent, 'I like to work with different frontend technologies and play video games.');
}
```

The second `Card` should display Tiffany's information.

```js
async () => {
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const secondCard = testElem.querySelectorAll('.card')[1];
  assert.equal(secondCard.querySelector('h2').textContent, 'Tiffany');
  assert.equal(secondCard.querySelector('.card-title').textContent, 'Engineering manager');
  assert.equal(secondCard.querySelector('p:last-child').textContent, 'I have worked in tech for 15 years and love to help people grow in this industry.');
  assert.match(code, /id\s*:\s*["']2["']/);
}
```

The third `Card` should display Doug's information.

```js
async () => {
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const thirdCard = testElem.querySelectorAll('.card')[2];
  assert.equal(thirdCard.querySelector('h2').textContent, 'Doug');
  assert.equal(thirdCard.querySelector('.card-title').textContent, 'Backend developer');
  assert.equal(thirdCard.querySelector('p:last-child').textContent, 'I have been a software developer for over 20 years and I love working with Go and Rust.');
  assert.match(code, /id\s*:\s*["']3["']/);
}
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
  { id: "1", name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
  --fcc-editable-region--
  
	--fcc-editable-region--
	];
  return (
    <div className="flex-container">
      {profiles.map((profile) => (
        <Card
          key={profile.id}
          name={profile.name}
          title={profile.title}
          bio={profile.bio}
        />
      ))}
    </div>
  );
}
```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
    { id: "1", name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
    { id: "2", name: 'Tiffany', title: "Engineering manager", bio: "I have worked in tech for 15 years and love to help people grow in this industry." },
    { id: "3", name: "Doug", title: "Backend developer", bio: "I have been a software developer for over 20 years and I love working with Go and Rust." }
  ];

  return (
    <div className="flex-container">
      {profiles.map((profile) => (
        <Card
          key={profile.id}
          name={profile.name}
          title={profile.title}
          bio={profile.bio}
        />
      ))}
    </div>
  );
}
```
