---
id: 68792530cde1deb18ea24996
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Now, inside your `map()` function, replace the `div` element with your `<Card />` component.

Next, pass the `name`, `title`, and `bio` props to the `<Card />` component.

# --hints--

Each `Card` component should have a `key` prop set to `profile.id`

```js
assert.match(code, /<\s*Card\s+[^>]*key\s*=\s*\{\s*profile\.id\s*}[^>]*>/);
```

Each `Card` components should have a `name` prop using `profile.name`.

```js
assert.match(code, /<\s*Card[^>]*name\s*=\s*\{\s*profile\.name\s*}[^>]*>/);
```

Each `Card` components should have a `title` prop using `profile.title`.

```js
assert.match(code, /<\s*Card[^>]*title\s*=\s*\{\s*profile\.title\s*}[^>]*>/);
```

Each `Card` components should have a `bio` prop using `profile.bio`.

```js
assert.match(code, /<\s*Card[^>]*bio\s*=\s*\{\s*profile\.bio\s*}[^>]*>/);
```

The `Card` component should render with correct data and proper syntax.

```js
const profiles = [
  	{ id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." }
	];

const cards = document.querySelectorAll('.card');
profiles.forEach((profile, i) => {
  assert.equal(cards[i]?.querySelector('h2')?.textContent, profile.name);
  assert.equal(cards[i]?.querySelector('.card-title')?.textContent, profile.title);
  assert.equal(cards[i]?.querySelector('p:not(.card-title)')?.textContent, profile.bio);
});
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
  	{ id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." }
	];
  return (
    <div className="flex-container">
--fcc-editable-region--
      {profiles.map((profile) => (
        <div key={profile.id}>
          {profile.name}
        </div>
      ))}
--fcc-editable-region--
    </div>
  );
}
```
