---
id: 687942c174eb0172ef4ded6a
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Now, let's add the other two profiles data to your `profiles` array in the `App` component. 

Use the data provided in the table below:

| Props  | First Card Values | Second Card Values |
|-------|------------------------------------------------------------|-----------------------------------------------------|
| `name`  | `"Tiffany"` | `"Doug"`|
| `title` | `"Engineering manager"` | `"Backend developer"` |
| `bio`   | `"I have worked in tech for 15 years and love to help people grow in this industry."` | `"I have been a software developer for over 20 years and I love working with Go and Rust."`|


This will demonstrate how powerful the `map()` function is, as all your cards will appear dynamically without you having to manually add individual `<Card />` components.

With that, your reusable profile card component project is complete!

# --hints--

The `profiles` array should contain all the correct profile data.

```js 
const cardElements = document?.querySelectorAll(".card");
const profiles = [
  { id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
  { id: '2', name: 'Tiffany', title: "Engineering manager", bio: "I have worked in tech for 15 years and love to help people grow in this industry." },
  { id: '3', name: "Doug", title: "Backend developer", bio: "I have been a software developer for over 20 years and I love working with Go and Rust." }
];

profiles.forEach((profile, index) => {
  const card = cardElements[index];
  assert.strictEqual(card?.querySelector("h2")?.textContent, profile.name);
  assert.strictEqual(card?.querySelector(".card-title")?.textContent, profile.title);
  assert.strictEqual(card?.querySelector("p:not(.card-title)")?.textContent, profile.bio);
});
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
  { id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
  --fcc-editable-region--
  
	--fcc-editable-region--
	];
  return (
    <div className="flex-container">
      {profiles.map((profile) => (
        <Card
          key={profile.id}
          name={profile.name}
          title={profile.title}
          bio={profile.bio}
        />
      ))}
    </div>
  );
}
```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
    { id: '1', name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." },
    { id: '2', name: 'Tiffany', title: "Engineering manager", bio: "I have worked in tech for 15 years and love to help people grow in this industry." },
    { id: '3', name: "Doug", title: "Backend developer", bio: "I have been a software developer for over 20 years and I love working with Go and Rust." }
  ];

  return (
    <div className="flex-container">
      {profiles.map((profile) => (
        <Card
          key={profile.id}
          name={profile.name}
          title={profile.title}
          bio={profile.bio}
        />
      ))}
    </div>
  );
}
```
