---
id: 68792530cde1deb18ea24996
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Now, inside your `map()` function, replace the `div` element with your `Card` component.

Next, pass the `name`, `title`, and `bio` props to the `Card` component.

# --hints--

You should replace the `div` element with the `Card` component in your `map()` function.

```js
assert.match(code, /profiles\.map\s*\(\s*\(?\s*profile\s*\)?\s*=>\s*(?:\(\s*)?<Card\b[^>]*\/>/);
```

Your `Card` component should have a `key` prop set to `profile.id`.

```js
assert.match(code, /<Card[^>]*key\s*=\s*\{\s*profile\.id\s*\}/);
```

Your `Card` component should have a `name` prop set to `profile.name`.

```js
assert.match(code, /<Card[^>]*name\s*=\s*\{\s*profile\.name\s*\}/);
```

Your `Card` component should have a `title` prop set to `profile.title`.

```js
assert.match(code, /<Card[^>]*title\s*=\s*\{\s*profile\.title\s*\}/);
```

Your `Card` component should have a `bio` prop set to `profile.bio`.

```js
assert.match(code, /<Card[^>]*bio\s*=\s*\{\s*profile\.bio\s*\}/);
```

Your `App` component should render one `Card` component with correct information.

```js
async () => {
  const testElem = await __helpers.prepTestComponent(window.index.App);
  const card = testElem.querySelector('.card');
  assert.exists(card);
  assert.equal(card.querySelector('h2').textContent, 'Mark');
  assert.equal(card.querySelector('.card-title').textContent, 'Frontend developer');
  assert.equal(card.querySelector('p:last-child').textContent, 'I like to work with different frontend technologies and play video games.');
}
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Card component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>

    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
     <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { App } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <App />
      );
    </script>
  </body>
</html>
```

```css
:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.flex-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: space-around;
  align-items: center;
}

.card {
  border: 5px solid var(--dark-orange);
  border-radius: 10px;
  width: 100%;
  padding: 20px;
  margin: 10px 0;
  background-color: var(--light-grey);
}

.card-title {
  border-bottom: 4px solid var(--dark-orange);
  width: fit-content;
}

@media (min-width: 768px) {
  .card {
    width: 300px;
  }
}
```

```jsx
export function Card({ name, title, bio }) {
  return (
    <div className="card">
      <h2>{name}</h2>
      <p className="card-title">{title}</p>
      <p>{bio}</p>
    </div>
  )
}

export function App() {
  const profiles = [
  	{ id: "1", name: "Mark", title: "Frontend developer", bio: "I like to work with different frontend technologies and play video games." }
	];
  return (
    <div className="flex-container">
--fcc-editable-region--
      {profiles.map((profile) => (
        <div key={profile.id}>
          {profile.name}
        </div>
      ))}
--fcc-editable-region--
    </div>
  );
}
```
