---
id: 615f357957e370510f21ea16
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

創建一個 `body` 選擇器，並將 `font-family` 設置爲 `Open Sans`，以 `sans-serif` 爲備用字體。

請記住，名稱中帶有空格的字體必須用 CSS 引號括起來。

# --hints--

應該有一個 `body` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 選擇器應該有一個 `font-family` 屬性設置爲 `"Open Sans", sans-serif`。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
