---
id: 615f3cafd794015aa9547a6d
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

請記住，`h1`、`h2` 和類似標籤的使用決定了 HTML 的語義結構。 但是，你可以調整這些元素的 CSS 來控制視覺流和層次結構。

創建 `h1` 規則並將 `font-weight` 屬性設置爲 `800`。 這將使 `h1` 文本更粗。

# --hints--

應該創建一個 `h1` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1` 選擇器應該有一個 `font-weight` 屬性設置爲 `800`。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
