---
id: 615f671b6d1919792745aa5d
title: 步驟 39
challengeType: 0
dashedName: step-39
---

# --description--

在 `.md` 元素下創建一個新的 `div` 元素。 給它一個 `class` 屬性設置爲 `daily-value sm-text`。 在這個新的 `div` 中，添加帶有文本 `% Daily Value *` 的 `p` 元素，並設置 `class` 屬性 `right bold`。

# --hints--

應該在 `.md` 元素之後創建一個新的 `div` 元素。

```js
assert(document.querySelector('.label')?.lastElementChild?.localName === 'div');
assert(document.querySelector('.label')?.lastElementChild?.previousElementSibling?.classList?.contains('divider'));
assert(document.querySelector('.label')?.lastElementChild?.previousElementSibling?.classList?.contains('md'));
```

新的 `div` 應該有一個 `class` 屬性設置爲 `daily-value sm-text`。

```js
assert(document.querySelector('.label')?.lastElementChild?.classList?.contains('daily-value'));
assert(document.querySelector('.label')?.lastElementChild?.classList?.contains('sm-text'));
```

新的 `div` 元素應該有一個 `p` 元素。

```js
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.localName === 'p');
```

新的 `p` 元素應包含文本 `% Daily Value *`。

```js
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.textContent === '% Daily Value *');
```

新的 `p` 元素應該有一個 `class` 屬性設置爲 `right bold`。

```js
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.classList?.contains('right'));
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.classList?.contains('bold'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
      <div class="calories-info">
        <p class="bold sm-text">Amount per serving</p>
        <h1>Calories <span class="right">230</span></h1>
      </div>
      <div class="divider md"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}

.md {
  height: 5px;
}

.sm-text {
  font-size: 0.85rem;
}

.calories-info h1 {
  margin: -5px -2px;
  overflow: hidden;
}

.calories-info span {
  font-size: 1.2em;
  margin-top: -7px;
}
```
