---
id: 615f4f9e4a40566b776a8f38
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

創建一個新的 `.large` 選擇器，並將其 `height` 屬性設置爲 `10px`。 同時創建一個 `.large, .medium` 選擇器，並將 `background-color` 屬性設置爲 `black`。

# --hints--

你應該有一個新的 `.large` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.large'));
```

你的 `.large` 選擇器應該將 `height` 屬性設置爲 `10px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.large')?.height === '10px');
```

你應該有一個新的 `.large, .medium` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.large, .medium'));
```

你的 `.large, .medium` 選擇器應該有一個 `background-color` 屬性設置爲 `black`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.large, .medium')?.backgroundColor === 'black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

--fcc-editable-region--

--fcc-editable-region--
```
