---
id: 615f51e4e5b24a6e80eccce1
title: 步驟 30
challengeType: 0
dashedName: step-30
---

# --description--

在你的 `.calories-info` 元素中，創建一個 `div` 元素。 給那個 `div` 元素一個 `class` 屬性設置爲 `left-container`。 在新創建的 `div` 元素中，創建一個帶有文本 `Amount per serving` 的 `h2` 元素。 給 `h2` 元素一個 `class` 屬性設置爲 `bold small-text`。

# --hints--

你應該在你的 `.calories-info` 元素中創建一個新的 `div` 元素。

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.localName === 'div');
```

你的新 `div` 元素應該有一個 `class` 屬性值爲 `left-container`。

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.classList.contains('left-container'));
```

你應該在你的 `.left-container` 元素中創建一個新的 `h2` 元素。

```js
assert(document.querySelector('.left-container')?.children?.[0]?.localName === 'h2');
```

你的新 `h2` 元素應該有一個 `class` 屬性設置爲 `bold small-text`。

```js
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('bold'));
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('small-text'));
```

你的新 `h2` 元素應包含文本 `Amount per serving`。

```js
assert(document.querySelector('.left-container')?.children?.[0]?.innerText === 'Amount per serving');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
--fcc-editable-region--
    <div class="calories-info">

    </div>
--fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}
```
