---
id: 615f34ecc1091b4fd5a8a484
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `head` 元素中，添加一個 `link` 元素，其中 `rel` 屬性設置爲 `stylesheet`，`href` 屬性設置爲 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`。

這將導入 `Open Sans` 字體系列，字體粗細值爲 `400`、`700` 和 `800`。

還要添加一個 `link` 元素來鏈接 `styles.css` 文件。

# --hints--

你的代碼應該有兩個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 2);
```

兩個 `link` 元素都應該將 `rel` 屬性設置爲 `stylesheet`。

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

`link` 元素之一應將 `href` 屬性設置爲 `./styles.css`。

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

`link` 元素之一應將 `href` 屬性設置爲 `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`。

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
