---
id: 615f3b091162165948e1cb33
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

如果使用瀏覽器的開發人員工具檢查 `.label` 元素，你可能會注意到它實際上是 288 像素寬而不是 270 像素。 這是因爲默認情況下，瀏覽器在確定元素大小時會包含邊框和內邊距。

要解決這個問題，創建一個 `*` 選擇器並給它一個 `border-box` 的 `box-sizing` 屬性來重置盒子模型。

# --hints--

應該創建一個 `*` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` 選擇器應該有一個 `box-sizing` 屬性設置爲 `border-box`。

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
