---
id: 64b171849f925b0773aa434c
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Create an list of numbers from 1 to 3 and assign it to the variable `xs`.

# --hints--

`xs` should be the list `[1,2,3]`.

```js
({ test: () => assert.deepEqual(Array.from(__userGlobals.get("xs")), [1,2,3])})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
xs = [1,2,3]
```
