---
id: 65df3afd233057f6a620a860
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

This shows how the helper functions can be used to analyze a learner's class definitions

# --instructions--

We can just find the class definition inside Python, return it to the JavaScript, and then check if it contains the string "pass".

If we need something more reliable, we can expand the helpers to check for pass statements specifically. This would be useful if say, we wanted to write classes containing, say, passengers.

# --hints--

`TreeNode` should not contain any `pass` statements.

```js
({ test: () => {
  const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
  assert.notInclude(pyClassStr, "pass")
}})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
  
class Trap:
    pass
--fcc-editable-region--
```

# --solutions--

```py
class TreeNode:
    def __init__(self, thing):
        self.thing = thing
        
class Trap:
    pass
  
```
