---
id: 6551ecf597038c36637f3759
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Змінні є важливою частиною Python та будь-якої мови програмування. Змінна — це назва, яка посилається на об’єкт або вказує на нього. Ви можете оголосити змінну, написавши назву змінної ліворуч від оператора присвоєння `=` та вказавши значення, яке потрібно присвоїти цій змінній, праворуч від оператора присвоєння:

```py
variable_name = value
```

Створіть змінну під назвою `number` та призначте значення `5` до нової змінної.

# --hints--

Оголосіть змінну під назвою `number`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^number\s*=/m)
```

Призначте значення `5` до змінної `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
