---
id: 6551f6f9ec6db270eb83f3db
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

На терміналі можна побачити, що надруковано `11`, оскільки `'Hello World'` містить 11 символів.

Ще одна корисна вбудована функція — `type()`, яка повертає тип даних змінної. Змініть виклик функції `print()`, щоб надрукувати тип даних змінної `text`.

# --hints--

Викличте функцію `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Передайте `text` до функції `type()`, написавши це в круглих дужках.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Передайте `type(text)` до функції `print()`, написавши це в круглих дужках. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Код не повинен містити `print(len(text))`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
