---
id: 6551fe3b1df7c9740f13f270
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Ключові аспекти найменування змінних в Python:

- Деякі слова є зарезервованими ключовими словами (наприклад, `for`, `while`, `True`). Вони мають особливе значення в Python, тому їх не можна використовувати для назв змінних.
- Назви змінних не можуть починатися з числа та можуть містити лише буквено-цифрові символи або підкреслення.
- Назви змінних чутливі до регістру, тобто `my_var` відрізняється від `my_Var` та `MY_VAR`.
- Часто назви змінних пишуть за допомогою `snake_case` (укр. зміїного_регістру), де пробіл замінюється підкресленням, а слова записані малими літерами.

Видаліть обидва виклики `print()` та оголосіть ще одну змінну під назвою `alphabet`. Призначте рядок `'abcdefghijklmnopqrstuvwxyz'` до цієї змінної.

# --hints--

Код не повинен містити `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Код не повинен містити `print(type(shift))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

Оголосіть змінну під назвою `alphabet`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^alphabet\s*=/m)
```

Призначте рядок `'abcdefghijklmnopqrstuvwxyz'` до змінної `alphabet`. Не забудьте, що потрібно використовувати або одинарні, або подвійні лапки.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
