---
id: 6551f5756c20146fc94f8675
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

You can access the number of characters in a string with the built-in `len()` function.

Modify your existing `print()` call by passing `len(text)` instead of `text[-1]`.

# --hints--


You should call the `len()` function.

```js
assert.match(code, /len\s*\(.*\)/)
```

You should pass the variable `text` to the `len()` function by including it between the parentheses.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

You should pass `len(text)` to the `print()` function by including it between the parentheses. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Код не повинен містити `print(text[-1])`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
