---
id: 6551f4259d66026ee28cb205
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Ви також можете отримати доступ до символів рядка, починаючи з кінця рядка. Останній символ має індекс `-1`, передостанній символ має індекс `-2` і так далі.

Тепер змініть наявний виклик `print()`, щоб надрукувати останній символ в рядку.

# --hints--

Ви все одно повинні викликати функцію `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Передайте `text[-1]` до функції `print()`, написавши це в круглих дужках. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
