---
id: 6551f3bfed58796ebb1268dc
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

На кожен символ рядка можна посилатися за допомогою числового індексу. Відрахунок індексу починається з нуля. Тобто, перший символ рядка має індекс `0`. Наприклад, в рядку `'Hello World'`: `'H'` має індекс `0`, `'e'` має індекс `1` і так далі.

До кожного символу рядка можна отримати доступ за допомогою дужкової нотації. Вам потрібно написати назву змінної, а потім квадратні дужки та додати індекс символу між дужками:

```py
text = 'Hello World'
r = text[8]
```

Тепер виведіть символ за індексом `6`, а не `text`.

# --hints--

Ви все одно повинні викликати функцію `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Передайте `text[6]` до функції `print()`, написавши це в круглих дужках. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
