---
id: 6551f5756c20146fc94f8675
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Ви можете отримати доступ до кількості символів у рядку за допомогою вбудованої функції `len()`.

Змініть наявний виклик `print()`, передавши `len(text)` замість `text[-1]`.

# --hints--


Викличте функцію `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Передайте змінну `text` до функції `len()`, написавши її в круглих дужках.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Передайте `len(text)` до функції `print()`, написавши це в круглих дужках. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Код не повинен містити `print(text[-1])`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
