---
id: 6551ef504d91936d2d4e54f8
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

*Аргумент* — це об’єкт або вираз, що передається до функції під час її виклику (в дужках):

```py
greet = 'Hello!'
print(greet)
```

Код у наведеному вище прикладі надрукує рядок `'Hello!'`, який є значенням змінної `greet`, переданої до `print()` як аргумент.

Надрукуйте змінну `text` на екрані, передавши змінну `text` як аргумент до функції `print()`.

# --hints--

Передайте `text` до функції `print()`, розмістивши назву змінної в дужках. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
