---
id: 6551f22653e6ce6e4f1d1370
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Modify your `print(shift)` call to print the data type of your `shift` variable.

# --hints--

You should keep your existing `print(type(text))` call.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

You should pass `shift` to the `type()` function.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

You should pass `type(shift)` to your `print()` function. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Код не повинен містити `print(shift)`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
