---
id: 6551fb531ce7f773071a0b84
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Як бачите, результатом друку `type(text)` є `<class 'str'>`, тобто змінна є рядком, що позначають як `str`.

Тепер перейдіть на новий рядок та створіть ще одну змінну під назвою `shift` і призначте до неї значення `3`.

# --hints--

Оголосіть змінну під назвою `shift`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^shift\s*=/m)
```

Призначте значення `3` до змінної `shift`.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
