---
id: 6551fe3b1df7c9740f13f270
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Key aspects of variable naming in Python are:

- Some words are reserved keywords (e.g. `for`, `while`, `True`). They have a special meaning in Python, so you cannot use them for variable names.
- Variable names cannot start with a number, and they can only contain alpha-numeric characters or underscores.
- Variable names are case sensitive, i.e. `my_var` is different from `my_Var` and `MY_VAR`.
- Finally, it is a common convention to write variable names using `snake_case`, where each space is replaced by an underscore character and the words are written in lowercase letters.

Remove both calls to `print()` and declare another variable called `alphabet`. Assign the string `'abcdefghijklmnopqrstuvwxyz'` to this variable.

# --hints--

Код не повинен містити `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Код не повинен містити `print(type(shift))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

Оголосіть змінну під назвою `alphabet`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^alphabet\s*=/m)
```

Призначте рядок `'abcdefghijklmnopqrstuvwxyz'` до змінної `alphabet`. Remember to use either single or double quotes to enclose the string.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
