---
id: 6551f22653e6ce6e4f1d1370
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Змініть виклик `print(shift)`, щоб надрукувати тип даних змінної `shift`.

# --hints--

Збережіть наявний виклик `print(type(text))`.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Передайте `shift` до функції `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Передайте `type(shift)` до функції `print()`. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Код не повинен містити `print(shift)`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
