---
id: 655220953ba90d80514d7ee2
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

Наразі функція `print()` приймає один аргумент `char`, але вона може приймати кілька аргументів, розділених комою.

Додайте другий аргумент до `print(char)`, щоб функція друкувала символ та його індекс в межах `alphabet`.

# --hints--

Додайте `index` як другий аргумент до наявного виклику `print(char)`. Не забудьте розділити аргументи комою.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
