---
id: 6552008c0d9d9075cbec9772
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Перший вид шифру, який ви створите, називається шифром *Цезаря*. Для цього потрібно взяти кожну літеру в повідомленні, знайти її позицію в алфавіті та замінити на літеру, яка розташована через 3 позиції.

Почніть із пошуку позиції першої літери в рядку. Один зі способів — виклик методу рядка `.find()` на рядку, який потрібно проаналізувати:

```py
text.find('W')
```

У прикладі вище, `'W'` у верхньому регістрі є символом, який потрібно знайти в межах рядка, що зберігається в змінній `text`. Метод поверне `6`, що є індексом символу `'W'` в межах рядка, який зберігається в змінній `text`.

В кінці коду викличте `.find()` на рядку `alphabet` та передайте `text[0]` до методу. Зауважте, що метод — це просто функція, яка належить об’єкту (про це ви дізнаєтеся більше в іншому проєкті).

# --hints--

Викличте метод `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Викличте метод `.find()` на змінній `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```

Викличте метод `.find()` на змінній `alphabet` та передайте `text[0]` до методу. Зверніть увагу на те, щоб розмістити виклик методу на початку рядка.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
