---
id: 65521fc818947e800bffe48a
title: Крок 27
challengeType: 20
dashedName: step-27
---

# --description--

Усередині циклу `for`, перш ніж друкувати поточний символ, оголосіть змінну під назвою `index` і призначте їй значення, яке повертає `alphabet.find(char)`.

# --hints--

Оголосіть нову змінну під назвою `index` на початку циклу `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Призначте `alphabet.find(char)` до нової змінної `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
