---
id: 65520e6f2b9678799977f24d
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Видаліть останній виклик `print()`. Потім замість `text[0]` передайте `text[0].lower()` до `.find()` і перегляньте результат.

# --hints--

Код не повинен містити `print(text.lower())`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Все одно надрукуйте змінну `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Оновіть виклик `alphabet.find(text[0])`, щоб використати `text[0].lower()` як аргумент. Зверніть увагу на те, щоб розмістити виклик методу на початку рядка.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
