---
id: 65520e6f2b9678799977f24d
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Remove the last `print()` call. Then, instead of `text[0]`, pass `text[0].lower()` to `.find()` and see the output.

# --hints--

Код не повинен містити `print(text.lower())`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

You should still print your `index` variable.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

You should update your `alphabet.find(text[0])` call to use `text[0].lower()` as the argument. Pay attention to place the method call at the beginning of the line.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
