---
id: 65525e359ca28d938baa82c5
title: Крок 40
challengeType: 20
dashedName: step-40
---

# --description--

Ви можете отримати той самий ефект `a = a + b`, використовуючи оператор додавання з присвоєнням:

```py
a += b
```

Оператор додавання з присвоєнням дозволяє додати значення до змінної, а потім призначити результат до неї.

Використайте оператор `+=`, щоб одночасно додати значення та призначити його до `encrypted_text`.

# --hints--

Використайте оператор додавання з присвоєнням, щоб додати `alphabet[new_index]` до поточного значення `encrypted_text` в тілі циклу.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
