---
id: 6552008c0d9d9075cbec9772
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Ви будете використовувати метод `.find()`, щоб знайти позицію кожної літери з повідомлення в алфавіті. Метод схожий на функцію, але належить до об’єкта.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

У прикладі вище, метод `.find()` *викликаний на* `sentence` (рядок, де відбувається пошук), а `'r'` (символ, який потрібно знайти) передано як аргумент. Виклик `sentence.find('r')` поверне `4`, що є індексом першого випадку, коли в `sentence` зустрічається `'r'`.

В кінці коду викличте `.find()` на `alphabet` та передайте `'z'` як аргумент до методу.

# --hints--

Викличте метод `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Викличте метод `.find()` на змінній `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Викличте метод `.find()` на змінній `alphabet` та передайте `'z'` до методу. Зверніть увагу на те, щоб розмістити виклик методу на початку рядка.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
