---
id: 65520f8282faf57a0db4f7fe
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Як бачите з результату, `'h'` в рядку `alphabet` має індекс `7`. Тепер вам потрібно знайти літеру за індексом `7` плюс значення `shift`. Для цього можна використати оператор додавання `+`, так само як і в математичному додаванні.

Змініть змінну `shifted`, щоб вона зберігала значення `alphabet` за індексом `index + shift`.

# --hints--

Призначте `alphabet[index + shift]` до змінної `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
