---
id: 65521203d7165c7b84b22ad4
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Цикл дозволяє систематично проходити послідовність елементів і виконувати дії над кожним з них.

У цьому випадку ви використаєте цикл `for`. Ось так ви можете ітерувати над `text`:

```py
for i in text:
```

`for` — це ключове слово, яке позначає тип циклу. `i` — це змінна, яка послідовно приймає значення елементів у `text`. Інструкція закінчується двокрапкою `:`.

Під рядком, де ви оголосили `alphabet`, напишіть цикл `for`, щоб ітерувати над `text`. Використайте `i` як змінну циклу.

Внаслідок в терміналі виникне помилка. Ви дізнаєтеся про це в наступному кроці.

# --hints--

Використайте ключове слово `for`, щоб створити цикл. Обов’язково розмістіть ключове слово `for` на початку рядка та поставте пробіл після ключового слова.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Напишіть змінну `i` після ключового слова `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Напишіть ключове слово `in` після `for i`. Не забудьте поставити пробіли навколо ключового слова `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Напишіть `text` після `for i in`. Не забудьте додати `:` вкінці.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Розмістіть цикл `for` під рядком коду `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
