---
id: 655208d59b131e7816f18c96
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Функція `print()` лише виводить дані на консолі, але функції та методи можуть повернути значення, яке можна використати в коді.

Тепер призначте `alphabet.find(text[0])` до змінної під назвою `index`. Таким чином `index` зберігатиме значення, яке повертає `alphabet.find(text[0])`.

# --hints--

Оголосіть змінну під назвою `index`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^index\s*=/m)
```

Призначте значення, яке повертає `alphabet.find(text[0])`, до змінної `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
