---
id: 65524d2a1a253b8bb5197ae2
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

`find` знову повертає `-1` для літер у верхньому регістрі та пробілу. Пробілом займемося пізніше.

А зараз, ітерацію `text` в циклі `for` змінимо на ітерацію `text.lower()`.

# --hints--

Змініть цикл `for`, щоб перебирати `text.lower()`, а не `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
