---
id: 655261b2e1f2c197093f3993
title: Крок 42
challengeType: 20
dashedName: step-42
---

# --description--

Наразі пробіли шифруються як `'c'`. Щоб зберегти вихідний інтервал в повідомленні, вам знадобиться умовна інструкція `if`. Вона складається з ключового слова `if`, умови та двокрапки `:`.

```py
if x != 0:
    print(x)
```

Умовою інструкції `if` в прикладі вище є `x != 0`. Код `print(x)` в межах тіла інструкції `if` виконується лише тоді, коли умова оцінюється як `True` (у цьому прикладі це означає, що `x` не дорівнює нулю).

Замініть `print(char == ' ')` у верхній частині циклу `for` на інструкцію `if`.  Умова цієї інструкції `if` має оцінюватись як `True`, якщо `char` є пробілом або `False`, якщо навпаки. Усередині тіла `if` надрукуйте рядок `'space!'`. Не забудьте зробити відступ для цього рядка.

# --hints--

Код не повинен містити `print(char == ' ')`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Замініть `print(char == ' ')` на інструкцію `if`, яка запускається, якщо `char == ' '`. Не використовуйте круглі дужки для інструкції `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Надрукуйте рядок `'space!'` в межах нової інструкції `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
