---
id: 65524b3aa6a1938a069a91ab
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

Далі надрукуйте `new_char` і перегляньте результат.

# --hints--

Надрукуйте змінну `new_char`.

```js
assert.match(code, /print\s*\(\s*new_char\s*\)/)
```

Надрукуйте змінну `new_char` в кінці тіла циклу.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*new_char\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
--fcc-editable-region--
```
