---
id: 655251308f31958d06cdf267
title: Крок 37
challengeType: 20
dashedName: step-37
---

# --description--

At the moment, the encrypted character is updated in every iteration. It would be better to store the encrypted string in a new variable. Before your `for` loop, declare a variable called `encrypted_text` and assign an empty string (`''`) to this variable.

# --hints--

Оголосіть змінну під назвою `encrypted_text` перед циклом `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Призначте порожній рядок до змінної `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
