---
id: 655251308f31958d06cdf267
title: Крок 37
challengeType: 20
dashedName: step-37
---

# --description--

Зараз зашифрований символ оновлюється на кожній ітерації. Було б краще зберегти зашифрований рядок у новій змінній. Перед циклом `for` оголосіть змінну під назвою `encrypted_text` та призначте до неї порожній рядок (`''`).

# --hints--

Оголосіть змінну під назвою `encrypted_text` перед циклом `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Призначте порожній рядок до змінної `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
