---
id: 65521203d7165c7b84b22ad4
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Repeating this process for each character in `text` can be tedious. Thankfully, you can simplify it using a loop. A loop allows you to systematically go through a sequence of elements and execute actions on each one.

In this case, you'll employ a `for` loop. Here's how you can iterate over `text`:

```py
for i in text:
```

`for` is the keyword denoting the loop type. `i` is a variable that sequentially takes the value of the elements in `text`. The statement ends with a colon, `:`.

Remove everything after the `alphabet` line. Then write a `for` loop to iterate over `text`.

# --hints--

You should create a `for` loop to iterate over `text`. Don't forget the final `:`.

```js
assert.match(code, /for\s+\w+\s+in\s+text\s*:/)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
