---
id: 655209a4a27dd37873c4cac3
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

`.find()` повертає індекс відповідного символу всередині рядка. Якщо символ не знайдено, повертається `-1`. Як бачите, перший символ в `text` (`'H'` у верхньому регістрі) не знайдено, оскільки `alphabet` містить літери лише в нижньому регістрі.

Ви можете перетворити рядок на його еквівалент у нижньому регістрі за допомогою методу `.lower()`. Додайте ще один виклик `print()`, щоб надрукувати `text.lower()` та подивіться на результат.

# --hints--

Код досі повинен містити `print(index)`. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Надрукуйте `text.lower()`. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
