---
id: 6553a755879b131a445e664c
title: Крок 45
challengeType: 20
dashedName: step-45
---

# --description--

Спробуйте призначити рядок `'Hello Zaira'` до змінної `text` та подивіться на результат на консолі.

# --hints--

Ви повинні мати змінну `text`.

```js
assert.match(code, /^text\s*=/m)
```

Призначте рядок `'Hello Zaira'` до змінної `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
