---
id: 6553f8c570f9982e013a8886
title: Крок 56
challengeType: 20
dashedName: step-56
---

# --description--

Наразі кожна літера завжди шифрується тією самою літерою, залежно від зазначеного зсуву. Що, якби зсув для літер був різним? Так було б набагато важче виконати розшифрування. Цей алгоритм називається шифром Віженера, де зсув для кожної літери визначається іншим текстом, який називається *ключем*.

Щоб розпочати перетворювати шифр Цезаря на шифр Віженера, видаліть два виклики функції.

# --hints--

Видаліть два виклики функції `caesar()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
