---
id: 6553f6b4ec51112d44d737c8
title: Крок 49
challengeType: 20
dashedName: step-49
---

# --description--

Додайте ще один виклик одразу перед викликом `print` та передайте `'plain text:', text` як аргументи до `print()`. Використайте той самий відступ.

# --hints--

Надрукуйте `'plain text:', text` після циклу `for`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
