---
id: 6553ed69ece88d29594748aa
title: Крок 51
challengeType: 20
dashedName: step-51
---

# --description--

Щоб виконати функцію, її потрібно викликати (або звернутися до неї), додавши пару круглих дужок після її назви:

```py
function_name()
```

У кінці коду викличте функцію `caesar`. Зверніть увагу на відступи.

# --hints--

Викличте функцію `caesar`. Обов’язково напишіть виклик функції на початку рядка.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
