---
id: 6553980e0527fa115c705646
title: Крок 43
challengeType: 20
dashedName: step-43
---

# --description--

Тепер замість того, щоб друкувати `'space!'`, використайте оператор додавання з присвоєнням, щоб додати пробіл (наразі зберігається в `char`) до поточного значення `encrypted_text`.

# --hints--

Код не повинен містити `print('space!')`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Використайте оператор `+=`, щоб додати `char` до поточного значення `encrypted_text` в межах нової інструкції `if`.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
