---
id: 6553995f412dd8122ed38e4a
title: Крок 44
challengeType: 20
dashedName: step-44
---

# --description--

Умовна інструкція також може мати умову `else`. Цю умову можна додати в кінець інструкції `if`, щоб виконати альтернативний код, якщо умова інструкції `if` хибна:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

Як можна побачити в результаті, коли ітерації циклу досягають пробілу, то його додають до зашифрованого рядка. Потім виконується код поза блоком `if` і до зашифрованого рядка додається `c`.

Щоб виправити це, додайте умову `else` після `encrypted_text += char` та зробіть відступ для всіх наступних рядків коду, крім виклику `print()`.

# --hints--

Створіть умову `else`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

Зробіть відступи для всіх рядків коду після умови `else`, окрім виклику `print()`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
