---
id: 6553f4f66099802c6ae94613
title: Крок 54
challengeType: 20
dashedName: step-54
---

# --description--

Наразі код не може виконатись, оскільки функція `caesar` визначена з двома параметрами (`message` та `offset`), але викликається без *аргументів*.

Надайте значення до `message` та `offset`, передавши `text` та `shift` як аргументи до виклику функції `caesar`.

# --hints--

Передайте `text` та `shift` як аргументи до виклику функції, розмістивши їх в круглих дужках. Не забудьте розділити аргументи комою.

```js
assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
