---
id: 6553f4f66099802c6ae94613
title: Step 53
challengeType: 20
dashedName: step-53
---

# --description--

Now, uncomment your `caesar()` call. You'll see that your code raises a `TypeError`, because the `caesar` function is defined with two parameters (`message` and `offset`), therefore it expects to be called with two *arguments*.

Calling `caesar()` without the required arguments stops the execution of the code.

Надайте значення до `message` та `offset`, передавши `text` та `shift` як аргументи до виклику функції `caesar`.

# --hints--

Передайте `text` та `shift` як аргументи до виклику функції, розмістивши їх в круглих дужках. Не забудьте розділити аргументи комою.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
# caesar()
--fcc-editable-region--
```
