---
id: 6553f6086add4b2cbb99fd78
title: Крок 55
challengeType: 20
dashedName: step-55
---

# --description--

У нижній частині коду, після наявного виклику `caesar(text, shift)`, знову викличте функцію.

Але цього разу передайте змінну `text` і ціле число `13` як аргументи.

# --hints--

Знову викличте функцію, цього разу передавши `text` та `13` як аргументи.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

Збережіть наявний виклик функції.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
