---
id: 6553a44b1801991847d8cc69
title: Крок 48
challengeType: 20
dashedName: step-48
---

# --description--

Потім змініть виклик `print()`, щоб надрукувати `'encrypted text:', encrypted_text` та розмістіть його поза циклом `for`, щоб зашифрований рядок друкувався один раз.

# --hints--

Надрукуйте `'encrypted text:', encrypted_text` після циклу `for`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

Код не повинен містити `print('char:', char, 'encrypted text:', encrypted_text)`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
