---
id: 6553f4f66099802c6ae94613
title: Крок 53
challengeType: 20
dashedName: step-53
---

# --description--

Наразі код видає `TypeError`, оскільки функцію `caesar` визначено з двома параметрами (`message` та `offset`), тому вона очікує бути викликаною з двома *аргументами*.

Виклик `caesar()` без необхідних аргументів зупиняє виконання коду.

Надайте значення до `message` та `offset`, передавши `text` та `shift` як аргументи до виклику функції `caesar`.

# --hints--

Передайте `text` та `shift` як аргументи до виклику функції, розмістивши їх в круглих дужках. Не забудьте розділити аргументи комою.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
