---
id: 6553a7d8d05cbb1ae335a665
title: Крок 46
challengeType: 20
dashedName: step-46
---

# --description--

Коли цикл досягає літери `Z`, сума `index + shift` перевищує довжину рядка `alphabet`.

У цьому випадку можна використати оператор ділення з остачею (`%`), щоб повернути остачу від ділення двох чисел. Наприклад, `5 % 2` дорівнює `1`, оскільки при діленні 5 на 2 ми отримаємо частку 2 та остачу 1.

Візьміть `index + shift` в круглі дужки та виконайте операцію ділення з остачею на `26`, що є довжиною `alphabet`.

# --hints--

Інструкція `else` має містити `new_index = (index + shift) % 26`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
