---
id: 655491bd5b98b813fa5bedca
title: Крок 57
challengeType: 20
dashedName: step-57
---

# --description--

Тепер змініть оголошення функції: змініть назву функції на `vigenere`, а другий параметр на `key`.

# --hints--

Змініть назву функції на `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

Функція `vigenere` має приймати `message` та `key` як параметри.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
