---
id: 6554d0332949b133a0b35eaa
title: Крок 67
challengeType: 20
dashedName: step-67
---

# --description--

Викличте функцію, передавши `text` та `custom_key` як аргументи. Збережіть значення, яке повертає виклик функції, у змінній під назвою `encryption`.

# --hints--

Викличте `vigenere`, передавши `text` та `custom_key` як аргументи.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

Ви повинні мати змінну `encryption`.

```js
({ test: () => assert(__userGlobals.has("encryption")) })
```

Призначте виклик функції до змінної `encryption`.

```js
({ test: () => assert.match(code, /encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():

        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]

    return encrypted_text
--fcc-editable-region--

--fcc-editable-region--
```
