---
id: 6554d25dc5ceaa354307a77e
title: Крок 69
challengeType: 20
dashedName: step-69
---

# --description--

Шифрування та розшифрування є протилежними процесами, а за допомогою певних налаштувань функція може робити і те, і те.

Додайте третій параметр під назвою `direction` до визначення функції. Також закоментуйте останні два рядки коду, щоб уникнути помилок на консолі.

# --hints--

Перетворіть останні два рядки коду на коментарі. Поставте `#` на початку кожного рядка.

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

Додайте `direction` як третій параметр функції.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    f = __locals.get("vigenere")
    sig = str(inspect.signature(f))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():

        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]

    return encrypted_text

encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
