---
id: 65549f90cf78131c96ebcf28
title: Крок 60
challengeType: 20
dashedName: step-60
---

# --description--

При написанні коду важлива читабельність. Коментарі служать ефективними примітками, які пояснюють логіку коду. Вони стають цінними, якщо через деякий час ви повертаєтесь до проєкту, а також допомагають колегам зрозуміти код.

Коментар в Python можна написати за допомогою `#`. Усе, що стоїть після `#`, не буде виконано.

Додайте коментар `Append space to the message` перед інструкцією `if`.

# --hints--

Додайте коментар `Append space to the message`. Не забудьте символ `#` напочатку.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--

        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
