---
id: 6554930320d70414e7b6acc6
title: Крок 58
challengeType: 20
dashedName: step-58
---

# --description--

Видаліть змінну `shift`. Потім оголосіть нову змінну під назвою `custom_key` та призначте цій змінній значення рядка `'python'`.

# --hints--

Видаліть змінну `shift` та її значення.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Оголосіть змінну під назвою `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Призначте рядок `'python'` до змінної `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
