---
id: 6554e2ee23bfd93f2c83640f
title: Крок 73
challengeType: 20
dashedName: step-73
---

# --description--

Тепер надрукуйте змінну `decryption`.

# --hints--

Надрукуйте нову змінну `decryption`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*decryption\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():

        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]

    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
--fcc-editable-region--
```
