---
id: 65553159615a8123b190ee43
title: Крок 86
challengeType: 20
dashedName: step-86
---

# --description--

Працює! Тепер ви почнете із зашифрованого повідомлення, яке потрібно розшифрувати.

Змініть значення `text` на рядок `'mrttaqrhknsw ih puggrur'`.

# --hints--

Ви досі повинні мати змінну `text`.

```js
({ test: () => assert(__userGlobals.has("text")) })
```

Змінна `text` повинна мати значення `'mrttaqrhknsw ih puggrur'`.

```js
({ test: () => assert.equal(__userGlobals.get("text"), "mrttaqrhknsw ih puggrur") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira!'
custom_key = 'python'
--fcc-editable-region--
def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)

encryption = encrypt(text, custom_key)
print(encryption)
decryption = decrypt(encryption, custom_key)
print(decryption)
```
