---
id: 655619327c7b364166f8dd6f
title: Крок 47
challengeType: 20
dashedName: step-47
---

# --description--

Якщо ви хочете використати додаткові символи в рядку `alphabet` (наприклад, цифри чи спеціальні символи), то ви дізнаєтесь, що потрібно вручну змінити правий операнд операції ділення з остачею.

Замініть `26` на `len(alphabet)`, щоб уникнути цієї проблеми.

# --hints--

Змініть значення `new_index`, замінивши `26` на `len(alphabet)`.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
  assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/));
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % 26
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
