---
id: 655a2a7210094920069b117c
title: Крок 41
challengeType: 20
dashedName: step-41
---

# --description--

Оператори порівняння дозволяють порівнювати два об’єкти на основі їхніх значень. Ви можете використати оператор порівняння, розмістивши його між об’єктами, які потрібно порівняти. Вони повертають *булеве* (або *логічне*) значення: `True` або `False`, залежно від правдивості виразу.

У Python є такі оператори порівняння:

<table>
  <thead>
    <tr>
      <th>Оператор</th>
      <th>Значення</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Дорівнює</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Не дорівнює</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Більше ніж</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Менше ніж</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Більше або дорівнює</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Менше або дорівнює</td>
    </tr>
  </tbody>
</table>

На початку тіла циклу надрукуйте результат порівняння `char` з пробілом (`' '`). Для цього використайте оператор рівності `==`.

# --hints--

Порівняйте `char` з пробілом, використавши оператор порівняння в циклі `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Надрукуйте результат порівняння `char` з пробілом в межах циклу `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Надрукуйте результат порівняння `char` з пробілом на початку циклу.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
