---
id: 65687d2f8c7ee27b0446cef3
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Код, який виконується на кожній ітерації (розміщений після `:`), становить тіло циклу. Цей код повинен мати відступ. У Python рекомендується використовувати 4 пробіли на рівень відступу. Цей рівень із відступом є блоком коду.

```py
for i in text:
    print(i)
```

Надайте циклу `for` тіло, додавши виклик до `print(i)`. Не забудьте зробити відступ для тіла циклу.

# --hints--

Додайте `print(i)` до тіла циклу `for`. Зверніть увагу на відступи.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
