---
id: 656886ec4e4d9b96ded3c565
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Змінна ітерації може мати будь-яку дійсну назву, але краще робити її змістовною.

Перейменуйте змінну `i` на `char`.

# --hints--

Замініть `i` на `char`.

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
