---
id: 656884439d4d7d929d58669d
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Python relies on indentation to indicate blocks of code. A colon at the end of a line is a signal that a new indented block of code will follow.

So, when no indented block is found after the final colon, the code execution stops and an `IndentationError` is thrown. Therefore, there's no output anymore.

Restore the indentation for your `print(i)` call.

# --hints--

You should indent `print(i)` inside your `for` loop body.

```js
({ test: () => assert.match(code, /for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
print(i)
--fcc-editable-region--
```
