---
id: 65688a50e6c998a21d8e41d3
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

Рядки незмінні, тобто їх не можна змінити після створення. Наприклад, ви можете подумати, що наступний код змінить значення `my_string` на рядок `'train'`, але це не так:

```py
my_string = 'brain'
my_string[0] = 't'
```

Перевірте це, використавши дужкову нотацію, щоб отримати першу літеру в `text` та спробуйте змінити її на будь-який символ. Ви побачите, що вивід зникне.

# --hints--

Код досі повинен містити `text = 'Hello World'`.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Отримайте доступ до першої літери в `text` за допомогою `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Використайте оператор `=`, щоб призначити будь-який символ до `text[0]`. Не забудьте, що потрібно використовувати або одинарні, або подвійні лапки.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
