---
id: 6568b7a40c7ba0ccbc4b4425
title: Крок 33
challengeType: 20
dashedName: step-33
---

# --description--

Як бачите, кожен символ перепризначеного рядка друкується всередині циклу.

Поверніться до вихідного рядка, видаливши перепризначення `text`.

# --hints--

Код не повинен містити `text = 'Albatross'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

Ви все одно повинні мати змінну `text` зі значенням `'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
