---
id: 656883980318048fb11a6e3b
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Тепер видаліть відступ і подивіться на результат.

# --hints--

Видаліть відступ з рядка, де викликаєте `print(i)`.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s*^print\s*\(\s*i\s*\)/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
