---
id: 656884439d4d7d929d58669d
title: Крок 25
challengeType: 20
dashedName: step-25
---

# --description--

В Python використовуються відступи, щоб вказувати блоки коду. Двокрапка в кінці рядка є сигналом того, що слідуватиме новий блок коду з відступом.

Якщо після останньої двокрапки не знайдено блоку з відступом, виконання коду припиняється та видається `IndentationError`. Тому нічого не виводиться.

Відновіть відступ для виклику `print(i)`.

# --hints--

Зробіть відступ для `print(i)` в тілі циклу `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text:\s*^\s+print\s*\(\s*i\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
print(i)
--fcc-editable-region--
```
