---
id: 65688b5a1655a7a6caede847
title: Step 31
challengeType: 20
dashedName: step-31
---

# --description--

When you try to change the individual characters of a string as you did in the previous step, you get a `TypeError`, which occurs when an object of inappropriate type is used in your code.

As you can see from the error message, strings do not support item assignment, because they are immutable. Однак змінній можна перепризначити інший рядок:

```py
message = 'Hello World'
message = 'Hello there!'
```

Видаліть рядок `text[0]` та перепризначте рядок `'Albatross'` до `text`.

# --hints--

Код не повинен містити `text[0] = 'C'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Перепризначте рядок `'Albatross'` до змінної `text`. Не змінюйте `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
