---
id: 65951b637ddec63611c482c0
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Ви можете використовувати вбудовану функцію `print()`, щоб вивести результат коду в терміналі.

Функції — це багаторазово використовувані блоки коду, які можна викликати (або запустити), коли вони вам потрібні. Щоб викликати функцію, вам просто потрібно написати пару круглих дужок після її назви. Незабаром ви дізнаєтеся більше про функції.

А зараз перейдіть на новий рядок і додайте порожній виклик до функції `print()`. Поки ви не побачите жодних результатів.

# --hints--

Код повинен містити `print()`. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
