---
id: 65952b70924edb449faf1950
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Declare a new variable named `shifted`. Use the bracket notation to access the value of `alphabet` at index `index` and assign it to your new variable.

# --hints--

You should declare a variable named `shifted`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^shifted\s*=/m)
```

Призначте `alphabet[index]` до змінної `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
