---
id: 659526f3e5424b3eac4b6aa7
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Тепер надрукуйте змінну `index` на консолі.

# --hints--

Надрукуйте змінну `index`. Зверніть увагу на те, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
