---
id: 65952b70924edb449faf1950
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Оголосіть нову змінну `shifted`. Використайте дужкову нотацію, щоб отримати доступ до значення змінної `alphabet` за індексом `index` та призначте його до нової змінної.

# --hints--

Оголосіть змінну під назвою `shifted`. Зверніть увагу на те, щоб розмістити назву змінної на початку рядка.

```js
assert.match(code, /^shifted\s*=/m)
```

Призначте `alphabet[index]` до змінної `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
