---
id: 6624e0ec0bf47017eed84b1f
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Перший вид шифру, який ви створите, називається шифром *Цезаря*.  Для цього потрібно взяти кожну літеру в повідомленні, знайти її позицію в алфавіті та замінити на літеру, яка розташована через 3 позиції.

Для цього ви будете використовувати метод `.find()`, про який згадувалось в попередньому кроці. Змініть наявний виклик `.find()`, передавши `text[0]` як аргумент (замість `'z'`).

# --hints--

Змініть наявний виклик `alphabet.find('z')`, передавши `text[0]` до методу.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
