---
id: 66605a64a9822f01ec5ae8ce
title: Крок 50
challengeType: 20
dashedName: step-50
---

# --description--

*Область* змінної в Python визначає, де можна отримати доступ до цієї змінної:

- змінні, визначені поза функцією, мають *глобальну* область, і до них можна отримати доступ з будь-якої частини коду;

- змінні, визначені в межах функції, є *локальними* для цієї функції, і до них неможливо отримати доступ поза нею.

Щоб побачити це в дії, спробуйте надрукувати змінну `alphabet` у кінці коду. Це видасть виняток `NameError`.

Ви повинні побачити повідомлення про помилку, яке вказує на те, що `alphabet` не визначено. Причина в тому, що змінну `alphabet` визначено в межах функції `caesar` та вона недоступна поза нею.

# --hints--

Спробуйте надрукувати змінну `alphabet` поза функцією `caesar`.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
