---
id: 6143862a5e54455d262c212e
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل `body`، قم بإنشاء عنصر `main`. ثم في هذا العنصر، أنشئ `section` مع `class` بقيمة `heading`.

# --hints--

يجب أن يكون لديك عنصر `main`.

```js
assert.exists(document.querySelector('main'));
```

عنصر `main` الخاص بك يجب أن يكون داخل عنصر `body` الخاص بك.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

يجب عليك إنشاء عنصر `section`.

```js
assert.exists(document.querySelector('section'));
```

عنصر `section` الخاص بك يجب أن يكون ضمن عنصر `main` الخاص بك.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

عنصر `section` الخاص بك يجب أن يحتوي على `class` بقيمة `heading`.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
