---
id: 614389f601bb4f611db98563
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

أسفل عنصر `.author`، قم بإنشاء عنصر `div` جديد مع class يدعي `social-icons`.

أضف خمسة عناصر `a` ضمن `div` الجديد، وإعطائهم خصائص `href` التالية.

- أول عنصر `a` يجب أن يحتوي علي `href` تم تعيينه إلى `https://www.facebook.com/freecodecamp`.
- ثاني عنصر `a` يجب أن يحتوي على `href` تم تعيينه إلى `https://twitter.com/freecodecamp`.
- ثالث عنصر `a` يجب ان يحتوي علي `href` تم تعيينه إلى `https://instagram.com/freecodecamp`.
- رابع عنصر `a` يجب أن يحتوي على `href` تم تعيينه إلى `https://www.linkedin.com/school/free-code-camp`.
- خامس عنصر `a` يجب أن يحتوي علي `href` تم تعيينه إلى `https://www.youtube.com/freecodecamp`.

# --hints--

يجب عليك إنشاء عنصر `div` جديد.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

عنصر `div` الجديد الخاص بك يجب أن يأتي بعد عنصر `.author` الخاص بك.

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

يجب أن يحتوي عنصر `div` الجديد على class يدعي `social-icons`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

يجب أن يحتوي عنصر `.social-icons` على خمسة عناصر `a`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

أول عنصر `a` يجب أن يحتوي علي `href` تم تعيينه إلى `https://www.facebook.com/freecodecamp`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

ثاني عنصر `a` يجب أن يحتوي على `href` تم تعيينه إلى `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

ثالث عنصر `a` يجب ان يحتوي علي `href` تم تعيينه إلى `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

رابع عنصر `a` يجب أن يحتوي على `href` تم تعيينه إلى `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

خامس عنصر `a` يجب أن يحتوي على `href` تم تعيينه إلى `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
