---
id: 614387cbefeeba5f3654a291
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

ضمن عنصر `header` ، أنشئ `div` مع `class` يسمي `author`.

ضمن هذا الـ `div` ، أنشئ عنصر `p` مع `class` بقيمة `author-name` وأعطيه النص `By freeCodeCamp`. قم بتغليف جزء `freeCodeCamp` في عنصر `a` مع تعيين `href` إلى `https://freecodecamp.org`، و `target` إلى `_blank`.


أسفل ذلك ، أضف عنصر `p` ثاني بـ class يسمي `publish-date` والنص `March 7, 2019`.

# --hints--

يجب عليك إضافة عنصر `div` جديد.

```js
assert.exists(document.querySelector('div'));
```

عنصر `div` الخاص بك يجب أن يأتي بعد عنصر `header` الخاص بك.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

عنصر `div` الخاص بك يجب أن يحتوي على `class` يسمي `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

يجب عليك إضافة عنصرين `p` جديدين.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

عناصران `p` الجديدان يجب ان يكوانا داخل عنصر `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

يجب أن يكون لعنصر `p` الجديد الخاص بك `class` بقيمة `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

يجب أن يحتوي العنصر الجديد الأول `p` على النص `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

يجب أن يحتوي عنصر `p` الجديد الثاني على `class` بقيمة `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

يجب أن يحتوي العنصر `p` الجديد الثاني على النص `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

يجب عليك إنشاء عنصر `a` جديد.

```js
assert.exists(document.querySelector('a'));
```

يجب أن يكون عنصر `a` الخاص بك ضمن العنصر الأول الجديد `p`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

عنصر `a` الخاص بك يجب أن يحتوي على `href` بقيمة `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

عنصر `a` الخاص بك يجب أن يحتوي على `target` بقيمة `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

عنصر `a` الخاص بك يجب أن يحيط بالنص `freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.textContent === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
