---
id: 614385513d91ae5c251c2052
title: الخطوة الثانية
challengeType: 0
dashedName: step-2
---

# --description--

اضف عنصر `title` يحتوي على النص `Magazine`، وعنصر `link` يحمل خط `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` الى ‏‏ورقة الأنماط (stylesheet), ثم `link` يحمل أدوات `https://use.fontawesome.com/releases/v5.8.2/css/all.css` الى ‏‏ورقة الأنماط (stylesheet), وأستخدم `link` يحمل `./styles.css` لربط ورقة الأنماط (stylesheet).

ستحمل ورقة نمط الخط (stylesheet) الخاصة بك ثلاثة خطوط منفصلة: `Anton`, و `Baskervville`, و `Raleway`.

# --hints--

الكود الخاص بك يجب أن يحتوي على ثلاث عناصر `link` مغلقة ذاتيا.

```js
assert(document.querySelectorAll('link').length === 3);
```

عنصر `link` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert(document.querySelectorAll('head > link').length === 3);
```

العناصر `link` الثلاث الخاص بك يجب أن تحتوي على سمة (attribute) الـ `rel` بقيمة `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

أحد عناصر الرابط (link) الخاص بك يجب أن يحتوى على `href` محدد إلى `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

واحد من الروابط (link) الخاص بك يجب أن يحتوي على `href` بقيمة `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

يجب أن يحتوي عنصر `link` من العناصر الخاص بك على خاصية `href` مع القيمة `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/)
```

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لمشروعك العنوان `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

تذكر أن الـحالة الحرف (casing) سواء كبيرا أو صغيرا (capital or small), والإملاء (spelling) مهمان للعنوان(title).

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
