---
id: 61437d575fb98f57fa8f7f36
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

ابدأ بمستند HTML المعتاد الخاص بك. أضف تعريف `DOCTYPE`، وعنصر `html` موضحاً لغة صفحة بالإنجليزية و عنصر `head` وعنصر `body`.

قم باضافة وسم `<meta>` مع الـ `charset` المناسبة و وسم `<meta>` للتجاوب مع شاشات المحمول داخل عنصر `head`.

# --hints--

يجب أن يحتوي الكود الخاص بك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` مع `>` بعد الـ type.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يكون لديك علامة فتح `<html>`، ويجب أن يحتوي على `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على closing tag.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك الـ opening tag الآتي `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك الـ closing tag الآتي `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك الـ opening tag الآتي `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك الـ closing tag الآتي `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

عناصر `head` و `body` يجب أن يكونا أشقاء (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

عنصر `head` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

عنصر `body` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

احد عناصر `meta` يجب أن يحتوي على `name` تم تعيينه إلى `viewport`، و `content` تم تعيينه إلى `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

عنصر `meta` الآخر يجب أن يحتوي على السمة `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
