---
id: 6143869bb45bd85e3b1926aa
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

داخل عنصر `.heading` الخاص بك، قم بإنشاء عنصر `header` مع `class` يسمي `hero`.

في هذا العنصر، أنشئ عنصر `img` مع `src` بقيمة `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` و `alt` تم تعيينه إلى `freecodecamp logo` و `class` تم تعيينه إلى `hero-img`.

السمة `loading` في عنصر `img` يمكن تعيينها إلى `lazy` لإخبار المتصفح بعدم جلب مورد الصورة إلى أن تكون مطلوبة (اي عندما يقوم المستخدم بالتمرير الي الصورة). كميزة إضافية ، لن يتم تحميل العناصر البطيئة التحميل حتى يتم تحميل العناصر غير البطيئة - وهذا يعني أن المستخدمين الذين لديهم اتصالات إنترنت بطيئة يمكنهم عرض محتوى صفحتك دون الحاجة إلى الانتظار حتى يتم تحميل الصور.

أعطِ عنصر `img` الجديد السمة `loading` التي تم تعيينها إلى `lazy`.

# --hints--

يجب إنشاء عنصر `header`.

```js
assert.exists(document.querySelector('header'));
```

عنصر `header` الخاص بك يجب أن يكون داخل عنصر `.heading` الخاص بك.

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

عنصر `header` الخاص بك يجب أن يحتوي على `class` يسمي `hero`.

```js
assert(document.querySelector('header')?.className === 'hero');
```

يجب إنشاء عنصر `img`.

```js
assert.exists(document.querySelector('img'));
```

عنصر `img` الخاص بك يجب أن يكون داخل عنصر `header` الخاص بك.

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

يجب أن يكون لعنصر `img` الخاص بك `src` بقيمة `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

يجب أن يكون لعنصر `img` الخاص بك السمة `alt` بقيمة `freecodecamp logo`.

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

يجب أن يحتوي عنصر `img` على السمة `loading` بقيمة `lazy`.

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

يجب أن يحتوي عنصر `img` الخاص بك على `class` بقيمة `hero-img`.

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
