---
id: 6165d3b702a5d92ad970b30c
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

بعد عنصر `img` الخاص بك، قم بإضافة عنصر `h1` مع `class` تم تعيينه إلى `hero-title` والنص المحدد إلى `OUR NEW CURRICULUM`، يتبعه عنصر `p` مع `class` تم تعيينه إلى `hero-subtitle` والنص المحدد إلى `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

يجب إنشاء عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

عنصر `h1` الخاص بك يجب أن يأتي بعد عنصر `img` الخاص بك.

```js
assert(document.querySelector('h1')?.previousElementSibling?.localName === 'img');
```

يجب أن يحتوي عنصر `h1` الخاص بك على `class` بقيمة `hero-title`.

```js
assert(document.querySelector('h1')?.className === 'hero-title');
```

يجب أن يحتوي العنصر `h1` الخاص بك علي النص `OUR NEW CURRICULUM`.

```js
assert(document.querySelector('h1')?.textContent === 'OUR NEW CURRICULUM');
```

يجب عليك إنشاء عنصر `p` جديد.

```js
assert.exists(document.querySelector('p'));
```

عنصر `p` الخاص بك يجب أن يأتي بعد عنصر `h1` الخاص بك.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

يجب أن يحتوي عنصر `p` الخاص بك على `class` بقيمة `hero-subtitle`.

```js
assert(document.querySelector('p')?.className === 'hero-subtitle');
```

عنصر `p` الخاص بك يجب أن يحتوي علي النص `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
