---
id: 6557709b0aee699a6a00528c
title: Hatua ya 12
challengeType: 20
dashedName: step-12
---

# --description--

Now, replace the existent keys with the strings `A` and `B` — one for each node. Then, replace each value with the string representing the node connected to the key.

# --hints--

Your dictionary should have an `A` key.

```js
({ test: () => assert(__pyodide.runPython(`
    "A" in my_graph
  `))
})
```

Your `A` key should have `B` as the value.

```js
({ test: () => assert(__pyodide.runPython(`
    my_graph["A"] == "B"
  `))
})
```

Your dictionary should have an `B` key.

```js
({ test: () => assert(__pyodide.runPython(`
    "B" in my_graph
  `))
})
```

Your `B` key should have `A` as the value.

```js
({ test: () => assert(__pyodide.runPython(`
    my_graph["B"] == "A"
  `))
})
```

Your dictionary should have two keys.

```js
({ test: () => assert(__pyodide.runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
