---
id: 5a23c84252665b21eecc8028
title: Stern-Brocot 序列
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

For this task, the Stern-Brocot sequence is to be generated by an algorithm similar to that employed in generating the <a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">Fibonacci sequence</a>.

<ol>
  <li>The first and second members of the sequence are both 1:</li>
    <ul><li>1, 1</li></ul>
  <li>首先考慮序列的第二個成員</li>
  <li>將序列的考慮成員與其先例相加，(1 + 1) = 2，並將其附加到序列的末尾：</li>
    <ul><li>1, 1, 2</li></ul>
  <li>將序列的考慮成員附加到序列的末尾：</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>考慮該系列的下一個成員（第三個成員，即 2）</li>
  <li>GOTO 3 </li>
    <ul>
      <li></li>
      <li> ──── 展開另一個循環，我們得到：────</li>
      <li></li>
    </ul>
  <li>將序列的考慮成員與其先例相加，(2 + 1) = 3，並將其附加到序列的末尾：</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>將序列的考慮成員附加到序列的末尾：</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>考慮該系列的下一個成員（第四個成員，即 1）</li>
</ol>

# --instructions--

創建一個函數，該函數返回 Stern-Brocot 序列中第一次遇到 $ n $ 的位置，該序列是使用上述方法生成的。 請注意，此序列使用基於 1 的索引。

# --hints--

`sternBrocot` 應該是一個函數。

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)` 應該返回一個數字。

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)` 應該返回 `3`。

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)` 應該返回 `5`。

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)` 應該返回 `11`。

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)` 應該返回 `19`。

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)` 應該返回 `39`。

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
