---
id: bad87fee1348bd9aed708826
title: Видаляйте елемент за допомогою jQuery
challengeType: 6
forumTopicId: 18262
dashedName: remove-an-element-using-jquery
---

# --description--

Видалимо елемент у форматі HTML із вашої сторінки за допомогою jQuery.

jQuery має функцію під назвою `.remove()`, яка повністю видалить HTML-елемент

Видаліть зі сторінки елемент `#target4` за допомогою функції `.remove()`.

# --hints--

Ви можете використовувати jQuery, щоб видалити елемент `target4` з вашої сторінки.

```js
assert(
  $('#target4').length === 0 && code.match(/\$\(["']#target4["']\).remove\(\)/g)
);
```

Використовуйте лише jQuery, щоб видалити цей елемент.

```js
assert(
  code.match(/id="target4/g) &&
    !code.match(/<!--.*id="target4".*-->/g) &&
    $('#right-well').length > 0
);
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("#target1").css("color", "red");
    $("#target1").prop("disabled", true);
    $("#target4").remove();
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
