---
id: 64e750d9e426ffe0a7c40dfa
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Asynchronous programming is a programming paradigm that allows tasks to run concurrently, without blocking the execution of other tasks.

In this challenge you would learn how concurrency works by building a prime number generator.

To start, import the `asyncio` library, which provides tools for working with asynchronous code.

# --hints--

You should use the `import` keyword to import the `asyncio` library.

```js
({ test: () =>  assert.equal(__userGlobals.get('import'), '{}') })
```

You should use the `asyncio` keyword.

```js
({ test: () =>  assert.equal(__userGlobals.get('asyncio'), '{}') })
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
