---
id: 64e750d9e426ffe0a7c40dfa
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

Asynchronous programming is a programming paradigm that allows tasks to run concurrently, without blocking the execution of other tasks.

In this challenge you would learn how concurrency works by building a prime number generator.

To start, import the `asyncio` library, which provides tools for working with asynchronous code.

# --hints--

You should use the `import` keyword followed by `asyncio` to import the `asyncio` library.

```js
assert.match(code, /import\s+asyncio/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
