---
id: 64f078a0ae487906f447a3a4
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Define an asynchronous function named `find_primes` that takes a start and an end value as input to find prime numbers in the given range.The syntax of defining async functions is `async def function_name(arg1, ...)`

# --hints--

Your function should be named `find_primes`.

```js
assert.match(code, /async\s+def\s+find_primes\s*\(\s*\)/);
```

Use the async keyword before defining the function.

```js
assert.match(code, /async\s+def\s+find_primes\s*\(\s*\)/);
```

`find_primes` should have two parameters `start` and `end`.

```js   
assert.match(code, /async\s+def\s+find_primes\s*\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--
import asyncio

```py
--fcc-editable-region--

--fcc-editable-region--
```

