---
id: 660ef31a5be625914a0102cd
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

When a variable is declared with the `let` keyword, you can <dfn>reassign</dfn> (or change the value of) that variable later on. In this example, the value of `programmer` is changed from `"Naomi"` to `"CamperChan"`.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Note that when reassigning a variable, you do **not** use the `let` keyword again.

After your `console.log`, assign the value `"World"` to your `character` variable.

# --hints--

You should use `character` twice in your code.

```js
assert.lengthOf(code.match(/character/g), 3);
```

You should use the assignment operator to reassign `character`.

```js
assert.lengthOf(code.match(/character\s*\=/g), 2);
```

You should use the string `"World"` in your code.

```js
assert.match(code, /("|')World\1/);
```

Your `character` variable should have the value `"World"`.

```js
assert.equal(character, "World");
```

Your reassignment should not use `let`.

```js
assert.isBelow(code.match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
