---
id: 660f0f980e98e8cf77f1ce31
title: Step 29
challengeType: 1
dashedName: step-29
---

# --description--

Change your `rows` declaration to be assigned an empty array again.

Also, change your `'Hello'` string to use double quotes again. Generally, it does not matter which of the two you prefer, but you will want to be consistent in that choice throughout your project.

# --hints--

Your `rows` array should be empty.

```js
assert.empty(rows);
```

Your `"Hello"` string should use double quotes.

```js
assert.match(code, /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
