---
id: 6610b8f6a98d25f4d485a94d
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

The default value of an uninitialized variable is `undefined`. This is a special data type that represents a value that does not have a definition yet.

You can still assign a value to an uninitialized variable. Here is an example:

```js
let uninitialized;
uninitialized = "assigned";
```

Assign the string `"Test"` to your `secondCharacter` variable below your declaration. Open the console to see how your log has changed.

# --hints--

You should not initialize `secondCharacter`. Remember that initialization means assigning a value when you declare the variable.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

You should use the assignment operator on `secondCharacter`.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

You should assign the string `"Test"` to your `secondCharacter` variable.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
