---
id: 66db529d37ad966480ebb633
title: Build a Lunch Picker Program
challengeType: 14
dashedName: build-a-lunch-picker-program
---

# --description--

In this lab, you will create a program that will help you decide what to eat for lunch. You will have two arrays, one for kids and one for adults, with different food items. You will use the `push` method to add items to the arrays and the `pop` and `shift` methods to remove items from the arrays. You will also use the `Math.random()` method to generate a random number that will be used to select a random meal from the arrays.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should create a `kidsMeal` array with one item, `"Chicken Nuggets"`, in it.
1. You should create an `adultsMeal` array with one item, `"Salad"`, in it.
1. You should use `push` to add three more food items being `"Pizza"`, `"Burgers"`, and `"Tacos"` to the `kidsMeal` array.
1. You should use `push` to add three more food items being `"Wraps"`, `"Hummus"`, and `"Rice"` to the `adultsMeal` array.
1. You should log the `kidsMeal` array to the console using the following format:
   
- `"Kids Menu: ", kidsMeal`.

1. You should log the `adultsMeal` array to the console using the following format:

- `"Adults Menu: ", adultsMeal`.

1. You should remove the last item from the `kidsMeal` array using `pop()`.
1. You should remove the first item from the `adultsMeal` array using `shift()`.
1. You should log the updated `kidsMeal` and `adultsMeal` array to the console using this format:

- `"Kids menu after removing the last item: ", kidsMeal`.
- `"Adults menu after removing the first item: ", adultsMeal`.

1.  Generate a random number that will be used to select a random meal from the `kidsMeal` array and store it in a variable called `randomIndexKids`.
1.  Generate a random number that will be used to select a random meal from the `adultsMeal` array and store it in a variable called `randomIndexAdults`.
1.  Log the following messages to the console:

- `"Today kids should eat:"` followed by the selected meal to the console.
- `"Today adults should eat:"` followed by the selected meal to the console.

# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let kidsMeal = ["Chicken Nuggets"];

let adultsMeal = ["Salad"];

kidsMeal.push("Pizza", "Burgers", "Tacos");

adultsMeal.push("Wraps", "Hummus", "Rice");

console.log("Kids Menu:", kidsMeal);
console.log("Adults Menu:", adultsMeal);

kidsMeal.pop();

adultsMeal.shift();

console.log("Kids menu after removing the last item: ", kidsMeal);
console.log("Adults Menu After Removing the first item:", adultsMeal);

let randomIndexKids = Math.floor(Math.random() * kidsMeal.length);
let randomIndexAdults = Math.floor(Math.random() * adultsMeal.length);

console.log("Today kids should eat:", kidsMeal[randomIndexKids]);
console.log("Today adults should eat:", adultsMeal[randomIndexAdults]);
```
