---
id: 66db529d37ad966480ebb633
title: Build a Lunch Picker Program
challengeType: 25
dashedName: build-a-lunch-picker-program
---

# --description--

In this lab, you'll build a program that helps in managing lunch options. You'll work with an array of lunches, add and remove items from the array, and randomly select a lunch option.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a mutable `lunches` variable and assign it an empty array that will be used to store lunch items.

2. You should create a function `addLunchToEnd` that takes an array as the first parameter and a string as the second parameter. The function should:
    - add the string to the end of the array
    - log the string `"[Lunch Item] added to the end of the lunch menu."`  to the console
    - return the updated array

3. You should create a function `addLunchToStart` that takes an array as the first parameter and a string as the second parameter. The function should:
    - add the string to the start of the array
    - log the string `"[Lunch Item] added to the start of the lunch menu."` to the console
    - return the updated array

4. You should create a function `removeLastLunch` that takes an array parameter and:
   - removes the last lunch item from the array
   - If successful, log the string `"[Lunch Item] removed from the end of the lunch menu."` to the console
   - If the array is empty, log the string `"No lunches to remove."` to the console
   - return the updated array

5. You should create a function `removeFirstLunch` that takes an array parameter and:
   - removes the first lunch item from the array
   - If successful, log the string `"[Lunch Item] removed from the start of the lunch menu."` to the console
   - If the array is empty, log the string `"No lunches to remove."` to the console
   - return the updated array

6. You should create a function `getRandomLunch` that takes an array parameter and:
   - selects a random lunch item from the array
   - If successful, log the string `"Randomly selected lunch: [Lunch Item]"` to the console
   - If the array is empty, log the string `"No lunches available."` to the console

7. You should create a function `showLunchMenu` that takes an array parameter and:
   - If there are lunch items in the array, log the string `"Menu items: [Lunch Item], [Lunch Item]...` to the console
   - If the array is empty, log the string `"The menu is empty."` to the console

# --hints--

You should declare a mutable variable `lunches` and assign it an empty array to store the lunch items.

```js
assert.isArray(lunches);
```

You should define a function `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

The function `addLunchToEnd` should take 2 arguments.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` should log the string `"Tacos added to the end of the lunch menu."`.

```js
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(lunches, "Pizza");
  assert.strictEqual(tempArr[0], "Pizza added to the end of the lunch menu.");
}finally{
  console.log = temp;
}
```

The function `addLunchToEnd` should return the correct array.

```js
lunches = ["Fries"]
assert.deepStrictEqual(addLunchToEnd(lunches, "Tacos"), ["Fries", "Tacos"]);
```

You should define a function `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

The function `addLunchToStart` should take two arguments.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` should log the string `"Sushi added to the start of the lunch menu."`.

```js
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  addLunchToStart(lunches, "Burger");
  assert.strictEqual(tempArr[0], "Burger added to the start of the lunch menu.");
}finally{
  console.log = temp;
}
```

The function `addLunchToStart` should return the correct array.

```js
lunches = ["Noodles"]
assert.deepStrictEqual(addLunchToStart(lunches, "Sushi"), ["Sushi", "Noodles"]);
```

You should define a function `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

The function `removeLastLunch` should take one argument.

```js
assert.lengthOf(removeLastLunch, 1);
```

When the input array is empty, the function `removeLastLunch` should log the string `"No lunches to remove."`.

```js
lunches = [];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  removeLastLunch(lunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
}finally{
  console.log = temp;
}
```

When `lunches = ["Stew", "Soup", "Toast"]`, `removeLastLunch(lunches)` should log the string `"Toast removed from the end of the lunch menu."`.

```js
lunches = ["Fries", "Tacos", "Noodles"];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  removeLastLunch(lunches);
  assert.strictEqual(tempArr[0], "Noodles removed from the end of the lunch menu.");
}finally{
  console.log = temp;
}
```

The function `removeLastLunch` should return the correct array.

```js
lunches = ["Pizza", "Burger"];
assert.deepStrictEqual(removeLastLunch(lunches), ["Pizza"]);
```

You should define a function `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch)
```

The function `removeFirstLunch` should take one argument.

```js
assert.lengthOf(removeFirstLunch, 1);
```

When the input array is empty, the function `removeFirstLunch` should log the string `"No lunches to remove."`.

```js
lunches = [];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(lunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
}finally{
  console.log = temp;
}
```

When `lunches = ["Salad", "Eggs", "Cheese"]`, `removeFirstLunch(lunches)` should log the string `"Salad removed from the start of the lunch menu."`.

```js
lunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(lunches);
  assert.strictEqual(tempArr[0], "Stew removed from the start of the lunch menu.");
}finally{
  console.log = temp;
}
```

The function `removeFirstLunch` should return the correct array.

```js
lunches = ["Pizza", "Burger"];
assert.deepStrictEqual(removeFirstLunch(lunches), ["Burger"]);
```

You should define a function `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

The function `getRandomLunch` should take one argument.

```js
assert.lengthOf(getRandomLunch, 1);
```

When the input array is empty, the function `getRandomLunch` should log the string `"No lunches available."`.

```js
lunches = [];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  getRandomLunch(lunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
}finally{
  console.log = temp;
}
```

When `lunches = ["Fish", "Fries", "Roast"]`, the function `getRandomLunch` should log a string in the format `"Randomly selected lunch: [Lunch Item]"`.

```js
lunches = ["Salad", "Eggs", "Cheese"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try{
  console.log = obj => tempArr.push(obj);

  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch(lunches);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${lunches[0]}`);

  // second item
  Math.random = () => 0.4;
  getRandomLunch(lunches);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${lunches[1]}`);

  // third item
  Math.random = () => 0.8;
  getRandomLunch(lunches);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${lunches[2]}`);

}finally{
  // restore Math.random
  Math.random = tempRandom;
  console.log = temp;
}
```

You should define a function `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

The function `showLunchMenu` should take one argument.

```js
assert.lengthOf(showLunchMenu, 1);
```

When the input array is empty, the function `showLunchMenu` should log the string `"The menu is empty."`.

```js
lunches = [];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  showLunchMenu(lunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
}finally{
  console.log = temp;
}
```

When `lunches = ["Greens", "Corns", "Beans"]`, the function `showLunchMenu` should log a string in the format `"Menu items: Greens, Corns, Beans"`.

```js
lunches = ["Pizza", "Burger", "Fries"];
const tempArr = [];
const temp = console.log;
try{
  console.log = obj => tempArr.push(obj);
  showLunchMenu(lunches);
  assert.strictEqual(tempArr[0], "Menu items: Pizza, Burger, Fries");
}finally{
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

lunches = addLunchToEnd(lunches, "Tacos");
lunches = addLunchToStart(lunches, "Sushi");
lunches = removeLastLunch(lunches);
lunches = removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
