---
id: 66db529d37ad966480ebb633
title: Build a Lunch Picker Program
challengeType: 14
dashedName: build-a-lunch-picker-program
---

# --description--

In this lab, you'll build a program that helps in managing lunch options. You'll work with an array of lunches, add and remove items from the array, and randomly select a lunch option.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a `lunches` variable and assign it an empty array that will be used to store lunch items.

2. You should create a function `addLunchToEnd` that takes a string parameter and adds it to the end of the `lunches` array and returns a string in the format:
     
   - `"[Lunch Item] added to the end of the lunch menu."`

3. You should create a function `addLunchToStart` that takes a string parameter and adds it to the start of the `lunches` array and returns a string in the format:  
   
   - `"[Lunch Item] added to the start of the lunch menu."`

4. You should create a function `removeLastLunch` that removes the last lunch item from the `lunches` array and:
   
   - If successful, returns a string in the format:  
     - `"[Lunch Item] removed from the lunch menu (last)."`
   - If the `lunches` array is empty, returns a string:  
     - `"No lunches to remove."`

5. You should create a function `removeFirstLunch` that removes the first lunch item from the `lunches` array and:
   
   - If successful, returns a string in the format:  
     - `"[Lunch Item] removed from the lunch menu (first)."`
   - If the `lunches` array is empty, returns a string:  
     - `"No lunches to remove."`

6. You should create a function `getRandomLunch` that selects a random lunch item from the `lunches` array and:
   
   - If successful, returns a string in the format:  
     - `"Randomly selected lunch: [Lunch Item]"`
   - If the `lunches` array is empty, returns a string:  
     - `"No lunches available."`

7. You should create a function `showLunchMenu` that:
   
   - If there are items in the `lunches` array, returns a string in the format:  
     - `"1. <Lunch Item>\n2. <Lunch Item>\n..."` (with a numbered list for all the lunch items in the array)
   - If the `lunches` array is empty, returns a string:  
     - `"The menu is empty."`

# --hints--

You should declare a `kidsMeal` array.

```js
assert.isArray(kidsMeal);
```

You should declare an `adultsMeal` array.

```js
assert.isArray(adultsMeal);
```

Your `kidsMeal` array should contain at least three items.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+kidsMeal\s*=\s*\[(?<items>.+?)\]/)
const kidsMealArray = matched.groups.items.split(",");
assert.isAtLeast(kidsMealArray.length, 3);
```

Your `kidsMeal` array should contain only non-empty string values.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+kidsMeal\s*=\s*\[(?<items>.+?)\]/)
const kidsMealArray = matched.groups.items.split(",");
assert.isAtLeast(kidsMealArray.length, 1);

for (let i = 0; i < kidsMealArray.length; i++) {
  const itemMatch = kidsMealArray[i].trim().match(/^(?<quote>'|"|`)(?<meal>.+)\1$/);
  assert.isNotNull(itemMatch);
  assert.notInclude(itemMatch.groups.meal, itemMatch.groups.quote)
}
```

Your `adultsMeal` array should contain at least three items.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+adultsMeal\s*=\s*\[(?<items>.+?)\]/)
const adultsMealArray = matched.groups.items.split(",");
assert.isAtLeast(adultsMealArray.length, 3);
```

Your `adultsMeal` array should contain only non-empty string values.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+adultsMeal\s*=\s*\[(?<items>.+?)\]/)
const adultsMealArray = matched.groups.items.split(",");
assert.isAtLeast(adultsMealArray.length, 1);


for (let i = 0; i < adultsMealArray.length; i++) {
  const itemMatch = adultsMealArray[i].trim().match(/^(?<quote>'|"|`)(?<meal>.+)\1$/);
  assert.isNotNull(itemMatch);
  assert.notInclude(itemMatch.groups.meal, itemMatch.groups.quote)
}
```

You should remove the last item from the `kidsMeal` array using `pop()`.

```js
assert.match(__helpers.removeJSComments(code), /kidsMeal\.pop/);
```

You should remove the first item from the `adultsMeal` array using `shift()`.

```js
assert.match(__helpers.removeJSComments(code), /adultsMeal\.shift/);
```

You should use `Math.random()` to generate a random number from `0` upto the length of the `kidsMeal` array that will be used to select a random meal from the `kidsMeal` array and store it in a variable called `randomIndexKids`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+randomIndexKids\s*=\s*Math\s*\.\s*floor\s*\(\s*Math\s*\.\s*random\s*\(\s*\)\s*\*\s*kidsMeal\s*\.\s*length\s*\)\s*?;?/);
```

You should use `Math.random()` to generate a random number from `0` upto the length of the `adultsMeal` array that will be used to select a random meal from the `adultsMeal` array and store it in a variable called `randomIndexAdults`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+randomIndexAdults\s*=\s*Math\s*\.\s*floor\s*\(\s*Math\s*\.\s*random\s*\(\s*\)\s*\*\s*adultsMeal\s*\.\s*length\s*\)\s*?;?/);
```

You should output `"Today kids should eat:", [kids' selected meal] ` to the console.

```js
const condition1 = /console\s*\.\s*log\s*\(\s*("|'|`)Today kids should eat:\1\s*,\s*kidsMeal\s*\[\s*randomIndexKids\s*\]\s*\)\s*?;?/.test(__helpers.removeJSComments(code));
const condition2 = /console\s*\.\s*log\s*\(\s*`Today kids should eat: \$\{\s*kidsMeal\s*\[\s*randomIndexKids\s*\]\s*}`\s*\)\s*?;?/.test(__helpers.removeJSComments(code));
assert.isTrue(condition1 || condition2);
```

You should output `"Today adults should eat:", [adult's selected meal] ` to the console.

```js
const condition1 = /console\s*\.\s*log\s*\(\s*("|'|`)Today adults should eat:\1\s*,\s*adultsMeal\s*\[\s*randomIndexAdults\s*\]\s*\)\s*?;?/.test(__helpers.removeJSComments(code));
const condition2 = /console\s*\.\s*log\s*\(\s*`Today adults should eat: \$\{\s*adultsMeal\s*\[\s*randomIndexAdults\s*\]\s*}`\s*\)\s*?;?/.test(__helpers.removeJSComments(code));
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the list (at the end)
function addLunchToEnd(newLunch) {
  lunches.push(newLunch);
  return `${newLunch} added to the end of the lunch menu.`;
}

// Add a new lunch to the start of the list
function addLunchToStart(newLunch) {
  lunches.unshift(newLunch);
  return `${newLunch} added to the start of the lunch menu.`;
}

// Remove the last lunch from the list
function removeLastLunch() {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    return `${removedLunch} removed from the lunch menu (last).`;
  } else {
    return "No lunches to remove.";
  }
}

// Remove the first lunch from the list
function removeFirstLunch() {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    return `${removedLunch} removed from the lunch menu (first).`;
  } else {
    return "No lunches to remove.";
  }
}

// Function to get a random lunch from the list
function getRandomLunch() {
  if (lunches.length === 0) {
    return "No lunches available.";
  }
  const randomIndex = Math.floor(Math.random() * lunches.length);
  const randomLunch = lunches[randomIndex];
  return `Randomly selected lunch: ${randomLunch}`;
}

// New function to display the lunches as a numbered menu using a for loop
function showLunchMenu() {
  if (lunches.length === 0) {
    return "The menu is empty.";
  }
  let menu = "Lunch Menu:\n";
  for (let i = 0; i < lunches.length; i++) {
    menu += `${i + 1}. ${lunches[i]}\n`;
  }
  return menu.trim();
}

addLunchToEnd("Tacos");
addLunchToStart("Sushi");
removeLastLunch();
removeFirstLunch();
getRandomLunch();
showLunchMenu();

```
