---
id: 66db529d37ad966480ebb633
title: Build a Lunch Picker Program
challengeType: 14
dashedName: build-a-lunch-picker-program
---

# --description--

In this lab, you'll build a program to assist in choosing lunch options. You'll work with two arrays: one for children's food choices and one for adults. You'll add and remove items from these arrays and then randomly select a meal from each. Finally, you'll log the results to the console.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should create a `kidsMeal` variable and assign it an array with at least three strings in it.
2. You should create an `adultsMeal` variable and assign it an array with at least three strings in it.
3. Once items have been added to the `kidsMeal`, you should use the appropriate array method to remove the last item.
4. Once items have been added to the `adultsMeal`, you should use the appropriate array method to remove the first item. 
5.  You should use the appropriate `Math` methods to generate a random number between `0` to the length of the `kidsMeal` array. Assign that result to a variable called `randomIndexKids`.
6. You should use the appropriate `Math` methods to generate a random number between `0` to the length of the `kidsMeal` array. Assign that result to a variable called `randomIndexAdults`.
7.  Log the following messages to the console:

- `"Today kids should eat:"` followed by the selected meal to the console.
- `"Today adults should eat:"` followed by the selected meal to the console.

# --hints--

You should declare a `kidsMeal` array.

```js
assert.isArray(kidsMeal);
```

You should declare an `adultsMeal` array.

```js
assert.isArray(adultsMeal);
```

Your `kidsMeal` array should contain at least three items.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+kidsMeal\s*=\s*\[(?<items>.+?)\]/)
const kidsMealArray = matched.groups.items.split(",");
assert.isAtLeast(kidsMealArray.length, 3);
```

Your `kidsMeal` array should contain only non-empty string values.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+kidsMeal\s*=\s*\[(?<items>.+?)\]/)
const kidsMealArray = matched.groups.items.split(",");
assert.isAtLeast(kidsMealArray.length, 1);

for (let i = 0; i < kidsMealArray.length; i++) {
  const itemMatch = kidsMealArray[i].trim().match(/^(?<quote>'|"|`)(?<meal>.+)\1$/);
  assert.isNotNull(itemMatch);
  assert.notInclude(itemMatch.groups.meal, itemMatch.groups.quote)
}
```

Your `adultsMeal` array should contain at least three items.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+adultsMeal\s*=\s*\[(?<items>.+?)\]/)
const adultsMealArray = matched.groups.items.split(",");
assert.isAtLeast(adultsMealArray.length, 3);
```

Your `adultsMeal` array should contain only non-empty string values.

```js
const matched = __helpers.removeJSComments(code).match(/(const|let)\s+adultsMeal\s*=\s*\[(?<items>.+?)\]/)
const adultsMealArray = matched.groups.items.split(",");
assert.isAtLeast(adultsMealArray.length, 1);


for (let i = 0; i < adultsMealArray.length; i++) {
  const itemMatch = adultsMealArray[i].trim().match(/^(?<quote>'|"|`)(?<meal>.+)\1$/);
  assert.isNotNull(itemMatch);
  assert.notInclude(itemMatch.groups.meal, itemMatch.groups.quote)
}
```

You should remove the last item from the `kidsMeal` array using `pop()`.

```js
assert.match(__helpers.removeJSComments(code), /kidsMeal\.pop/);
```

You should remove the first item from the `adultsMeal` array using `shift()`.

```js
assert.match(__helpers.removeJSComments(code), /adultsMeal\.shift/);
```

You should use `Math.random()` to generate a random number from `0` upto the length of the `kidsMeal` array that will be used to select a random meal from the `kidsMeal` array and store it in a variable called `randomIndexKids`.

```js
// Check that Math.random() is used
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);

// Ensure that randomIndexKids is not null
assert.isNotNull(randomIndexKids);

// Ensure that randomIndexKids is at least 0
assert.isAtLeast(randomIndexKids, 0);

// Ensure that randomIndexKids is at most kidsMeal.length - 1 (not hardcoded)
assert.match(__helpers.removeJSComments(code), /Math\.floor\(Math\.random\(\) \* kidsMeal\.length\)/);

// Ensure that randomIndexKids does not exceed kidsMeal.length dynamically
assert.isAtMost(randomIndexKids, kidsMeal.length - 1);
```

You should use `Math.random()` to generate a random number from `0` upto the length of the `adultsMeal` array that will be used to select a random meal from the `adultsMeal` array and store it in a variable called `randomIndexAdults`.

```js
// Check that Math.random() is used
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);

// Ensure that randomIndexAdults is not null
assert.isNotNull(randomIndexAdults);

// Ensure that randomIndexAdults is at least 0
assert.isAtLeast(randomIndexAdults, 0);

// Ensure that randomIndexAdults is at most adultsMeal.length - 1 (not hardcoded)
assert.match(__helpers.removeJSComments(code), /Math\.floor\(Math\.random\(\) \* adultsMeal\.length\)/);

// Ensure that randomIndexAdults does not exceed adultsMeal.length dynamically
assert.isAtMost(randomIndexAdults, adultsMeal.length - 1);
```

You should output `"Today kids should eat:", [kids' selected meal] ` to the console.

```js
const condition1 = /console\.log\(\s*["']Today\s+kids\s+should\s+eat:["'],\s+kidsMeal\[randomIndexKids\]\);?/gm.test(__helpers.removeJSComments(code));
const condition2 = /console\.log\(\s*`Today\s+kids\s+should\s+eat:\s*\${kidsMeal\[randomIndexKids\]}\s*`\);?/gm.test(__helpers.removeJSComments(code));
assert.isTrue(condition1 || condition2);

```

You should output `"Today adults should eat:", [adult's selected meal] ` to the console.

```js
const condition1 = /console\.log\(\s*["']Today\s+adults\s+should\s+eat:["'],\s+adultsMeal\[randomIndexAdults\]\);?/gm.test(__helpers.removeJSComments(code));
const condition2 = /console\.log\(\s*`Today\s+adults\s+should\s+eat:\s*\${adultsMeal\[randomIndexAdults\]}\s*`\);?/gm.test(__helpers.removeJSComments(code));
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let kidsMeal = ["Chicken Nuggets","Pizza", "Burgers", "Tacos"];

let adultsMeal = ["Salad", "Wraps", "Hummus", "Rice"];

kidsMeal.pop();

adultsMeal.shift();


let randomIndexKids = Math.floor(Math.random() * kidsMeal.length);
let randomIndexAdults = Math.floor(Math.random() * adultsMeal.length);

console.log("Today kids should eat:", kidsMeal[randomIndexKids]);

console.log("Today adults should eat:", adultsMeal[randomIndexAdults]);

```
