---
id: 6733b5e6f210c1bd1afdf8b3
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

The remaining plant objects have been added for you. Now it's time to create your catalog.

Like regular JavaScript objects, `Map` objects contain a collection of key-value pairs. To create a `Map` object, you can use the `Map` constructor:

```js
const myMap = new Map();
```

Declare a `catalog` variable and assign it a `Map` object.

# --hints--

You should have a variable named `catalog`.

```js
assert.exists(catalog);
```

You should assign a `Map` object to your `catalog` variable.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
