---
id: 6734a2385c52c54d916a8b86
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Add another key-value pair to your catalog. Use your `prettyPolly` variable as the key and make the value an object with `small`, `medium`, and `large` properties whose values are integers again.

# --hints--

You should set your `prettyPolly` object as a key of your `catalog` map.

```js
assert.isTrue(catalog.has(prettyPolly));
```

The value of the `prettyPolly` key in the `catalog` map should be an object with the properties `small`, `medium`, and `large`.

```js
const b = catalog.get(prettyPolly)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

The `small` property of the object associated with the `prettyPolly` key in the `catalog` map should be a positive integer.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

The `medium` property of the object associated with the `prettyPolly` key in the `catalog` map should be a positive integer.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

The `large` property of the object associated with the `prettyPolly` key in the `catalog` map should be a positive integer.

```js
const b = catalog.get(prettyPolly)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

You should still have the `ballerina` key in your `catalog` map.

```js
assert.isTrue(catalog.has(ballerina));
```

Your `catalog` map should contain two elements.

```js
assert.lengthOf(catalog, 2);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
