---
id: 6734a606e5968a5a19b77171
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Keep adding the remaining plants to your catalog as you did in the previous steps.

# --hints--

You should set your `willowVale` object as a key of your `catalog` map.

```js
assert.isTrue(catalog.has(willowVale));
```

The value of the `willowVale` key in the `catalog` map should be an object with the properties `small`, `medium`, and `large`, each with the value of a positive integer.

```js
const b = catalog.get(willowVale)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

You should set your `hidcote` object as a key of your `catalog` map.

```js
assert.isTrue(catalog.has(hidcote));
```

The value of the `hidcote` key in the `catalog` map should be an object with the properties `small`, `medium`, and `large`, each with the value of a positive integer.

```js
const b = catalog.get(hidcote)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

You should set your `imperialGem` object as a key of your `catalog` map.

```js
assert.isTrue(catalog.has(imperialGem));
```

The value of the `imperialGem` key in the `catalog` map should be an object with the properties `small`, `medium`, and `large`, each with the value of a positive integer.

```js
const b = catalog.get(imperialGem)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

You should set your `royalCrown` object as a key of your `catalog` map.

```js
assert.isTrue(catalog.has(royalCrown));
```

The value of the `royalCrown` key in the `catalog` map should be an object with the properties `small`, `medium`, and `large`, each with the value of a positive integer.

```js
const b = catalog.get(royalCrown)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

You should still have the `ballerina` key in your `catalog` map.

```js
assert.isTrue(catalog.has(ballerina));
```

You should still have the `prettyPolly` key in your `catalog` map.

```js
assert.isTrue(catalog.has(prettyPolly));
```

Your `catalog` map should contain six elements.

```js
assert.lengthOf(catalog, 6);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set(ballerina, { small: 20, medium: 15, large: 12 });
catalog.set(prettyPolly, { small: 31, medium: 14, large: 24 });

--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
