---
id: 5dc174fcf86c76b9248c6eb2
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Los elementos HTML tienen etiquetas de apertura como `<h1>` y etiquetas de cierre como `</h1>`.

El texto de un elemento va entre sus etiquetas de apertura y cierre.

Encuentra el elemento `h1` y cambia su texto a:

`CatPhotoApp`

# --hints--

El texto `CatPhotoApp` debe estar presente en el código. Tal vez deberías revisar tu ortografía.

```js
assert(code.match(/catphotoapp/i));
```

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/h1\>/));
```

Tienes más de un elemento `h1`. Elimina el elemento `h1` extra.

```js
assert(document.querySelectorAll('h1').length === 1);
```

El texto de tu elemento `h1` deber ser `CatPhotoApp`. Probablemente olvidaste añadir el texto, tienes un error tipográfico o no está entre las etiquetas de apertura y cierre del elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Parece que está utilizando una extensión del navegador que está modificando la página. Asegúrese de desactivar todas las extensiones del navegador.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

