---
id: 5dc17d3bf86c76b9248c6eb4
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

El elemento `p` es usado para crear un párrafo de texto en sitios webs. Crea un elemento `p` debajo de tu elemento `h2` y dale el siguiente texto:

`See more cat photos in our gallery.`

# --hints--

El elemento `p` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('p'));
```

El elemento `p` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen una `/` después del carácter `<`.

```js
assert(code.match(/<\/p\>/));
```

El texto de su elemento `p` debe ser `See more cat photos in our gallery.` Probablemente no has añadido el texto o tienes un error tipográfico.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

El elemento `p` debe estar debajo del elemento `h2`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

