---
id: 5dc17dc8f86c76b9248c6eb5
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Los comentarios te permiten dejar mensajes sin afectar lo que se muestra en el navegador. Te permite crear código que no se ejecutará. Un comentario en HTML comienza con `<!--`, contiene tantas líneas de texto como quieras, y termina con `-->`. Por ejemplo, el comentario`<!-- TODO: Remove h1 -->` contiene el texto `TODO: Remove h1`.

Añade un comentario sobre el elemento `p` con este texto:

`TODO: Add link to cat photos`

# --hints--

Tu comentario debe iniciar con `<!--`. Faltan uno o más de los caracteres que definen el inicio de un comentario.

```js
assert(code.match(/<!--/));
```

Tu comentario debe terminar con `-->`.  Faltan uno o más de los caracteres que definen el final de un comentario.

```js
assert(code.match(/-->/));
```

Tu código no debe tener caracteres extra de apertura/cierre para crear comentarios. Tienes los caracteres `<!--` o `-->` extra mostrándose en el navegador.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Tu comentario debe contener el texto `TODO: Add link to cat photos`.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

Tu comentario debe estar sobre el elemento `p`. Los tienes en el orden incorrecto.

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

