---
id: 5dc174fcf86c76b9248c6eb2
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Los elementos HTML tienen una etiqueta de apertura y una de cierre con contenido entre ellas.

Esta es la sintáxis básica:

```html
<openingTag>content</closingTag>
```

Primero aprenderás sobre el elemento `h1`. El elemento `h1` es un elemento de encabezado y se utiliza para indicar el encabezado principal de una página web.

```html
<h1>This is a main heading</h1>
```

Cambia el texto del elemento `h1` más abajo de `Hello World` a `CatPhotoApp` y observa el cambio en la vista previa del navegador.

Cuando hayas terminado, presiona el botón "Comprueba tu código" para ver si es correcto.

# --hints--

El texto `CatPhotoApp` debe estar presente en el código. Tal vez quieras revisar tu ortografía.

```js
assert(code.match(/catphotoapp/i));
```

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del caracter `<`.

```js
assert(code.match(/<\/h1\>/));
```

Tienes más de un elemento `h1`. Elimina el elemento `h1` extra.

```js
assert(document.querySelectorAll('h1').length === 1);
```

El texto de tu elemento `h1` debe ser `CatPhotoApp`. O has omitido el texto, tienes un error tipográfico, o no está entre las etiquetas de apertura y cierre del elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Parece que estás utilizando una extensión del navegador que está modificando la página. Asegúrate de desactivar todas las extensiones del navegador.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

