---
id: 5dc2385ff86c76b9248c6eb7
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 tiene diferentes elementos que ayudan a diferenciar diferentes tipos de contenido. Estos elementos hacen tu código HTML más fácil de leer y ayudan con el Posicionamiento en buscadores (Search Engine Optimization - SEO) y accesibilidad.

Identifica la sección principal de esta página añadiendo la etiqueta de apertura `<main>` después del elemento `h1` y una etiqueta de cierre `</main>` después del elemento `p`.

# --hints--

Tu elemento `main` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('main'));
```

Tu elemento `main` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen una `/` después del carácter `<`.

```js
assert(code.match(/<\/main\>/));
```

La etiqueta de apertura del elemento `main` debe estar debajo del elemento `h1`. Los tienes en el orden incorrecto.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling?.tagName, "H1");
```

La etiqueta de apertura del elemento `main` debe estar encima del elemento `h2`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('main,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H2'));
```

La etiqueta de cierre del elemento `main` debe estar debajo del elemento `p`. Los tienes en el orden incorrecto.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/click here to view more cat photos/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
--fcc-editable-region--
  </body>
</html>
```

