---
id: 5dc24165f86c76b9248c6ebc
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Todos los elementos `img` deben tener un atributo `alt`. El texto del atributo `alt` es utilizado por lectores de pantalla para mejora la accesibilidad y es mostrado en caso de que la imagen falle en cargar. Por ejemplo, `<img src="cat.jpg" alt="A cat">` tiene un atributo `alt` con el texto `A cat`.

Dentro del elemento `img`, añade un atributo `alt` con este texto:

`A cute orange cat lying on its back`

# --hints--

Tu código debe tener un elemento `img`. Eliminaste el elemento `img` de uno de los pasos anteriores.

```js
assert(document.querySelector('img'));
```

Tu elemento `img` no tiene un atributo `alt`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

El atributo `alt` del elemento `img` tiene un valor diferente a 'A cute orange cat lying on its back'. Asegúrate de que el valor del atributo `alt` esté entre comillas.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

