---
id: 5dc24614f86c76b9248c6ebd
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Puedes hacer un enlace que te llevará a otra página con el elemento anchor (`a`). Por ejemplo, `<a href='https://freecodecamp.org'></a>` te llevará a `freecodecamp.org`.

Agrega un elemento anchor después del párrafo que te lleve a `https://freecatphotoapp.com`. Por ahora, el enlace no aparecerá en la vista previa.

# --hints--

El elemento (`a`) debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('a'));
```

El elemento anchor (`a`) debe tener una etiqueta de cierre. Una etiqueta de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/a\>/));
```

Tu elemento anchor (`a`) debe estar debajo del elemento `p`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Tu elemento anchor (`a`) no tiene un atributo `href`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Tu elemento anchor (`a`) te debe de llevar a `https://freecatphotoapp.com`. Probablemente no has añadido la URL o tienes un error tipográfico.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Aunque le has dado el enlace correcto al parámetro `href` del elemento anchor ('a'), se recomienda siempre poner el valor de un atributo entre comillas.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

