---
id: 5dc23f9bf86c76b9248c6eba
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Puedes añadir imágenes a tu sitio web utilizando el elemento `img`. Los elementos `img` tienen una etiqueta de apertura sin una etiqueta de cierre. La etiqueta de un elemento que no necesita etiqueta de cierre se conoce como <dfn>etiqueta de autocierre</dfn> (self-closing tag).

Añade un elemento `img` debajo del elemento `p`. Por ahora no se mostrará ninguna imagen en el navegador.

# --hints--

Tu elemento `img` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Tu elemento `img` no debe tener una etiqueta de cierre. Una etiqueta de cierre tiene una `/` después del carácter `<`.

```js
assert(!code.match(/<\/img\>/));
```

Solo debes tener un elemento `img`. Elimina los que estén de más.

```js
assert(document.querySelectorAll('img').length === 1);
```

Tu elemento `img` debe estar debajo del elemento `p`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

