---
id: 5dc23991f86c76b9248c6eb8
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

En el paso anterior, anidaste un elemento `h2`, un comentario y un elemento `p` dentro de un elemento `main`. Un elemento anidado es hijo de su elemento padre. Debe tener una indentación o sangría de dos espacios más que su elemento padre para mejorar la legibilidad, de esta forma:

```html
<ul> 
  <li> Child Element 1 </li>
  <li> Child Element 2 </li>
</ul>
```

Añade dos espacios delante de los tres elementos hijos del elemento `main` para que tu código HTML sea más legible.

# --hints--

Tu código debe tener un elemento `h2` con el texto `Cat Photos`. Puede que lo hayas eliminado accidentalmente, faltan las etiquetas de apertura y de cierre o el texto ha sido modificado.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Tu elemento `h2` debe estar debajo de la etiqueta de apertura del elemento `main` y la etiqueta de apertura debe tener 6 espacios de separación desde inicio de la línea.

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s{6}<h2>/));
```

Tu código debe tener un comentario. Eliminaste el comentario de un paso anterior.

```js
assert(code.match(/<!--.*-->/));
```

El texto del comentario debe ser `TODO: Add link to cat photos`. No cambies el texto o los espacios del comentario.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Tu comentario debe estar debajo del elemento `h2` y tener 6 espacios desde el inicio de la línea.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Tu código debe tener un elemento `p`. Eliminaste el elemento `p` de un paso anterior.

```js
assert(document.querySelector('p'));
```

El texto del elemento `p` debe ser `Click here to view more cat photos.` No cambies el texto, el espaciado, o la puntuación del elemento `p`.

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

Tu elemento `p` debe estar debajo del comentario y su etiqueta de apertura debe tener 6 espacios desde el inicio de la línea.

```js
assert(code.toLowerCase().match(/-->\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

