---
id: 5dfa22d1b521be39a3de7be0
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Convierte las palabras `cat photos` ubicadas dentro del elemento `p` en un enlace usando el mismo valor para el atributo `href` que el enlace debajo del elemento `p`. El elemento `p` debe seguir mostrando el mismo texto en el navegador, pero las palabras `cat photos` ahora deben ser un enlace. Asegúrate de eliminar el elemento `a` con el texto `cat photos` en la linea debajo del elemento `p`.

# --hints--

Tu código solo debe contener un elemento anchor (`a`). Elimina cualquier elemento anchor extra.

```js
assert(document.querySelectorAll('a').length === 1);
```

Tu elemento anchor (`a`) debe estar anidado dentro del elemento `p`.

```js
assert($('p > a').length);
```

El texto del enlace debe ser `cat photos`. Probablemente no has añadido el texto o tienes un error tipográfico.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com' &&
    nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Después de anidar el elemento achor (`a`), el único contenido visible en el navegador del elemento `p` debe ser `Click here to view more cat photos.` Comprueba el texto, el espaciado y la puntuación del elemento `p` y del elemento anchor anidado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

