---
id: 5dfa22d1b521be39a3de7be0
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Ahora puedes ver las palabras `cat photos` subrayadas al lado de la imagen en la vista previa. Ese es tu enlace; no dudes en darle un clic.

En el texto del elemento `p`, convierte las palabras `cat photos` en un enlace a `https://freecatphotoapp.com` anidando estas palabras dentro de un nuevo elemento de anclaje.

# --hints--

Deberías anidar un nuevo elemento (`a`) dentro del elemento `p`.

```js
assert($('p > a').length);
```

El valor del `href` del enlace debe ser `https://freecatphotoapp.com`. Has omitido el valor del `href` o tienes un error de escritura.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

El texto del enlace debe ser `cat photos`. Has omitido el texto o tienes un error de escritura.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Después de anidar el elemento de anclaje (`a`), el único contenido del elemento `p` visible en el navegador debe ser `Click here to view more cat photos.` Compruebe el texto, el espaciado o la puntuación tanto del elemento `p` como del elemento ancla anidado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

