---
id: 5dfa2407b521be39a3de7be1
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Añade un atributo `target` con el valor `_blank` en la etiqueta de apertura, del elemento anchor (`a`) para que el enlace abra en una nueva pestaña.

# --hints--

Tu elemento `p` debe tener un elemento anchor (`a`) anidado, con el texto `cat photos`. Es posible que lo hayas eliminado o que tengas un error tipográfico.

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Tu elemento anchor (`a`) no tiene un atributo `target`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
assert(document.querySelector('a').hasAttribute('target'));
```

El valor del atributo `target` debe ser `_blank`. Probablemente no has añadido el valor o tienes un error tipográfico. Recuerda que los valores de los atributos deben estar entre comillas.

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

