---
id: 5dfa22d1b521be39a3de7be0
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Puedes convertir cualquier texto en un enlace, como el texto dentro de un elemento `p`.

``` html
<p>I think <a href="https://www.freecodecamp.org">freeCodeCamp</a> is great.</p>
```

En el texto de su elemento `p`, convierta las palabras `fotos de gatos` en un enlace agregando etiquetas de anclaje de apertura y cierre (`a`) alrededor de estas palabras.. Luego establece el atributo `href` en `https://freecatphotoapp.com`


# --hints--

Deberías anidar un nuevo elemento (`a`) dentro del elemento `p`.

```js
const nestedAnchor = document.querySelector(`p > a`);
assert.isNotNull(nestedAnchor)
```

El valor del enlace `href` debe ser `https://freecatphotoapp.com`. Has omitido el valor `href` o tienes un error tipográfico.

```js
assert(
  document.querySelector('p > a').href === 'https://freecatphotoapp.com/'
);
```

El texto del enlace debe ser `cat photos`. Has omitido el texto o tienes un error de escritura.

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isTrue(innerContent.trim() === 'cat photos');
```

El texto dentro de tu elemento anchor tiene espacio en blanco extra inicial o final. El único espacio en el texto del elemento anchor debería estar entre la palabra `cat` y la palabra `photos`.

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

Después de anidar el elemento anchor (`a`), el único contenido visible del elemento `p` en el navegador debería ser `See more cat photos in our gallery.` Verifica el texto, el espaciado o la puntuación tanto del elemento `p` como del elemento anchor anidado.

```js
assert.match(code, /<p>see more <a[^>]*>cat photos<\/a> in our gallery\.?<\/p>/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

