---
id: 5dfa22d1b521be39a3de7be0
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

En el paso anterior has convertido las palabras `link to cat pictures` en un enlace colocándolas entre la apertura y el cierre de la etiqueta (`a`). Puedes hacer lo mismo con las palabras dentro de un elemento, como un elemento de `p`.

En el texto del elemento `p`, convierte las palabras `cat photos` en un enlace a `https://freecatphotoapp.com` añadiendo etiquetas anchor (`a`) de apertura y cierre alrededor de dichas palabras.


# --hints--

Deberías anidar un nuevo elemento (`a`) dentro del elemento `p`.

```js
assert($('p > a').length);
```

El valor del enlace `href` debe ser `https://freecatphotoapp.com`. Has omitido el valor `href` o tienes un error tipográfico.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

El texto del enlace debe ser `cat photos`. Has omitido el texto o tienes un error de escritura.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Después de anidar el elemento anchor (`a`), el único contenido del elemento `p` visible en el navegador debe ser `See more cat photos in our gallery.` Verifique el texto, el espaciado o la puntuación tanto del elemento `p` como del elemento de anclaje anidado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

