---
id: 5dfb6250eacea3f48c6300b2
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

After the unordered list, add a new image with a `src` attribute value set to:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Y un atributo `alt` con el valor:

`A slice of lasagna on a plate.`

# --hints--

Debe haber un elemento `img` después de la etiqueta de cierre del elemento `</ul>`.

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

La nueva imagen no tiene un atributo `alt`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

La nueva imagen debe tener un atributo `alt` con el valor `A slice of lasagna on a plate.` Asegurate de que el valor del atributo `alt` está entre comillas.

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

The new image does not have a `src` attribute. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

The new image should have a `src` value of `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Asegúrate de que el valor del atributo `src` esté entre comillas.

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Aunque el atributo `src` de la nueva imagen tenga la URL correcta, se recomienda siempre poner entre comillas el valor de un atributo.

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

