---
id: 5dfb6a35eacea3f48c6300b4
title: Paso 24
challengeType: 0
dashedName: step-24
---

# --description--

Un elemento (`figcaption`), se utiliza para añadir una descripción o leyenda para describir una imagen anidada en un elemento `figure`. Por ejemplo, `<figcaption>A cute cat</figcaption>` añade la leyenda `A cute cat`.

Después de anidar la imagen en el elemento `figure`, añade un elemento `figcaption` con el texto:

`Cats love lasagna.`

# --hints--

Tu elemento `figcaption` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

Tu elemento `figcaption` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

El elemento `figcaption` debe estar anidado en el elemento `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

El elemento `img` de lasaña debe estar anidado en el elemento `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

El elemento `figcaption` anidado en el elemento `figure` debe estar debajo del elemento `img`. Los tienes en el orden equivocado.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

El texto de tu elemento `figcaption` debe ser `Cats love lasagna.` Has omitido el texto o tienes un error tipográfico.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

