---
id: 5ef9b03c81a63668521804d7
title: Paso 36
challengeType: 0
dashedName: step-36
---

# --description--

El atributo `action` indica, dónde deben enviarse los datos del formulario. Por ejemplo, `<form action="/submit-url"></form>` le dice al navegador que los datos del formulario deben ser enviados a la ruta `/submit-url`.

Al elemento `form`, añádele un atributo `action` con el valor `https://freecatphotoapp.com/submit-cat-photo`.

# --hints--

Tu elemento `form` debe tener una etiqueta de apertura y una etiqueta de cierre. Puede que te falte una o dos de las etiquetas requeridas o las tienes en el orden incorrecto.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Tu elemento `form` anidado en el último elemento `section` debe estar debajo del elemento `h2`. Tienes el elemento `h2` y el elemento `form` en el orden incorrecto.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Tu elemento `form` no debe tener contenido. Elimina cualquier elemento HTML o texto que tengas entre las etiquetas del elemento `form`.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

Tu elemento `form` no tiene un atributo `action`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Tu formulario `form` debe tener un atributo `action` con el valor `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Aunque el atributo `action` tenga la URL correcta, se recomienda siempre poner entre comillas el valor de un atributo.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

