---
id: 5ef9b03c81a63668521804d8
title: Paso 37
challengeType: 0
dashedName: step-37
---

# --description--

El elemento `input` te permite recolectar datos desde un formulario web de diferentes formas. Al igual que los elementos `img`, los elementos `input` son de <dfn>auto-cierre</dfn> y no necesitan etiquetas de cierre.

Anida un elemento `input` dentro del elemento `form`.

# --hints--

Tu elemento `form` debe tener una etiqueta de apertura y una etiqueta de cierre en el orden correcto. Puede que te falte una o dos de las etiquetas requeridas o las tienes en el orden incorrecto.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

La etiqueta de apertura de tu elemento `form` solo debe tener un atributo `action`. Elimina cualquier otra cosa que hayas escrito en ella.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Debes crear un elemento `input`. Revisa la sintaxis.

```js
assert(document.querySelector('input'));
```

Tu elemento `input` debe tener una etiqueta de apertura, pero no una etiquete de cierre.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Tu elemento `input` debe estar anidado dentro del elemento `form`.

```js
assert(document.querySelector('form > input'));
```

Tu elemento `form` solo debe contener el elemento `input`. Elimina cualquier elemento HTML o texto que tengas entre las etiquetas del elemento `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">

        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

