---
id: 5ef9b03c81a63668521804d4
title: Paso 32
challengeType: 0
dashedName: step-32
---

# --description--

El elemento `strong` se utiliza para indicar que una parte de un texto es importante o urgente.

En el elemento `figcaption` que acabas de añadir, indica que la palabra `hate` tiene mayor importancia, envolviéndola con un elemento `strong`.

# --hints--

Tu elemento `strong` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('strong'));
```

Tu elemento `strong` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/strong\>/));
```

Tu elemento `strong` debe envolver la palabra `hate` en el texto `Cats hate other cats.` Has omitido el texto o tienes un error tipográfico.

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.toLowerCase() === 'hate'
);
```

El texto del `figcaption` debe ser `Cats hate other cats.` Compruebe que no hay errores tipográficos y que los espacios necesarios están presentes alrededor de las etiquetas de apertura y cierre del elemento `strong`.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

