---
id: 5ef9b03c81a63668521804d6
title: Paso 35
challengeType: 0
dashedName: step-35
---

# --description--

Ahora añadirás un formulario web para recolectar información de los usuarios.

Después del título `Cat Form`, añade un elemento `form`.

# --hints--

Tu elemento `form` debe tener una etiqueta de apertura y una etiqueta de cierre. Puede que te falte una o dos de las etiquetas requeridas o las tienes en el orden incorrecto.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Las etiquetas del elemento `form` no están en el orden correcto.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

El elemento `form` anidado en el último elemento `section` debe estar debajo del elemento `h2`. Tienes el elemento `h2` y el elemento `form` en el orden incorrecto.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Tu elemento `form` no debe tener contenido. Elimina cualquier elemento HTML o texto que tengas entre las etiquetas del elemento `form`.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

