---
id: 5ef9b03c81a63668521804d0
title: Paso 25
challengeType: 0
dashedName: step-25
---

# --description--

Enfatiza la palabra `love` en el elemento `figcaption` envolviendola en un elemento énfasis `em`.

# --hints--

Tu elemento énfasis `em` debería tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Tu elemento énfasis `em` debería tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/em\>/));
```

Has eliminado el elemento `figcaption` o le falta una etiqueta de apertura o de cierre.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Tu elemento énfasis `em` debería rodear el texto `love`. Probablemente no has añadido el texto o tienes un error tipográfico.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

El texto del elemento `figcaption` debe ser `Cats love lasagna`. Comprueba que no haya errores tipográficos y que los espacios entre las etiquetas del elemento `em` son los correctos.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

