---
id: 5ef9b03c81a63668521804dc
title: Paso 44
challengeType: 0
dashedName: step-44
---

# --description--

Puedes usar botones de radio (radio button) para preguntas en las que quieres que el usuario solo te dé una respuesta a partir de múltiples opciones.

Aquí hay un ejemplo de un radio button con la opción `cat`: `<input type="radio"> cat`. Recuerda que los elementos `input` son de <dfn>auto-cierre</dfn>, no necesitan etiqueta de cierre.

Antes de la casilla de entrada de texto (text input), añade un radio button con la opción:

`Indoor`

# --hints--

Debes crear un elemento `input` para tu botón de radio. Revisa la sintaxis.

```js
assert($('form > input').length >= 2);
```

Tu elemento `input` debe tener una etiqueta de apertura, pero no una etiquete de cierre.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

Debes tener solo un elemento `input` agregado para tu botón de radio. Elimina los que estén de más.

```js
assert($('form > input').length === 2);
```

Tu nuevo elemento `input` debe estar arriba del elemento `input` ya existente que tiene un atributo `type` con el valor `text`. Los tienes en el orden incorrecto.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Tu nuevo elemento `input` no tiene un atributo `type`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura.

```js
assert($('input')[0].hasAttribute('type'));
```

Your new `input` element should have only one `type` attribute. Remove any extras.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Your new `input` element should have a `type` attribute with the value `radio`. You have either omitted the value or have a typo. Remember that attribute values should be surrounded with quotation marks.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Although you have set the new `input` element's `type` attribute to `radio`, it is recommended to always surround the value of an attribute with quotation marks.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

The `radio` button's `Indoor` text should be located after it instead of before it.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

The text `Indoor` should be located directly to the right of your `radio` button. You have either omitted the text or have a typo.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

