---
id: 5ef9b03c81a63668521804db
title: Paso 41
challengeType: 0
dashedName: step-41
---

# --description--

Para evitar que un usuario entregue su formulario incompleto, necesitas añadir el atributo `required` al elemento `input`. El atributo `required` no necesita ningún valor. Simplemente necesitas añadir la palabra `required` al elemento `input`, asegurándote de que hay espacios entre ella y los otros atributos.

# --hints--

Has eliminado el elemento `input` o la sintaxis es incorrecta. Todos los valores de los atributos deben estar entre comillas.

```js
assert($('input').length);
```

Tu elemento `form` solo debe contener el elemento `input`. Elimina cualquier elemento HTML o texto que tengas entre las etiquetas del elemento `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Tu elemento `input` debe tener el atributo `required`. Recuerda, simplemente añade la palabra `required` a la etiqueta del elemento `input`.

```js
assert($('input')[0].hasAttribute('required'));
```

No debes añadir ningún valor al atributo `required`.

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

