---
id: 5f05a1d8e233dff4a68508d8
title: Paso 47
challengeType: 0
dashedName: step-47
---

# --description--

Crea otro botón de radio debajo del primero. Anídalo dentro de un elemento `label` con `Outdoor` como texto del `label`. Dale al botón de radio un atributo `id` con `outdoor` como valor.

# --hints--

Solo el `input` de botón de radio `Indoor` original debe tener un `id` establecido en `indoor`. Puede reiniciar el paso para recuperar el HTML original si es necesario.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

No debes hacer ningún cambio en el botón de radio `Indoor`. Puede reiniciar el paso para recuperar el HTML original si es necesario.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Debes agregar exactamente un nuevo elemento `input` anidado en un nuevo elemento `label`. Asegurate de que tu nuevo `label` tiene una etiqueta de apertura y de cierre.

```js
assert(document.querySelectorAll('label input').length === 2);
```

No debes agregar ningún atributo a ninguna de las etiquetas de apertura del `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

Tu nuevo `input` debe tener un atributo `id` con el valor `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

Tu nuevo `input` debe tener un atributo `type` con el valor `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

No debes agregar ningún elemento nuevo que no sea un `input` anidado en un `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

No debe haber texto a la izquierda de tu nuevo `input`.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

El texto de la etiqueta de su nuevo botón de radio debe ser exactamente `Outdoor`. Esto incluye las mayúsculas.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

El nuevo botón de radio y la etiqueta deben estar inmediatamente debajo/después del botón de radio `Indoor` y la etiqueta. No debe haber ninguna otra etiqueta entre ellos.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

