---
id: 5f07be6ef7412fbad0c5626b
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Antes de añadir nuevo contenido, deberías utilizar un elemento `section`, para separar el contenido de cat photos, del contenido que añadiremos después.

Anida en un elemento `section` los elementos `h2`, el comentario, `p`, y anchor (`a`).

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Tu elemento `section` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

El elemento `section` debe estar entre las etiquetas de apertura y cierre del elemento `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Los elementos `h2`, comentario, `p` y anchor (`a`) deben estar entre las etiquetas de apertura y cierre del elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

El elemento `h1` no debe estar anidado en el elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

