---
id: 5f07be6ef7412fbad0c5626b
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Antes de añadir nuevo contenido, deberías utilizar un elemento `section`, para separar el contenido de cat photos, del contenido que añadiremos después.

Toma todos los elementos actualmente ubicados dentro del elemento `main` y anidarlos dentro de un elemento `section`.

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Tu elemento `section` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert(code.match(/<\/section\>/));
```

El elemento `section` debe estar entre las etiquetas de apertura y cierre del elemento `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

Los elementos `h2`, comentario, `p` y anchor (`a`) deben estar entre las etiquetas de apertura y cierre del elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElems = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElems.length === 3);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
    </main>
--fcc-editable-region--
  </body>
</html>
```

