---
id: 62dabe2ef403a12d5d295273
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Ahora que has convertido el texto `cat photos` dentro del elemento `p` en un enlace, no necesitas el segundo enlace debajo del elemento `p`. Elimina todo el elemento de ancla debajo del elemento `p`.

# --hints--

Your code should only have one paragraph (`p`) element.

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

Tu código solo debe tener un elemto de ancla (`a`).

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

Do not make any changes to the elements you created in the previous step.

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

No debería estar el texto `cat photos` debajo del elemento `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
