---
id: 62dabe2ef403a12d5d295273
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Ahora que has convertido el texto `cat photos` dentro del elemento `p` en un enlace, no necesitas el segundo enlace debajo del elemento `p`. Elimina todo el elemento de ancla debajo del elemento `p`.

# --hints--

Tu código solo debe tener un elemto de ancla (`a`).

```js
assert(document.querySelectorAll('a').length === 1);
```

El elemento `p` todavía debe tener el elemento de ancla (`a`) dentro de él.

```js
assert($('p > a').length);
```

El valor del enlace `href` debe ser `https://freecatphotoapp.com`. Has omitido el valor `href` o tienes un error tipográfico.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

El texto del enlace debe ser `cat photos`. Has omitido el texto o tienes un error tipográfico.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

El contenido del elemento `p` visible en el navegador debe ser `Click here to view more cat photos.` Comprobar el texto, el espacio, o la puntuación del elemento `p` y el elemento ancla anidado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

No debería estar el texto `cat photos` debajo del elemento `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more <a href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
