---
id: 6140cbeec34e970dfe75e710
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Unda kichaguzi cha kipengele chako cha `.wheel`. Anza kwa kuweka `border` kuwa `2px solid black`, `border-radius` kuwa `50%` na `margin-left` kuwa `50px`.

# --hints--

Unapaswa kuwa na kichagui cha `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Kichaguzi chako cha `.wheel` kinapaswa kuwa na sifa ya `border` iliyowekwa kuwa `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Kichaguzi chako cha `.wheel` kinafaa kuwa na sifa ya `border-radius` iliyowekwa kuwa `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Kichaguzi chako cha `.wheel` kinafaa kuwa na sifa ya `margin-left` iliyowekwa kuwa `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
