---
id: 6140c7e645d8e905819f1dd4
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Anza na msingi wa kanuni wa kawaida (boilerplate). Ongeza tamko lako la `DOCTYPE`, kipengele chako cha `html` chenye lugha iliyowekwa kwa Kiingereza, vipengele vyako vya `head` na `body`.

Ongeza kipengee chako cha `meta` kwa `charset` sahihi, kipengele chako cha `title` na `link` kwa `./styles.css` faili.

Weka `title` kuwa `Ferris Wheel`.

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Unapaswa kufafanua aina ya hati kuwa `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Unapaswa kufunga tamk la `DOCTYPE` kwa `>` baada ya aina.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Kipengele chako cha `html` kinafaa kuwa na lebo ya ufunguzi yenye sifa ya `lang` ya `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert.match(code, /<\/html\s*>/);
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Unapaswa kuwa na tagi ya `head` ya kufungua.

```js
assert.match(code, /<head\s*>/i);
```

Unapaswa kuwa na tagi ya kufunga ya `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Unapaswa kuwa na tagi ya `body` ya kufungua.

```js
assert.match(code, /<body\s*>/i);
```

Unapaswa kuwa na tagi ya kufunga ya `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Msimbo wako unapaswa kuwa na kipengele cha `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya `charset` yenye thamani ya `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Mradi wako unapaswa kuwa na kichwa cha `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Kumbuka, ukubwa wa herufi na tahajia ni jambo la muhimu katika kichwa.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `rel` yenye thamani `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
