---
id: 6140cd32d018ed0f600eefce
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Weka sifa ya `position` ya kichaguzi `.wheel` kuwa `absolute`. Weka `height` na `width` zote kuwa `55vw`.

# --hints--

Kichaguzi chako cha `.wheel` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

Kichaguzi chako cha `.wheel` kinafaa kuwa na sifa ya `height` iliyowekwa kuwa `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

Kichaguzi chako cha `.wheel` kinafaa kuwa na sifa ya `width` iliyowekwa kuwa `55vw`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
