---
id: 6140d10d50636e14695013b2
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Unda kichaguzi ili kulenga kipengele chako cha pili cha `.line`. Weka sifa ya `transform` kuwa `rotate(60deg)`.

Kumbuka kwamba sifa ya `transform` hukuruhusu kudhibiti umbo la kipengele. Katika hali hii, kwa kutumia thamani ya `rotate(60deg)` kutazungusha kipengele kuzunguka uhakika wake wa `transform-origin` kwa nyuzi 60 kisaa.

# --hints--

Unapaswa kuunda kichaguzi cha `.line:nth-of-type(2)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

Kichaguzi chako cha `.line:nth-of-type(2)` kinafaa kuwa na sifa ya `transform` iliyowekwa kuwa `rotate(60deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
