---
id: 6140d0069049f5139d78da40
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

Sifa ya `transform-origin` inatumika mahali ambapo mabadiliko ya CSS yanatumika. Kwa mfano, unapofanya `rotate` (ambayo utafanya baadaye katika mradi huu), `transform-origin` huamua ni sehemu gani kipengele kinazungushwa.

Kipe kichaguzi cha `.line` sifa ya `transform-origin` ya `0% 0%`. Hii itaondoa sehemu ya asili kwa `0%` kutoka kushoto na `0%` kutoka juu, na kuiweka kwenye kona ya juu kushoto ya kipengele.

# --hints--

Kichaguzi chako cha `.line` kinafaa kuwa na sifa ya `transform-origin` iliyowekwa kuwa `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
