---
id: 6140d1a351e88f159ed54fca
title: Hatua ya 10
challengeType: 0
dashedName: step-10
---

# --description--

Kwa kutumia mchoro sawa, tengeneza kichaguzi tofauti cha `.line` cha tatu, `.line` ya nne, `.line` ya tano, na `.line` ya sita.

Weka sifa ya `transform` ya `.line` ya tatu kuwa `rotate(120deg)`, ya nne `rotate(180deg)`, ya tano `rotate(240deg)`, na ya sita `rotate(300deg)`.

# --hints--

Unapaswa kuunda kichaguzi cha `.line:nth-of-type(3)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Kichaguzi chako cha `.line:nth-of-type(3)` kinafaa kuwa na sifa ya `transform` iliyowekwa kuwa `rotate(120deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform === 'rotate(120deg)');
```

Unapaswa kuunda kichaguzi cha `.line:nth-of-type(4)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Kichaguzi chako cha `.line:nth-of-type(4)` kinafaa kuwa na sifa ya `transform` iliyowekwa kuwa `rotate(180deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform === 'rotate(180deg)');
```

Unapaswa kuunda kichaguzi cha `.line:nth-of-type(5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Kichaguzi chako cha `.line:nth-of-type(5)` kinafaa kuwa na sifa ya `transform` iliyowekwa kuwa `rotate(240deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform === 'rotate(240deg)');
```

Unapaswa kuunda kichaguzi cha `.line:nth-of-type(6)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Kichaguzi chako cha `.line:nth-of-type(6)` kinafaa kuwa na sifa ya `transform` iliyowekwa kuwa `rotate(300deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform === 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
