---
id: 68af0687ef34c76c28ffa547
title: "Challenge 30: Unique Characters"
challengeType: 28
dashedName: challenge-30
---

# --description--

Given a string, determine if all the characters in the string are unique.

- Uppercase and lowercase letters should be considered different characters.

# --hints--

`allUnique("abc")` should return `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` should return `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` should return `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` should return `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` should return `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` should return `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` should return `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
