---
id: 68c1a929005bf54d342aa8d5
title: "Challenge 58: Space Week Day 4: Placeholder"
challengeType: 28
dashedName: challenge-58
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
