---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Placeholder"
challengeType: 28
dashedName: challenge-59
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
