---
id: 68c497f3aaefc9fd9f1b0e26
title: "Challenge 62: Placeholder"
challengeType: 28
dashedName: challenge-62
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
