---
id: 6912601784c1e8fc20e472cf
title: Maskify Numerals
challengeType: 28
dashedName: maskify-numerals
---

# --description--

ATM machines often mask all numbers, besides the last 4 numbers of a credit/debit card. Write a function that masks all numbers with the character "#", while leaving the last 4 numbers revealed.

# --hints--

`maskify("12345678910")` should return `"#######8910"`.

```js
assert.equal(maskify("12345678910"), "#######8910")
```

`maskify("12345")` should return `"#2345"`.

```js
assert.equal(maskify("12345"), "#2345")
```

`maskify("123")` should return `"123"`


```js
assert.equal(maskify("123"), "123")
```


# --seed--

## --seed-contents--

```js
function maskify (str) {

return str
}
```

# --solutions--

```js
function maskify(str) {
	const strSplit = str.split('')
	const strSlice = strSplit.slice(-4)
	const strSliceJoin = strSlice.join('')
	const stor = []
	
	for (let i = 0; i < strSplit.length - 4; i++) {
	strSplit[i] = '#'
	stor.push(strSplit[i])
}
	const eleJoin = stor.concat(strSliceJoin)
	
    const strFilter = eleJoin.filter((element) => element !== ',')
	const stringify = strFilter.toString()
	const stringifyJoin = stringify.split('').join('')
	
	return stringifyJoin.replace(/,/g, "")
}

    ```
    