---
id: 691b559495c5cb5a37b9b486
title: "Challenge 126: Capitalize It"
challengeType: 28
dashedName: challenge-126
---

# --description--

Given a string title, return a new string formatted in title case using the following rules:

- Capitalize the first letter of each word.
- Make all other letters in each word lowercase.
- Words are always separated by a single space.

# --hints--

`titleCase("hello world")` should return `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` should return `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` should return `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` should return `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
