---
id: 6925e2068081f40f549ced1c
title: "Challenge 138: Sum of Divisors"
challengeType: 28
dashedName: challenge-138
---

# --description--

Given a positive integer, return the sum of all its divisors.

- A divisor is any integer that divides the number evenly (the remainder is `0`).
- Only count each divisor once.

For example, given `6`, return `12` because the divisors of `6` are `1`, `2`, `3`, and `6`, and the sum of those is `12`.

# --hints--

`sumDivisors(6)` should return `12`.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` should return `14`.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` should return `56`.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` should return `224`.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` should return `806`.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` should return `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
