---
id: 6939b873185d8e00d453563b
title: "Challenge 156: Odd or Even?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Given a positive integer, return `"Odd"` if it's an odd number, and `"Even"` if it's even.

# --hints--

`oddOrEven(1)` should return `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` should return `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` should return `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` should return `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` should return `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
