---
id: 6939b873185d8e00d453563d
title: "Challenge 158: Array Swap"
challengeType: 28
dashedName: challenge-158
---

# --description--

Given an array with two values, return an array with the values swapped.

For example, given `["A", "B"]` return `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` should return `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` should return `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` should return `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` should return `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
