---
id: 697a49e9860d24853adef67f
title: "Challenge 194: Placeholder"
challengeType: 28
dashedName: challenge-194
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
