---
id: 697a49e9860d24853adef67e
title: "Challenge 193: Placeholder"
challengeType: 28
dashedName: challenge-193
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
