---
id: 697a49e6ff50d756c9b69365
title: "Challenge 188: Placeholder"
challengeType: 28
dashedName: challenge-188
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
