---
id: 697a49e6ff50d756c9b6935d
title: "Challenge 180: 2026 Winter Games Day 1: Opening Day"
challengeType: 28
dashedName: challenge-180
---

# --description--

Today marks the start of the 2026 Winter Games. The next 17 days will bring you coding challenges inspired by them.

For the first one, you are given a two-letter country code and need return the flag emoji for that country.

Use this list:

| Country | Code | Flag |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"🇦🇱"` |
| Andorra                    |            `"AD"` | `"🇦🇩"` |
| Argentina                  |            `"AR"` | `"🇦🇷"` |
| Armenia                    |            `"AM"` | `"🇦🇲"` |
| Australia                  |            `"AU"` | `"🇦🇺"` |
| Austria                    |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                 |            `"AZ"` | `"🇦🇿"` |
| Belgium                    |            `"BE"` | `"🇧🇪"` |
| Benin                      |            `"BJ"` | `"🇧🇯"` |
| Bolivia                    |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina     |            `"BA"` | `"🇧🇦"` |
| Brazil                     |            `"BR"` | `"🇧🇷"` |
| Bulgaria                   |            `"BG"` | `"🇧🇬"` |
| Canada                     |            `"CA"` | `"🇨🇦"` |
| Chile                      |            `"CL"` | `"🇨🇱"` |
| China                      |            `"CN"` | `"🇨🇳"` |
| Colombia                   |            `"CO"` | `"🇨🇴"` |
| Croatia                    |            `"HR"` | `"🇭🇷"` |
| Cyprus                     |            `"CY"` | `"🇨🇾"` |
| Czech Republic             |            `"CZ"` | `"🇨🇿"` |
| Denmark                    |            `"DK"` | `"🇩🇰"` |
| Ecuador                    |            `"EC"` | `"🇪🇨"` |
| Eritrea                    |            `"ER"` | `"🇪🇷"` |
| Estonia                    |            `"EE"` | `"🇪🇪"` |
| Finland                    |            `"FI"` | `"🇫🇮"` |
| France                     |            `"FR"` | `"🇫🇷"` |
| Georgia                    |            `"GE"` | `"🇬🇪"` |
| Germany                    |            `"DE"` | `"🇩🇪"` |
| Great Britain              |            `"GB"` | `"🇬🇧"` |
| Greece                     |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau              |            `"GW"` | `"🇬🇼"` |
| Haiti                      |            `"HT"` | `"🇭🇹"` |
| Hong Kong                  |            `"HK"` | `"🇭🇰"` |
| Hungary                    |            `"HU"` | `"🇭🇺"` |
| Iceland                    |            `"IS"` | `"🇮🇸"` |
| India                      |            `"IN"` | `"🇮🇳"` |
| Iran                       |            `"IR"` | `"🇮🇷"` |
| Ireland                    |            `"IE"` | `"🇮🇪"` |
| Israel                     |            `"IL"` | `"🇮🇱"` |
| Italy                      |            `"IT"` | `"🇮🇹"` |
| Jamaica                    |            `"JM"` | `"🇯🇲"` |
| Japan                      |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                 |            `"KZ"` | `"🇰🇿"` |
| Kenya                      |            `"KE"` | `"🇰🇪"` |
| Kosovo                     |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                 |            `"KG"` | `"🇰🇬"` |
| Latvia                     |            `"LV"` | `"🇱🇻"` |
| Lebanon                    |            `"LB"` | `"🇱🇧"` |
| Liechtenstein              |            `"LI"` | `"🇱🇮"` |
| Lithuania                  |            `"LT"` | `"🇱🇹"` |
| Luxembourg                 |            `"LU"` | `"🇱🇺"` |
| Madagascar                 |            `"MG"` | `"🇲🇬"` |
| Malaysia                   |            `"MY"` | `"🇲🇾"` |
| Malta                      |            `"MT"` | `"🇲🇹"` |
| Mexico                     |            `"MX"` | `"🇲🇽"` |
| Moldova                    |            `"MD"` | `"🇲🇩"` |
| Monaco                     |            `"MC"` | `"🇲🇨"` |
| Mongolia                   |            `"MN"` | `"🇲🇳"` |
| Montenegro                 |            `"ME"` | `"🇲🇪"` |
| Morocco                    |            `"MA"` | `"🇲🇦"` |
| Netherlands                |            `"NL"` | `"🇳🇱"` |
| New Zealand                |            `"NZ"` | `"🇳🇿"` |
| Nigeria                    |            `"NG"` | `"🇳🇬"` |
| North Macedonia            |            `"MK"` | `"🇲🇰"` |
| Norway                     |            `"NO"` | `"🇳🇴"` |
| Pakistan                   |            `"PK"` | `"🇵🇰"` |
| Philippines                |            `"PH"` | `"🇵🇭"` |
| Poland                     |            `"PL"` | `"🇵🇱"` |
| Portugal                   |            `"PT"` | `"🇵🇹"` |
| Puerto Rico                |            `"PR"` | `"🇵🇷"` |
| Romania                    |            `"RO"` | `"🇷🇴"` |
| San Marino                 |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia               |            `"SA"` | `"🇸🇦"` |
| Serbia                     |            `"RS"` | `"🇷🇸"` |
| Singapore                  |            `"SG"` | `"🇸🇬"` |
| Slovakia                   |            `"SK"` | `"🇸🇰"` |
| Slovenia                   |            `"SI"` | `"🇸🇮"` |
| South Africa               |            `"ZA"` | `"🇿🇦"` |
| South Korea                |            `"KR"` | `"🇰🇷"` |
| Spain                      |            `"ES"` | `"🇪🇸"` |
| Sweden                     |            `"SE"` | `"🇸🇪"` |
| Switzerland                |            `"CH"` | `"🇨🇭"` |
| Thailand                   |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago          |            `"TT"` | `"🇹🇹"` |
| Turkey                     |            `"TR"` | `"🇹🇷"` |
| Ukraine                    |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates       |            `"AE"` | `"🇦🇪"` |
| United States              |            `"US"` | `"🇺🇸"` |
| Uruguay                    |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                 |            `"UZ"` | `"🇺🇿"` |
| Venezuela                  |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` should return `"🇦🇱"`.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` should return `"🇦🇩"`.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` should return `"🇦🇷"`.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` should return `"🇦🇲"`.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` should return `"🇦🇺"`.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` should return `"🇦🇹"`.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` should return `"🇦🇿"`.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` should return `"🇧🇪"`.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` should return `"🇧🇯"`.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` should return `"🇧🇴"`.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` should return `"🇧🇦"`.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` should return `"🇧🇷"`.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` should return `"🇧🇬"`.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` should return `"🇨🇦"`.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` should return `"🇨🇱"`.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` should return `"🇨🇳"`.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` should return `"🇨🇴"`.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` should return `"🇭🇷"`.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` should return `"🇨🇾"`.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` should return `"🇨🇿"`.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` should return `"🇩🇰"`.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` should return `"🇪🇨"`.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` should return `"🇪🇷"`.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` should return `"🇪🇪"`.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` should return `"🇫🇮"`.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` should return `"🇫🇷"`.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` should return `"🇬🇪"`.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` should return `"🇩🇪"`.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` should return `"🇬🇧"`.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` should return `"🇬🇷"`.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` should return `"🇬🇼"`.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` should return `"🇭🇹"`.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` should return `"🇭🇰"`.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` should return `"🇭🇺"`.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` should return `"🇮🇸"`.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` should return `"🇮🇳"`.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` should return `"🇮🇷"`.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` should return `"🇮🇪"`.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` should return `"🇮🇱"`.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` should return `"🇮🇹"`.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` should return `"🇯🇲"`.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` should return `"🇯🇵"`.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` should return `"🇰🇿"`.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` should return `"🇰🇪"`.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` should return `"🇽🇰"`.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` should return `"🇰🇬"`.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` should return `"🇱🇻"`.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` should return `"🇱🇧"`.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` should return `"🇱🇮"`.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` should return `"🇱🇹"`.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` should return `"🇱🇺"`.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` should return `"🇲🇬"`.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` should return `"🇲🇾"`.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` should return `"🇲🇹"`.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` should return `"🇲🇽"`.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` should return `"🇲🇩"`.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` should return `"🇲🇨"`.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` should return `"🇲🇳"`.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` should return `"🇲🇪"`.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` should return `"🇲🇦"`.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` should return `"🇳🇱"`.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` should return `"🇳🇿"`.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` should return `"🇳🇬"`.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` should return `"🇲🇰"`.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` should return `"🇳🇴"`.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` should return `"🇵🇰"`.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` should return `"🇵🇭"`.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` should return `"🇵🇱"`.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` should return `"🇵🇹"`.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` should return `"🇵🇷"`.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` should return `"🇷🇴"`.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` should return `"🇸🇲"`.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` should return `"🇸🇦"`.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` should return `"🇷🇸"`.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` should return `"🇸🇬"`.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` should return `"🇸🇰"`.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` should return `"🇸🇮"`.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` should return `"🇿🇦"`.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` should return `"🇰🇷"`.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` should return `"🇪🇸"`.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` should return `"🇸🇪"`.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` should return `"🇨🇭"`.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` should return `"🇹🇭"`.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` should return `"🇹🇹"`.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` should return `"🇹🇷"`.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` should return `"🇺🇦"`.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` should return `"🇦🇪"`.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` should return `"🇺🇸"`.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` should return `"🇺🇾"`.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` should return `"🇺🇿"`.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` should return `"🇻🇪"`.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
