---
id: 697a49e6ff50d756c9b69363
title: "Challenge 186: 2026 Winter Games Day 7: Speed Skating"
challengeType: 28
dashedName: challenge-186
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
