---
id: 698a1a863194f1f4e63f6460
title: "Challenge 209: Placeholder"
challengeType: 28
dashedName: challenge-209
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
