---
id: 698a1a863194f1f4e63f645f
title: "Challenge 208: Placeholder"
challengeType: 28
dashedName: challenge-208
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
