---
id: 699c8e045ee7cb94ed2322d9
title: "Challenge 218: Evenly Divisible"
challengeType: 28
dashedName: challenge-218
---

# --description--

Given two integers, determine if you can evenly divide the first one by the secone one.

# --hints--

`isEvenlyDivisible(4, 2)` should return `true`.

```js
assert.isTrue(isEvenlyDivisible(4, 2));
```

`isEvenlyDivisible(7, 3)` should return `false`.

```js
assert.isFalse(isEvenlyDivisible(7, 3));
```

`isEvenlyDivisible(5, 10)` should return `false`.

```js
assert.isFalse(isEvenlyDivisible(5, 10));
```

`isEvenlyDivisible(48, 6)` should return `true`.

```js
assert.isTrue(isEvenlyDivisible(48, 6));
```

`isEvenlyDivisible(3186, 9)` should return `true`.

```js
assert.isTrue(isEvenlyDivisible(3186, 9));
```

`isEvenlyDivisible(4192, 11)` should return `false`.

```js
assert.isFalse(isEvenlyDivisible(4192, 11));
```

# --seed--

## --seed-contents--

```js
function isEvenlyDivisible(a, b) {

  return a;
}
```

# --solutions--

```js
function isEvenlyDivisible(a, b) {
  return a % b === 0;
}
```
