---
id: 69a890af247de743333bd4d0
title: "Challenge 227: Cooldown Time"
challengeType: 28
dashedName: challenge-227
---

# --description--

Given two timestamps, the first representing when a user finished an exam, and the second representing the current time, determine whether the user can take an exam again.

- Both timestamps will be given the format: `"YYYY-MM-DDTHH:MM:SS"`, for example `"2026-03-25T14:00:00"`. Note that the time is 24-hour clock.
- A user must wait at least 48 hours before retaking an exam.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` should return `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` should return `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` should return `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` should return `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  const finishDate = new Date(finishTime);
  const currentDate = new Date(currentTime);
  const diffMs = currentDate - finishDate;
  const cooldownMs = 48 * 60 * 60 * 1000;

  return diffMs >= cooldownMs;
}
```
