---
id: 69c5f3d787b1725d5f00c8bd
title: "Challenge 259: FizzBuzz Explosion"
challengeType: 28
dashedName: challenge-259
---

# --description--

Given an integer, return the number of steps it takes to turn the word `"fizzbuzz"` into a string with at least the given number of `"z"`'s using the following rules:

- Start with the string `"fizzbuzz"`.
- Each step, apply the standard FizzBuzz rules using the letter position in the string (the first `"f"` is position 1).
  - If the letter position is divisible by 3, replace the letter with `"fizz"`
  - If it's divisible by 5, replace the letter with `"buzz"`
  - If it's divisible by 3 and 5, replace the letter with `"fizzbuzz"`

So after 1 step, `"fizzbuzz"` turns into `"fifizzzbuzzfizzzz"`, which has 9 `"z"`'s.

# --hints--

`explodeFizzbuzz(9)` should return `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` should return `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` should return `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` should return `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` should return `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` should return `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` should return `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` should return `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
