---
id: 69c6ff713a52713463aa7929
title: "Challenge 263: Binary Crossword"
challengeType: 28
dashedName: challenge-263
---

# --description--

Given a character, determine if its 8-bit binary representation can be found in the following grid, horizontally or vertically in either direction:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

For example, `"A"` has the binary representation `01000001`, which appears in the first row from left to right.

# --hints--

`isInCrossword("I")` should return `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` should return `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` should return `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` should return `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` should return `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` should return `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` should return `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` should return `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
