---
id: 69c6ff713a52713463aa7926
title: "Challenge 260: Word Score"
challengeType: 28
dashedName: challenge-260
---

# --description--

Given a word, return its score using a standard letter-value table:

| Letter | Value |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Upper and lowercase letters have the same value.

# --hints--

`getWordScore("hi")` should return `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` should return `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` should return `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` should return `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
