---
id: 69e2383af7832c8032603b94
title: "Challenge 278: Coffee Order Parser"
challengeType: 28
dashedName: challenge-278
---

# --description--

Given a string for a coffee order, identify any menu items and return a formatted order.

Use the following menu items and prices:

| Item | Price | 
| - | - |
| `"cold brew"` | $4.50 |
| `"oat latte"` | $5.00 |
| `"cappuccino"` | $4.75 |
| `"espresso"` | $3.00 |
| `"vanilla syrup"` | $0.75 |
| `"caramel drizzle"` | $0.60 |
| `"extra shot"` | $0.50 |
| `"oat milk"` | $0.75 |
| `"cream"` | $0.75 |

Return a string with the matched items joined by `" + "`, followed by a colon and space (`": "`), and the total price.

For example, given `"I'd like an oat latte with vanilla syrup and an extra shot please."`, return `"oat latte + vanilla syrup + extra shot: $6.25"`

Items should appear in the order they appear in the menu and the total price should always have two decimal places.

# --hints--

`formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please.")` should return `"oat latte + vanilla syrup + extra shot: $6.25"`.

```js
assert.equal(formatCoffeeOrder("I'd like an oat latte with vanilla syrup and an extra shot please."), "oat latte + vanilla syrup + extra shot: $6.25");
```

`formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk.")` should return `"cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85"`.

```js
assert.equal(formatCoffeeOrder("Give me a cappuccino with caramel drizzle, vanilla syrup, and some oat milk."), "cappuccino + vanilla syrup + caramel drizzle + oat milk: $6.85");
```

`formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot.")` should return `"cold brew + extra shot + cream: $5.75"`.

```js
assert.equal(formatCoffeeOrder("Can I get a cold brew with some cream and an extra shot."), "cold brew + extra shot + cream: $5.75");
```

`formatCoffeeOrder("Just an espresso please.")` should return `"espresso: $3.00"`.

```js
assert.equal(formatCoffeeOrder("Just an espresso please."), "espresso: $3.00");
```

`formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle.")` should return `"oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60"`.

```js
assert.equal(formatCoffeeOrder("I'll take an oat latte with cream and an extra shot, and some vanilla syrup and caramel drizzle."), "oat latte + vanilla syrup + caramel drizzle + extra shot + cream: $7.60");
```

# --seed--

## --seed-contents--

```js
function formatCoffeeOrder(order) {

  return order;
}
```

# --solutions--

```js
function formatCoffeeOrder(order) {
  const menu = [
    { name: "cold brew", price: 4.50 },
    { name: "oat latte", price: 5.00 },
    { name: "cappuccino", price: 4.75 },
    { name: "espresso", price: 3.00 },
    { name: "vanilla syrup", price: 0.75 },
    { name: "caramel drizzle", price: 0.60 },
    { name: "extra shot", price: 0.50 },
    { name: "oat milk", price: 0.75 },
    { name: "cream", price: 0.75 },
  ];

  const found = menu.filter(item => order.toLowerCase().includes(item.name));
  const total = found.reduce((sum, item) => sum + item.price, 0);
  return `${found.map(item => item.name).join(" + ")}: $${total.toFixed(2)}`;
}
```
