---
id: 69f8c998d78ad3171a0713b9
title: "Challenge 286: Open Issues"
challengeType: 28
dashedName: challenge-286
---

# --description--

Given an array of issue numbers and another array of pull request (PR) numbers, return an array of issues that remain open after all PRs have been merged.

- A PR closes an issue if their digits are a rotation of each other. For example, issue 123 would be closed by PR 231 or 312.
- A PR does not close an issue with the exact same number. For example, PR 123 does not close issue 123. So an issue with all the same number can't get closed.
- Either number may have leading zeros stripped. For example, PR 201 would close issue 12 (012, a rotation of 201). Similarily, issue 201 would be closed by PR 12.

Return the remaining open issues in the order they were given.

# --hints--

`getOpenIssues([123, 234], [231])` should return `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` should return `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` should return `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` should return `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` should return `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
