---
id: 69f8c998d78ad3171a0713bb
title: "Challenge 288: Secret Number"
challengeType: 28
dashedName: challenge-288
---

# --description--

Given a secret number and a guess, determine if the guess is correct.

Return:

- `"higher"` if the secret number is higher than the guess.
- `"lower"` if the secret number is lower than the guess.
- `"you got it!"` if the guess is correct.

# --hints--

`guessNumber(50, 30)` should return `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` should return `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` should return `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` should return `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` should return `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` should return `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
