---
id: 69f90c1329a94b37e2a2086d
title: "Challenge 294: Parentheses Combinations"
challengeType: 28
dashedName: challenge-294
---

# --description--

Given an integer, `n`, return the number of valid combinations of `n` pairs of parentheses.

- A valid combination is a string where every opening parentheses has a corresponding closing parentheses, and no closing parentheses appears before its matching opening parentheses.

For example, given `2`, there are 2 valid combinations:

```json
(())
()()
```

# --hints--

`getCombinations(2)` should return `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` should return `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` should return `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` should return `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` should return `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
