---
id: 6a0dcc730cb92a616f86f0c4
title: "Challenge 300: Placeholder"
challengeType: 28
dashedName: challenge-300
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
