---
id: 6a0dcc730cb92a616f86f0bf
title: "Challenge 295: Schema Validator Part 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Given an object (JavaScript) or dictionary (Python), determine if it matches the following schema:

```json
{
  username: string
}
```

- Extra keys are allowed

# --hints--

`isValidSchema({ username: "bob" })` should return `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` should return `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` should return `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` should return `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` should return `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
