---
id: 6a0dcd03ee4e68698080ef6c
title: "Challenge 308: Placeholder"
challengeType: 28
dashedName: challenge-308
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
