---
id: 6a0dcd03ee4e68698080ef6a
title: "Challenge 306: HTML Content Extractor"
challengeType: 28
dashedName: challenge-306
---

# --description--

Given a string of HTML code, remove all the tags and return the plain content.

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
