---
id: 6a15cadf5f240d05a2649556
title: "Challenge 312: Streaming Cost"
challengeType: 28
dashedName: challenge-312
---

# --description--

Given an array representing movies in the cart of your streaming service, and a string for your subscription tier, return the total cost of the movies.

Each item in the cart is an object with a `"format"` (`"HD"` or `"4K"`) and a `"type"` (`"rent"` or `"buy"`). Their costs are:

| | `"rent"` | `"buy"` |
| - | - | - |
| `"HD"` | $3.99 | $12.99 |
| `"4K"` | $5.99 | $19.99 |

Apply the following subscription tier discounts:

- `"none"`: full price
- `"basic"`: 10% off
- `"premium"`: 25% off

Return the total cost rounded to two decimal places in the format `"$D.CC"`.

# --hints--

`getStreamingBill([{ format: "HD", type: "rent" }], "none")` should return `"$3.99"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }], "none"), "$3.99");
```

`getStreamingBill([{ format: "HD", type: "buy" }], "premium")` should return `"$9.74"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "buy" }], "premium"), "$9.74");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic")` should return `"$18.87"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "HD", type: "rent" }, { format: "HD", type: "buy" }], "basic"), "$18.87");
```

`getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium")` should return `"$44.98"`.

```js
assert.equal(getStreamingBill([{ format: "4K", type: "buy" }, { format: "4K", type: "buy" }, { format: "4K", type: "buy" }], "premium"), "$44.98");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none")` should return `"$42.96"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "none"), "$42.96");
```

`getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic")` should return `"$38.66"`.

```js
assert.equal(getStreamingBill([{ format: "HD", type: "rent" }, { format: "4K", type: "rent" }, { format: "HD", type: "buy" }, { format: "4K", type: "buy" }], "basic"), "$38.66");
```

# --seed--

## --seed-contents--

```js
function getStreamingBill(cart, subscription) {

  return cart;
}
```

# --solutions--

```js
function getStreamingBill(cart, subscription) {
  const pricing = {
    HD: { rent: 3.99, buy: 12.99 },
    "4K": { rent: 5.99, buy: 19.99 },
  };

  const discounts = {
    none: 1,
    basic: 0.9,
    premium: 0.75,
  };

  const multiplier = discounts[subscription];

  const total = cart.reduce((sum, item) => {
    const price = pricing[item.format][item.type];
    return sum + price * multiplier;
  }, 0);

  return `$${total.toFixed(2)}`;
}
```
