---
id: 6a15cadf5f240d05a264955a
title: "Challenge 316: Leet Speak"
challengeType: 28
dashedName: challenge-316
---

# --description--

Given a lowercase string, return it translated into leet speak by replacing the letters below with their leet substitutions:

| Letter | Leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `i` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- Characters with no substitution are left unchanged.

# --hints--

`makeLeet("cool")` should return `"c00l"`.

```js
assert.equal(makeLeet("cool"), "c00l");
```

`makeLeet("leet")` should return `"l337"`.

```js
assert.equal(makeLeet("leet"), "l337");
```

`makeLeet("hacker")` should return `"h4ck3r"`.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` should return `"5473ll173"`.

```js
assert.equal(makeLeet("satellite"), "5473ll173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` should return `"4bcd3fgh1jklmn0pqr57uvwxyz"`.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3fgh1jklmn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return arg;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", i: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
