---
id: 6a15cadf5f240d05a264955a
title: "Challenge 316: 1337 Speak"
challengeType: 28
dashedName: challenge-316
---

# --description--

Given a lowercase string, return it translated into leet speak by replacing the letters below with their leet substitutions:

| Letter | Leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `g` | `9` |
| `i` | `1` |
| `l` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- Characters with no substitution are left unchanged.

# --hints--

`makeLeet("cool")` should return `"c001"`.

```js
assert.equal(makeLeet("cool"), "c001");
```

`makeLeet("leet")` should return `"1337"`.

```js
assert.equal(makeLeet("leet"), "1337");
```

`makeLeet("hacker")` should return `"h4ck3r"`.

```js
assert.equal(makeLeet("hacker"), "h4ck3r");
```

`makeLeet("satellite")` should return `"547311173"`.

```js
assert.equal(makeLeet("satellite"), "547311173");
```

`makeLeet("abcdefghijklmnopqrstuvwxyz")` should return `"4bcd3f9h1jk1mn0pqr57uvwxyz"`.

```js
assert.equal(makeLeet("abcdefghijklmnopqrstuvwxyz"), "4bcd3f9h1jk1mn0pqr57uvwxyz");
```

# --seed--

## --seed-contents--

```js
function makeLeet(str) {

  return str;
}
```

# --solutions--

```js
function makeLeet(str) {
  const map = { a: "4", e: "3", g: "9", i: "1", l: "1", o: "0", t: "7", s: "5" };
  return str.split("").map(char => map[char] || char).join("");
}
```
