---
id: 6a15cadf5f240d05a264955c
title: "Challenge 318: DNA Mutations"
challengeType: 28
dashedName: challenge-318
---

# --description--

Given two DNA strands of equal length, return an array of indexes where the strands differ (mutations).

- DNA strands are strings made up of the characters `"A"`, `"T"`, `"C"`, and `"G"`
- Return the indexes in ascending order
- If there are no mutations, return an empty array

# --hints--

`detectMutations("ATCG", "ATGG")` should return `[2]`.

```js
assert.deepEqual(detectMutations("ATCG", "ATGG"), [2]);
```

`detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC")` should return `[4, 9, 11]`.

```js
assert.deepEqual(detectMutations("ATGCGTACGTTAGC", "ATGCATACGATTGC"), [4, 9, 11]);
```

`detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG")` should return `[]`.

```js
assert.deepEqual(detectMutations("GATCTAGCTAGGCTAGCTAG", "GATCTAGCTAGGCTAGCTAG"), []);
```

`detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG")` should return `[15, 16, 17, 18]`.

```js
assert.deepEqual(detectMutations("TCAGATCATGGCTAGCTACGATCAGCTAGCATGCATATCGACTG", "TCAGATCATGGCTAGAGCTGATCAGCTAGCATGCATATCGACTG"), [15, 16, 17, 18]);
```

`detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT")` should return `[1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]`.

```js
assert.deepEqual(detectMutations("ACGTCAGTACGCACATGACCATTGACATA", "AACGTCAGTACGCACATGACCATTGACAT"), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 20, 21, 23, 24, 25, 26, 27, 28]);
```

# --seed--

## --seed-contents--

```js
function detectMutations(strand1, strand2) {

  return strand1;
}
```

# --solutions--

```js
function detectMutations(strand1, strand2) {
  return strand1
    .split("")
    .reduce((mutations, base, i) => {
      if (base !== strand2[i]) mutations.push(i);
      return mutations;
    }, []);
}
```
