---
id: 6a15cadf5f240d05a2649559
title: "Challenge 315: Summer Solstice"
challengeType: 28
dashedName: challenge-315
---

# --description--

Today is the summer solstice, the longest day of the year in the Northern Hemisphere and the shortest in the Southern. Given a latitude, return a string representing daytime and nighttime hours.

- The latitude will be between 90 (north pole) and -90 (south pole), inclusive
- The number of daytime hours = 12 + (latitude / 90) * 12
- Round the daytime hours to the nearest even number

Return a 24-character string using `"☀️"` for daytime hours and `"🌑"` for nighttime hours, where:

- Each character represents one hour, starting at midnight (hour 0)
- Sunrise and sunset fall symmetrically around noon

For example, a latitude of `0` (the equator) has 12 hours of daylight, so sunrise is at 6:00 AM and sunset is at 6:00 PM. Return: `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

# --hints--

`getDaytimeHours(0)` should return `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(90)` should return `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`.

```js
assert.equal(getDaytimeHours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️");
```

`getDaytimeHours(-90)` should return `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(-33)` should return `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑");
```

`getDaytimeHours(66.5)` should return `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`.

```js
assert.equal(getDaytimeHours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑");
```

`getDaytimeHours(40)` should return `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑");
```

`getDaytimeHours(-50)` should return `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
assert.equal(getDaytimeHours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑");
```

# --seed--

## --seed-contents--

```js
function getDaytimeHours(latitude) {

  return latitude;
}
```

# --solutions--

```js
function getDaytimeHours(latitude) {
  const daylightHours = Math.round((12 + (latitude / 90) * 12) / 2) * 2;
  const sunrise = 12 - daylightHours / 2;
  const sunset = 12 + daylightHours / 2;

  let result = "";
  for (let hour = 0; hour < 24; hour++) {
    result += hour >= sunrise && hour < sunset ? "☀️" : "🌑";
  }

  return result;
}
```
