---
id: 6a15cadf5f240d05a264955e
title: "Challenge 320: Blood Bank"
challengeType: 28
dashedName: challenge-320
---

# --description--

Given an array of the inventory at a blood bank and an array of patient blood type requests, return a string in the format `"X of Y patients served"`. Where `X` is the maximum number of patients that can receive blood from the bank's inventory, and `Y` is the total number of patients.

Each entry in both arrays is one of the following blood types: `"AB"`, `"A"`, `"B"`, or `"O"`.

Compatibility rules:

- `"AB"` can receive from any blood type.
- `"A"` can receive from `"A"` and `"O"`.
- `"B"` can receive from `"B"` and `"O"`.
- `"O"` can only receive from `"O"`.

Duplicate entries in the given arrays represent quantity.

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` should return `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` should return `"3 of 5 patients served"`.

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` should return `"4 of 5 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` should return `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` should return `"8 of 11 patients served"`.

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` should return `"13 of 14 patients served"`.

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
