---
id: 6a19b1062d1b153d8ac76d73
title: "Challenge 325: Lucky Number"
challengeType: 28
dashedName: challenge-325
---

# --description--

Given a string of a person's first and last name, calculate their lucky number using the following rules:

- First and last names are separated by a space
- Find the vowel and consonant count for each name
- Multiply the smaller vowel and consonant counts by each other and then by the length of the smaller name
- Do the same for the two larger counts and the larger name
- Subtract the smaller value from the larger one to get their lucky number

If the final value is zero (`0`), return `13`.

# --hints--

`getLuckyNumber("John Doe")` should return `21`.

```js
assert.equal(getLuckyNumber("John Doe"), 21);
```

`getLuckyNumber("Olivia Lewis")` should return `52`.

```js
assert.equal(getLuckyNumber("Olivia Lewis"), 52);
```

`getLuckyNumber("James Wilson")` should return `18`.

```js
assert.equal(getLuckyNumber("James Wilson"), 18);
```

`getLuckyNumber("Elizabeth Hernandez")` should return `81`.

```js
assert.equal(getLuckyNumber("Elizabeth Hernandez"), 81);
```

`getLuckyNumber("Mike Walker")` should return `32`.

```js
assert.equal(getLuckyNumber("Mike Walker"), 32);
```

`getLuckyNumber("Chloe Perez")` should return `13`.

```js
assert.equal(getLuckyNumber("Chloe Perez"), 13);
```

# --seed--

## --seed-contents--

```js
function getLuckyNumber(name) {

  return name;
}
```

# --solutions--

```js
function getLuckyNumber(name) {
  const vowels = "aeiouAEIOU";

  function nameStats(str) {
    const vowelCount = str.split("").filter(c => vowels.includes(c)).length;
    const consonantCount = str.split("").filter(c => c.match(/[a-zA-Z]/) && !vowels.includes(c)).length;
    return { vowelCount, consonantCount, length: str.length };
  }

  const [first, last] = name.split(" ");
  const firstStats = nameStats(first);
  const lastStats = nameStats(last);

  const largerVowels = Math.max(firstStats.vowelCount, lastStats.vowelCount);
  const smallerVowels = Math.min(firstStats.vowelCount, lastStats.vowelCount);
  const largerConsonants = Math.max(firstStats.consonantCount, lastStats.consonantCount);
  const smallerConsonants = Math.min(firstStats.consonantCount, lastStats.consonantCount);
  const largerLength = Math.max(firstStats.length, lastStats.length);
  const smallerLength = Math.min(firstStats.length, lastStats.length);

  const largerValue = largerVowels * largerConsonants * largerLength;
  const smallerValue = smallerVowels * smallerConsonants * smallerLength;
  const result = largerValue - smallerValue;

  return result === 0 ? 13 : result;
}
```
