---
id: 6a19b1062d1b153d8ac76d72
title: "Challenge 324: Duplicate Character Count"
challengeType: 28
dashedName: challenge-324
---

# --description--

Given two strings, return a count of characters from the second string that can be found in the first.

- Duplicate characters in the second string are counted separately.

# --hints--

`duplicateCharacterCount("aloha", "hei")` should return `1`.

```js
assert.equal(duplicateCharacterCount("aloha", "hei"), 1);
```

`duplicateCharacterCount("jambo", "bonjour")` should return `4`.

```js
assert.equal(duplicateCharacterCount("jambo", "bonjour"), 4);
```

`duplicateCharacterCount("hello", "hola")` should return `3`.

```js
assert.equal(duplicateCharacterCount("hello", "hola"), 3);
```

`duplicateCharacterCount("ola", "hej")` should return `0`.

```js
assert.equal(duplicateCharacterCount("ola", "hej"), 0);
```

`duplicateCharacterCount("ciao", "konnichiwa")` should return `5`.

```js
assert.equal(duplicateCharacterCount("ciao", "konnichiwa"), 5);
```

`duplicateCharacterCount("merhaba", "xin chao")` should return `2`.

```js
assert.equal(duplicateCharacterCount("merhaba", "xin chao"), 2);
```

`duplicateCharacterCount("hello world", "hello to everyone around the world")` should return `26`.

```js
assert.equal(duplicateCharacterCount("hello world", "hello to everyone around the world"), 26);
```

# --seed--

## --seed-contents--

```js
function duplicateCharacterCount(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function duplicateCharacterCount(str1, str2) {
  return str2.split("").filter(char => str1.includes(char)).length;
}
```
