---
id: 6a1d9f98e819ed70a0e994da
title: "Challenge 328: Kaprekar's Routine"
challengeType: 28
dashedName: challenge-328
---

# --description--

Given a 4-digit number, return the number of times you need to apply Kaprekar's routine until reaching 6174.

Kaprekar's routine works as follows:

- Arrange the digits in descending order to form the largest number
- Arrange the digits in ascending order to form the smallest number (pad with leading zeros if necessary)
- Subtract the smaller from the larger
- Repeat with the new number

# --hints--

`kaprekar(1234)` should return `3`.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` should return `6`.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` should return `4`.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` should return `7`.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` should return `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
