---
id: 6a1d9f98e819ed70a0e994d9
title: "Challenge 327: Database Migration"
challengeType: 28
dashedName: challenge-327
---

# --description--

Given two database objects, return the second object with any missing properties from the first filled in.

- Fields that already exist in the record should not be overwritten.

# --hints--

`migrateRecord({ username: "", posts: 0 }, { verified: true })` should return `{ username: "", posts: 0, verified: true }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { verified: true })` should return `{ username: "", posts: 0, verified: true }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { verified: true })` should return `{ username: "", posts: 0, verified: true }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

`migrateRecord({ username: "", posts: 0 }, { verified: true })` should return `{ username: "", posts: 0, verified: true }`.

```js
assert.deepEqual(migrateRecord({ username: "", posts: 0 }, { verified: true }), { username: "", posts: 0, verified: true });
```

# --seed--

## --seed-contents--

```js
function migrateRecord(schema, record) {

  return schema;
}
```

# --solutions--

```js
function migrateRecord(schema, record) {
  return { ...schema, ...record };
}
```
