---
id: 6a1d9f98e819ed70a0e994e1
title: "Challenge 335: Five Dice"
challengeType: 28
dashedName: challenge-335
---

# --description--

Given an array of five dice with values 1-6, return the best possible hand.

Here are the hands ranked lowest to highest:

| Hand | Description |
|------|-------------|
| `"no pair"` | No pair or better |
| `"pair"` | Two dice with the same value |
| `"two pair"` | Two different pairs |
| `"three of a kind"` | Three dice with the same value |
| `"small straight"` | Four consecutive values |
| `"large straight"` | Five consecutive values |
| `"full house"` | Three of a kind and a pair |
| `"four of a kind"` | Four dice with the same value |
| `"five of a kind"` | All five dice with the same value |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` should return `"five of a kind"`.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` should return `"four of a kind"`.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` should return `"large straight"`.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` should return `"three of a kind"`.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` should return `"pair"`.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` should return `"no pair"`.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` should return `"small straight"`.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` should return `"full house"`.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` should return `"two pair"`.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
