---
id: 6a2037a68a0bc2aef0005fff
title: "Challenge 337: Tally Counter"
challengeType: 28
dashedName: challenge-337
---

# --description--

Given a string of tally marks, return the total count represented.

- Each pipe `"|"` represents one count.
- Every fifth mark is represented as a forward slash `"/"`, completing a group of five (`"||||/"`).
- Groups are separated by a space.

# --hints--

`getTallyCount("||||")` should return `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` should return `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` should return `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` should return `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` should return `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
