---
id: 6a2037a68a0bc2aef0006000
title: "Challenge 338: Pet Age Calculator"
challengeType: 28
dashedName: challenge-338
---

# --description--

Given a pet type and age in human years, return the equivalent age in pet years using the following conversion table:

| Pet | Multiplier |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` should return `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` should return `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` should return `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` should return `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` should return `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` should return `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` should return `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
