---
id: 6a2037a68a0bc2aef0006002
title: "Challenge 340: Pig Latin Converter"
challengeType: 28
dashedName: challenge-340
---

# --description--

Given a string, convert it to Pig Latin using the following rules:

- If a word begins with a vowel (`"a"`, `"e"`, `"i"`, `"o"`, or `"u"`), add `"way"` to the end. For example, `"universe"` converts to `"universeway"`.
- If a word begins with one or more consonants, move them to the end and add `"ay"`. For example, `"hello"` converts to `"ellohay"`.
- Preserve the case of the first letter. For example, `"Hello"` converts to `"Ellohay"`.

# --hints--

`pigLatin("universe")` should return `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` should return `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` should return `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` should return `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` should return `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` should return `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
