---
id: 6a22d77ddf034bc4e35b1d58
title: "Challenge 351: Pronic Number"
challengeType: 28
dashedName: challenge-351
---

# --description--

Given a number, determine whether it is a pronic number.

A pronic number is the product of two consecutive integers. For example, 6 is pronic because 2 * 3 = 6.

# --hints--

`isPronic(6)` should return `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` should return `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` should return `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` should return `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` should return `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` should return `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
