---
id: 6a26df95efa55a2524399746
title: "Challenge 365: The Last Challenge"
challengeType: 28
dashedName: challenge-365
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function placeholder(arg) {

  return arg;
}
```

# --solutions--

```js
function placeholder(arg) {

  return arg;
}
```
