---
id: 673b567e3ba535dda140d278
title: Build a Voting System Using Maps and Sets
challengeType: 14
dashedName: build-a-voting-system-using-maps-and-sets
---

# --description--

In this lab, you will build a voting system that would use `Map` to create a poll and `Set` to prevent duplicate voting.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should initialize the `poll` variable to a new `Map` object.

2. You should have a function `addOption` that accepts a parameter `option`.

3. In the `addOption` function:

   - If the `option` does not already exist in the poll, it should be added to the poll with an empty `Set` as its value to track voters. You should also return the message `Option "<option>" added to the poll.`

   - If the `option` already exists, it should return the message `Option "<option>" already exists.`.

4. You should have a function `vote` that accepts two parameters, `option`(the option to vote for) and `voterId` (a unique ID for the voter).

5. In the `vote` function:

   - If the `option` does not exist in the poll, the function should return the message `Option "<option>" does not exist.`;

   - If the `option` exists, the function should check if the `voterId` has already voted for this `option`.

   - If the voter has already voted, the function should display the message `Voter <voterId> has already voted for "<option>".`

   - If the voter has not voted, `voterId` should be added to the `Set` of `voters` for this option. The function should return the message  `Voter <voterId> voted for "<option>".`

6. You should have a function `displayResults` that returns the poll results in the following format:

```js
Poll Results:
Option A: N votes
Option B: N votes
.
.
```

# --hints--


You should have a `poll` variable initialized to a new `Map` object.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

You should define a function `addOption` that accepts a parameter `option`.

```js
assert.isFunction(addOption);
assert.match(addOption.toString(), /\(\s*option\s*\)/);
```

You should define a function `vote` that accepts two parameters, `option` and `voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

You should define a function `displayResults` to display the poll results.

```js
assert.isFunction(displayResults);
```

You should ensure each voting option maps to a `Set` object.

```js
poll.clear();

addOption("Turkey");
assert.instanceOf(poll.get("Turkey"), Set);
```

`addOption("Turkey")` should return `Option "Turkey" added to the poll.`

```js
poll.clear();

assert.equal(addOption("Turkey"), 'Option "Turkey" added to the poll.');
```

Adding an empty option should return `"Option cannot be empty."`

```js
assert.equal(addOption(""), 'Option cannot be empty.');
```

When `"Turkey"` is already added, `addOption("Turkey")` should return `Option "Turkey" already exists.`

```js
poll.clear();

addOption("Turkey");
assert.equal(addOption("Turkey"), 'Option "Turkey" already exists.');
```

When `Turkey` exists in the voting options, `vote("Turkey", "traveler1")` should return `Voter traveler1 voted for "Turkey".`

```js
poll.clear();

addOption("Turkey");
assert.equal(vote("Turkey", "traveler1"), 'Voter traveler1 voted for "Turkey".');
```

`vote` should update the `Set` of voters for an option.

```js
poll.clear();

addOption("Turkey");
vote("Turkey", "traveler1");
assert.isTrue(poll.get("Turkey").has("traveler1"));
```

When `traveler1` tries to vote for `Turkey` again, `vote("Turkey", "traveler1")` should return `Voter traveler1 has already voted for "Turkey".`

```js
poll.clear();

addOption("Turkey");
vote("Turkey", "traveler1");
assert.equal(vote("Turkey", "traveler1"), 'Voter traveler1 has already voted for "Turkey".');
```

Duplicate votes should not increase the size of the `Set`.

```js
poll.clear();

addOption("Turkey");
vote("Turkey", "traveler1");
const voterSet = poll.get("Turkey");
const initialSize = voterSet.size;

vote("Turkey", "traveler1");
assert.equal(voterSet.size, initialSize);
```

When `Morocco` is not in the voting options, `vote("Morocco", "traveler2")` should return `Option "Morocco" does not exist.`

```js
poll.clear();

assert.equal(vote("Morocco", "traveler2"), 'Option "Morocco" does not exist.');
```

A unique option should be able to receive multiple votes.

```js
poll.clear();
  
addOption("Turkey");
assert.equal(vote("Turkey", "traveler1"), 'Voter traveler1 voted for "Turkey".');
assert.equal(vote("Turkey", "traveler2"), 'Voter traveler2 voted for "Turkey".');
assert.equal(poll.get("Turkey").size, 2);
```

`displayResults()` should iterate over the `Map` and return the results in the correct format.

```js
poll.clear();

assert.match(displayResults.toString(), /poll\.(entries|forEach|keys)/);
addOption("Turkey");
addOption("Morocco");

vote("Turkey", "traveler1");
vote("Turkey", "traveler2");
vote("Morocco", "traveler3");

assert.equal(displayResults(), "Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === "") {
    return "Option cannot be empty.";
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}


function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId); 
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
    let results = "Poll Results:\n";
    for (let [option, voters] of poll.entries()) {
        results += `${option}: ${voters.size} votes\n`;
    }
    return results.trim();
}
```
