---
id: 673b567e3ba535dda140d278
title: Build a Voting System Using Maps and Sets
challengeType: 14
dashedName: build-a-voting-system-using-maps-and-sets
---

# --description--

In this lab, you will build a voting system that would use `Map` to create a poll and `Set` to prevent duplicate voting.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should initialize the `poll` variable to a new `Map` object.

2. You should have a function `addOption` that accepts a parameter `option`.

3. In the `addOption` function:

   - If the `option` does not already exist in the poll, it should be added to the poll with an empty `Set` as its value to track voters. You should also return the message `Option "<option>" added to the poll.`

   - If the `option` already exists, it should return the message `Option "<option>" already exists.`.

4. You should have a function `vote` that accepts two parameters, `option`(the option to vote for) and `voterId` (a unique ID for the voter).

5. In the `vote` function:

   - If the `option` does not exist in the poll, the function should return the message `Option "<option>" does not exist.`;

   - If the `option` exists, the function should check if the `voterId` has already voted for this `option`.

   - If the voter has already voted, the function should display the message `Voter <voterId> has already voted for "<option>".`

   - If the voter has not voted, `voterId` should be added to the `Set` of `voters` for this option. The function should return the message  `Voter <voterId> voted for "<option>".`

6. You should have a function `displayResults` that returns the poll results in the following format:

```js
Poll Results:
Option A: N votes
Option B: N votes
.
.
```

# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
// Initialize the poll with a Map
const poll = new Map();

function addOption(option) {
  if (!poll.has(option)) {
    poll.set(option, new Set()); // Each option has a Set to track unique voters
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId); // Add voterId to the Set for this option
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
    let results = "Poll Results:\n";
    for (let [option, voters] of poll.entries()) {
        results += `${option}: ${voters.size} votes\n`;
    }
    return results.trim();
}

// Example usage

// add options
console.log(addOption("Turkey"));
console.log(addOption("Morocco"));
console.log(addOption("Qatar"));

// vote
console.log(vote("Turkey", "traveler1"));
console.log(vote("Morocco", "traveler2"));
console.log(vote("Turkey", "traveler2")); // A second unique voter for Option A
console.log(vote("Qatar", "traveler3"));
console.log(vote("Turkey", "traveler1")); // Duplicate vote attempt for voter1 on Option A

console.log(displayResults());
```
