---
id: 65560ffdb7d05d248e012280
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside the `TreeNode` class, replace `pass` with an `__init__` method so that you can initialize the attributes of the object. Don't add any parameters and use the `pass` keyword to avoid an error.

# --hints--

Визначте метод `__init__` в межах класу `TreeNode`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\([^(]*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
