---
id: 65560f9380be92226084ef46
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви створите бінарне дерево пошуку (БДП). БДП — це структура даних, у якій кожен вузол має не більше двох дочірніх елементів, причому ліва дитина містить значення, менші за батьківський вузол, а права дитина містить значення, більші за батьківський вузол, що дозволяє ефективні операції пошуку та сортування.

Ось так виглядає бінарне дерево пошуку:

<img class="img-responsive center-block" alt="бінарне дерево пошуку" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px;" />

Почніть із визначення порожнього класу `TreeNode`. Клас `TreeNode` представляє вузол у бінарному дереві пошуку. Використайте ключове слово `pass`, щоб заповнити тіло класу та уникнути помилок.

# --hints--

Використайте ключове слово `class`, щоб оголосити порожній клас під назвою `TreeNode`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло класу.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get("TreeNode")
      inspect.isclass(a)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

