---
id: 65560ffdb7d05d248e012280
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Замініть `pass` на метод `__init__` в межах класу `TreeNode`, щоб ініціалізувати атрибути об’єкта. Не додавайте жодних параметрів і використайте ключове слово `pass`, щоб уникнути помилки.

# --hints--

Визначте метод `__init__` в межах класу `TreeNode`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\([^(]*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
