---
id: 65561022956c1024e7184add
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Метод `__init__` приймає два параметри: `self` (представляє екземпляр класу, який створюється) та `key` (значення, яке буде зберігатися у вузлі). Додайте ці два параметри до методу `__init__()`.

# --hints--

Додайте параметр `self` до методу.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Додайте параметр `key` до методу.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
