---
id: 655a4ffc762d117470b94e3b
title: Крок 37
challengeType: 20
dashedName: step-37
---

# --description--

Якщо жодна з попередніх умов не виконується, то вузол має і ліві, і праві дочірні елементи.

Щоб вибрати нащадка, потрібно знайти мінімальне значення в правому піддереві. Найменше значення буде порядковим нащадком поточного вузла.

Пізніше ви визначите допоміжний метод під назвою `_min_value`, який знаходить найменше значення в заданому піддереві.

Наразі додайте виклик `_min_value` після блоку `elif`, передавши `node.right` як аргумент. Призначте його до ключа поточного вузла.

# --hints--

Призначте `self._min_value(node.right)` до `node.key` після блоку `elif`.

```js
assert.match(code, /node\.key\s*=\s*self\._min_value\(\s*node\.right\s*\)/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
--fcc-editable-region--
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left                
--fcc-editable-region--
```
