---
id: 655a50f992ba7177aff2b718
title: Крок 38
challengeType: 20
dashedName: step-38
---

# --description--

Після знаходження мінімального значення вам потрібно буде рекурсивно видалити вузол з мінімальним значенням із правого піддерева.

Цей крок гарантує, що вузол з мінімальним значенням буде видалено з дерева, зберігаючи властивість бінарного дерева пошуку.

# --hints--

Рекурсивно викличте метод `_delete` з аргументами `node.right` та `node.key`.

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

Призначте значення, яке повертає виклик `_delete()`, до правої дитини поточного вузла.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left
--fcc-editable-region--
            node.key = self._min_value(node.right)
--fcc-editable-region--
```
