---
id: 655a13aa8af2510de51f2e1c
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Метод `insert` буде викликаний користувачем. Окрім параметра `self`, йому також знадобиться параметр `key`. Цей параметр буде ключовим значенням, яке буде вставлено в бінарне дерево пошуку.

Додайте `key` як другий параметр до визначення функції.

# --hints--

Метод `insert` має містити параметр `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

Метод `insert` має містити параметр `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
    def insert(self):
        pass
--fcc-editable-region--
```
