---
id: 655a13744e82580d7ee3073d
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Next, you need to define a mechanism to insert nodes in the tree. For that, define an empty `insert` method within the `BinarySearchTree` class and give it a `self` parameter.

# --hints--

Визначте метод `insert` з параметром `self` в межах класу `BinarySearchTree`. Remember the `pass` keyword.

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None


--fcc-editable-region--
```
