---
id: 655a1564f3aa8210938cdf68
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Add an `elif` conditional statement that checks if `key` > `node.key`.

# --hints--

You should add an `elif` conditional statement to check if `key > node.key`. Не забудьте використати ключове слово `pass`.

```js
({
  test: () =>
    assert.match(
      code,
      /^(\s+)if\s+node\s+is\s+None\s*:.*?\1elif\s+key\s*>\s*node\.key\s*:/ms
    )
});
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
--fcc-editable-region--
        if key < node.key:
            node.left = self._insert(node.left, key)
--fcc-editable-region--
```
