---
id: 655a10e4a620fc091ba43b3d
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Видаліть `pass` в межах методу `__init__` та призначте значення параметра `key` до атрибута вузла `key` за допомогою `self.key`.

Це означає, що атрибут `key` екземпляра `TreeNode` буде встановлено на значення, передане під час створення об’єкта.

# --hints--

Призначте значення параметра `key` до атрибута вузла `key` за допомогою `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
