---
id: 655a151cc6041f0ff7d24ded
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Now you need to recursively traverse the tree and insert the values using the principle for binary trees:

- Values smaller than the key are placed in the left subtree
- Values greater than the key are placed in the right subtree

After your existing conditional statement, write another `if` statement to check if `key` is strictly less than `node.key`.

# --hints--

You should write another if statement to check if `key` is strictly less than `node.key`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s*:.*?\1if\s+key\s*<\s*node\.key\s*:/ms) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)                    
--fcc-editable-region--
```
