---
id: 655a47a9404d856743c7f529
title: Крок 32
challengeType: 20
dashedName: step-32
---

# --description--

Замініть `pass` в межах блоку `if` на виклик методу `_delete`, передавши ліву дитину поточного вузла та `key` як аргументи. Призначте виклик функції до лівого вузла.

# --hints--

Викличте метод `_delete` з аргументами `node.left` та `key`.

```js
assert.match(code, /self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

Призначте результат виклику `_delete()` до лівої дитини поточного вузла (`node.left`).

```js
assert.match(code, /node\.left\s*=\s*self\._delete\(\s*node\.left\s*,\s*key\s*\)/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            pass
--fcc-editable-region--
```
