---
id: 655a46aa8e10c26218c5034c
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Inside the `BinarySearchTree` class, define a new helper method called `_delete`  that takes three parameters: `self`, `node`, and `key`.

# --hints--

Визначте метод `_delete` в межах класу `BinarySearchTree` з параметрами `self`, `node` та `key`. Не забудьте використати ключове слово `pass`.

```js
assert.match(code, /def\s+_delete\(\s*self\s*,\s*node\s*,\s*key\s*\)/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
