---
id: 655a153a6b362d103e125028
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

If `key` < `node.key` returns `True`, then the new node should be placed in the left subtree.

Delete `pass` and recursively call the `_insert` method on the left child of the current node. Then, assign the result to the `left` attribute of the current node.

# --hints--

You should call the `self._insert` method passing `node.left` and `key` as the arguments.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

You should assign the result of your `self._insert()` call to the `left` attribute of the current node.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            pass
--fcc-editable-region--
```
