---
id: 655a153a6b362d103e125028
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Якщо `key` < `node.key` повертає `True`, то новий вузол потрібно розмістити в лівому піддереві.

Видаліть `pass` та рекурсивно викличте метод `_insert` на лівій дитині поточного вузла. Потім призначте результат до атрибута поточного вузла `left`.

# --hints--

Викличте метод `self._insert`, передавши `node.left` та `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Призначте результат виклику `self._insert()` до атрибута поточного вузла `left`.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            pass
--fcc-editable-region--
```
