---
id: 655a54cd88e33b8646c67e16
title: Крок 46
challengeType: 20
dashedName: step-46
---

# --description--

Within the `inorder_traversal` method, start the in-order traversal by calling the helper method `_inorder_traversal` and pass the BST root and the `result` list as the arguments.

This will start the traversal from the root of the binary search tree (`self.root`), and the `result` list will be passed to accumulate the keys during the traversal.

# --hints--

You should call `_inorder_traversal` and pass `self.root` and `result` as the arguments.

```js
({ test: () => assert.match(code, /self\._inorder_traversal\(\s*self\.root\s*,\s*result\s*\)/) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)
        return node

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

--fcc-editable-region--
    def inorder_traversal(self):
        result = []
--fcc-editable-region--
```
