---
id: 655a227e57aabb25d1f9c987
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Тепер настав час попрацювати над пошуком.

Визначте метод під назвою `search` у класі `BinarySearchTree`. Надайте методу `search` два параметри: `self` та `key`.


# --hints--

Визначте метод під назвою `search`. Пам’ятайте про ключове слово `pass`.

```js
({ test: () => assert.match(code, /def\s+search/) });
```

Метод `search` має приймати два параметри: `self` та `key`.

```js
({ test: () => assert.match(code, /def\s+search\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--

--fcc-editable-region-- 
```
