---
id: 655a51a705c97d7a9294ab2a
title: Крок 40
challengeType: 20
dashedName: step-40
---

# --description--

Якщо вузол, який ви хочете видалити, має 2 дочірні елементи, то з них потрібно вибрати вузол для заміни. Метод порядкового нащадка вибирає найменший елемент з правого піддерева та розміщує цей елемент на місці видаленого вузла.

Визначте метод `_min_value` і надайте йому два параметри: `self` та `node`.

# --hints--

Визначте метод `_min_value` з параметрами `self` та `node`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert.match(code, /def\s+_min_value\(\s*self\s*,\s*node\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right

            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)
        return node
--fcc-editable-region--

--fcc-editable-region--
```
