---
id: 655a142ae611b30e5df0ac16
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Now, inside the `insert` method, you need to call another method `_insert()` that performs the actual insertion. You are going to define the `_insert` method soon.

Delete `pass` and assign `self._insert(self.root, key)` to `self.root`.

Note that:

- `self.root` passes the root node of the tree as the first argument. This is the starting point for the insertion process.
- `key`: passes the `key` value you want to insert as the second argument.

# --hints--

You should recursively call the `_insert()` method using `self._insert()`

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "insert");
    const {function_body} = ins;    
    assert(function_body.match(/self\._insert\s*\([^(]*\)/));
  }
})
```

Передайте `self.root` та `key` до виклику `_insert()`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "insert");
    const {function_body} = ins;    
    assert(function_body.match(/self\._insert\s*\(\s*self\.root\s*,\s*key\s*\)/));
  }
})
```

You should assign the return value of your `_insert()` call to `self.root`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}self\.root\s*=\s*self\._insert\s*\(\s*self\.root\s*,\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
    def insert(self,key):
        pass
--fcc-editable-region--
```
