---
id: 655a4614304cd36031cb4e75
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

The deletion operation might result in a new root (for example, if the node to be deleted is the current root).

To handle this case, assign the result of the `_delete` call to `self.root`.

# --hints--

Призначте результат виклику `_delete()` до `self.root`.

```js
({ test: () => assert.match(code, /self\.root\s*=\s*self\._delete\(\s*self\.root\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
--fcc-editable-region--
    def delete(self, key):
        self._delete(self.root, key)
--fcc-editable-region--
```
