---
id: 655a158c7e80c810f6eff1e1
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Замініть `pass` на виклик методу `_insert` на правій дитині поточного вузла в межах умови `elif`. Призначте результат до атрибута поточного вузла `right`.

# --hints--

Викличте метод `self._insert()`, передавши `node.right` та `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Призначте результат виклику `self._insert()` до атрибута поточного вузла `right` (`node.right`).

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
--fcc-editable-region--
        elif key > node.key:
            pass
--fcc-editable-region--
```
