---
id: 655a48da2c25656d2f7bab4d
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

Inside the `else` block, replace `pass` with an `if` statement to check if the left child of the current node is `None`.

When `node.left` is `None`, there is no left child. Therefore, return the right child from the new `if` block as a replacement.

# --hints--

You should write an `if` statement that checks if `node.left` is `None`.

```js
assert.match(code, /if\s+node\.left\s*is\s*None/);
```

You should return the right child of the current node.

```js 
assert.match(code, /return\s+node\.right/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
--fcc-editable-region--
        else:
            pass
--fcc-editable-region--
```
