---
id: 655a194276dfa11460f7b5e3
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Тепер ви визначите базовий випадок для рекурсивного пошуку. Напишіть інструкцію `if`, яка перевіряє дві умови:

- якщо `node` є `None`: це вказує на те, що пошук досягнув кінця гілки, не знайшовши ключа;
- якщо `node.key == key`: це означає, що ключ знайдено у поточному вузлі.

Поєднайте дві умови за допомогою оператора `or` та поверніть поточний вузол у блоці `if`.

# --hints--

Напишіть інструкцію `if`, яка перевіряє, чи `node` має значення `None` або `node.key` дорівнює `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Поверніть `node` з блоку `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        pass
--fcc-editable-region--
```
