---
id: 655a13744e82580d7ee3073d
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Далі вам потрібно визначити механізм вставки вузлів у дерево. Для цього визначте порожній метод `insert` у класі `BinarySearchTree` та надайте йому параметр `self`.

# --hints--

Визначте метод `insert` з параметром `self` в межах класу `BinarySearchTree`. Пам’ятайте про ключове слово `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None


--fcc-editable-region--
```
