---
id: 655a48da2c25656d2f7bab4d
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

Замініть `pass` в межах блоку `else` на інструкцію `if`, яка перевірятиме, чи ліва дитина поточного вузла є `None`.

Якщо `node.left` є `None`, то лівої дитини немає. В такому випадку поверніть праву дитину з нового блоку `if` як заміну.

# --hints--

Напишіть інструкцію `if`, яка перевіряє, чи `node.left` є `None`.

```js
assert.match(code, /if\s+node\.left\s*is\s*None/);
```

Поверніть праву дитину поточного вузла.

```js 
assert.match(code, /return\s+node\.right/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
--fcc-editable-region--
        else:
            pass
--fcc-editable-region--
```
