---
id: 655a151cc6041f0ff7d24ded
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Тепер вам потрібно рекурсивно пройти по дереву та вставити значення за принципом бінарних дерев:

- значення, менші за ключ, розміщуються в лівому піддереві
- значення, більші за ключ, розміщуються в правому піддереві

Після наявної умовної інструкції напишіть ще одну інструкцію `if`, щоб перевірити, чи `key` строго менший за `node.key`.

# --hints--

Напишіть ще одну інструкцію `if`, щоб перевірити, чи `key` строго менший за `node.key`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s*:.*?\1if\s+key\s*<\s*node\.key\s*:/ms) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)                    
--fcc-editable-region--
```
