---
id: 655a1188d6cfbb0a3ec63c57
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Inside the `__init__` method, initialize the `left` and `right` attributes of the node to `None`. This is because when a node is first created, it doesn't have any left or right children. Не забудьте використати ключове слово `self`.

# --hints--

You should assign `None` to the `left` attribute of the node using `self.left`.

```js
({ test: () => assert.match(code, /^\s{8}self\.left\s*=\s*None/m) })
```

You should assign `None` to the `right` attribute of the node using `self.right`.

```js
({ test: () => assert.match(code, /^\s{8}self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
