---
id: 655a12c2ab0bcc0c2ba30e16
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Within the `BinarySearchTree` class, replace `pass` with an `__init__` method and add a `self` parameter to this method.

# --hints--

Визначте метод `__init__` та додайте до нього параметр `self`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
