---
id: 655a11eb4e54b60acd6bd641
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Create another empty class called `BinarySearchTree` that represents a binary search tree.

# --hints--

Ви повинні мати клас під назвою `BinarySearchTree`. Не забудьте використати ключове слово `pass`.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      bst = __locals.get("BinarySearchTree")
      inspect.isclass(bst)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--

--fcc-editable-region--
```
