---
id: 655a489d83b1996bd537b153
title: Крок 34
challengeType: 20
dashedName: step-34
---

# --description--

Умовні інструкції, які ви написали в попередніх кроках, дійсні для вузлів, які мають 0 або 1 дочірній елемент. Майбутні умови будуть дійсними, якщо є 2 дочірні елементи.

Для цього додайте умову `else` до умовної інструкції.

# --hints--

Додайте умову `else`. Не забудьте використати ключове слово `pass`.

```js
assert.match(code, /else\s*:/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)            
--fcc-editable-region--
```
