---
id: 655a173e5b8adc13b761ed74
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Now, define the `_search` method with three parameters, namely `self`,`node` and `key`.

# --hints--

Оголосіть метод `_search`. Note the leading underscore and remember to use `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Метод `_search` має приймати три параметри: `self`, `node` та `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--

--fcc-editable-region--
```
