---
id: 655a167ea3e96512bf1343ce
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Видаліть `pass` в межах методу `search` та викличте приватний допоміжний метод `_search`, використавши аргументи нижче. Ви визначите `_search` у наступних кроках.

- `self.root`: це корінь бінарного дерева пошуку. Пошук починається з кореня.
- `key`: це значення, яке користувач хоче знайти у бінарному дереві пошуку.

Внутрішньо, `search` доручає пошук приватному методу `_search`, який виконує рекурсивний пошук в межах бінарного дерева пошуку.

# --hints--

Викличте метод `_search` в межах методу `search`.

```js
({ test: () => assert.match(code, /self\._search\([^(]*\)/) });
```

Викличте метод `_search`, передавши `self.root` та `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._search\(\s*self\.root\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--
    def search(self, key):
        pass
--fcc-editable-region--
```
