---
id: 655a132b56e0160cfeca08d4
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Видаліть `pass` з методу `__init__` та встановіть значення атрибута екземпляра під назвою `root` на `None`.

Атрибут `root` представляє кореневий вузол бінарного дерева пошуку. Оскільки під час створення нового об’єкта `BinarySearchTree` він є конструктором, то спочатку дерево порожнє, тому атрибут `root` має значення `None`.

# --hints--

Встановіть значення атрибута `root` на `None` за допомогою `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
