---
id: 65c4f2b7178afed8e88f782f
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Якщо `key` < `node.key` оцінюється як `True`, то новий вузол потрібно розмістити в лівому піддереві.

Видаліть `pass` та рекурсивно викличте метод `_insert` з лівою дитиною як першим аргументом і `key` як другим аргументом. Призначте результат до атрибута поточного вузла `left`.

# --hints--

Видаліть ключове слово `pass` з блоку `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Викличте метод `self._insert`, передавши `node.left` та `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Призначте результат виклику `self._insert()` до атрибута поточного вузла `left`.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
