---
id: 65c4f2851645e8d84d14f5d1
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер вам потрібно рекурсивно пройти по дереву та вставити значення за принципом бінарних дерев:

- значення, менші за ключ, розміщуються в лівому піддереві
- значення, більші за ключ, розміщуються в правому піддереві

Після наявної умовної інструкції напишіть ще одну інструкцію `if`, щоб перевірити, чи `key` менший за `node.key`.

# --hints--

Напишіть ще одну інструкцію `if`, щоб перевірити, чи `key` менший за `node.key`. Не забудьте використати ключове слово `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

