---
id: 65c4f300da28d8da361bfa93
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Замініть `pass` в межах умови `elif` на виклик до методу `_insert`, першим аргументом якого є права дитина поточного вузла, а другим — `key`. Призначте результат до атрибута поточного вузла `right`.

# --hints--

Видаліть ключове слово `pass` з блоку `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Викличте метод `self._insert()`, передавши `node.right` та `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Призначте результат виклику `self._insert()` до атрибута поточного вузла `right` (`node.right`).

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
