---
id: 65c4f07da204f4d2f325fbd3
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Призначте значення `None` до атрибутів вузла `left` та `right` в межах методу `__init__`. Причина в тому, що коли вузол створюється вперше, він не має лівих чи правих дітей. Не забудьте використати ключове слово `self`.

# --hints--

Призначте `None` до атрибута вузла `left` за допомогою `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Призначте `None` до атрибута вузла `right` за допомогою `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
