---
id: 65c4f0b82db314d3fcc0b8e1
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Замініть `pass` на метод `__init__` в межах класу `BinarySearchTree` та додайте до цього методу параметр `self`.

# --hints--

Видаліть ключове слово `pass` з класу `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


Визначте метод `__init__` та додайте до нього параметр `self`. Не забудьте використати ключове слово `pass` в межах тіла методу.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
