---
id: 65c4eb814cc977c95cd7df0e
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви створите бінарне дерево пошуку (БДП). БДП — це структура даних, у якій кожен вузол має не більше двох дочірніх елементів, причому ліва дитина містить значення, менші за батьківський вузол, а права дитина містить значення, більші за батьківський вузол, що дозволяє ефективні операції пошуку та сортування.

Ось так виглядає бінарне дерево пошуку:

<img class="img-responsive center-block" alt="бінарне дерево пошуку" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px;" />

Для початку визначте порожній клас `TreeNode`. Клас `TreeNode` представляє вузол у бінарному дереві пошуку. Використайте ключове слово `pass`, щоб заповнити тіло класу та уникнути помилок.

# --hints--

Використайте ключове слово `class`, щоб оголосити порожній клас під назвою `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
