---
id: 65c4ef8463f869d0eaf87c67
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Замініть `pass` на метод `__init__` в межах класу `TreeNode`, щоб ініціалізувати атрибути об’єкта. Поки не додавайте жодних параметрів.

# --hints--

Видаліть ключове слово `pass` з класу `TreeNode` та перемістіть його до методу `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Визначте метод `__init__` в межах класу `TreeNode`. Не забудьте використати ключове слово `pass`.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
