---
id: 65c4f02dbd3557d23f12db62
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Видаліть `pass` в межах методу `__init__` та призначте значення параметра `key` до атрибута вузла `key` за допомогою `self.key`.

Це означає, що атрибут `key` екземпляра `TreeNode` буде встановлено на значення, передане під час створення об’єкта.

# --hints--

Видаліть інструкцію `pass` з методу `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Призначте значення параметра `key` до атрибута вузла `key` за допомогою `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
