---
id: 65c4f37ba6cbcfdc77e24165
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Метод `insert` буде викликаний користувачем. Окрім параметра `self`, йому також знадобиться параметр `key`. Цей параметр буде ключовим значенням, яке буде вставлено в бінарне дерево пошуку.

Додайте `key` як другий параметр до визначення функції.

# --hints--

Метод `insert` має містити параметр `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

Метод `insert` має містити параметр `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
