---
id: 65c63da2ed6769a10e141341
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Тепер ви визначите базовий випадок для рекурсивного пошуку. Видаліть наявне ключове слово `pass` та напишіть інструкцію `if`, яка перевіряє дві умови:

- чи `node` є `None`: це вказує на те, що пошук досягнув кінця гілки, не знайшовши ключа;
- чи `node.key == key`: це означає, що ключ знайдено у поточному вузлі.

Поєднайте дві умови за допомогою оператора `or` та поверніть поточний вузол у блоці `if`.

# --hints--

Видаліть ключове слово `pass` з методу `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Напишіть інструкцію `if`, яка перевіряє, чи `node` має значення `None` або `node.key` дорівнює `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Поверніть `node` з блоку `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
