---
id: 65c6409418806da8d0636ffc
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Функціональності `insert` та `search` готові, час їх протестувати.

Зверніть увагу, що наразі вузли не відсортовані, а лише вставлені. Над сортуванням ви працюватимете пізніше, використовуючи серединний обхід.

Екземпляр класу в Python можна створити ось так:

```python
object_name = ClassName()
```

Створіть екземпляр класу `BinarySearchTree` поза визначеннями класів та призначте його до змінної `bst`.


# --hints--

Створіть екземпляр класу `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Призначте новий екземпляр `BinarySearchTree` до змінної `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
