---
id: 65c644829cfb63acf3479d09
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

Вставте вузли зі списку `nodes` до екземпляра `bst`, ітерувавши над списком та викликавши метод `insert` на кожному вузлі зі списку. Використайте `node` як змінну ітерації.

# --hints--

Ви повинні мати цикл `for node in nodes`.

```js
assert.match(code, /for\s+node\s+in\s+nodes/);
```

Цикл `for` повинен містити `bst.insert(node)`.

```js
assert.match(code, /bst\.insert/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

--fcc-editable-region--
```
