---
id: 65c63a4c4da62e9ae18e321a
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Тепер настав час попрацювати над пошуком. Так само, як ви створили допоміжний метод `_insert` для методу `insert`, вам потрібно створити допоміжний метод `_search` для методу `search`.

Визначте метод `_search` з трьома параметрами, а саме `self`, `node` та `key`.

# --hints--

Оголосіть метод `_search`. Зверніть увагу на знак підкреслення напочатку та не забудьте використати `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Метод `_search` має приймати три параметри: `self`, `node` та `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
