---
id: 65c63e6962159fa372ecdbec
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Якщо друга інструкція `if` не `True`, то цільовий ключ більший або дорівнює ключу поточного вузла.

Якщо цільовий ключ більший за ключ поточного вузла в бінарному дереві пошуку, то пошук продовжується в правому піддереві.

Поверніть результат виклику методу `_search`, аргументами якого є права дитина поточного вузла та ключ, після блоку `if`.

# --hints--

Рекурсивно викличте метод `_search` за допомогою `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Передайте два аргументи до виклику методу `_search`: `node.right` та `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Поверніть результат виклику методу `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
