---
id: 65c63eb01b9563a4b8a046b6
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Тепер визначте метод під назвою `search` в межах класу `BinarySearchTree`. Надайте методу `search` два параметри: `self` та `key`.

# --hints--

Визначте метод під назвою `search`. Пам’ятайте про ключове слово `pass`.

```js
({ test: () => assert.match(code, /def\s+search/) });
```

Метод `search` має приймати два параметри: `self` та `key`.

```js
({ test: () => assert.match(code, /def\s+search\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

--fcc-editable-region--

--fcc-editable-region--
```
