---
id: 65c63fd9b387f0a6c5bf6a72
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Видаліть `pass` в межах методу `search` та викличте допоміжний метод `_search`, використавши аргументи нижче.

- `self.root`: це корінь бінарного дерева пошуку. Пошук починається з кореня.
- `key`: це значення, яке користувач хоче знайти у бінарному дереві пошуку.

Внутрішньо, `search` доручає пошук допоміжному методу `_search`, який виконує рекурсивний пошук в межах бінарного дерева пошуку.

# --hints--

Видаліть ключове слово `pass` з методу `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Викличте метод `_search` в межах методу `search`.

```js
({ test: () => assert.match(code, /self\._search\([^(]*\)/) });
```

Викличте метод `_search`, передавши `self.root` та `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._search\(\s*self\.root\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def search(self, key):
        pass
--fcc-editable-region--
```
