---
id: 65c645b838e7deb080fc25e0
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Тепер вузли вставлено. Щоб перевірити, чи все правильно, можна знайти вузол в дереві.

Поза межами циклу `for`, знайдіть вузол `80` в екземплярі `bst` та додайте його до виклику `print`. Першим аргументом функції `print` має бути `'Search for 80:'`.

# --hints--

Надрукуйте результат виклику `bst.search(80)`. Першим аргументом інструкції `print` має бути `'Search for 80:'`.

```js
assert.match(code, /^print\(('|")Search for 80:('|"),\s*bst\.search\(80\)/gm);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

--fcc-editable-region--
```
