---
id: 65ca0e8eb2c9c215269d6a66
title: Крок 54
challengeType: 20
dashedName: step-54
---

# --description--

Метод `inorder_traversal` відповідає за виконання серединного обходу бінарного дерева пошуку. Він повертає ключі вузлів у відсортованому порядку. Для цього він буде використовувати допоміжний метод `_inorder_traversal`.

Створіть метод `inorder_traversal`, який приймає `self` як параметр.

# --hints--

Визначте метод `inorder_traversal`, який приймає `self` як параметр. Не забудьте використати `pass`.

```js
({ test: () => assert.match(code, /def\s+inorder_traversal\(\s*self\s*\)/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   

            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   

        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

    def _inorder_traversal(self, node, result):
        if node:
            self._inorder_traversal(node.left, result)
            result.append(node.key)
            self._inorder_traversal(node.right, result)
--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
