---
id: 6553f4f66099802c6ae94613
title: 步骤 50
challengeType: 20
dashedName: step-50
---

# --description--

Give `message` and `offset` values, by passing `text` and `shift` as *arguments* to the `caesar` function call.

# --hints--

You should pass `text` and `shift` as the arguments to your function call.

```js
({ test: () =>  assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar()
--fcc-editable-region--
```
