---
id: 65549561463f0016876e852c
title: 步骤 55
challengeType: 20
dashedName: step-55
---

# --description--

Since your key is shorter than the text that you will have to encrypt, you will need to repeat it to generate the whole encrypted text. At the beginning of your function body, declare a `key_index` variable and set it to zero.

# --hints--

You should declare a variable called `key_index`.

```js
({ 
  test: () => {
    const vigenere = __helpers.python.getDef(code, "vigenere");
    const {function_body} = vigenere;    
    assert(function_body.match(/key_index\s*=/));
  }
})
```

You should assign `0` to your `key_index` variable.

```js
({ 
  test: () => {
    const vigenere = __helpers.python.getDef(code, "vigenere");
    const {function_body} = vigenere;    
    assert(function_body.match(/key_index\s*=\s*0/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
