---
id: 6555e547c18a2b1a7b795bd8
title: 步骤 90
challengeType: 20
dashedName: step-90
---

# --description--

Add two newline characters to your last `print()` call: one at the beginning and the other one at the end of the string.

# --hints--

You should modify the `print(f'Decrypted text: {decryption}')` call by adding a `\n` character at the beginning of the string and another one at the end.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')\\nDecrypted\stext:\s\{\s*decryption\s*\}\\n\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'\nEncrypted text: {text}')
print(f'Key: {custom_key}')
decryption = decrypt(text, custom_key)
print(f'Decrypted text: {decryption}')
--fcc-editable-region--
```
