---
id: 61437d575fb98f57fa8f7f36
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

從你的標準 HTML 模板開始。 添加 `DOCTYPE` 聲明、語言設置爲英語的 `html` 元素、 `head` 和 `body` 元素。

在 `head` 元素中添加一個具有適當 `charset` 的 `<meta>` 標籤和一個用於移動響應的 `<meta>`標籤。

# --hints--

你的代碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該把文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `DOCTYPE` 聲明中的類型後面加上一個 `>`。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你應該有一個 `<html>` 開始標籤，它應該有 `lang` 值爲 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

確保 `html` 元素有一個閉合標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 聲明應該位於 HTML 的開始行。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你需要一個起始標籤 `<head>`

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個閉合標籤 `</head>`

```js
assert(code.match(/<\/head\s*>/i));
```

您應該有一個 `<body>` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個 `</body>` 閉合標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一個 `meta` 元素應將 `name` 設置爲 `viewport`，並將 `content` 設置爲 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應該將 `charset` 屬性設置爲 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
