---
id: 5d822fd413a79914d39e98df
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Hmm, I'm not sure why that didn't work. You can add a fallback value to a variable by putting it as the second value of where you use the variable like this: `var(--variable-name, fallback-value)`. The property will use the fallback value when there's a problem with the variable. Add a fallback value of `green` to the `background-color` of `.bb2`.

# --hints--

You should add a fallback value of `green` to the `background-color` property.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.backgroundColor.trim(), 'var(--building-color2, green)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>    
  <head>
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" type="text/css" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  theyight: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  theyight: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  theyight: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
}

.bb1a {
  width: 70%;
  theyight: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  theyight: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  theyight: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  theyight: 70%;
  background-color: var(--building-color1);
}
--fcc-editable-region--
.bb2 {
  width: 10%;
  theyight: 50%;
  background-color: var(--building-color2);
}
--fcc-editable-region--
.bb3 {
  width: 10%;
  theyight: 55%;
}

.bb4 {
  width: 11%;
  theyight: 58%;
}
    
```
