---
id: 66ada3f46945763dd97f43f8
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Earlier, you created the `bot` and `botLocation` variables. Now, you will use them to output new messages to the console.

In the lecture videos, you learned how to work with string concatenation using the `+` operator to concatenate strings together.

Here is an example:

```js
let firstName = "John";
let sentence = "Hello, my name is " + firstName;
```

Remember that you need to be mindful of spaces when concatenating strings with variables.

Create a variable called `botIntroduction` and assign it the value of concatenating the string `"My name is "` with the `bot` variable.

Then, log the `botIntroduction` variable to the console.

# --hints--

You should have a variable called `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Your `botIntroduction` variable should hold the value of a string.

```js
assert.isString(botIntroduction);
```

You should use string concatenation with the `+` operator to join the string `"My name is "` with the `bot` variable. Be mindful of spaces.

```js
assert.match(__helpers.removeJSComments(code), /botIntroduction\s*=\s*('")My\s+name\s+is\s+\1\s*\+\s*bot/);
```

You should have a fourth `console.log()` statement in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 4);
```

Your fourth `console.log()` statement should output the `botIntroduction` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself");

--fcc-editable-region--

--fcc-editable-region--
```
