---
id: 66ad9c633fa26d3c1475eae3
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now, it is time to add another message from the bot. 

Add another `console` statement to the code that logs the message `"Allow me to introduce myself."`.

# --hints--

You should have a `console` statement.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Your `console` statement should have the message `"Allow me to introduce myself."`. Don't forget the quotes.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
