---
id: 66ad8d150264db3926eccfeb
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In the previous lecture videos, you learned about the `let` keyword and how to declare variables in JavaScript. 

Remember that a variable is a container for storing data values and declaring a variable means creating a new variable in the program.

Here is an example of declaring a variable using the `let` keyword:

```js
let greeting;
```

Use the `let` keyword to declare a variable called `bot`.

# --hints--

You should have the `let` keyword in your code.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

You should declare a variable called `bot`.

```js
assert.isNotNull(bot);
```

You should use the `let` keyword to declare the variable.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you");
--fcc-editable-region--

--fcc-editable-region--
```
