---
id: 66adc5919c1853448f119d11
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Now it looks like the bot wants to change their nickname.

Using reassignment, assign the string `"awesomeTeacherBot"` to the `bot` variable.

# --hints--

You should assign the string `"awesomeTeacherBot"` to the `bot` variable.

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
