---
id: 66ad8ab945d266383f318cbf
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Now you should see the first message from the bot in the console. 

It is time to add a second message from the bot.

Add another `console.log` statement to output the message `"I am excited to talk to you."` to the console.

# --hints--

You should have a `console.log()` statement in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

Your `console` statement should have the message `"I am excited to talk to you."`. Don't forget the quotes.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
