---
id: 66ada3f46945763dd97f43f8
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Earlier, you created the `bot` and `botLocation` variables. Now, you will use them to output new messages to the console.

In the lecture videos, you learned how to work with string concatenation using the `+` operator to concatenate strings together like this:

```js
let firstName = "John";

// result: "Hello, my name is John."
"Hello, my name is " + firstName + ".";
```

Remember that you need to be mindful of spaces when concatenating strings with variables.

Create a variable called `botIntroduction`.

Then use string concatenation with the `+` operator to join the string `"My name is "` followed by the `bot` variable followed by a period (`.`). 

Assign this value to the `botIntroduction` variable.

Then, log the `botIntroduction` variable to the console.

# --hints--

You should have a variable called `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Your `botIntroduction` variable should hold the value of a string.

```js
assert.isString(botIntroduction);
```

You should use string concatenation with the `+` operator to join the string `"My name is "` with the `bot` variable followed by a period (`.`). Be mindful of spaces.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

You should have a `console` statement in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Your `console` statement should output the `botIntroduction` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
