---
id: 61537c5f81f0cf325b4a854c
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Füge ein `header`-Element innerhalb des `body`-Elements hinzu und weise ihm eine Klasse von `header` zu.

Erstelle im `header` einen `h1` mit `css flexbox photo gallery` als Text.

# --hints--

Du solltest ein `header`-Element innerhalb deines `body`-Elements haben.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Your `header` element should have a `class` with `header` as the value.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Dein `header`-Element sollte ein `h1`-Element haben.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Dein `h1`-Element sollte den Text `css flexbox photo gallery` haben.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
