---
id: 61537bb9b1a29430ac15ad38
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Füge in deinem `head`-Element ein `title`-Element mit dem Text `Photo Gallery` und ein `link`-Element hinzu, um deine `styles.css`-Datei zur Seite hinzuzufügen.

# --hints--

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Dein Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein Projekt sollte einen Titel namens `Photo Gallery` tragen.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
