---
id: 6153893900438b4643590367
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Richte deinen `.header`-Text mittig aus. Schreibe den Text in Großbuchstaben unter der Verwendung der `text-transform`Eigenschaft mit dem Wert `uppercase`.

Weise ihm auf allen Seiten ein Padding von `32px` zu. Setze die Farbwerte des Hintergrundes auf `#0a0a23` und des Textes auf `#fff`.

Füge eine `border-bottom` mit `4px solid #fdb347` als Wert hinzu.

# --hints--

Du solltest einen `.header`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Dein `.header`-Selektor sollte eine `text-align`-Eigenschaft auf `center` als Wert setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Dein `.header`-Selektor sollte eine `text-transform`-Eigenschaft auf `uppercase` als Wert setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Dein `.header`-Selektor sollte eine `padding`-Eigenschaft auf `32px` als Wert setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Dein `.header`-Selektor sollte eine `background-color`-Eigenschaft auf `#0a0a23` als Wert setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Dein `.header`-Selektor sollte eine `color`-Eigenschaft auf `#fff` als Wert setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Dein `.header`-Selektor sollte eine `border-bottom`-Eigenschaft auf `4px solid #fdb347` als Wert setzen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
