---
id: 6153a04847abee57a3a406ac
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Deine Bilder benötigen etwas Abstand dazwischen.

Die zusammenfassende `gap`-CSS-Eigenschaft legt die Lücken fest, die auch als Zwischenraum zwischen den Zeilen und Spalten bekannt sind. Die `gap`-Eigenschaft und ihre `row-gap` und `column-gap`-Untereigenschaften bieten diese Funktionalität dem Flex-, Gitter- und Mehrspalten-Layout. Du wendest die Eigenschaft auf das Containerelement an.

Gib deinem `.gallery`-Flexcontainer eine `gap`-Eigenschaft mit `16px` als Wert.

# --hints--

Dein `.gallery`-Selektor sollte eine `gap`-Eigenschaft mit `16px` als Wert haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.gap === '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
}
```
