---
id: 61537d86bdc3dd343688fceb
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Gib als Nächstes jedem `img` ein `src`-Attribut entsprechend seiner Reihenfolge im Dokument. Das erste `img` sollte einen `src` von `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` haben. Der Rest sollte der gleiche sein, außer das `1` durch die Nummer ersetzt werden soll, die `img` im Dokument hat.

# --hints--

Alle neun der `img`-Elemente sollten ein `src`-Attribut haben.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

Dein neues `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Dein zweites `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` als `src`-Attribut gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Dein drittes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` als `src`-Attribut gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Dein viertes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Dein fünftes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Dein sechstes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Dein siebtes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Dein achtes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Dein neuntes `img`-Element sollte `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` als `src`-Attributwert gesetzt haben.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
