---
id: 6153897c27f6334716ee5abe
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Align your `.header` text in the center. Make the text uppercase using the `text-transform` property with `uppercase` as the value.

Gib ihm auf allen Seiten ein Padding von `32px`. Setze die Farbwerte des Hintergrundes auf `#0a0a23` und des Textes auf `#fff`.

Füge einen `border-bottom` mit `4px solid #fdb347` als Wert hinzu.


# --hints--

Du solltest einen `.header`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Your `.header` selector should have a `text-align` property set to `center` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign === 'center');
```

Your `.header` selector should have a `text-transform` property set to `uppercase` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform === 'uppercase');
```

Your `.header` selector should have a `padding` property set to `32px` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.padding === '32px');
```

Your `.header` selector should have a `background-color` property set to `#0a0a23` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor === 'rgb(10, 10, 35)');
```

Your `.header` selector should have a `color` property set to `#fff` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.color === 'rgb(255, 255, 255)');
```

Your `.header` selector should have a `border-bottom` property set to `4px solid #fdb347` as the value.

```js
console.log(new __helpers.CSSHelp(document).getStyle('.header'))
assert(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom === '4px solid rgb(253, 179, 71)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
