---
id: 61537485c4f2a624f18d7794
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Beginne mit deinem standardmäßigen HTML-Boilerplate-Code. Füge eine `DOCTYPE`-Deklaration, ein `html`-Element, ein `head`-Element und ein `body`-Element hinzu.

Add the `lang` attribute to the opening `<html>` tag with `en` set as the value.

# --hints--

Dein Code sollte die `DOCTYPE`-Referenz beinhalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

You should define `html` as the document type.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag haben.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Your opening `<html>` tag should have the `lang` attribute with `en` as the value.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Das `head`- und `body`-Element sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
