---
id: 615380dff67172357fcf0425
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Deine Bilder sind jetzt zu groß.

Create a `.gallery img` selector to target them. Give them all a `width` of `100%` and a `max-width` of `350px`.

Setze ebenfalls die `height`-Eigenschaft auf `300px`, damit die Bilder eine einheitliche Größe haben.

# --hints--

Du solltest einen `.gallery img`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

Dein `.gallery img`-Selektor sollte eine `width`-Eigenschaft mit `100%` als Wert enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

Dein `.gallery img`-Selektor sollte eine `max-width`-Eigenschaft mit `350px` als Wert enthalten.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

Dein `.gallery img`-Selektor sollte eine `height`-Eigenschaft mit `300px` als Wert enthalten.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
