---
id: 61537c9eecea6a335db6da79
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Below your `.header` element, create a new `div` element and assign it a `class` of `gallery`. Dieses `div`-Element dient als Container für die Galeriebilder.

Erstelle innerhalb deines `.gallery`-Elements neun `img`-Elemente.

# --hints--

Du solltest ein `div`-Element in deinem `body`-Element erstellen.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Your new `div` element should have a `class` with `gallery` set as the value.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Dein neues `div`-Element sollte nach deinem `.header`-Element stehen.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Dein `.gallery`-Element sollte neun `img`-Elemente haben.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
