---
id: 6153897c27f6334716ee5abe
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Flexbox ist ein eindimensionales CSS-Layout, das steuern kann, wie Elemente innerhalb eines Containers angeordnet und ausgerichtet sind.

Um es zu verwenden, weise einem Element eine `display`-Eigenschaft mit `flex` als Wert zu. Das wird das Element zu einem <em>flex container</em> machen. Alle direkt untergeordneten Elemente eines Flex-Containers werden <em>flex items</em> genannt.

Erstelle einen `.gallery`-Selektor und mache ihn zu einem Flex-Container.



# --hints--

Du solltest einen `.gallery`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Dein `.gallery`-Selektor sollte eine `display`-Eigenschaft mit `flex` als Wert enthalten.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
