---
id: 615f171d05def3218035dc85
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Entferne die Margin aus deinem body-Element, setze die `font-family` auf `sans-serif` und weise ihm eine `background-color` mit dem Wert `#f5f6f7` zu.


# --hints--

Du solltest einen `body`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Dein `body`-Selektor sollte eine `margin`-Eigenschaft auf 0 als Wert gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginTop, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginBottom, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginLeft, '0px');

assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.marginRight, '0px');
```

Dein `body`-Selektor sollte eine `font-family`-Eigenschaft auf `sans-serif` als Wert gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === 'sans-serif');
```

Dein `body`-Selektor sollte eine `background-color`-Eigenschaft auf `#f5f6f7` als Wert gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}

```
