---
id: 6493bc0d99879635209565aa
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Das `border-box`-Größenmodell bewirkt das Gegenteil von `content-box`. Die Gesamtbreite des Elements, einschließlich des Paddings und Rands, entspricht der ausdrücklich festgelegten Breite. Der Inhalt des Elements wird verkleinert, um Platz für das Padding und den Rand zu schaffen.

Ändere die `box-sizing`-Eigenschaft zu `border-box`. Beachte, wie deine blauen Bildränder jetzt in deinen roten Galerierahmen passen.

# --hints--


Your * selector should have a `box-sizing` property set to `border-box` as the value.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
--fcc-editable-region--
 * {
  box-sizing: content-box;
}
--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
