---
id: 6494da0daf5df5197963671d
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Beachte, wie der blaue Bildrand über den roten Galerierand hinausragt. Dies ist auf die Art und Weise zurückzuführen, wie Browser die Größe von Containerelementen berechnen.

Die `box-sizing`-Eigenschaft wird verwendet, um dieses Verhalten festzulegen. Standardmäßig wird das `content-box`-Modell verwendet. With this model, when an element has a specific width, that width is calculated based only on the element's content. Padding- und Rand-Werte werden der Gesamtbreite hinzugefügt, so dass das Elemente wächst, um diese Werte aufzunehmen.

Versuche, `box-sizing` auf `content-box` mit dem globalen `*`-Selektor festzulegen. Zu diesem Zeitpunkt wirst du keine Änderungen sehen, da du den Standardwert verwendest.

# --hints--

Du solltest einen `*`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Dein `*`-Selektor sollte eine `box-sizing`-Eigenschaft enthalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Dein `*`-Selektor sollte eine `box-sizing`-Eigenschaft auf `content-box` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
