---
id: 64dd076a1b33c86d84b4232f
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Um die Größe deiner Elemente besser zu veranschaulichen, kann es hilfreich sein, einen Rand hinzuzufügen.

Weise deinem `.gallery`-Element eine `width` von `50%` und eine `border` zu, die auf `5px solid red` gesetzt ist.

Weise deinen `img`-Elementen dann eine `width` von `100%`, ein `padding` von `5px` und eine `border` mit dem Wert `5px solid blue` zu.

# --hints--

Du solltest einen `.gallery`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Dein `.gallery`-Selektor sollte eine `width`-Eigenschaft enthalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Dein `.gallery`-Selektor sollte eine `width`-Eigenschaft auf `50%` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Dein `.gallery`-Selektor sollte eine `border`-Eigenschaft enthalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Dein `.gallery`-Selektor sollte eine `border`-Eigenschaft auf `5px solid red` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Du solltest einen `img`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Dein `img`-Selektor sollte eine `width`-Eigenschaft enthalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Dein `img`-Selektor sollte eine `width`-Eigenschaft auf `100%` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Dein `img`-Selektor sollte eine `padding`-Eigenschaft enthalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Dein `img`-Selektor sollte die `padding`-Eigenschaft auf `5px` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Dein `img`-Selektor sollte eine `border`-Eigenschaft enthalten.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Dein `img`-Selektor sollte eine `border`-Eigenschaft auf `5px solid blue` gesetzt haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
