---
id: 64dd136e4e8b0894f9c70d86
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Nun, da du deinen `box-sizing`-Ansatz gefunden hast, kannst du das von dir hinzugefügte CSS bereinigen, um die Änderungen zu sehen.

Entferne deine `.gallery`- und `img`-Selektoren und alle darin enthaltenen Regeln.

# --hints--

Du solltest keinen `.gallery`-Selektor haben.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Du solltest keinen `img`-Selektor haben.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('img'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>


```

```css
 * {
  box-sizing: border-box;
}

--fcc-editable-region--
.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
--fcc-editable-region--
```
