---
id: 66c051d13a6a20255a963695
title: Build a Stylized To-Do list
challengeType: 14
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have an `h1` element with the title `My To-Do List`.
2. You should have an unordered list with the class `todo-list`.
3. Inside the unordered list, you should have 4 list items. Each list item should have a checkbox, a task label and two links that point to a valid URL. 
4. The first link in each list item should have the text `Edit` and the second link should have the text `Delete`. The links should also have the class `edit` and `delete` respectively.
5. The list items should have a bottom margin of `15px` and a line height of `1.6`.
6. Each list item should have squared bullets.
7. The links should be `blue` initially and have a default of no underline. 
8. When you hover over the links, the font color should change to `red`.
9. Visited links should be `purple`.
10. When a link is focused, there should be a `blue` border around it.
11. When a link is being clicked, its color should change to `green`.
    
# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title></title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="container">
        <h1>My To-Do List</h1>
        <ul class="todo-list">
            <li>
                <input type="checkbox" id="task1">
                <label for="task1">Explore gaming keyboards</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">Store link</a>
                    </li>
                </ul>
            </li>

            <li>
                <input type="checkbox" id="task2">
                <label for="task2">Finish the report</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">Request Access</a>
                    </li>
                </ul>

            </li>

            <li>
                <input type="checkbox" id="task3">
                <label for="task3">View Phone's Warranty</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">View Receipts</a>

                    </li>
                </ul>
            </li>

            <li>
                <input type="checkbox" id="task4">
                <label for="task4">Check Processor Specs</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">View Model Number</a>

                    </li>
                </ul>
            </li>
        </ul>
    </div>
</body>

</html>
```

```css
body {
    font-family: Arial, sans-serif;
    background-color: #f0f0f0;
    margin: 0;
    padding: 20px;
}

.container {
    max-width: 400px;
    margin: 0 auto;
    background-color: #fff;
    padding: 20px;
    border-radius: 10px;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
    text-align: center;
    color: #333;
}

.todo-list {
    list-style-type: none;
    padding: 0;
}

.todo-list li {
    margin-bottom: 15px;
    line-height: 1.6;
    padding: 10px;
    background-color: #e5f1e5;
    border-radius: 10px;
    position: relative;
}

.todo-list li a {
    margin-left: 10px;
    text-decoration: none;
}

.todo-list li a.sub-item-link:link {
    color: blue;
}

.todo-list li a.sub-item-link:visited {
    color: purple;
}

.todo-list li a.sub-item-link:hover {
    color: red;
}

.todo-list li a.sub-item-link:focus {
    outline: 2px solid blue;
    outline-offset: 2px;
}

.todo-list li a.sub-item-link:active {
    color: green;
}

.sub-item {
    list-style-type: square;
}
```
