---
id: 66c051d13a6a20255a963695
title: Build a Stylized To-Do list
challengeType: 14
dashedName: build-a-stylized-to-do-list
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. **Do not copy this demo project**.

**User Stories:**

1. You should have an unordered list with the class `todo-list`.
1. Inside the unordered list, you should have 4 list items.
1. Inside each list item, there should be:
   
- An `input` element with the type `checkbox` and `id` set to a unique value.
- A `label` element with the `for` attribute set to the corresponding `input` element's `id`.
- An unordered list with the class `sub-item`.
- A list item with an `anchor` element with the class `sub-item-link` and a valid `href` within the `sub-item` unordered list.
  
1. There shouldn't be any list styles applied to the unordered lists.
2. The anchor elements shouldn't have any default text decoration.
3. You should set the text color of the links to a colour of your choice.
4. When the links are visited, the color should change to another colour of your choice.
5. When the links are hovered over, the color should change to another colour of your choice.
6. When the links are focused, there should be a coloured outline around the link.
7. When the links are clicked, the color should change to another colour of your choice.
    
# --hints--

You should have an unordered list with the class `todo-list`.

```js
assert.exists(document.querySelector('.todo-list'));
```

You should have 4 `input` elements with the type `checkbox`.

```js
const checkboxes = document.querySelectorAll('.todo-list li input[type="checkbox"]');

assert.strictEqual(checkboxes.length, 4);
```

All of the `input` elements should have an `id`.

```js
const inputs = document.querySelectorAll('input');
for (let input of inputs) {
    assert.notEmpty(input?.id);
}
```

After the `label` elements, there should be an unordered list with the class `sub-item`.

```js
assert.exists(document.querySelector('.todo-list li ul.sub-item'));
```

The `li` inside the `ul` with the class `sub-item` should have an anchor element with the class `sub-item-link`.

```js
assert.exists(document.querySelector('.todo-list li ul.sub-item li a.sub-item-link'));
```

All of the anchor elements should have a valid `href`.

```js
const links = $('a');
assert(!!links.length);
for (let link of links) {
    assert.isAbove(link.href.length, 0);
}
```

All of the `anchor` elements should have a text.

```js
const links = $('a');
assert(!!links.length);
for (let link of links) {
    assert.isAbove(link.innerText.length, 0);
}
```

All of the `label` elements should have a `for` attribute.

```js
const labels = document.querySelectorAll('label');
for (let label of labels) {
  assert.notEmpty(label?.htmlFor);
}
```

All of the `label` elements should have some text.

```js
const labels = $('label');
assert(!!labels.length);
for (let label of labels) {
    assert.isAbove(label.innerText.length, 0);
}
```

Each `a` element should have a `target` attribute with the value of `_blank`.

```js
const anchors = $('a');
assert(!!anchors.length);
for (let anchor of anchors) {
    assert.equal(anchor.target, '_blank');
}
```

The links on the page should have no underline by default.

```js
const links = document.querySelectorAll("a");
[...links].forEach((link) => {
    assert(getComputedStyle(link).textDecoration.startsWith("none"));
});
```

The links should change color when hovered over.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:link')?.color || new __helpers.CSSHelp(document).getStyle('a:link')?.color);
```

The links should change color when they are being clicked.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:active')?.color || new __helpers.CSSHelp(document).getStyle('a:active')?.color);
```

The links should have an outline when focused.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:focus')?.outline || new __helpers.CSSHelp(document).getStyle('a:focus')?.outline);
```

The links should change color once visited.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sub-item-link:visited')?.color || new __helpers.CSSHelp(document).getStyle('a:visited')?.color);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Styled To-Do List</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled To-Do List</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="container">
        <h1>My To-Do List</h1>
        <ul class="todo-list">
            <li>
                <input type="checkbox" id="task1">
                <label for="task1">Explore gaming keyboards</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">Store link</a>
                    </li>
                </ul>
            </li>

            <li>
                <input type="checkbox" id="task2">
                <label for="task2">Finish the report</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">Request Access</a>
                    </li>
                </ul>

            </li>

            <li>
                <input type="checkbox" id="task3">
                <label for="task3">View Phone's Warranty</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">View Receipts</a>

                    </li>
                </ul>
            </li>

            <li>
                <input type="checkbox" id="task4">
                <label for="task4">Check Processor Specs</label>
                <ul class="sub-item">
                    <li>
                        <a href="#" target="_blank" class="sub-item-link">View Model Number</a>

                    </li>
                </ul>
            </li>
        </ul>
    </div>
</body>

</html>
```

```css
body {
    font-family: Arial, sans-serif;
    background-color: #f0f0f0;
    margin: 0;
    padding: 20px;
}

.container {
    max-width: 400px;
    margin: 0 auto;
    background-color: #fff;
    padding: 20px;
    border-radius: 10px;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

h1 {
    text-align: center;
    color: #333;
}

.todo-list {
    list-style-type: none;
    padding: 0;
}

.todo-list li {
    margin-bottom: 15px;
    line-height: 1.6;
    padding: 10px;
    background-color: #e5f1e5;
    border-radius: 10px;
    position: relative;
}

.todo-list li a {
    margin-left: 10px;
    text-decoration: none;
}

.sub-item-link:link {
    color: blue;
}

.sub-item-link:visited {
    color: purple;
}

.sub-item-link:hover {
    color: red;
}

.sub-item-link:focus {
    outline: 2px solid blue;
    outline-offset: 2px;
}

.sub-item-link:active {
    color: green;
}

.sub-item {
    list-style-type: square;
}
```
