---
id: 587d7b89367417b2b2512b4b
title: Деструктуроване присвоєння для присвоєння змінних з масивів
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6 робить деструктуризацію масивів такою ж легкою, як і деструктуризацію об’єктів.

Ключовою відмінністю між оператором розширення та деструктурованим масивом є те, що оператор розширення розпаковує весь вміст масиву в список, розділений комою. Тому ви не можете вибрати, які елементи потрібно присвоїти до змінних.

А деструктуризація масиву дозволяє це зробити:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

Консоль показуватиме значення `a` та `b` як `1, 2`.

Змінній `a` присвоєно перше значення масиву, а `b` присвоєно друге значення масиву. Ми також можемо отримати доступ до значення за будь-яким індексом масиву за допомогою деструктуризації, використовуючи коми для досягнення потрібного індексу:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

Консоль показуватиме значення `a`, `b` та `c` як `1, 2, 5`.

# --instructions--

Використайте деструктуроване присвоєння, щоб поміняти місцями значення `a` та `b` (`a` отримує значення, збережене в `b`, а `b` отримує значення, збережене в `a`).

# --hints--

Значенням `a` повинне стати `6`.

```js
assert(a === 6);
```

Значенням `b` повинне стати `8`.

```js
assert(b === 8);
```

Ви повинні використати деструктуризацію масиву, щоб поміняти значення `a` та `b`.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(code));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
