---
id: 66db57ad34c7089b9b41bfd6
title: Build a Lightbox Viewer
challengeType: 14
dashedName: build-a-lightbox-viewer
demoType: onClick
---

# --description--

A lightbox displays a larger version of an image when clicked and shadows the rest of the page. This project will create a lightbox gallery that displays full-size images when a thumbnail is clicked. For each image, two versions are provided: a thumbnail and a full-size image. The full-size image is the same as the thumbnail, but without the `-thumbnail` suffix.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `div` with a class of `gallery` within your `body`.
1. Within the `.gallery` element, you should have three image thumbnails, each with a class of `gallery-item`. You should use the following links for thumbnail images:
   
<!-- would be replaced by cdn -->

- `./images/stonehenge-thumbnail.jpg`
- `./images/storm-thumbnail.jpg`
- `./images/trees-thumbnail.jpg`

1. You should have a `div` with a class of `lightbox` within your `body`.
1. You should have a `span` with an `id` of `close` within your `.lightbox` element. You can use `&times;` as its text if you want.
1. You should have a `img` with an `id` of `lightbox-image` within your `.lightbox` element.
1. Your `.lightbox` element should cover the entire viewport, have a semi-transparent background, and center its content. Initially, its `display` property should be set to `none` to hide it.
1. When you click one of your `.gallery-item` elements, the `.lightbox` element’s `display` property should be set to `flex` to make it, and the two elements within it, visible.
1. When you click one of your `.gallery-item` elements, the `#lightbox-image` element’s `src` should be set to a full-size version of the image clicked by removing `-thumbnail` from the image’s `src` attribute. The full-size images are located at the following links:

<!-- would be replaced by cdn -->
- `./images/stonehenge.jpg`
- `./images/storm.jpg`
- `./images/trees.jpg`

1. When your `.lightbox` element is visible and you click the `.close` button or the `.lightbox` element, the `.lightbox` elements `display` should be set back to `none`.
1.  You should use separate functions `closeLightbox` and `openLightbox` to handle the opening and closing of the lightbox.

**Note:** Be sure to link your stylesheet and the JavaScript file in your HTML.

# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Lightbox Viewer</title>
</head>

<body>

</body>

</html>
```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Image Gallery with Lightbox</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="title"><h1>Museum Gallery</h1></div>

    <div class="gallery">
        <img src="./images/stonehenge-thumbnail.jpg" alt="Image 1" class="gallery-item">
         <img src="./images/storm-thumbnail.jpg" alt="Image 2" class="gallery-item">
        <img src="./images/trees-thumbnail.jpg" alt="Image 3" class="gallery-item">
    </div>

    <div id="lightbox" class="lightbox">
        <span id="close" class="close">&times;</span>
        <img id="lightbox-image" src="" alt="Lightbox Image">
    </div>

    <script src="script.js"></script>
</body>

</html>
```

```css
body {
    margin: 0;
    padding: 0;
    background-color: #f0f0f0;
    color: #333;
}

.title {
    margin-top: 100px;
    display: flex;
    flex-wrap: wrap;
    justify-content: center;
    background-color: rgb(211, 221, 224);
   
}

.gallery {
    display: flex;
    flex-wrap: wrap;
    gap: 10px;
    justify-content: center;
    padding: 20px;
    
}

.gallery-item {
    width: 250px;
    height: 300px;
    object-fit: cover;
    cursor: pointer;
    transition: transform 0.2s;
}

.gallery-item:hover {
    transform: scale(1.05);
}

.lightbox {
    display: none;
    position: fixed;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background-color: rgba(0, 0, 0, 0.8);
    justify-content: center;
    align-items: center;
    z-index: 1000;
}

.lightbox img {
    max-width: 90%;
    max-height: 90%;
}

.close {
    position: absolute;
    top: 20px;
    right: 20px;
    font-size: 2em;
    color: white;
    cursor: pointer;
}
```

```js
// Get references to the elements
const galleryItems = document.querySelectorAll(".gallery-item");
const lightbox = document.getElementById("lightbox");
const lightboxImage = document.getElementById("lightbox-image");
const closeButton = document.getElementById("close");

// Function to open the lightbox
function openLightbox(src) {
  lightboxImage.src = src;
  lightbox.style.display = "flex"; // Show the lightbox
}

// Function to close the lightbox
function closeLightbox() {
  lightbox.style.display = "none"; // Hide the lightbox
}

// Add event listener to the close button
closeButton.addEventListener("click", closeLightbox);

// Add event listener to the lightbox background
  lightbox.addEventListener("click", (e) => {
  if (e.target === lightbox) {
    closeLightbox(); // Close lightbox if clicking outside the image
  }
});

// Add event listeners to gallery items
galleryItems.forEach((item) => {
  item.addEventListener("click", () => {
    openLightbox(item.src.replace("-thumbnail", "")); // Open lightbox with the full image
  });
});
```
