---
id: 65577236b056379d5dbc7000
title: Paso 18
challengeType: 20
dashedName: step-18
---

# --description--

Now you are going to start developing the algorithm to calculate the shortest path between each node in your new graph.

Declare an empty function called `shortest_path` and don't forget the `pass` keyword.

# --hints--

You should have a `shortest_path` function.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    f = __locals.get("shortest_path")
    inspect.isfunction(f)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
