---
id: 6579dd49fa8a8e1fd06b85a9
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Graphs are data structures representing relations between pairs of elements. These elements, called *nodes*, can be real-life objects, entities, points in space or others. The connections between the nodes are called the *edges*.

For example, a graph can be used to represent two points in the space, `A` and `B`, connected by a path. A graph like this will be made of two nodes connected by an edge.

Rename the `copper` dictionary into `my_graph`. This will represent the graph to test your algorithm.

# --hints--

You should rename your `copper` dictionary into `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Your `my_graph` variable should be a dictionary.

```js
({ test: () => assert(__pyodide.runPython(`
    my_graph = __locals.get("my_graph")
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
