---
id: 6579cd5f6dd62c189e53ddbb
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Now that you got the basic aspects of dictionaries, you can proceed to build the shortest path algorithm.

Delete every line of code after the declaration of the `copper` dictionary.

# --hints--

You should delete the lines after the declaration of your dictionary.

```js
({ test: () => assert.isFalse( /copper\s*\[.*?\]|del|print\s*\(.*?\)/ms.test(code)) })
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']
print(copper)
--fcc-editable-region--
```
