---
id: 619691693bc14b0e528f5a20
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Bestimme das `.ground`-Element und setze seine `width` so, dass es die gesamte Beite des Ansichtsfensters einnimmt. Setze dann die `height` auf `400px`.

# --hints--

Du solltest den `.ground`-Selektor verwenden.

```js
assert.match(code, /\.ground\s*\{/);
```

Du solltest dem `.ground`-Element eine `width` von `100vw` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

Du solltest dem `.ground`-Element eine `height` von `400px` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
