---
id: 619665c9abd72906f3ad30f9
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Du wirst einen glücklichen Flappy Penguin bauen und in diesem Prozess CSS-Transformationen und Animationen weiter erkunden.

Beginne mit deinem grundlegenden HTML-Boilerplate. Füge die `DOCTYPE`-Deklaration ein, das `html`-Element mit der Sprache Englisch, die passenden `meta`-Tags, ein `head`, `body` und `title`-Element. Verlinke auch dein Stylesheet mit der Seite.

# --hints--

Dein Code sollte eine `<!DOCTYPE html>`-Deklaration haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Dein Code sollte ein `html`-Element haben.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein abschließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/gi));
```

Dein Code sollte innerhalb des `html`-Elements ein `head`-Element haben.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Dein Code sollte innerhalb des `html`-Elements ein `body`-Element haben.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Dein `head`-Element sollte vor deinem `body`-Element stehen.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Du solltest zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Ein `meta`-Element sollte `name` auf `viewport` gesetzt haben und `content` sollte auf `width=device-width, initial-scale=1.0` gesetzt sein.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset`-Attribut auf `UTF-8` eingestellt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Dein Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein `title` sollte etwas Text haben.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Dein Code sollte ein `link`-Element haben.

```js
assert.exists(document.querySelector('link'));
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
