---
id: 61967e74a8e3690ab6292daa
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Target the `body` element to set the `background` to a linear gradient angled 45 degrees clockwise, starting at `rgb(118, 201, 255)` and ending at `rgb(247, 255, 222)`.

# --hints--

Du solltest den `body`-Element-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Du solltest die `background`-Eigenschaft im `body`-Selektor verwenden.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Du solltest den `background` auf `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))` setzen.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
