---
id: 613297a923965e0703b64796
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

你可能已經熟悉 `meta` 元素；它用於指定有關頁面的信息，例如標題、描述、關鍵詞和作者。

給你的頁面一個 `meta` 元素，然後給它的屬性 `charset` 一個合適的值。

`charset` 屬性指定了頁面的字符編碼，`UTF-8` 是現在大多數瀏覽器支持的唯一編碼。

# --hints--

你應該在 `head` 元素中創建一個 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你應該給 `meta` 元素添加一個 `charset` 屬性，值爲 `UTF-8`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
