---
id: 61329b210dac0b08047fd6ab
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

繼續學習 `meta` 元素，`viewport` 定義告訴瀏覽器如何呈現頁面。 包括改善移動設備上的視覺可訪問性，並改進 _SEO_（搜索引擎優化）。

添加帶有 `content` 屬性的 `viewport` 定義，詳細說明頁面的 `width` 和 `initial-scale`。

# --hints--

你應該在 `head` 中創建另一個 `meta` 元素。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

你應該給 `meta` 元素添加一個 `name` 屬性，值爲 `viewport`。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

你應該給 `meta` 添加一個 `content` 屬性，值爲 `width=device-width, initial-scale=1`。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
