---
id: 61329d52e5010e08d9b9d66b
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

另一個關於可訪問性和 SEO 的重要 `meta` 元素是 `description` 定義。 `content` 屬性的值被搜索引擎用來提供你的頁面的描述。

添加一個 `meta` 元素，並將 `name` 屬性設置爲 `description`，併爲其添加一個有用的 `content` 屬性。

# --hints--

你應該給 `head` 添加一個新的 `meta` 元素。

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

你應該給 `meta` 添加一個值爲 `description` 的 `name` 屬性。

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

你應該給 `meta` 添加一個 `content` 屬性。

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

`content` 屬性的值不應超過 165 個字符。 _這是 Google 接受的最大描述長度。_

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
