---
id: 6133d11ef548f51f876149e3
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

導航是無障礙的核心部分，屏幕閱讀器依靠你提供你的頁面結構。 這是通過語義 HTML 元素完成的。

將 `header` 元素和 `main` 元素添加到你的頁面。

`header` 元素將用於介紹頁面，以及提供導航菜單。

`main` 元素將包含頁面的核心內容。

# --hints--

你應該將給 `body` 元素添加一個 `header` 元素。

```js
assert.exists(document.querySelector('body > header'));
```

你應該給 `body` 元素添加一個 `main` 元素。

```js
assert.exists(document.querySelector('body > main'));
```

`header` 元素應該在 `main` 元素之前。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
