---
id: 613e275749ebd008e74bb62e
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

_SVG_（可縮放矢量圖形）的一個有用之處是它包含一個 `path` 屬性，該屬性允許在不影響圖像分辨率的情況下縮放圖像。

`img` 當前是默認尺寸，這個尺寸太大。 可以使用它的 `id` 作爲選擇器來縮放圖像，並將 `width` 設置爲 `max(100px, 18vw)`。

# --hints--

你應該使用 `#logo` 選擇器來定位 `img` 元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

你應該給 `img` 元素添加一個值爲 `max(100px, 18vw)` 的 `width`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
