---
id: 613e2546d0594208229ada50
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

在 `header` 中，通過嵌套一個 `img` 元素、`h1` 元素和 `nav` 元素來提供有關頁面的上下文。

`img` 元素應該指向 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`，並且有一個 `id` 爲 `logo`。

`h1` 元素應該包含文本 `HTML/CSS Quiz`。

# --hints--

你應該將 `img` 元素添加到 `header` 元素。

```js
assert.exists(document.querySelector('header > img'));
```

你應該將 `h1` 元素添加到 `header` 元素。

```js
assert.exists(document.querySelector('header > h1'));
```

你應該將 `nav` 元素添加到 `header` 元素。

```js
assert.exists(document.querySelector('header > nav'));
```

你應該按照依次添加 `img` 元素、`h1` 元素和 `nav` 元素。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

你應該給 `img` 元素添加一個 `src` 屬性，值爲 `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

你應該給 `img` 元素添加一個 `id` 屬性，值爲 `logo`。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

你應該給 `h1` 元素添加文本 `HTML/CSS Quiz`。

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
