---
id: 61408f155e798909b6908712
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

要在頁面上啓用導航，請添加一個包含以下三個列表項的無序列表：

- `INFO`
- `HTML`
- `CSS`

列表項文本應包含在錨標記中。

# --hints--

你應該在 `nav` 中嵌套一個 `ul` 元素。

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

你應該在 `ul` 元素中嵌套三個 `li` 元素。

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

你應該在每個 `li` 元素中嵌套一個 `a` 元素。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

你應該爲第一個 `a` 元素添加文本 `INFO`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

你應該爲第二個 `a` 元素添加文本 `HTML`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

你應該爲第三個 `a` 元素添加文本 `CSS`。

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
