---
id: 6140883f74224508174794c0
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

使 `header` 佔用其父容器的全寬，設置它的 `height` 爲 `50px`，並將屬性 `background-color` 設置爲 `#1b1b32`。 然後設置 `display` 屬性時使用 _Flexbox_。

# --hints--

你應該使用 `header` 元素選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

你應該把 `header` 元素 `width` 屬性的值設置爲 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

你應該把 `header` 的 `height` 屬性設置爲 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

你應該把 `header` 元素 `background-color` 屬性的值設置爲 `#1b1b32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

你應該把 `header` 元素 `display` 屬性設置爲 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
