---
id: 6140827cff96e906bd38fc2b
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

如 <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp 樣式指南</a>中所述，logo 應保持 `35 / 4` 的長寬比，並且文本週圍有 padding。

首先，將 `background-color` 設置爲 `#0a0a23`，以便你能看見 logo。 然後，使用 `aspect-ratio` 屬性將所需的長寬比設置爲 `35 / 4`。 最後，在周圍添加 `0.4rem` 的 `padding`。

# --hints--

你應該給 `#logo` 添加一個值爲 `#0a0a23` 的 `background-color` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

你需要使用 `aspect-ratio` 屬性。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

你不應該使用 `height` 屬性。

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

你不應該修改`width`屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

你需要給 `img` 設置一個值爲 `35 / 4` 的 `aspect-ratio` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

你需要給 `img` 設置一個值爲 `0.4rem` 的 `padding` 的屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(100px, 18vw);

}
--fcc-editable-region--

```
