---
id: 6141fabd6f75610664e908fd
title: 步驟 14
challengeType: 0
dashedName: step-14
---

# --description--

由於這是一個測試題，你需要給用戶提供一個表單以便他們提交答案。 你可以使用 `section` 元素在語義上分隔表單中的內容。

在 `main` 元素中，創建一個包含三個嵌套 `section` 元素的表單。 然後，使用正確的方法將表單提交到 `https://freecodecamp.org/practice-project/accessibility-quiz`。

# --hints--

你應該在 `main` 元素中嵌套一個 `form` 元素。

```js
assert.exists(document.querySelector('main > form'));
```

你應該在 `form` 元素中嵌套三個 `section` 元素。

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

你應該給 `form` 元素一個 `action` 屬性。

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

你應該將 `action` 屬性的值設置爲 `https://freecodecamp.org/practice-project/accessibility-quiz`。

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

你應該給 `form` 元素一個 `method` 屬性。

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

你應該給 `form` 元素一個 `method` 屬性，其值爲 `post`。

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
