---
id: 614202874ca576084fca625f
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

Every `region` role requires a label, which helps screen reader users understand the purpose of the region. One method for adding a label is to add a heading element inside the region and then reference it with the `aria-labelledby` attribute.

Add the following `aria-labelledby` attributes to the `section` elements:

- `student-info`
- `html-questions`
- `css-questions`

然後，在每個 `section` 元素中，嵌套一個 `h2` 元素，其 `id` 與相應的 `aria-labelledby` 屬性匹配。 給每個 `h2` 添加適當的文本內容。

# --hints--

你應該給第一個 `section` 元素添加值爲 `student-info` 的 `aria-labelledby` 屬性。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

你應該給第二個 `section` 元素添加值爲 `html-questions` 的 `aria-labelledby` 屬性。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

你應該給第三個 `section` 元素添加值爲 `css-questions` 的 `aria-labelledby` 屬性。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

你應該在第一個 `section` 元素中嵌套一個 `h2` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

你應該在第二個 `section` 元素中嵌套一個 `h2` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

你應該在第三個 `section` 元素中嵌套一個 `h2` 元素。

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

你應該給第一個 `h2` 元素添加 `id` 爲 `student-info`。

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

你應該給第二個 `h2` 元素添加 `id` 爲 `html-questions`。

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

你應該給第三個 `h2` 元素添加 `id` 爲 `css-questions`。

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

你應該給第一個 `h2` 元素添加適當的文本內容。 _提示：我會使用 `Student Info`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[0]?.textContent?.length, 1);
```

你應該給第二個 `h2` 元素添加適當的文本內容。 _提示：我會使用 `HTML`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[1]?.textContent?.length, 1);
```

你應該給第三個 `h2` 元素添加適當的文本內容。 _提示：我會使用 `CSS`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[2]?.textContent?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
