---
id: 6568c0a5edddc3daa65d20b2
title: Крок 40
challengeType: 20
dashedName: step-40
---

# --description--

Now, you need to check if a given number is not equal to the number in the specified column of the current row.

For that, first, iterate over the rows of the 2D list `self.board` using a `for` loop in the range `0` to `8`. Використайте `row` як змінну ітерації.

# --hints--

`valid_in_col` має містити `for row in range(9)`.

```js
const tCode = code.replace(/\r/g, '');
const valid = __helpers.python.getDef(tCode, "valid_in_col");
const {function_body} = valid;
assert.match(function_body, /for\s+row\s+in\s+range\(\s*9\s*\)/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

            if index < 8:
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'
                else:
                    board_string += middle_lines
            else:
                board_string += lower_lines

        return board_string

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

--fcc-editable-region--
    def valid_in_col(self, col, num):

--fcc-editable-region--
```
