---
id: 65688f737b0ef396bf0c22d6
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер ви перейдете до створення самої сітки 9×9.

Вхідна головоломка виглядатиме так:

```py
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]
```

Отримана сітка виглядатиме так:

<img class="img-responsive" alt="сітка судоку" src="https://cdn.freecodecamp.org/curriculum/python/sample-board.png" style="background-color: white; height:300px; width:300px; padding: 10px;" />


Визначте метод `__str__` в межах класу `Board`. Також додайте параметр `self`. Цей метод автоматично викликається, якщо використати функцію `str()` на екземплярі класу або використати `print()` з об’єктом.

# --hints--

Метод повинен мати назву `__str__`.

```js
assert.match(code, /def\s+__str__\s*\(/);
```

Додайте параметр `self` до методу.

```js
assert.match(code, /def\s+__str__\s*\(\s*self\s*\)\s*:/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--

--fcc-editable-region--
```
