---
id: 6568bc19f3418dc8a8821187
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Тепер потрібно обробити останній ряд всієї сітки.

`lower_lines` представляє нижній край всієї сітки.

Створіть блок `else`, щоб додати рядок `lower_lines` до `board_string`, якщо умова зовнішньої інструкції `if` є хибною.

# --hints--

Створіть блок `else` в межах зовнішньої інструкції `if`. Зверніть увагу на відступи.

```js
({ test: () => assert.match(code, /else\s*:/m) })
```

Зовнішня умова `else` повинна містити `board_string += lower_lines`.

```js
({ test: () => assert.match(code, /board_string\s*\+=\s*lower_lines/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

--fcc-editable-region--
            if index < 8:
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'
                else:
                    board_string += middle_lines

--fcc-editable-region--
```
