---
id: 6568a242a3e1efc22b07274d
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

`board_string` поступово створюється, оскільки цикл ітерує над кожним рядом, створюючи повне ASCII-представлення сітки судоку.

Додайте змінений рядок `row_empty` до `board_string`.

# --hints--

Зовнішній цикл `for` повинен містити `board_string += row_empty`.

```js
({ test: () => assert.match(code, /board_string\s*\+=\s*row_empty/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

--fcc-editable-region--
            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')

--fcc-editable-region--
```
