---
id: 6568bc19f3418dc8a8821187
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Now, you need to handle the last row of the entire board.

`lower_lines` represents the bottom border of the entire sudoku board.

Create an `else` block to append the `lower_lines` string to `board_string` when the outer `if` condition is false.

# --hints--

You should create an `else` block for the outermost `if` statement. Pay attention to the indentation.

```js
({ test: () => assert.match(code, /else\s*:/m) })
```

You should have `board_string += lower_lines` within the outermost `else` clause.

```js
({ test: () => assert.match(code, /board_string\s*\+=\s*lower_lines/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

--fcc-editable-region--
            if index < 8:
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'
                else:
                    board_string += middle_lines

--fcc-editable-region--
```
