---
id: 6568904b83a2f29878578146
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Initialize a `board_string` variable with the content of `upper_lines`. This will be the starting point for building the entire visual representation of the sudoku board.

# --hints--

Метод `__str__` повинен містити `board_string = upper_lines`.

```js
({ test: () =>
  {
    const  str = __helpers.python.getDef(code.replace(/\r/g, ''), "__str__");
    const {function_body} = str;    
    assert(function_body.match(/board_string\s*=\s*upper_lines/));
  }
})
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
--fcc-editable-region--
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'

--fcc-editable-region--
```
