---
id: 656873ffdc638f7e290f60de
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви дізнаєтесь про класи та об’єкти, створивши розв’язувач судоку.

Клас в Python — це основа для створення об’єктів. Об’єкти, створені з класу, є екземплярами цього класу. Клас можна створити за допомогою такого синтаксису:

```js
class ClassName:
```

Спочатку ви створите дошку 9×9 за допомогою класів, а потім заповните її значеннями головоломки.

Почніть зі створення класу `Board`.

# --hints--

Клас повинен мати назву `Board`.

```js
assert.match(code, /class\s+Board\s*\:/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
