---
id: 6568bbc8c3bda1c773e23cf1
title: Крок 25
challengeType: 20
dashedName: step-25
---

# --description--

Тепер потрібно обробити інші ряди. Якщо внутрішня умова оцінюється як `False` (тобто поточний ряд не є останнім рядом квадрата 3×3), додайте рядок `middle_lines` до `board_string`. Напишіть це в блоці `else`.

Пригадайте, що `middle_lines` представляє середні краї сітки та включає горизонтальні перетини.

# --hints--

Додайте блок `else` до інструкції `if`.

```js
({ test: () => assert.match(code, /else\s*:/m) })
```

Блок `else` повинен містити `board_string += middle_lines`.

```js
({ test: () => assert.match(code, /board_string\s*\+=\s*middle_lines/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

            if index < 8:
--fcc-editable-region--
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'

--fcc-editable-region--
```
