---
id: 656874efd5102b81815c8ef7
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

A new instance of a class is created by using the function notation: `ClassName()`. The instantiation creates an empty object. Classes can have methods, which are like local functions for each instance. Methods are declared as follows:

```python
class ClassName:
    def method_name():
        pass
```

The `__init__` method is a special method that allows you to instantiate an object to a customized state. When a class implements an `__init__` method, `__init__` is automatically called upon instantiation.

Create an `__init__` method inside your `Board` class.

# --hints--

Метод повинен мати назву `__init__`. Don't add any parameters.

```js
assert.match(code, / +def\s+__init__\s*\(\s*\)\:/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:

--fcc-editable-region--
```
