---
id: 6568991b4d4874a4d5271337
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Далі ви розділите кожен ряд на три сегменти, щоб правильно представити квадрати 3×3.

Створіть вкладений цикл `for`, щоб ітерувати над кожним сегментом ряду. Використайте `square_no`, `part` як змінну ітерації та функцію `enumerate()`. Наразі залиште виклик `enumerate()` порожнім.

# --hints--

Внутрішній цикл повинен мати `square_no` як лічильник та `part` як елемент з ітерованого об’єкта.

```js
assert.match(code, /for\s+square_no\s*,\s*part/)
```

Наявний цикл `for` повинен містити `for square_no, part in enumerate()`.

```js
assert.match(code, /for\s+square_no\s*,\s*part\s+in\s+enumerate\s*\(\s*\)\:/)
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

--fcc-editable-region--
        for index, line in enumerate(self.board):
          row_list = []

--fcc-editable-region--
```
