---
id: 656873ffdc638f7e290f60de
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you will learn about classes and objects by building a sudoku puzzle solver.

In Python, a class is a blueprint for creating objects. Objects created from a class are instances of that class. You can create a class using this syntax:

```js
class ClassName:
```

First, you will create a 9x9 board by using classes and then populate it with the puzzle values.

Begin by creating a `Board` class.

# --hints--

Клас повинен мати назву `Board`.

```js
assert.match(code, /class\s+Board\s*\:/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
