---
id: 65689ad61dfa81ab9ffafc86
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Далі ви створите рядкове представлення ряду з пробілами між елементами.

Для цього створіть рядок `row` поза межами тіла внутрішнього циклу `for`. Призначте до нього відформатований рядок `f'║ {" ".join(row_list)} ║\n'`, щоб об’єднати елементи з `row_list`, розділені пробілом.

# --hints--

Призначте відформатований рядок `f'║ {" ".join(row_list)} ║\n'` до змінної під назвою `row`.

```js
({ test: () => assert.match(code, /row\s*\=\s*f("|')║\s\{(?!\1)("|')\s\2\.join\s*\(\s*row_list\s*\)\s*\}\s║\\n\1/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
            row_list = []
--fcc-editable-region--
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)

                if square_no != 3:
                    row_list.append('║')

--fcc-editable-region--
```
