---
id: 65688f22703200963a85dfb7
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Призначте значення параметра `board` (що передається під час створення екземпляра класу `Board`) до змінної екземпляра під назвою `board`, використавши `self.board` в межах методу `__init__`.

`self.board` посилається на атрибут екземпляра класу `board`. Це змінна, яка належить до об’єкта, створеного з класу `Board`.

# --hints--

Метод `__init__` повинен містити `self.board = board`.

```js
assert.match(code, /self\.board\s*=\s*board/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):

--fcc-editable-region--
```
