---
id: 65688f22703200963a85dfb7
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Inside the `__init__` method, assign the value of the `board` parameter (which is passed when creating an instance of the `Board` class) to an instance variable named `board` using `self.board`.

`self.board` refers to the board attribute of the instance of the class. It's a variable that belongs to the object created from the `Board` class.

# --hints--

Метод `__init__` повинен містити `self.board = board`.

```js
assert.match(code, /self\.board\s*=\s*board/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):

--fcc-editable-region--
```
