---
id: 65689b055e6f49ac6f82d3cf
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Якщо ви передаєте вхідну сітку головоломки, то `0` використовується для порожніх клітинок.

Для кращого візуального представлення замініть порожні клітинки в ряді на пробіл, використавши метод `replace`.

Метод `replace()` приймає два аргументи: перший — символ, який потрібно замінити, а другий — символ, на який потрібно замінити.

Після заміни призначте результат до змінної `row_empty`.

# --hints--

Замініть всі `0` в ряді на пробіли за допомогою методу `replace`.

```js
({ test: () => assert.match(code, /row\.replace\s*\(\s*("|')0\1\s*,\s*("|') \2\s*\)/) })
```

Зовнішній цикл `for` повинен містити `row_empty = row.repalce('0', ' ')`.

```js
({ test: () => assert.match(code, /row_empty\s*=\s*row\.replace\s*\(\s*("|')0\1\s*,\s*("|') \2\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)

                if square_no != 3:
                    row_list.append('║')

--fcc-editable-region--
            row = f'║ {" ".join(row_list)} ║\n'

--fcc-editable-region--
```
