---
id: 65689a748de8fbaa00c5617e
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Створіть інструкцію `if` в межах внутрішнього циклу, щоб перевірити, чи поточний сегмент (`square_no`) не є останнім (тобто не дорівнює `3`)

# --hints--

Перевірте, чи поточний сегмент (`square_no`) не дорівнює `3`.

```js
assert.match(code, /if\s+square_no\s*!=\s*3/m)
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
            row_list = []

--fcc-editable-region--
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)

--fcc-editable-region--
```
