---
id: 6568beebba98a3d1f26f6bf8
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

In the `try` block, attempt to find the index of the first occurrence of `0` in the current row using `contents.index(0)`. Збережіть результати в змінній `col`.

# --hints--

Use the index of the first occurrence of `0` in the current row using `contents.index(0)`.

```js
assert.match(code, /contents\.index\(\s*0\s*\)/)
```

Блок `try` повинен містити `col = contents.index(0)`.

```js
assert.match(code, /col\s*=\s*contents\.index\(\s*0\s*\)/)
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

            if index < 8:
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'
                else:
                    board_string += middle_lines
            else:
                board_string += lower_lines

        return board_string

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
--fcc-editable-region--
            try:

--fcc-editable-region--
```
