---
id: 65688efcc78c9495e73acfc9
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Додайте два параметри до методу `__init__`. Порядок має значення:

- `self`: це посилання на екземпляр класу. Цей параметр прийнято називати self.
- `board`: очікується, що параметр board буде передано під час створення екземпляра класу `Board`.

# --hints--

Додайте параметри `self` та `board` до методу.

```js
assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*board\s*\):/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__():

--fcc-editable-region--
```
