---
id: 6568917528820d99236ad811
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер вам потрібно переглянути кожен ряд сітки судоку.

Перелічування — це зручний спосіб відстежувати елемент і його позицію в списку. Функція `enumerate()` — це вбудована функція в Python, яка приймає ітерований об’єкт (наприклад, список, кортеж або рядок) та повертає ітератор, який створює кортежі, що містять індекси та відповідні значення з ітерованого об’єкта.

Ініціюйте цикл `for`, щоб ітерувати над кожним рядом (`line`) на сітці судоку (`self.board`).

Використайте перелічування, щоб отримати індекс (`index`) та вміст (`line`) всіх рядів.

Загальний синтаксис буде таким:

```js
for x, y in enumerate(parameter):
```

# --hints--

Метод `__str__` повинен містити `for index, line in enumerate(self.board):`.

```js
assert.match(
  code,
  /for\s+(\w+)\s*,\s*(?!\1)(\w+)\s+in\s+enumerate\s*\(\s*self\.board\s*\)\s*:/m
);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
--fcc-editable-region--

--fcc-editable-region--
```
