---
id: 656899f4214ee6a881bc8649
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Extend the `row_list` with the elements of the `row_square` string.

# --hints--

You should have `row_list.extend(row_square)` within the innermost `for` loop.

```js
({ test: () => assert.match(code, /row_list\.extend\s*\(\s*row_square\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []

--fcc-editable-region--
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)

--fcc-editable-region--
```
