---
id: 65688f93a1b6e9970f710f62
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Щоб створити верхній край сітки, створіть змінну `upper_lines` і призначте до неї значення `f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'`.

Цей рядок представляє верхній край сітки судоку у візуально привабливому стилі ASCII. Він використовує спеціальні символи Unicode, щоб намалювати краї та перетини.

# --hints--

Призначте значення `f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'` до змінної під назвою `upper_lines`.

```js
({ test: () => assert.match(code, /upper_lines\s*=\s*f("|')\\n╔═══\{\s*(?=[^\1])("|')╤═══\2\s*\*\s*2\s*\}\{\s*\2╦═══\2\s*\}\{\s*\2╤═══\2\s*\*\s*2\s*\}\{\s*\2╦═══\2\s*\}\{\s*\2╤═══\2\s*\*\s*2\s*\}╗\\n\1/m) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

--fcc-editable-region--
    def __str__(self):

--fcc-editable-region--
```
