---
id: 6568994faf481da5d37bfa40
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Тепер ви повинні створити три відрізки, які потрібно передати до функції `enumerate`.

Використайте зрізи списку, щоб створити три списки однакової довжини, які представляють сегмент `line` кожного квадрата 3×3, і передайте їх до виклику `enumerate()`. Додайте `start = 1`, щоб почати перелічування з `1`, а не `0`.

# --hints--

Внутрішній цикл `for` повинен містити `enumerate([line[:3], line[3:6], line[6:]], start=1)`.

```js
assert.match(code, /\[\s*line\s*\[\s*:3\s*\]\s*,\s*line\s*\[\s*3\s*:\s*6\s*\]\s*,\s*line\s*\[\s*6\s*:\s*\]\s*\]\s*,\s*start\s*=\s*1/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

--fcc-editable-region--
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate():
--fcc-editable-region--
```
