---
id: 656874efd5102b81815c8ef7
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Новий екземпляр класу створюється за допомогою нотації функції `ClassName()`. Створення екземпляра створює порожній об’єкт. Класи можуть мати методи, які є локальними функціями для кожного екземпляра. Методи оголошуються наступним чином:

```python
class ClassName:
    def method_name():
        pass
```

Метод `__init__` — це спеціальний метод, який дозволяє створити екземпляр об’єкта в налаштованому стані. Коли клас реалізує метод `__init__`, то `__init__` автоматично викликається під час створення екземпляра.

Створіть метод `__init__` в межах класу `Board`.

# --hints--

Метод повинен мати назву `__init__`. Не додавайте жодних параметрів.

```js
assert.match(code, / +def\s+__init__\s*\(\s*\)\:/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:

--fcc-editable-region--
```
