---
id: 656896ffecbf07a2d3402a93
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Inside the loop, initialize an empty list `row_list` to store the elements of a single row in the sudoku board.

# --hints--

Ви повинні мати змінну `row_list`.

```js
({ test: () =>
  {
    const str = __helpers.python.getDef(code.replace(/\r/g, ''), "__str__");
    const {function_body} = str;    
    assert(function_body.match(/row_list\s*=/));
  }
})
```

`row_list` має бути порожнім списком.

```js
({ test: () =>
  {
    const str = __helpers.python.getDef(code.replace(/\r/g, ''), "__str__");
    const {function_body} = str;    
    assert(function_body.match(/row_list\s*=\s*\[\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

--fcc-editable-region--
        for index, line in enumerate(self.board):

--fcc-editable-region--
```
