---
id: 6569dffeee007f26d2b56d46
title: Крок 52
challengeType: 20
dashedName: step-52
---

# --description--

Within the `Board` class, create another method `is_valid`. It would take three parameters:

- `self` (representing the instance of the class),
- `empty` (a tuple representing the row and column indices of an empty cell)
- `num` (representing the number to be checked).

This method checks if a given number is a valid choice for an empty cell in the sudoku board by validating its compatibility with the row, column, and 3x3 square of the specified empty cell.

# --hints--

The method name should be `is_valid`.

```js
const tCode = code.replace(/\r/g, '');
const board = __helpers.python.getBlock("\n" + tCode, "class Board");
const valid = __helpers.python.getDef(board.block_body, "is_valid");
assert.exists(valid);
```

Метод має приймати три параметри: `self`, `empty` та `num`. Order matters

```js
const tCode = code.replace(/\r/g, '');
const valid = __helpers.python.getDef(tCode, "is_valid");
const {function_parameters} = valid;
assert.match(function_parameters, /self\s*,\s*empty\s*,\s*num/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('║')

            row = f'║ {" ".join(row_list)} ║\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

            if index < 8:
                if index % 3 == 2:
                    board_string += f'╠═══{"╪═══"*2}{"╬═══"}{"╪═══"*2}{"╬═══"}{"╪═══"*2}╣\n'
                else:
                    board_string += middle_lines
            else:
                board_string += lower_lines

        return board_string

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(
            self.board[row][col] != num
            for row in range(9)
        )

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start=(col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--

--fcc-editable-region--
```
