---
id: 660699aabc59c532f2d556e5
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Щоб викликати метод екземпляра, потрібно використати точкову нотацію:

```py
instance_name.method_name()
```

Де `instance_name` — це екземпляр або об’єкт, а `method_name` — це метод, який потрібно викликати.

Викличте метод `spam` об’єкта `gameboard`.

# --hints--

Викличте метод `spam` об’єкта `gameboard`, використавши `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')

gameboard = Board()
--fcc-editable-region--
```
