---
id: 66069e5759b800364707988e
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Повертаючись до методу `__init__`: він вимагає додаткового параметра, що представляє головоломку, яку потрібно розв’язати.

Додайте другий параметр під назвою `board` до методу `__init__` та виправте створення екземпляра `gameboard`, передавши до нього список `puzzle` (як ви б передали аргумент до виклику функції).

# --hints--

Метод `__init__` повинен мати два параметри в такому ж порядку: `self` та `board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

Змінна `gameboard` повинна мати значення `Board(puzzle)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
