---
id: 6606c14182435d4bab0de2ee
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Далі ви будете створювати метод, який перевіряє, чи можна вставити цифру у вказаний стовпець сітки (перевіряючи, чи в цьому стовпці немає даної цифри).

В межах класу `Board` створіть метод під назвою `valid_in_col` та надайте йому три параметри: `self`, `col` та `num`. Де `col` та `num` відповідно представляють індекс стовпця та цифру, яку потрібно перевірити.

# --hints--

Створіть новий метод під назвою `valid_in_col` в межах класу `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

Метод `valid_in_col` повинен мати три параметри в такому ж порядку: `self`, `col` та `num`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--

--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
