---
id: 66069167b3307b2f4067b22b
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Новий екземпляр класу створюється за допомогою нотації функції, що передбачає додавання пари круглих дужок до назви класу.

Створіть екземпляр класу `Board` поза визначенням класу та призначте його до змінної під назвою `gameboard`.

# --hints--

Оголосіть змінну `gameboard` поза класом `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Змінна `gameboard` повинна мати значення `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
