---
id: 6606927d010be4300a4e5330
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Створення екземпляра створює порожній об’єкт. Але класи можуть мати методи, які є локальними функціями для кожного екземпляра. У межах класу методи оголошуються наступним чином:

```py
class ClassName:
    def method_name():
        pass
```

Замініть `pass` в межах класу `Board` на порожній метод `spam`.

# --hints--

Визначте метод під назвою `spam` в межах класу `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass

gameboard = Board()
--fcc-editable-region--
```
