---
id: 66069b992c1c5e3451f3deb0
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Створення екземпляра створює порожній об’єкт. Метод `__init__` — це спеціальний метод, який дозволяє створити екземпляр об’єкта в налаштованому стані. Коли клас реалізує метод `__init__`, то `__init__` автоматично викликається під час створення екземпляра.

Видаліть метод `spam` з класу `Board` та замініть його на метод `__init__`, який містить параметр `self`.

# --hints--

Клас `Board` не повинен містити метод `spam`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Визначте метод `__init__` в класі `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Метод `__init__` повинен мати параметр `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')

gameboard = Board()
--fcc-editable-region--
```
