---
id: 66068fb0bfddba2b7977eb60
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви дізнаєтесь про класи та об’єкти, створивши розв’язувач судоку.

Клас в Python — це основа для створення об’єктів. Об’єкти, створені з класу, є екземплярами цього класу. Клас можна створити за допомогою такого синтаксису:

```py
class ClassName:
    pass
```

Де `class` — це ключове слово, необхідне для визначення класу, а `ClassName` — це ім’я класу, написане за конвенцією у *ПаскалевомуРегістрі*.

Почніть зі створення класу `Board`.

# --hints--

Створіть клас під назвою `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
