---
id: 660699119472f332798860ad
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер замініть `pass` на виклик `print` та передайте до нього рядок `'Spam!'`.

# --hints--

Видаліть `pass` та надрукуйте рядок `'Spam!'` у методі `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass

gameboard = Board()
--fcc-editable-region--
```
