---
id: 6606bc4e5535c0484990ccd5
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Далі ви будете працювати над методом, який перевіряє, чи можна вставити задану цифру у конкретний ряд сітки.

В межах класу `Board` створіть метод під назвою `valid_in_row` та надайте йому три параметри: `self`, `row` та `num`. Де `self` представляє екземпляр класу, а `row` та `num` відповідно є індексом ряду та цифрою, яку потрібно перевірити.


# --hints--

Створіть новий метод під назвою `valid_in_row` в межах класу `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

Метод `valid_in_row` повинен мати три параметри в такому ж порядку: `self`, `row` та `num`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--

--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
