---
id: 6606cc754a8834509cd0afb6
title: Крок 36
challengeType: 20
dashedName: step-36
---

# --description--

Ще одна річ, яку потрібно перевірити, — чи можна вставити цифру в квадрат 3×3.

В межах класу `Board` створіть метод під назвою `valid_in_square` з чотирма параметрами: `self`, `row`, `col` та `num`. Де `row`, `col` та `num` відповідно представляють індекс ряду, індекс стовпця та цифру, яку потрібно перевірити.

# --hints--

Створіть новий метод під назвою `valid_in_square` в межах класу `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_square")`)) })
```

Метод `valid_in_square` повинен мати чотири параметри в такому ж порядку: `self`, `row`, `col` та `num`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_square").has_args("self, row, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))
--fcc-editable-region--

--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
