---
id: cf1111c1c11feddfaeb3bdef
title: إضافة رقمين باستخدام JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` هو نوع من أنواع البيانات في لغة JavaScript و هو يمثل بيانات رقمية.

حاول إضافة رقمين باستخدام JavaScript.

تستخدم JavaScript رمز `+` كمشغل الإضافة عند وضعه بين رقمين.

**مثال:**

```js
const myVar = 5 + 10;
```

لدي `myVar` الآن القيمة `15`.

# --instructions--

غير `0` بحيث يكون المجموع `20`.

# --hints--

يجب أن `sum` يساوي `20`.

```js
assert(sum === 20);
```

يجب عليك استخدام المعامل `+`.

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
