---
id: cf1231c1c11feddfaeb5bdef
title: ضرب رقمين باستخدام JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

يمكننا أيضًا ضرب رَقْم في آخر.

يستخدم JavaScript رمز `*` لضرب رقمين.

**مثال**

```js
const myVar = 13 * 13;
```

لدي `myVar` القيمة `169` الآن.

# --instructions--

غيّر `0` بحيث يساوي الناتج (product) `80`.

# --hints--

يجب أن يساوي المتغير `product` قيمة 80.

```js
assert(product === 80);
```

يجب عليك استخدام المشغل `*`.

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'product = '+z;})(product);
```

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
