---
id: bd7993c9c69feddfaeb7bdef
title: اضرب رقمين عشريين بواسطة JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/ce2GeHq'
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

في JavaScript، يمكنك أيضًا إجراء حسابات بأرقام عشرية، تماما مثل الأرقام صحيحة.

دعونا نضرب عددين عشريين معًا لنحصل على حاصل ضربهما.

# --instructions--

غيّر `0.0` بحيث يساوي الناتج `5.0`.

# --hints--

يجب أن يساوي المتغير `product` قيمة `5.0`.

```js
assert(product === 5.0);
```

يجب عليك استخدام المشغل `*`

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(y){return 'product = '+y;})(product);
```

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
