---
id: 642206e054eef81b5e3092ed
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Du musst eine Zahl vor dem Text `dollars` abgleichen. While you could write out `0|1|2` and so on, regular expressions have a feature that makes this easier.

Eine <dfn>Zeichenklasse</dfn> wird durch eckige Klammern definiert und stimmt mit jedem Zeichen innerhalb der Klammern überein. Zum Beispiel stimmt `[aeiou]` mit jedem Zeichen in der `aeiou`-Liste überein. Du kannst auch einen Bereich von Zeichen definieren, der mit einem Bindestrich abgeglichen werden soll. Beispielsweise stimmt `[a-z]` mit jedem Zeichen von `a` bis `z` überein.

Füge eine Zeichenklasse hinzu, die die Ziffern `0` bis `9` mit deinem `dollarRegex`-Ausdruck verbindet - beachte, dass die Ziffer vor dem Wort `dollars` stehen muss, und zwischen der Ziffer und dem Wort ein Leerzeichen sein sollte.

# --hints--

Deine `dollarRegex` sollte eine Zeichenklasse verwenden.

```js
assert.match(dollarRegex.source, /\[.*\]/);
```

Deine Zeichenklasse sollte `0-9` sein.

```js
assert.match(dollarRegex.source, /\[0-9\] dollars/);
```

Deine `dollarRegex` sollte mit `1 dollars` übereinstimmen.

```js
assert.match('1 dollars', dollarRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width" />
    <title>Learn Regular Expressions by Building a Spam Filter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <header class="main-text">
      <h1 class="title">Is this Spam?</h1>
      <p class="description">
        Enter a phrase to check if it would be marked as spam or not.
      </p>
    </header>

    <main>
      <label class="message-label" for="message-input">Message: </label>
      <textarea
        placeholder="Enter message here"
        value=""
        type="text"
        name="message"
        id="message-input"
        rows="10"
        cols="40"
      ></textarea>
      <button class="btn" id="check-message-btn" type="button">
        Check message
      </button>
      <p id="result"></p>
    </main>

    <footer class="footer">&copy; freeCodeCamp</footer>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
  color: var(--light-grey);
}

body,
#message-input:placeholder-shown {
  text-align: center;
}

textarea {
  max-width: 90%;
}

.main-text {
  margin: 25px 0;
}

.title {
  font-size: 2.5rem;
}

.description {
  margin-top: 15px;
  font-size: 1.4rem;
}

.message-label {
  display: block;
  margin-bottom: 20px;
  font-size: 1.5rem;
}

#message-input:placeholder-shown,
textarea {
  font-size: 1.1rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 200px;
  margin: 10px auto;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result {
  font-size: 2rem;
  margin: 20px 0;
}

.footer {
  margin-top: 10px;
}
```

```js
const messageInput = document.getElementById("message-input");
const result = document.getElementById("result");
const checkMessageButton = document.getElementById("check-message-btn");

const helpRegex = /please help|assist me/i;
--fcc-editable-region--
const dollarRegex = /dollars/i;
--fcc-editable-region--

const denyList = [helpRegex, dollarRegex];

const isSpam = (msg) => denyList.some((regex) => regex.test(msg));

checkMessageButton.addEventListener("click", () => {
  if (messageInput.value === "") {
    alert("Please enter a message.");
    return;
  }

  result.textContent = isSpam(messageInput.value)
    ? "Oh no! This looks like a spam message."
    : "This message does not seem to contain any spam.";
  messageInput.value = "";
});
```
