---
id: 615f3cafd794015aa9547a6d
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Ricorda che l'uso di `h1`, `h2` e tag simili determina la struttura semantica dell'HTML. Tuttavia, è possibile regolare il CSS di questi elementi per controllarne il flusso visivo e la gerarchia.

Crea una regola `h1` e imposta la proprietà `font-weight` su `800`. Questo renderà il tuo testo `h1` più spesso.

# --hints--

Dovresti creare un selettore `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Il selettore `h1` dovrebbe avere una proprietà `font-weight` con il valore `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
