---
id: 615f4f9e4a40566b776a8f38
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

Crea un nuovo selettore `.lg` e assegnagli una proprietà `height` con valore di `10px`. Crea anche un selettore `.lg, .md` e imposta la proprietà `background-color` con valore di `black`.

# --hints--

Dovresti avere un nuovo selettore `.lg`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.lg'));
```

Il tuo selettore `.lg` dovrebbe avere una proprietà `height` impostata su `10px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.lg')?.height === '10px');
```

Dovresti avere un nuovo selettore `.lg, .md`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.lg, .md'));
```

Il tuo selettore `.lg, .md` dovrebbe avere una proprietà `background-color` impostata su `black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.lg, .md')?.backgroundColor === 'black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

--fcc-editable-region--

--fcc-editable-region--
```
