---
id: 615f3b091162165948e1cb33
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Se ispezioni l'elemento `.label` con gli strumenti da sviluppatore del tuo browser, potrai vedere che in realtà è largo 288 pixel invece di 270. Questo perché, come impostazione predefinita, il browser include il bordo e il padding quando determina la dimensione di un elemento.

Per risolvere questa situazione, ripristina il box model creando un selettore `*` e assegnandogli una proprietà `box-sizing` con il valore `border-box`.

# --hints--

Dovresti creare un selettore `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

Il selettore `*` dovrebbe avere una proprietà `box-sizing` con il valore `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
