---
id: 615f2abbe7d18d49a1e0e1c8
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Ti abbiamo fornito il testo stardard HTML di base.

Crea un elemento `h1` all'interno dell'elemento `body` e assegnagli il testo `Nutrition Facts`.

# --hints--

Dovresti aggiungere un nuovo elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

L'elemento `h1` dovrebbe essere all'interno dell'elemento `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

L'elemento `h1` dovrebbe avere il testo `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
