---
id: 615f34ecc1091b4fd5a8a484
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

All'interno del tuo elemento `head`, aggiungi un elemento `link` con l'attributo `rel` impostato a `stylesheet<code>` e l'attributo `href` impostato a `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Questo importerà la famiglia di caratteri `Open Sans`, con spessori di `400`, `700`, e `800`.

Aggiungi anche un elemento `link` per collegare il tuo file `styles.css`.

# --hints--

Dovresti avere due elementi `link`.

```js
assert(code.match(/<link/g)?.length === 2);
```

Entrambi gli elementi `link` dovrebbero avere un attributo `rel` con valore di `stylesheet`.

```js
assert(code.match(/<link[\s\S]*?rel=('|"|`)stylesheet\1/)?.length === 2);
```

Uno degli elementi `link` dovrebbe avere un attributo `href` con valore di `./styles.css`.

```js
assert(code.match(/<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Uno degli elementi `link` dovrebbe avere attributo `href` con valore di `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>

  </head>
  <body>
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </body>
</html>
--fcc-editable-region--
```

```css

```
