---
id: 615f34ecc1091b4fd5a8a484
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

All'interno dell'elemento `head`, aggiungi un elemento `link` con l'attributo `rel` impostato su `stylesheet` e l'attributo `href` con il valore `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

In questo modo, importerai la famiglia di caratteri `Open Sans`, con spessori di `400`, `700`, e `800`.

Aggiungi anche un elemento `link` per collegare il file `styles.css`.

# --hints--

Dovresti avere due elementi `link` auto-concludenti.

```js
assert(document.querySelectorAll('link').length === 2);
```

Entrambi gli elementi `link` dovrebbero avere un attributo `rel` con il valore `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno degli elementi `link` dovrebbe avere un attributo `href` con il valore `./styles.css`.

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Uno degli elementi `link` dovrebbe avere attributo `href` con il valore `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
