---
id: 615f378014c2da526a109c81
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Il carattere è un po' piccolo. Crea un selettore `html` e aumenta la dimensione del carattere a `16px`.

# --hints--

Dovresti avere un selettore `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

Il selettore `html` dovrebbe avere una proprietà `font-size` con il valore `16px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.fontSize === '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

body {
  font-family: 'Open Sans', sans-serif;
}
```
