---
id: 615f2d4150fe0d4cbd0f2628
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Sotto il tuo elemento `h1`, aggiungi un elemento `p` con il testo `8 servings per container`.

# --hints--

Dovresti aggiungere un nuovo elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Il tuo elemento `p` dovrebbe essere all'interno del tuo elemento `body`.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

Il tuo elemento `p` dovrebbe venire dopo il tuo elemento `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Il tuo elemento `p` dovrebbe avere il testo `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
  </head>
  <body>
    <h1>Nutrition Facts</h1>

  </body>
</html>
--fcc-editable-region--
```

```css

```
