---
id: 615f34ecc1091b4fd5a8a484
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

All'interno dell'elemento `head`, aggiungi un elemento `link` con l'attributo `rel` impostato su `stylesheet` e l'attributo `href` con il valore `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

In questo modo, importerai la famiglia di caratteri `Open Sans`, con spessori di `400`, `700`, e `800`.

Aggiungi anche un elemento `link` per collegare il file `styles.css`.

# --hints--

Your code should have two `link` elements.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Your `link` elements should be void elements, they should not have a end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Your two `link` elements should be inside the `head` element.

```js
const headContentRegex = /(?<=<head\s*>)(?:.|\s*)*?(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Entrambi gli elementi `link` dovrebbero avere un attributo `rel` con il valore `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

One of your `link` elements should have an `href` attribute set to `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Uno degli elementi `link` dovrebbe avere attributo `href` con il valore `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
