---
id: 615f34948891834dd77655a6
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Aggiungi un secondo elemento `p` con il testo `Serving size 2/3 cup (55g)`.

# --hints--

Dovresti avere un secondo elemento `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Il secondo elemento `p` dovrebbe essere all'interno dell'elemento `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Il secondo elemento `p` dovrebbe essere all'interno dell'elemento `p`.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Il secondo elemento `p` dovrebbe avere il testo `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
