---
id: 615f671b6d1919792745aa5d
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--

Crea un nuovo elemento `div` sotto l'elemento `.medium`. Assegnagli un attributo `class` impostato su `daily-value small-text`. All'interno di questo nuovo `div`, crea un elemento `p` con il testo `% Daily Value *` e imposta l'attributo `class` con il valore `bold right`.

# --hints--

Dovresti creare un nuovo elemento `div` dopo l'elemento `.medium`.

```js
assert(document.querySelector('.label')?.lastElementChild?.localName === 'div');
assert(document.querySelector('.label')?.lastElementChild?.previousElementSibling?.classList?.contains('divider'));
assert(document.querySelector('.label')?.lastElementChild?.previousElementSibling?.classList?.contains('medium'));
```

Il nuovo `div` dovrebbe avere un attributo `class` con il valore `daily-value small-text`.

```js
assert(document.querySelector('.label')?.lastElementChild?.classList?.contains('daily-value'));
assert(document.querySelector('.label')?.lastElementChild?.classList?.contains('small-text'));
```

Il nuovo elemento `div` dovrebbe avere un elemento `p`.

```js
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.localName === 'p');
```

Il nuovo elemento `p` dovrebbe avere il testo `% Daily Value *`.

```js
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.textContent === '% Daily Value *');
```

Il nuovo elemento `p` dovrebbe avere un attributo `class` con il valore `bold right`.

```js
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.classList?.contains('bold'));
assert(document.querySelector('.label')?.lastElementChild?.firstElementChild?.classList?.contains('right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
        <p>Calories</p>
      </div>
      <span>230</span>
    </div>
    --fcc-editable-region--
    <div class="divider medium"></div>

    --fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

.medium {
  height: 5px;
}

.small-text {
  font-size: 0.85rem;
}

.calories-info {
  display: flex;
  justify-content: space-between;
  align-items: flex-end;
}

.calories-info h2 {
  margin: 0;
}

.left-container p {
  margin: -5px -2px;
  font-size: 2em;
  font-weight: 700;
}

.calories-info span {
  margin: -7px -2px;
  font-size: 2.4em;
  font-weight: 700;
}
```
