---
id: 69b1f12e7941b5947294426c
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--
 
In this workshop, you will build a scatterplot of animals based on their weight and lifespan. 

Start by creating an `svg` variable and assign it a D3 selection of the `body` element. Then, append an `svg` element to the `body` and give it a class of `graph`.

# --hints--

You should have a variable called `svg`.

```js
assert.isDefined(svg);
```

Your `svg` variable should be a D3 selection of the `body` element with an `svg` element appended to it.

```js
assert.equal(svg.node().tagName, "svg");
```

Your `svg` element should have a class of `graph`.

```js
assert.isTrue(svg.classed("graph"));
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
