---
id: 69b1f12e7941b5947294426c
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

After previous lesson, now we know importance of Data and Data Visualisation. In this Workshop, you will build a scatterplot of animals based on their `Weight` and `Lifespan`. Let's get fimiliar with D3.js library.

Now we will create a empty container, which is the primary step to create any plot, graph or chart.

```js
const svg = d3.select("body")
  .append("svg")
  .attr("width", 600)
  .attr("height", 400)
  .style("border", "1px solid #000000");
```

Here, a `svg` element is appended to the HTML `body`, with attribute `Width`, `Height` and `Border` (style).

```js
   .attr("class", "graph")
```

Now create a `svg` with `class` : `graph`, and append it to the `body`. After doing it properly you'll see a empty box in the in preview.

# --hints--

Select the `body`, append a `svg` element, which has `class`:`graph` atrribute.

```js
// assert.match(
//   code,
//   /const\s+svg\s*=\s*d3\s*\.select\(\s*"body"\s*\)\s*\.append\(\s*"svg"\s*\)\s*\.attr\(\s*"class"\s*,\s*"graph"\s*\)\s*;?/
// );
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
