---
id: 69baa985f829ef676ddb3e8e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

In the previous step, you created a container for your graph. Now you will need to create some variables to hold the dimensions for your graph.

Create three variables called `margin`, `maxHeight` and `maxWidth`.

Your `margin` variable should have a value of an object with the following properties:

```md
top: 60,
right: 60,
bottom: 50,
left: 60
```

Your `maxHeight` variable should have a value of `500` and your `maxWidth` variable should have a value of `800`.

# --hints--

You should have a variable called `margin`.

```js
assert.isDefined(margin);
```

Your `margin` variable should be an object.

```js
assert.isObject(margin);
```

Your `margin` object should have a `top` property with a value of `60`.

```js
assert.equal(margin.top, 60);
```

Your `margin` object should have a `right` property with a value of `60`.

```js
assert.equal(margin.right, 60);
```

Your `margin` object should have a `bottom` property with a value of `50`.

```js
assert.equal(margin.bottom, 50);
```

Your `margin` object should have a `left` property with a value of `60`.

```js
assert.equal(margin.left, 60);
```

You should have a variable called `maxHeight` with a value of `500`.


```js
assert.equal(maxHeight, 500);
```

You should have a variable called `maxWidth` with a value of `800`.

```js
assert.equal(maxWidth, 800);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
```
