---
id: 69baa985f829ef676ddb3e8e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

In step-1, you have built a container . Now, you will create *3* constant as *margin*, *maxHeight* and *maxWidth*.

*margin* should be object with top=60, right=60, bottom=50 and left=60 attributes.  
*maxHeight* and *maxWidth* is maximum height and width limit for `svg` container.  

Create these constants according to the above description.


# --hints--


You should create a `const` named *margin*.


```js
assert.isObject(margin);
```


You should create a `const` named *maxHeight*.


```js
assert.isDefined(maxHeight);
```

You should create a `const` named *maxWidth*.


```js
assert.isDefined(maxWidth);
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
```
