---
id: 69baa985f829ef676ddb3e8e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

**Let's make me container responsive**  
In step-1, you have built a container. But that was static container, now you will make it responsive for different `windowWidth`. For this just add a Event listener on the window, call a fuction that resizes that container for each window resize, simple.

```js
const maxHeight = 500;
const maxWidth = 800;

function resizeFunction() {
  const windowWidth = window.innerWidth;

  let width = Math.min(maxWidth, windowWidth * 0.9);  
  let height = Math.min(maxHeight, window.innerHeight * 0.7);

  d3.select("container")
    .attr("width", width)
    .attr("height", height);
}

resizeFunction();
window.addEventListener("resize", resizeFunction);
```
  
Now, create resizing a function named *drawScatterplot*, that will update *width* and *height* of svg container. 


# --hints--


You should create a function *drawScatterplot*


```js
assert.match(
  code,
  /function\s+drawScatterplot\s*\(\s*\)\s*\{[\s\S]*\}/
);
```


Select "svg" and modify *width* and *height*


```js
assert.match(
  code,
  /d3\s*\.select\(\s*"svg"\s*\)\s*\.attr\(\s*"width"\s*,[\s\S]*?\)\s*\.attr\(\s*"height"\s*,[\s\S]*?\)\s*;?/
);
```

Call the *drawScatterplot* and add event listener for later call on window resize


```js
assert.match(
  code,
  /drawScatterplot\s*\(\s*\)\s*;\s*window\s*\.addEventListener\(\s*"resize"\s*,\s*drawScatterplot\s*\)\s*;?/
);
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

--fcc-editable-region--

--fcc-editable-region--
```
