---
id: 69bba2f81c07d5d804cc4399
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

So far, you have created a responsive `svg` container, Now, Let's draw red circle inside.

```js
const svg_circle = d3
  .select("body")
  .append("svg")
  .attr("width", '800px')
  .attr("height", '500px')
  .style("border", '2px solid black'); 

// appends a circle element at (cx,cy):(0,0) left-top position 
svg_circle
  .append("circle")
  .attr("cx", 0)   // cx goes Left to Right
  .attr("cy", 0)   // cy goes Top to Buttom
  .attr("r", 50)
  .style("fill", "red");
```

Now append a `circle` element into the `svg` you have created earlier. You will see **1/4th** of a red circle, rest of the circle is outside of the circle.  
Then you can play with different `(cx,cy)` center location.  

It gives clear understanding of co-ordinate system of `svg`.  

# --hints--


Append `circle` with `cx`, `cy`, `r` and `fill` attribute.


```js
assert.match(code, /svg\s*\.append\(\s*"circle"\s*\)/);;
```

Add center *(cx,cy)* to the circle

```js
assert.match(code, /\.attr\(\s*"cx"\s*,\s*\d+\s*\)\s*\.attr\(\s*"cy"\s*,\s*\d+\s*\)/);
```

Add radius *(r)* to circle

```js
assert.match(code, /\.attr\(\s*"r"\s*,\s*\d+\s*\)/);
```

Add fill attribute to the circle

```js
assert.match(code, /\.style\(\s*"fill"\s*,[\s\S]*?\)/);
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  --fcc-editable-region--

  --fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
