---
id: 69bbd3f66dfa39e277252772
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Now, you will learn about scale generator of d3 library.   
There several types of [scale generators](https://d3js.org/d3-scale) in D3 library. Here, you will focus in two scale generator: 

  1. scaleOrdinal()
  2. scaleLinear()

In this step, you will create a category to color mapping function using the `scaleOrdinal` function.
`scaleOrdinal` maps categories to corresponding values like colors or labels. In this case, you will map animal categories to color values.

Here is an example:

```js
const color = d3.scaleOrdinal(["color1", "color2", "color3"], ["red", "green", "blue"]);
color("color1") // "red"
```

All elements of the first array `["color1", "color2", "color3"]` get mapped with each element of the second array `["red", "green", "blue"]`.   

Create an ordinal scale named `color` that maps the animal types `["Mammal", "Bird", "Reptile"]` to the colors `["red", "green", "blue"]`.

# --hints--

You should create a variable called `color` and assign it `d3.scaleOrdinal()`.

```js
assert.match(
  code,
  /(const|let)\s+color\s*=\s*d3\s*\.scaleOrdinal\s*\([\s\S]*\)\s*;?/
);
```

Your scale should map `["Mammal", "Bird", "Reptile"]` to `["red", "green", "blue"]`.

```js
assert.match(
  code,
  /d3\s*\.scaleOrdinal\s*\(\s*\[\s*["']Mammal["']\s*,\s*["']Bird["']\s*,\s*["']Reptile["']\s*\]\s*,\s*\[\s*["']red["']\s*,\s*["']green["']\s*,\s*["']blue["']\s*\]\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

--fcc-editable-region--

--fcc-editable-region--

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  
  svg
    .append("circle")
    .attr("cx", 0)   
    .attr("cy", 0)  
    .attr("r", 50)
    .style("fill", "red");
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);

```
