---
id: 69bbe55e7a490ec46081ab51
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

In this step, you will create a `g` element in `svg` and transform it by applying `margin.left` along the x‑axis and `margin.top` along the y‑axis.

```js
const svg = d3
  .select("body")
  .append("svg")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);
```

To keep free margin space both horizontally and vertically, the `g` element has been translated.

Now, add a `g` element inside the `svg` container and translate it by `margin.left` on the x‑axis and `margin.top` on the y‑axis.

# --hints--

You should append a group inside the `svg`.

```js
assert.match(
  code,
  /(const|let)\s+svg\s*=\s*d3\s*\.select\s*\(\s*["']body["']\s*\)\s*\.append\s*\(\s*["']svg["'][\s\S]*\)\s*\.append\s*\(\s*["']g["']\s*\)/
);
```

Your `g` element should be traslated.

```js
assert.match(
  code,
  /\.attr\s*\(\s*["']transform["']\s*,\s*`translate\(\$\{margin\.left\},\$\{margin\.top\}\)`\s*\)/
);
```



# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  --fcc-editable-region--

  --fcc-editable-region--


function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);


  const color = d3.scaleOrdinal(["r", "g", "b"], ["red", "green", "blue"]);
  const x = d3.scaleLinear([0, 100], [0, width]); 
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
