---
id: 69bbd3f66dfa39e277252772
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

**Let's get familier with `scale generator's` in D3**  
There several types of [scale generators](https://d3js.org/d3-scale) in D3 library. Here we will focus in to two scale generator: 

  1. scaleOrdinal()
  2. scaleLinear()

**scaleOrdinal()**

```js
const color = d3.scaleOrdinal(["a", "b", "c"], ["red", "green", "blue"]);
color("a") // "red"
```

Here all element first array `["a", "b", "c"]` get mapped with each element of second array `["red", "green", "blue"]`. 

**scaleLinear()**  
You know about Linear equation, right! This scale is like a replica of linear equations.

```js
const f = d3.scaleLinear([0, 100], [0, 1000])
```

It maps 0 -> 0  
        1 -> 10  
        2 -> 20  
        .  
        .  
        50 -> 500  
        .  
        .  
        100 -> 1000  
Similarly for : 

```js
const f = d3.scaleLinear([0, 100], [1000, 0])
```

It maps 0 -> 1000  
        1 -> 900  
        2 -> 800  
        .  
        .  
        50 -> 500  
        .  
        .  
        100 -> 0  

That's how linear scale maps a continous renge to another continous renge.  
Now, 

  1. create a ordinal scale for `type to color` named as color. ["r", "g", "b"] -> ["red", "green", "blue"]
  2. create a linear scale x that maps (0,100) -> (0,width) of content.



# --hints--

Create ordinal scale as `color`

```js
assert.match(
  code,
 /const\s+color\s*=\s*d3\.scaleOrdinal\(\s*\[\s*"r"\s*,\s*"g"\s*,\s*"b"\s*\]\s*,\s*\[\s*"red"\s*,\s*"green"\s*,\s*"blue"\s*\]\s*\);\s*/
);
```

Create a Linear scale as `x`


```js
assert.match(
  code,
  /const\s+x\s*=\s*d3\.scaleLinear\(\s*\[\s*0\s*,\s*100\s*\]\s*,\s*\[\s*0\s*,\s*width\s*\]\s*\)/
);
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  
  svg
    .append("circle")
    .attr("cx", 0)   
    .attr("cy", 0)  
    .attr("r", 50)
    .style("fill", "red");

  --fcc-editable-region--

  --fcc-editable-region--

  console.log(color("Mammal"));    
  console.log(color("Bird"));    
  console.log(color("Reptile"));    
  console.log(color("Amphibian"));   

  console.log(); 
  console.log("Total width: " + width);   
  console.log(x(0));    
  console.log(x(25));   
  console.log(x(50));   
  console.log(x(75));   
  console.log(x(100));  
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);

```
