---
id: 69bbe55e7a490ec46081ab51
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

`g` element stands for group, It’s a container that lets you organize multiple shapes, paths, or other SVG elements together 
so you can apply transformations, styles, or attributes to them collectively. Here we will use `g` only for grouping 
everything in the scatterplot together.

```js
const svg = d3
  .select("body")
  .append("svg")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);
```

Here, `svg` is inside of `body`, then similarly `g` is appended inside `svg`. This group will contain shapes for our plot.
Here, the `g` has arrtibute of transform, it means it will start at (margin.left,margin.top) coordinate
from top-left point of the svg container.

Now, add a `g` element inside svg container we have, make sure it is starting from (margin.left,margin.top) co-ordinate relative to the top-left corner.

# --hints--

Append a group inside the `svg` and transform it.

```js
assert.match(
  code,
  /const\s+svg\s*=\s*d3\s*\.select\("body"\)\s*\.append\("svg"\)\s*\.attr\("class",\s*"graph"\)\s*\.append\("g"\)\s*\.attr\("transform",\s*`translate\(\$\{margin\.left\},\$\{margin\.top\}\)`\);/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  --fcc-editable-region--

  --fcc-editable-region--


function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);


  const color = d3.scaleOrdinal(["r", "g", "b"], ["red", "green", "blue"]);
  const x = d3.scaleLinear([0, 100], [0, width]); 
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
