---
id: 69bc4fe22f78c853b35eed76
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

**Let's draw the x-axis**  
To start creating the scatterplot, you need two axes that are x-axis and y-axis. X-axis will be mapped to width, 0 to max value on x-axis will be
mapped to width (0,width). 

```js
const x = d3
  .scaleLinear()
  .domain([0, 100])
  .range([0, width]);
```

we have learnt about Linear scale in previous scale, now using this scale we can create xAxix by axis generator **axisBottom()**  

```js
const xAxis = d3.axisBottom(x);
```

But this axis generator doesn't draws anything, now it has to be drawn inside a group `g`, then to be appended to the `svg` we have.  


```js
svg
  .append("g")
  .attr("transform", `translate(0,${height})`)
  .call(xAxis)              // draws the axis from axis generator
  .append("text")           // This text element for adding a label under x-axis
  .attr("x", width / 2)     // center horizontally
  .attr("y", 40)            // push below the axis line
  .style("text-anchor", "middle")
  .style("fill", "black")   // ensure visible color
  .style("font-size", windowWidth > 700 ? "12px" : "16px")
  .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");
```

Now draw x-axis in our context and check the preview.    
  
**Note:** One thing that if you now resize preview tab, there multiple overriding line one above of other. This happens because in each *drawScatterplot()* call new `g` are appened in the parent `svg`, so before drawing a new x-axis, we have to clear previous draws.

```js
  svg.selectAll("*").remove();
```


# --hints--

Create a d3 bottom-axis named `xAxis` using  **axisBottom()** axis generator

```js
assert.match(
  code,
  /const\s+xAxis\s*=\s*d3\.axisBottom\(x\);?/
);
```

Create a `group`, create the axis and append `text` element with position attribute.


```js
assert.match(
  code,
   /svg\s*\.append\("g"\)[\s\S]*?\.call\(xAxis\)[\s\S]*?\.append\("text"\)[\s\S]*?\.attr\("x",\s*.+\)[\s\S]*?\.attr\("y",\s*.+\)/
);
```



# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  const x = d3
    .scaleLinear()
    .domain([0, 100])
    .range([0, width]);
  
  --fcc-editable-region--

  --fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
