---
id: 69bc4fe22f78c853b35eed76
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--
  
For scatterplot, you need x-axis and y-axis. X-axis will be mapped to width, 0 to max value width on x-axis will be
mapped to width (0,width) using linear scale (scaleLinear()) you have created earlier.  

```js
const x = d3
  .scaleLinear()
  .domain([0, 100])
  .range([0, width]);
```

To create axis-x, you need a axis generator. For creating bottom axis you will use axisBottom() generator. Generator will
take linear scale as reference for the scale.

```js
const xAxis = d3.axisBottom(x);
```

But this axis generator doesn't draws anything, now it has to be drawn inside a group `g`, this task will be done using 
call metheod of `svg`. call() method creates the axis using a axis generator.

```js
svg
  .call(xAxis)              // draws the axis from axis generator
```

Now, you should draw x-axis using previously created `x` scale.

# --hints--

You should create a d3 bottom-axis generator named `xAxis`.

```js
assert.match(
  code,
  /const\s+xAxis\s*=\s*d3\.axisBottom\(\s*x\s*\);?/
);
```

You should draw `xAxis` inside the svg.

```js
assert.match(
  code,
   /svg\s*\.call\s*\(\s*xAxis\s*\)/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  const x = d3
    .scaleLinear()
    .domain([0, 100])
    .range([0, width]);
  
  --fcc-editable-region--

  --fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
