---
id: 69bd7592bfedc186d58068c2
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--
 
As of now you were working on a example that had domain of x ε [1,100], and range y ε [1,10]. Now, you have a modification that for smaller windowWidth, you want to plot in opposite manner.  
Means domain of x ε [1,10] and range y ε [1,100].  
  
To do this just add a check on windowWidth, change scale *x* and *y* accordinly.

```js
let x, y, xAxis, yAxis;

if (windowWidth > 700) {
  // Horizontal layout: Weight X, Lifespan Y
  x = d3
    .scaleLinear()
    .domain([0,100])
    .range([0, width]);
  y = d3
    .scaleLinear()
    .domain([0, 10])
    .range([height, 0]);
  xAxis = d3.axisBottom(x);
  yAxis = d3.axisLeft(y);
} else {
  // Vertical layout for mobile: Weight Y, Lifespan X (rotated)
  x = d3
    .scaleLinear()
    .domain([0, 10])
    .range([0, width]);
  y = d3
    .scaleLinear()
    .domain([0, 100])
    .range([height, 0]);
  xAxis = d3.axisBottom(x);
  yAxis = d3.axisLeft(y);
}
```

Add this responsiveness in our function and check the preview.

# --hints--

Add a if condition that *switch* the scales based window width.

```js
assert.match(
  code,
  /if\s*\(\s*windowWidth\s*>\s*\d+(\.\d+)?\s*\)\s*\{[\s\S]*\}[\s\S]*else\s*\{[\s\S]*\}/m
);
```

Create the scales and generate axes

```js
assert.strictEqual(typeof x, "function", "x should be a scale function");
assert.strictEqual(typeof y, "function", "y should be a scale function");
assert.strictEqual(typeof xAxis, "function", "xAxis should be an axis generator");
assert.strictEqual(typeof yAxis, "function", "yAxis should be an axis generator");

```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const color = d3.scaleOrdinal(["a", "b", "c"], ["red", "green", "blue"]);

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();

  
  
  --fcc-editable-region--

  --fcc-editable-region--


  svg
    .append("g")
    .attr("transform", `translate(0,${height})`)
    .call(xAxis)
    .append("text")           
    .attr("x", width / 2)    
    .attr("y", 40)           
    .style("text-anchor", "middle")
    .style("fill", "black")   
    .style("font-size", windowWidth > 700 ? "16px" : "12px")
    .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");

  svg
    .append("g")
    .call(yAxis)
    .append("text")
    .attr("transform", "rotate(-90)")
    .attr("x", -height / 2)             
    .attr("y", -margin.left + 15)        
    .style("text-anchor", "middle")      
    .style("fill", "black")              
    .style("font-size", windowWidth > 700 ? "16px" : "12px")
    .text(windowWidth > 700 ? "Lifespan (years)" : "Weight (kg)");


  svg
    .append("circle")
    .style("fill","transparent")
    .style("stroke","black")
    .attr("r", 15)
    .attr("cx", x(0))
    .attr("cy", y(0));
    
  svg
    .append("circle")
    .style("fill","transparent")
    .style("stroke","black")
    .attr("r", 15)
    .attr("cx", x(0))
    .attr("cy", y(10));

  svg
    .append("circle")
    .style("fill","transparent")
    .style("stroke","black")
    .attr("r", 15)
    .attr("cx", x(100))
    .attr("cy", y(0));
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
