---
id: 69be5d6aff9618eb5bd51dc2
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--
 
Our x scale here mapping [0,100] to width.

```js
x = d3
  .scaleLinear()
  .domain([0,100])
  .range([0, width]);
```

But this `100` is irrevelant to us. Here, we want [0,max_weight_of_animals] to be mapped [0,width].
We can find max_weight_of_animals, using *d3.max()* funtion.

```js
let Array1 = [10, 20, 30, 40, 50, 60];
let Array2 = [1, 2];
let Array3 = [0, 1.5, 6.8];
let Array4 = [.8, .08, .008];
let Array_object = [{"name":"jon doe" , "age": 25}, {"name":"jon doe" , "age": 30}];

// Calling to d3.max() function
let A,B,C,D,E;
A = d3.max(Array1);    //60
B = d3.max(Array2);    // 2
C = d3.max(Array3);    // 6.8
D = d3.max(Array4);    //0,8
E = d3.max(Array_object, (each_element) => each_element.age);   // 30
```

Now, map x scale to [0,max_weight_of_animals] and y scale to [0,max_lifespan_of_animals].

```js
x = d3
  .scaleLinear()
  .domain([0, d3.max(Array, (each_element) => each_element.Weight)])
  .range([0, width]);
y = d3
  .scaleLinear()
  .domain([0, d3.max(Array, (each_element) => each_element.Lifespan)])
  .range([height, 0]);
```



# --hints--

*x* scales domain will be [0,max_weight_of_animals]

```js
assert.match(
  code,
  /x\s*=\s*d3\s*\.scaleLinear\(\)\s*\.domain\(\s*\[\s*0\s*,\s*d3\.max\(data\s*,\s*\(\s*each_element\s*\)\s*=>\s*each_element\.Weight\s*\)\s*\]\s*\)\s*\.range\(\s*\[\s*0\s*,\s*width\s*\]\s*\)/m
);
```


*y* scales domain will be [0,max_lifespan_of_animals]


```js
assert.match(
  code,
  /y\s*=\s*d3\s*\.scaleLinear\(\)\s*\.domain\(\s*\[\s*0\s*,\s*d3\.max\(data\s*,\s*\(\s*each_element\s*\)\s*=>\s*each_element\.Lifespan\s*\)\s*\]\s*\)\s*\.range\(\s*\[\s*height\s*,\s*0\s*\]\s*\)/m
);
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

const color = d3
  .scaleOrdinal()
  .domain(["Mammal", "Bird", "Reptile", "Amphibian"])
  .range(["#1f77b4", "#ff7f0e", "#2ca02c", "#d62728"]);



function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();

  let x, y, xAxis, yAxis;

  d3.json("https://cdn.freecodecamp.org/curriculum/workshop-build-a-zoo-animal-scatterplot/data.json").then((data) => {

    --fcc-editable-region--

    --fcc-editable-region--

    xAxis = d3.axisBottom(x);
    yAxis = d3.axisLeft(y);

    svg
      .append("g")
      .attr("transform", `translate(0,${height})`)
      .call(xAxis)
      .append("text")           
      .attr("x", width / 2)    
      .attr("y", 40)           
      .style("text-anchor", "middle")
      .style("fill", "black")   
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");

    svg
      .append("g")
      .call(yAxis)
      .append("text")
      .attr("transform", "rotate(-90)")    
      .attr("x", -height / 2)             
      .attr("y", -margin.left + 15)        
      .style("text-anchor", "middle")      
      .style("fill", "black")              
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Lifespan (years)" : "Weight (kg)");
  });
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
