---
id: 69be5704f659249afe2d9e0c
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

**Let's create color Ordinal scales for animal types**  
Before we have built a ordinal scale Like: 

```js
const color = d3.scaleOrdinal(["a", "b", "c"], ["red", "green", "blue"]);
color("a") // "red"
```

Now, create a ordinal scale that maps
["Mammal", "Bird", "Reptile", "Amphibian"] to ["#1f77b4", "#ff7f0e", "#2ca02c", "#d62728"].  

Here, all of the animal type in *data.json* has been mapped to defferent color. Examine them using : 

```js
console.log(color("Amphibian"));
```

# --hints--

Create a Ordinal scale based on named `color`

```js
assert.match(
  code,
  /(?:const\s+)?color\s*=\s*d3\s*[\s\S]*\.scaleOrdinal\([\s\S]*\)\s*/m
);
```

Map animal to color in correct order.

```js
assert.match(
  code,
  /\[\s*"Mammal"\s*,\s*"Bird"\s*,\s*"Reptile"\s*,\s*"Amphibian"\s*\][\s\S]*\[\s*"#1f77b4"\s*,\s*"#ff7f0e"\s*,\s*"#2ca02c"\s*,\s*"#d62728"\s*\]/m);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;
 

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

--fcc-editable-region--

--fcc-editable-region--


function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();

  let x, y, xAxis, yAxis;

  d3.json("https://cdn.freecodecamp.org/curriculum/workshop-build-a-zoo-animal-scatterplot/data.json").then((data) => {


    if (windowWidth > 700) {
      x = d3
        .scaleLinear()
        .domain([0,100])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, 10])
        .range([height, 0]);
      xAxis = d3.axisBottom(x);
      yAxis = d3.axisLeft(y);
    } else {
      x = d3
        .scaleLinear()
        .domain([0, 10])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, 100])
        .range([height, 0]);
      xAxis = d3.axisBottom(x);
      yAxis = d3.axisLeft(y);
    }
  
    svg
      .append("g")
      .attr("transform", `translate(0,${height})`)
      .call(xAxis)
      .append("text")           
      .attr("x", width / 2)    
      .attr("y", 40)           
      .style("text-anchor", "middle")
      .style("fill", "black")   
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");

    svg
      .append("g")
      .call(yAxis)
      .append("text")
      .attr("transform", "rotate(-90)")    
      .attr("x", -height / 2)             
      .attr("y", -margin.left + 15)        
      .style("text-anchor", "middle")      
      .style("fill", "black")              
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Lifespan (years)" : "Weight (kg)");
  });
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
