---
id: 69c0f14cd948e0c56b9208c1
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--
 
Now, you will add a event listener to each circle, that will show some information about that animal on *mouseover*.

```js
const tooltip = d3
  .select("body")
  .append("div")
  .attr("class", "tooltip")
  .style("opacity", 0);
```

On that purpose, a tooltip `div` is created with class="tooltip", this div will be shown different position with necessary information.

```js
svg
  .append("circle")
  .attr("r", 12)
  .attr("cx", 100)
  .attr("cy", 100)
  .style("fill", "pink")
  .on("mouseover", (event, d) => {
    tooltip.transition().duration(200).style("opacity", 0.9);
    tooltip
      .html(
        `${d.Name} (${d.Type})<br/>Weight: ${d.Weight} kg<br/>Lifespan: ${d.Lifespan} yrs`,
      )
      .style("left", event.pageX + "px")
      .style("top", event.pageY - 28 + "px");
  })
  .on("mouseout", () =>
    tooltip.transition().duration(200).style("opacity", 0),
  );
```

*mouseover* on any circle triggers a callback function that makes changes on tooltip div. It changes of tooltip position to (event.pageX, event.pageY - 28), here 
(event.pageX, event.pageY) is the position from event is triggered.
  
As our circles are a selection, it will also receive a data for each element. (here *d*)  

Now, add this tooltip to the circle selection.

# --hints--

Add a *mouseover* event listener in each `circle` with callback function.

```js
assert.match(
  code,
/\.on\(\s*"mouseover"\s*,\s*\(event\s*,\s*d\)\s*=>\s*\{[\s\S]*\}\s*\)/m
);
```

Add a *mouseout* event listener in each `circle` with callback function.

```js
assert.match(
  code,
/\.on\(\s*"mouseout"\s*,\s*\(\)\s*=>[\s\S]*\)\s*;/m
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

const color = d3
  .scaleOrdinal()
  .domain(["Mammal", "Bird", "Reptile", "Amphibian"])
  .range(["#1f77b4", "#ff7f0e", "#2ca02c", "#d62728"]);

const tooltip = d3
  .select("body")
  .append("div")
  .attr("class", "tooltip")
  .style("opacity", 0);

  

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();

  let x, y, xAxis, yAxis;

  d3.json("https://cdn.freecodecamp.org/curriculum/workshop-build-a-zoo-animal-scatterplot/data.json").then((data) => {

    if (windowWidth > 700) {

      x = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Weight) + 100])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Lifespan) + 10])
        .range([height, 0]);
    } else {

      x = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Lifespan) + 10])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Weight) + 100])
        .range([height, 0]);
    }

    xAxis = d3.axisBottom(x);
    yAxis = d3.axisLeft(y);

    svg
      .append("g")
      .attr("transform", `translate(0,${height})`)
      .call(xAxis)
      .append("text")           
      .attr("x", width / 2)    
      .attr("y", 40)           
      .style("text-anchor", "middle")
      .style("fill", "black")   
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");

    svg
      .append("g")
      .call(yAxis)
      .append("text")
      .attr("transform", "rotate(-90)")    
      .attr("x", -height / 2)             
      .attr("y", -margin.left + 15)        
      .style("text-anchor", "middle")      
      .style("fill", "black")              
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Lifespan (years)" : "Weight (kg)");

    svg
      .selectAll(".dot")
      .data(data)
      .enter()
      .append("circle")
      .attr("class", "dot")
      .attr("r", windowWidth > 700 ? 12 : 8)
      .attr("cx", (d) => (windowWidth > 700 ? x(d.Weight) : x(d.Lifespan)))
      .attr("cy", (d) => (windowWidth > 700 ? y(d.Lifespan) : y(d.Weight)))
      .style("fill", (d) => color(d.Type))
      --fcc-editable-region--

      --fcc-editable-region--

    const legend = svg
            .selectAll(".legend")
            .data(color.domain())
            .enter()
            .append("g")
            .attr("class", "legend")
            .attr("transform", (d, i) => `translate(0,${i * 20})`);
          
    legend
      .append("rect")
      .attr("x", width - 18)
      .attr("width", 18)
      .attr("height", 18)
      .style("fill", color);

    legend
      .append("text")
      .attr("x", width - 24)
      .attr("y", 9)
      .text((d) => d)
      .attr("dy", ".35em")
      .style("text-anchor", "end")
      .style("font-size", windowWidth > 700 ? "16px" : "12px");
  });
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```






# --solutions--


```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```


```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```


```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

const color = d3
  .scaleOrdinal()
  .domain(["Mammal", "Bird", "Reptile", "Amphibian"])
  .range(["#1f77b4", "#ff7f0e", "#2ca02c", "#d62728"]);

const tooltip = d3
  .select("body")
  .append("div")
  .attr("class", "tooltip")
  .style("opacity", 0);

  

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();

  let x, y, xAxis, yAxis;

  d3.json("https://cdn.freecodecamp.org/curriculum/workshop-build-a-zoo-animal-scatterplot/data.json").then((data) => {

    if (windowWidth > 700) {

      x = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Weight) + 100])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Lifespan) + 10])
        .range([height, 0]);
    } else {

      x = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Lifespan) + 10])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Weight) + 100])
        .range([height, 0]);
    }

    xAxis = d3.axisBottom(x);
    yAxis = d3.axisLeft(y);

    svg
      .append("g")
      .attr("transform", `translate(0,${height})`)
      .call(xAxis)
      .append("text")           
      .attr("x", width / 2)    
      .attr("y", 40)           
      .style("text-anchor", "middle")
      .style("fill", "black")   
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");

    svg
      .append("g")
      .call(yAxis)
      .append("text")
      .attr("transform", "rotate(-90)")    
      .attr("x", -height / 2)             
      .attr("y", -margin.left + 15)        
      .style("text-anchor", "middle")      
      .style("fill", "black")              
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Lifespan (years)" : "Weight (kg)");

    svg
      .selectAll(".dot")
      .data(data)
      .enter()
      .append("circle")
      .attr("class", "dot")
      .attr("r", windowWidth > 700 ? 12 : 8)
      .attr("cx", (d) => (windowWidth > 700 ? x(d.Weight) : x(d.Lifespan)))
      .attr("cy", (d) => (windowWidth > 700 ? y(d.Lifespan) : y(d.Weight)))
      .style("fill", (d) => color(d.Type))
      .on("mouseover", (event, d) => {
        tooltip.transition().duration(200).style("opacity", 0.9);
        tooltip
          .html(
            `${d.Name} (${d.Type})<br/>Weight: ${d.Weight} kg<br/>Lifespan: ${d.Lifespan} yrs`,
          )
          .style("left", event.pageX + "px")
          .style("top", event.pageY - 28 + "px");
      })
      .on("mouseout", () =>
        tooltip.transition().duration(200).style("opacity", 0),
      );


    const legend = svg
            .selectAll(".legend")
            .data(color.domain())
            .enter()
            .append("g")
            .attr("class", "legend")
            .attr("transform", (d, i) => `translate(0,${i * 20})`);
          
    legend
      .append("rect")
      .attr("x", width - 18)
      .attr("width", 18)
      .attr("height", 18)
      .style("fill", color);

    legend
      .append("text")
      .attr("x", width - 24)
      .attr("y", 9)
      .attr("dy", ".35em")
      .style("text-anchor", "end")
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text((d) => d); 
  });
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
