---
id: 69c047e8f9be3bc18162080a
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

You have a array of object carrying information about different animals, like: 

```js
{
  Name: 'Blue Whale',
  Type: 'Mammal',
  Weight: 120000,
  Lifespan: 90 
}
```

Now, to plot them in the plane, we are assuming (x,y) = (Weight,Lifespan).  
For, this let's understand some D3.js methods: 

```js
svg
  .selectAll(".dot")
  .data(data)
  .enter()
  .append("circle")
  .attr("class", "dot")
  .attr("r", 12)
  .attr("cx", (d) =>  x(d.Weight) )
  .attr("cy", (d) => y(d.Lifespan) )
  .style("fill", (d) => color(d.Type))
```

*selectAll()* selects all element with class="dot", as there is no element with class="dot", so this selection is
empty. Though this *selectAll* act as indication of group, because we are going to bind this group with *data* array 
we have by *.data(data)* and *enter()* function just telling new element not present the selection currently those are to be entered.


# --hints--

Create selection with *class = "dot"* , bind it with data array, enter new elements.

```js
assert.match(
  code,
  /svg\s*\.selectAll\(\s*"\.dot"\s*\)\s*\.data\(\s*data\s*\)\s*\.enter\(\)/m
);
```

Append `circle` element with *class = "dot"*

```js
assert.match(
  code,
  /\.append\(\s*"circle"\s*\)\s*\.attr\(\s*"class"\s*,\s*"dot"\s*\)/m
);
```

Add *radius*, *center* and *color* to the `circle`


```js
assert.match(
  code,
/\.attr\(\s*"r"\s*,[\s\S]*\)\s*\.attr\(\s*"cx"\s*,\s*\(d\)\s*=>[\s\S]*\)\s*\.attr\(\s*"cy"\s*,\s*\(d\)\s*=>[\s\S]*\)\s*\.style\(\s*"fill"\s*,\s*\(d\)\s*=>[\s\S]*\)/m
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

const color = d3
  .scaleOrdinal()
  .domain(["Mammal", "Bird", "Reptile", "Amphibian"])
  .range(["#1f77b4", "#ff7f0e", "#2ca02c", "#d62728"]);

  

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();

  let x, y, xAxis, yAxis;

  d3.json("https://cdn.freecodecamp.org/curriculum/workshop-build-a-zoo-animal-scatterplot/data.json").then((data) => {

    if (windowWidth > 700) {

      x = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Weight) + 100])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Lifespan) + 10])
        .range([height, 0]);
    } else {

      x = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Lifespan) + 10])
        .range([0, width]);
      y = d3
        .scaleLinear()
        .domain([0, d3.max(data, (d) => d.Weight) + 100])
        .range([height, 0]);
    }

    xAxis = d3.axisBottom(x);
    yAxis = d3.axisLeft(y);

    svg
      .append("g")
      .attr("transform", `translate(0,${height})`)
      .call(xAxis)
      .append("text")           
      .attr("x", width / 2)    
      .attr("y", 40)           
      .style("text-anchor", "middle")
      .style("fill", "black")   
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Weight (kg)" : "Lifespan (years)");

    svg
      .append("g")
      .call(yAxis)
      .append("text")
      .attr("transform", "rotate(-90)")    
      .attr("x", -height / 2)             
      .attr("y", -margin.left + 15)        
      .style("text-anchor", "middle")      
      .style("fill", "black")              
      .style("font-size", windowWidth > 700 ? "16px" : "12px")
      .text(windowWidth > 700 ? "Lifespan (years)" : "Weight (kg)");

    --fcc-editable-region--

    --fcc-editable-region--
  });
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
