---
id: 69cfe4d5d21fc3ac05a0e8b5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Now, You should apply width and height to the `svg` selection. Remember you have declared *width* and *height* for the content aria in last step. So, *width* of `svg` selection will be *content_aria_width + left_margin + right_margin*.  
  
Similarly, height of `svg` selection will be *content_aria_height + top_margin + bottom_margin*.  
Here below is a example for selecting in `svg` using d3 *select()* method and updating it's *width* and *height* attribute.     

```js
d3.select("svg")
  .attr("width", width)
  .attr("height", height);
```

Start selecting the `svg` you have appended into `body`, then update it's *width* and *height*.

# --hints--


Add Width attribute to svg selection.

```js
assert.match(code, /d3\.select\("svg"\)[\s\S]*\.attr\("width",\s*width[\s\S]*\)/);
```

Add Left and Right margin to the width.

```js
assert.match(code, /width\s*\+\s*margin\.left\s*\+\s*margin\.right/);
```

Add Height attribute to svg selection.

```js
assert.match(code, /d3\.select\("svg"\)[\s\S]*\.attr\("height",\s*height[\s\S]*\)/);
```

Add Left and Right margin to the width.

```js
assert.match(code, /height\s*\+\s*margin\.top\s*\+\s*margin\.bottom/);
```



# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

  
function drawScatterplot() {
  const windowWidth = window.innerWidth;
  const windowHeight = window.innerHeight;

  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

--fcc-editable-region--

--fcc-editable-region--
}
```
