---
id: 69cfd1aa0676f24eff0473da
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

At this point you will create two variable `width` and `height`. This two variable will be the *width* and *height* of content aria inside `svg` container you made earlier.  
You should keep width as 90% of windowWidth and height 70% of windowHeight.  
Additionally, *width* and *height* should not exceed *maxWidth* and *maxHeigth*.

```js
let width = Math.min(maxWidth, windowWidth * 0.9);
```

As `width` should be the width of content aria, You should reduce left and right margin from width. Example

```js
let width = Math.min(maxWidth, windowWidth * 0.9) - left_margin - right_margin;
```

Now, declare two variable `width` and `height` for content aria, where you will draw the scatterplot.


# --hints--

You should declare a variable `width` equal to min of maxWidth and 90% of window width

```js
assert.match(code, /let\s+width\s*=\s*Math\.min\(maxWidth,\s*windowWidth\s*\*\s*0\.9\)/);
```

Left and Right margin should be reduced from `width`

```js
assert.match(code, /-\s*margin\.left\s*-\s*margin\.right/);
```

You should declare a variable `height` equal to min of maxHeight and 70% of window height

```js
assert.match(code, /let\s+height\s*=\s*Math\.min\(maxHeight,\s*window\.innerHeight\s*\*\s*0\.7\)/);
```

Top and Bottom margin should be reduced from `height`

```js
assert.match(code, /-\s*margin\.top\s*-\s*margin\.bottom/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

  
function drawScatterplot() {
  const windowWidth = window.innerWidth;
  const windowHeight = window.innerHeight;

  --fcc-editable-region--

  --fcc-editable-region--
}
```
