---
id: 69cf7491752dfb71c55eb447
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Your scatterplot graph will need to know the width and height of the browser window to determine how much space it has to draw the graph.

Information for the width and height of browser window can be retrieved from the global `window` object in JavaScript.  

```js
const widthInsideBrowser = window.innerWidth;
const heightInsideBrowser = window.innerHeight;
```

Inside your `drawScatterplot` function, declare two variables called `windowWidth` and `windowHeight`. Assign `windowWidth` to the inner width of the browser window and `windowHeight` to the inner height of the browser window.

# --hints--

You should have a variable called `windowWidth` that is equal to inner window width inside the `drawScatterplot` function.

```js
assert.match(
  code,
  /(const|let)\s+windowWidth\s*=\s*window\s*\.innerWidth\s*;?/
);
```

You should have a variable called `windowHeight` that is equal to inner window height inside the `drawScatterplot` function.

```js
assert.match(
  code,
  /(const|let)\s+windowHeight\s*=\s*window\s*\.innerHeight\s*;?/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

function drawScatterplot() {
  --fcc-editable-region--
 
  --fcc-editable-region--
}
```
