---
id: 69cfd1aa0676f24eff0473da
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

The next step is to determine the height of the content area inside the `svg` container.  

Below your `width` variable, declare another variable called `height`. 

Assign the value `Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom` to the `height` variable. 

This represents the minimum of `maxHeight` and `70%` of the inner height of the browser window, minus the top and bottom margins.

# --hints--

You should have a `height` variable.

```js
assert.match(code, /(const|let)\s+height\s*=/);
```

You should assign `Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom` to your `height` variable.

```js
assert.match(code, /(const|let)\s+height\s*=\s*Math\.min\s*\(\s*maxHeight\s*,\s*window\.innerHeight\s*\*\s*0\.7\s*\)\s*-\s*margin\.top\s*-\s*margin\.bottom/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

  
function drawScatterplot() {
  const windowWidth = window.innerWidth;
  const windowHeight = window.innerHeight;

  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;

  --fcc-editable-region--

  --fcc-editable-region--
}
```
