---
id: 69cfe8a2c376f8d915f96462
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

In this step, you will update your scatterplot with the changing window size.
Start by calling the `drawScatterplot` function.

Then below that function call, add an event listener to the `window` object.  The first argument for the `addEventListener()` will be the `"resize"` event and the second argument will be a reference to the `drawScatterplot` function.

An example of addEventListener could be : 

```js
function drawScatterplot() {
  console.log("Scatterplot drawn!");
}

window.addEventListener("resize", drawScatterplot);
```

# --hints--

You should call the `drawScatterplot` function.

```js
assert.match(code, /drawScatterplot\(\s*\)\s*;/);
```

Your first argument for the `addEventListener()` should be the `"resize"` event and the second argument should be a reference to the `drawScatterplot` function.

```js
assert.match(code, /window\.addEventListener\(\s*"resize",\s*drawScatterplot\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

  
function drawScatterplot() {
  const windowWidth = window.innerWidth;
  const windowHeight = window.innerHeight;

  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

}

--fcc-editable-region--

--fcc-editable-region--
```
