---
id: 69cf7491752dfb71c55eb447
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In this step, you will create a function named *drawScatterplot* and inside this function you should create two variable *windowWidth* and *windowHeight*.

Here, *windowWidth* and *windowHeight* should be width and height of browser window. Information about width and height of *browser window* can be retreived from global *window* of JavaScript.  

```js
const width_inside_browser = window.innerWidth;
const height_inside_browser = window.innerHeight;
```

Now, create *drawScatterplot* function and declare variable *windowWidth* and *windowHeight*.

# --hints--

You should create a function `drawScatterplot`

```js
assert.isFunction(drawScatterplot);
```

You should create a variable `windowWidth` equal to inner window width.

```js
assert.match(
  code,
  /const\s+windowWidth\s*=\s*window\s*\.innerWidth\s*;?/
);
```

You should create a variable `windowHeight` equal to inner window height.

```js
assert.match(
  code,
  /const\s+windowHeight\s*=\s*window\s*\.innerHeight\s*;?/
);
```


# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = { top: 60, right: 60, bottom: 50, left: 60 };
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")


--fcc-editable-region--

--fcc-editable-region--
```
