---
id: 69e6e202a8e7c61e893ce690
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Your `drawScatterplot` function will need to determine the width of the content area inside the `svg` container.

Declare a variable called `width` and assign it `Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right`. This represents the minimum of `maxWidth` and `90%` of the `windowWidth`, minus the left and right margins.

# --hints--

You should have a `width` variable.

```js
assert.match(code, /(const|let)\s+width\s*=/);
```

You should assign `Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right` to your `width` variable.

```js
assert.match(code, /(const|let)\s+width\s*=\s*Math\.min\s*\(\s*maxWidth\s*,\s*windowWidth\s*\*\s*0\.9\s*\)\s*-\s*margin\.left\s*-\s*margin\.right/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  const windowHeight = window.innerHeight;

  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
