---
id: 69f4597f48697f7e92ddc24b
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

In this step, you will learn about the 2nd scale generator `scaleLinear()`.  
*scaleLinear()*  
In the scale, you can map continous domain to a continous range. To understand the above code look carefully to the below diagrams, that map domain value 0 to 100 with range value 0 to 1000.

```js
const f = d3.scaleLinear([0, 100], [0, 1000])
f(10);    // 100 
```

Inside d3.scaleLinear() method, the first array refers the domain and second array refers the range.

It maps 0 -> 0  
        1 -> 10  
        2 -> 20  
        .  
        .  
        50 -> 500  
        .  
        . 
        99 -> 990
        100 -> 1000  

 
 Similarly, you can map inversely, as example domain value 0 to 100 with range value 1000 to 0. 

```js
const f = d3.scaleLinear([0, 100], [1000, 0])
f(10)     //900
```

It maps 0 -> 1000  
        1 -> 990  
        2 -> 980  
        .  
        .  
        50 -> 500  
        .  
        .  
        99 -> 10
        100 -> 0  

Now, with this understanding create a linear scale `x` that maps domain 0 to 100 with range 0 to `width` of content.

# --hints--

You should create a Linear scale `x`.

```js
assert.match(
  code,
  /const\s+x\s*=\s*d3\.scaleLinear\([\s\S]*\)/
);
```

Your `x` scale should have domain value 0 to 100 and range value 0 to width.

```js
assert.match(
  code,
  /d3\s*\.scaleLinear\s*\(\s*\[\s*0\s*,\s*100\s*\]\s*,\s*\[\s*0\s*,\s*width\s*\]\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

const color = d3.scaleOrdinal(["Mammal", "Bird", "Reptile"], ["red", "green", "blue"]);


function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  
  svg
    .append("circle")
    .attr("cx", 0)   
    .attr("cy", 0)  
    .attr("r", 50)
    .style("fill", "red");

--fcc-editable-region--

--fcc-editable-region--

  console.log(color("Mammal"));    
  console.log(color("Bird"));    
  console.log(color("Reptile"));    
  console.log(color("Amphibian"));   

  console.log(); 
  console.log("Total width: " + width);   
  console.log(x(0));    
  console.log(x(25));   
  console.log(x(50));   
  console.log(x(75));   
  console.log(x(100));  
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);

```
