---
id: 69f4597f48697f7e92ddc24b
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Unlike `scaleOrdinal()`, which maps categories to values, `scaleLinear()` is used to map continuous numeric values. 

Here is an example of converting values from `0` to `100` into pixel positions between `0` and `width` so elements can be positioned correctly across the SVG.

```js
const f = d3.scaleLinear([0, 100], [0, width])
```

Create a linear scale `x` that map value `0` to `100` into pixel positions `0` to `width` continously.

# --hints--

You should created a variable called `x` and assign it `d3.scaleLinear()`

```js
assert.match(
  code,
  /(const|let)\s+x\s*=\s*d3\.scaleLinear\([\s\S]*\)/
);
```

Your `x` scale should have domain values from `0` to `100` and range values from `0` to `width`.

```js
assert.match(
  code,
  /d3\s*\.scaleLinear\s*\(\s*\[\s*0\s*,\s*100\s*\]\s*,\s*\[\s*0\s*,\s*width\s*\]\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")

const color = d3.scaleOrdinal(["Mammal", "Bird", "Reptile"], ["red", "green", "blue"]);


function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  

  --fcc-editable-region--

  --fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);

```
