---
id: 69f50d9c3022e3c419abec2e
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

In this step, you will a text label below x-axis.  
You will append a `text` element inside the same group x-axis is residing. Then you will place the text in middle of width of x-axis, and some 40px
below of that group.

```js
svg
  .append("g")
  .append("text")           
  .attr("x", width / 2)    
  .attr("y", 40) 
```

Then you can add some `css` styling to text, like "text-anchor" as "middle", "fill" as "black", "font-size" as "12px" and then set the text.

```js
svg
  .append("g")
  .append("text")                    
  .style("text-anchor", "middle")
  .style("fill", "black")   
  .style("font-size", "12px")
  .text("x-axis");
```

Now add a `text` element in the middle and below x-axis as "x-axis".

# --hints--

You should add a `text` element into the `g` where xAxis residing.

```js
assert.match(
  code,
  /svg\s*\.append\s*\(\s*["']g["']\s*\)[\s\S]*\.call\s*\(\s*xAxis\s*\)\s*\.append\s*\(\s*["']text["']\s*\)/
);
```

You should position the `text` at the middle of the width and 40px below the `g`.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.attr\s*\(\s*["']x["']\s*,\s*width\s*\/\s*2\s*\)\s*\.attr\s*\(\s*["']y["']\s*,\s*40\s*\)/
);
```

You should set style of `fill` attribute as `black`.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.style\s*\(\s*["']fill["']\s*,\s*["']black["']\s*\)/
);
```

You should add "x-axis" as 'text' attribute in the `text` element.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.text\s*\(\s*["']x-axis["']\s*\)/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  const x = d3
    .scaleLinear()
    .domain([0, 100])
    .range([0, width]);

  const xAxis = d3.axisBottom(x);
  
  svg
    .append("g")
    .attr("transform", `translate(0,${height})`)
    .call(xAxis)
--fcc-editable-region--

--fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
