---
id: 69f5003e6808e654306cae6c
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

In last step, you have drawn the x-axis. But x-axis is placed in the top of the `g` element. 
In this step, you will place the x-axis in right place.

For this you will wrap the axis inside a `g` element, then you will move the group to the bottom.

```js
svg
  .append("g")
  .call(xAxis)
```

Above code shows the axis is drawn inside an `g` element. To translate the group to the bottom, you can use the transform attribute of `g` element.

```js
svg
  .append("g")
  .attr("transform", `translate(0,${height})`)   // it translates 'height' unit along the y-axis
```

Now, you should place the `xAxis` in bottom of `svg`.

# --hints--

You should draw `xAxis` inside a `g` element.

```js
assert.match(
  code,
  /svg\s*\.append\s*\(\s*["']g["']\s*\)[\s\S]*\.call\s*\(\s*xAxis\s*\)/
);
```

You should translate the `g` to the bottom.

```js
assert.match(
  code,
  /svg\s*\.append\s*\(\s*["']g["']\s*\)\s*\.attr\s*\(\s*["']transform["']\s*,\s*`translate\(0,\$\{height\}\)`\s*\)/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);

  const x = d3
    .scaleLinear()
    .domain([0, 100])
    .range([0, width]);

  const xAxis = d3.axisBottom(x);
  
  svg
    .call(xAxis)
  
--fcc-editable-region--

--fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
