---
id: 6a09870b04bf7553d96a5265
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

In this step, you will create a linear scale for y-axis.  
Scale `y` will map value `0` to `10` into pixel positions `height` to `0` continously.

```js
const y = d3.scaleLinear([0, 10],[height, 0])
```

Create a linear scale `y` that map value `0` to `100` into pixel positions `height` to `0` continously.

# --hints--

You should created a variable called `y` and assign it `d3.scaleLinear()`

```js
assert.match(
  code,
  /(const|let)\s+y\s*=\s*d3\.scaleLinear\([\s\S]*\)/
);
```

Your `y` scale should have domain values from `0` to `100` and range values from `height` to `0`.

```js
assert.match(
  code,
  /d3\s*\.scaleLinear\s*\(\s*\[\s*0\s*,\s*100\s*\]\s*,\s*\[\s*height\s*,\s*0\s*\]\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js
const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
  
  svg.selectAll("*").remove();
  
  const x = d3.scaleLinear([0, 100],[0, width])
  const xAxis = d3.axisBottom(x);

  --fcc-editable-region--

  --fcc-editable-region--

  
  svg
    .append("g")
    .attr("transform", `translate(0,${height})`)
    .call(xAxis)
    .append("text")           
    .attr("x", width / 2)    
    .attr("y", 40)           
    .style("text-anchor", "middle")
    .style("fill", "black")   
    .style("font-size", "12px")
    .text('x-axis');
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
