---
id: 6a09900acc7ad946bcb39a00
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

In this step, you will add a vertical text label for the y‑axis.  
The `text` should be rotated 90° anti‑clockwise, and its transformed `(x, y)` position should be `(-height / 2, -margin.left + 15)`.

An example of adding rotated `text` below:  

```js
svg
  .append("text")
  .attr("transform", "rotate(-90)")    // rotate the label
  .attr("x", height)              // center vertically (negative because of rotation)
  .attr("y", margin.left)        // move left of the axis line
  .text("y-axis");
```

Now, add a vertical text label for the y‑axis.  

# --hints--

You should append a `text` element to the `svg` for the y‑axis.

```js
assert.match(
  code,
  /svg[\s\S]*\.call\s*\(\s*yAxis\s*\)[\s\S]*\.append\s*\(\s*["']text["']\s*\)/
);
```

You should rotate the `text` by ‑90 degrees.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.attr\s*\(\s*["']transform["']\s*,\s*["']rotate\(-90\)["']\s*\)/
);
```

You should position the `text` at `-height / 2` on the x‑axis.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.attr\s*\(\s*["']x["']\s*,\s*-\s*height\s*\/\s*2\s*\)/
);
```

You should position the `text` at `-margin.left + 15` on the y‑axis.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.attr\s*\(\s*["']y["']\s*,\s*-\s*margin\.left\s*\+\s*15\s*\)/
);
```

You should add `"y-axis"` as the text content of the `text` element.

```js
assert.match(
  code,
  /\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.text\s*\(\s*["']y-axis["']\s*\)/
);
```

You should set style of `fill` attribute as `black`.

```js
assert.match(
  code,
  /svg[\s\S]*\.call\s*\(\s*yAxis\s*\)[\s\S]*\.append\s*\(\s*["']text["']\s*\)[\s\S]*\.style\s*\(\s*["']fill["']\s*,\s*["']black["']\s*\)/
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Zoo Animal Scatterplot</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <script src="script.js"></script>
  </body>
</html>
```

```css
.graph {
  display: block;
  margin: auto;
  border: 1px solid #000;
  background-color: white;
}

body {
  font: 10px sans-serif;
}

.axis path,
.axis line {
  fill: none;
  stroke: #000;
  shape-rendering: crispEdges;
}

.dot {
  stroke: #000;
  opacity: 0.8;
}

div.tooltip {
  position: absolute;
  padding: 10px;
  font: 12px sans-serif;
  background: lightsteelblue;
  border-radius: 8px;
  pointer-events: none;
}
```

```js

const margin = {
  top: 60,
  right: 60,
  bottom: 50,
  left: 60,
};
const maxHeight = 500;
const maxWidth = 800;

const svg = d3
  .select("body")
  .append("svg")
  .attr("class", "graph")
  .append("g")
  .attr("transform", `translate(${margin.left},${margin.top})`);

function drawScatterplot() {
  const windowWidth = window.innerWidth;
  let width = Math.min(maxWidth, windowWidth * 0.9) - margin.left - margin.right;
  let height = Math.min(maxHeight, window.innerHeight * 0.7) - margin.top - margin.bottom;

  d3.select("svg")
    .attr("width", width + margin.left + margin.right)
    .attr("height", height + margin.top + margin.bottom);
    
  svg.selectAll("*").remove();

  const x = d3.scaleLinear([0, 100],[0, width]);
  const xAxis = d3.axisBottom(x);

  const y = d3.scaleLinear([0, 100],[height, 0]);
  const yAxis = d3.axisLeft(y);

  svg
  .append("g")
  .attr("transform", `translate(0,${height})`)
  .call(xAxis)
  .append("text")           
  .attr("x", width / 2)    
  .attr("y", 40)           
  .style("text-anchor", "middle")
  .style("fill", "black")   
  .style("font-size", "12px")
  .text('x-axis');

  svg
    .append("g")
    .call(yAxis)
  --fcc-editable-region--

  --fcc-editable-region--
}

drawScatterplot();
window.addEventListener("resize", drawScatterplot);
```
