---
id: 5d5a813321b9e3db6c106a46
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

To keep track of the player's experience points, we've declared a variable called `xp` and assigned it the starting value of 0.

Create another variable to keep track of health and start it at 100.

# --hints--

See description above for instructions.

```js
assert(health === 100);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
