---
id: 62a3a7e4f1060e2fc5ffb34b
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Create another variable called `currentWeapon` and set it to `0`.

When a variable name has multiple words, the convention in JavaScript is to use what's called <dfn>camelCase</dfn>. The first word is lowercase, and the first letter of every following word is uppercase.

```js
let thisIsCamelCase;
```

# --hints--

You should use `let` to declare a variable called `currentWeapon`.

```js
assert.match(code, /let\s+currentWeapon/i);
```

You should use camelCase to name your variable.

```js
assert.match(code, /currentWeapon/);
```

Your `currentWeapon` variable should be set to `0`.

```js
assert.equal(currentWeapon, 0);
```

You should initialize your variable to `0`.

```js
assert.match(code, /let\s+currentWeapon\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
