---
id: 62a39f5a5790eb27c1e5d4bf
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Your view has been switched to your new `script.js` file. Remember that you can use the tabs above to switch between files.

Add your `console.log("Hello World");` line to this file, and see it appear in your console.

# --hints--

You should have a `console.log("Hello World");` line in your code.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
