---
id: 62a3a0a3c0a4b32915d26a6e
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Remove your `console.log("Hello World");` line to begin writing your project code.

In JavaScript, a <dfn>variable</dfn> is used to hold a value. To use a variable, you must first <dfn>declare</dfn> it. For example, to declare a variable called `camperbot`, you would write:

```js
let camperbot;
```

The `let` keyword tells JavaScript you are declaring a variable. Declare a variable called `xp`.

# --hints--

You should not have a `console.log("Hello World");` line in your code.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

You should use the `let` keyword to declare your variable.

```js
assert.match(code, /let/);
```

You should declare a variable named `xp`.

```js
assert.match(code, /let\s+xp/);
```

You should not assign a value to your variable.

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

Don't forget the semi-colon at the end of the line.

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
