---
id: 62a3b49686792938718b90d3
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

The variables you have assigned have all had values that are numbers. JavaScript has multiple different data types. The next one you will use is the <dfn>string</dfn>. Strings are used to store things like words or text. Strings are surrounded with double quotes, single quotes, or backticks. Here is an example of declaring a variable with a string:

```js
let developer = "Naomi";
```

Assign the `inventory` variable to have the value of `stick`.

# --hints--

You should set `inventory` to the string `stick`.

```js
assert.equal(inventory, "stick");
```

You should initialize your `inventory` variable with the string `stick`.

```js
assert.match(code, /let\s+inventory\s*=\s*('|"|`)stick\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
