---
id: 5dc17d3bf86c76b9248c6eb4
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Kipengele cha `p` kinatumika kuunda aya ya maandishi kwenye tovuti. Unda kipengele cha `p` chini ya kipengele chako cha `h2` na ukipe maandishi yafuatayo:

`See more cat photos in our gallery.`

# --hints--

Kipengele chako cha `p` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Kipengele chako cha `p` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/p\>/));
```

Maandishi ya kipengele chako cha `p` yanapaswa kuwa `See more cat photos in our gallery.` Aidha umesahau maandishi, au kuna makosa ya herufi.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

Kipengele chako cha `p` kinapaswa kuwa chini ya kipengele chako cha `h2`. Unazo katika mpangilio mbaya.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

