---
id: 5dc17dc8f86c76b9248c6eb5
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Kutoa maoni hukuruhusu kuacha ujumbe bila kuathiri onyesho la kivinjari. Pia hukuruhusu kufanya msimbo kutofanya kazi. Maoni katika HTML huanza na `<!--`, yana idadi yoyote ya mistari ya maandishi, na kuishia na `-->`.

Here is an example of a comment with the `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Ongeza maoni juu ya kipengele cha `p` na maandishi haya:

`TODO: Add link to cat photos`

# --hints--

Maoni yako yanapaswa kuanza na `<!--`. Unakosa herufi moja au zaidi zinazofafanua mwanzo wa maoni.

```js
assert(code.match(/<!--/));
```

Maoni yako yanafaa kuishia na `-->`.  Unakosa herufi moja au zaidi zinazofafanua mwisho wa maoni.

```js
assert(code.match(/-->/));
```

Msibo wako haupaswi kuwa na vibambo vya ziada vya kufungua/kufunga maoni. Una `<!--` au `-->` ya ziada inayoonyeshwa kwenye kivinjari.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Maoni yako yanapaswa kuwa na maandishi `TODO: Add link to cat photos`.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

Maoni yako yanapaswa kuwa juu ya kipengele cha `p`. Unazo katika mpangilio mbaya.

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

