---
id: 5dc1798ff86c76b9248c6eb3
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Vipengee vya vichwa vya `h1` kupitia `h6` hutumika kuashiria umuhimu wa maudhui yaliyo chini yao. Kadiri nambari inavyopungua ndivyo umuhimu unavyoongezeka, kwa hivyo vipengele vya `h2` vina umuhimu mdogo kuliko vipengee vya `h1`.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Tumia kipengele kimoja cha `h1` pekee kwa kila ukurasa na uweke vichwa vya umuhimu wa chini chini ya vichwa vya umuhimu wa juu.

Chini ya kipengele cha `h1`, ongeza kipengele cha `h2` na maandishi haya:

`Cat Photos`

# --hints--

Kipengele chako cha `h1` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/h1\>/));
```

Unapaswa kuwa na kipengele kimoja cha `h1` tu. Ondoa ziada yoyote.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Maandishi ya kipengele cha `h1` yanapaswa kuwa 'CatPhotoApp'. Aidha umesahau maandishi, au kuna makosa ya herufi.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Kipengele chako cha `h2` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Kipengele chako cha `h2` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/h2\>/));
```

Maandishi ya kipengele cha `h2` yanapaswa kuwa 'Cat Photos'. Weka maandishi ya `Cat Photos` pekee kati ya tagi za kufungua na kufunga za`h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Kipengele chako cha `h2` kinapaswa kuwa chini ya kipengele cha `h1`. Kipengele cha `h1` kina umuhimu mkubwa na lazima kiwe juu ya kipengele cha `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

