---
id: 5dc174fcf86c76b9248c6eb2
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Vipengele vya HTML vina tagi za ufunguzi kama `<h1>` na tagi za kufunga kama `</h1>`.

Maandishi ya kipengele huenda kati ya tagi zake za kufungua na kufunga.

Tafuta kipengele cha `h1` na ubadilishe maandishi yake kuwa:

`CatPhotoApp`

# --hints--

Maandishi `CatPhotoApp` yanapaswa kuwepo kwenye msimbo. Huenda ukahitaji kuangalia tahajia yako.

```js
assert(code.match(/catphotoapp/i));
```

Kipengele chako cha `h1` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/h1\>/));
```

Una zaidi ya kipengele kimoja cha `h1`. Ondoa kipengele cha ziada cha `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Maandishi ya kipengele cha `h1` yanapaswa kuwa `CatPhotoApp`. Aidha umeacha maandishi, umeandika makosa, au haiko kati ya tagi za kufungua na kufunga za kipengele cha `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Inaonekana unatumia kiendelezi cha kivinjari ambacho kinarekebisha ukurasa. Hakikisha umezima viendelezi vyote vya kivinjari.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

