---
id: 5dc23f9bf86c76b9248c6eba
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Unaweza kuongeza picha kwenye tovuti yako kwa kutumia kipengele cha `img`. Vipengele vya `img` vina tagi inayofungua bila tagi ya kufunga. Tagi ya kipengele bila tagi ya kufunga inajulikana kama <dfn>self-closing tag</dfn>.

Ongeza kipengele cha `img` chini ya kipengele cha `p`. Katika hatua hii, hakuna picha itaonyeshwa kwenye kivinjari.

# --hints--

Kipengele chako cha `img` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Kipengele chako cha `img` hakipaswi kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(!code.match(/<\/img\>/));
```

Unapaswa kuwa na kipengele kimoja tu cha `img`. Ondoa ziada yoyote.

```js
assert(document.querySelectorAll('img').length === 1);
```

Kipengele chako cha `img` kinapaswa kuwa chini ya kipengele cha `p`. Unazo katika mpangilio mbaya.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

