---
id: 5dc24073f86c76b9248c6ebb
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>attributes</dfn> ni maneno maalum yanayotumika ndani ya tagi ya ufunguzi wa kipengele ili kudhibiti tabia ya kipengele. Sifa ya `src` katika kipengele cha `img` inabainisha URL ya picha (ambapo picha iko).

Huu hapa ni mfano wa kipengele cha `img` chenye sifa ya `src` inayoelekeza kwenye nembo ya freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Inside the existing `img` element, add a `src` attribute with this URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `img`. Huenda umeondoa kipengele cha `img` au hujazingira sifa ya `src` kwa nukuu.

```js
assert(document.querySelector('img'));
```

Your `img` element should have a `src` attribute. Aidha umesahau sifa au umeandika makosa. Hakikisha kuna nafasi kati ya jina la kipengele na jina la sifa.

```js
assert(document.querySelector('img').src);
```

Kipengele cha `img` cha sifa yako ya `src` inapaswa kuwekwa kuwa '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Aidha umesahau URL au umeandika makosa. Kesi ya URL ni muhimu.

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Ingawa umeweka kipengele `img` cha `src` kwa URL sahihi, inashauriwa kuzunguka thamani ya sifa kila wakati kwa alama za nukuu.

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

Your `img` element should be closed either with `>` or `/>`.

```js
assert(/<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

