---
id: 5dfa22d1b521be39a3de7be0
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

You can turn any text into a link, such as the text inside of a `p` element.

``` html
<p>I think <a href="https://www.freecodecamp.org/">freeCodeCamp</a> is great.</p>
```

In the text of your `p` element, turn the words `cat photos` into a link by adding opening and closing anchor (`a`) tags around these words. Then set the `href` attribute to `https://freecatphotoapp.com/`


# --hints--

Unapaswa kuweka kipengee kipya cha (`a`) ndani ya kipengele cha `p`.

```js
assert($('p > a').length);
```

Thamani ya `href` ya kiungo inapaswa kuwa `https://freecatphotoapp.com`. Umeacha thamani ya `href` au umeandika makosa.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Maandishi ya kiungo yanapaswa kuwa `cat photos`. Aidha umesahau maandishi, au kuna makosa ya herufi.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Baada ya kuweka kipengee cha nanga (`a`), kipengele pekee cha `p` kinachoonekana kwenye kivinjari kinapaswa kuwa `See more cat photos in our gallery.` Mara mbili angalia maandishi, nafasi, au uakifishaji wa `p` na kipengele cha nanga kilichowekwa.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

