---
id: 5dfa37b9eacea3f48c6300b0
title: Hatua ya 20
challengeType: 0
dashedName: step-20
---

# --description--

To create an unordered list of items, you can use the `ul` element.

Baada ya kipengele cha `h3` chenye maandishi ya `Things cats love:`, ongeza orodha isiyopangwa ya kipengele (`ul`). Fahamu kuwa hakuna kitakachoonyeshwa kwa wakati huu.

# --hints--

Kipengele chako cha `ul` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('ul'));
```

Kipengele chako cha `ul` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/ul>/));
```

Kipengele cha `ul` kinafaa kuwa juu ya tagi ya kufunga ya kipengele cha pili cha `section`.

```js
const secondSectionLastElemNode = $('main > section')[1].lastElementChild;
assert(secondSectionLastElemNode.nodeName === 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

