---
id: 5dfa2407b521be39a3de7be1
title: Hatua ya 14
challengeType: 0
dashedName: step-14
---

# --description--

To open links in a new tab, you can use the `target` attribute on the anchor (`a`) element.

The `target` attribute specifies where to open the linked document. `target="_blank"` opens the linked document in a new tab or window.

Here is the basic syntax for an `a` element with a `target` attribute:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Ongeza sifa ya `target` yenye thamani `_blank` kwenye tagi ya ufunguzi ya kipengele (`a`), ili kiungo kifunguke katika kichupo kipya.

# --hints--

Kipengele chako cha `p` kinapaswa kuwa na kipengee cha nanga cha (`a`) chenye maandishi `cat photos`. Huenda umeifuta au kuna makosa ya kuandika.

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Kipengele chako cha nanga (`a`) hakina sifa ya `target`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
assert(document.querySelector('a').hasAttribute('target'));
```

Thamani ya sifa ya `target` inapaswa kuwa `_blank`. Umeacha thamani au umeandika makosa. Kumbuka kwamba thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

