---
id: 5dfa22d1b521be39a3de7be0
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Katika hatua ya awali, uligeuza maneno `link to cat pictures` kuwa kiungo kwa kuyaweka kati ya tagi za kufungua na kufunga (`a`). Unaweza kufanya vivyo hivyo kwa maneno ndani ya kipengele, kama vile kipengele cha `p`.

Katika maandishi ya kipengele chako cha `p`, geuza maneno `cat photos` kuwa kiungo hadi `https://freecatphotoapp.com` kwa kuongeza nanga ya ufunguzi na kufunga ya tagi (`a`) karibu na maneno haya.


# --hints--

Unapaswa kuweka kipengee kipya cha (`a`) ndani ya kipengele cha `p`.

```js
assert($('p > a').length);
```

Thamani ya `href` ya kiungo inapaswa kuwa `https://freecatphotoapp.com`. Umeacha thamani ya `href` au umeandika makosa.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Maandishi ya kiungo yanapaswa kuwa `cat photos`. Aidha umesahau maandishi, au kuna makosa ya herufi.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Baada ya kuweka kipengee cha nanga (`a`), kipengele pekee cha `p` kinachoonekana kwenye kivinjari kinapaswa kuwa `See more cat photos in our gallery.` Mara mbili angalia maandishi, nafasi, au uakifishaji wa `p` na kipengele cha nanga kilichowekwa.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

