---
id: 5dfb655eeacea3f48c6300b3
title: Hatua ya 23
challengeType: 0
dashedName: step-23
---

# --description--

Kipengele cha `figure` kinawakilisha maudhui yanayojitosheleza na kitakuruhusu kuhusisha picha na maelezo mafupi.

Weka picha uliyoongeza hivi punde ndani ya kipengele cha `figure`.

# --hints--

Kipengele chako cha `figure` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

Kipengele chako cha `figure` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/figure\>/));
```

Kunapaswa kuwa na kipengele cha `figure` juu ya tagi ya kufunga ya kipengele cha pili cha `section`.

```js
assert($('section')[1].lastElementChild.nodeName === 'FIGURE');
```

Kipengee cha lasagna cha `img` kinapaswa kuwekwa katika kipengele cha `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

