---
id: 5ef9b03c81a63668521804e3
title: Hatua ya 58
challengeType: 0
dashedName: step-58
---

# --description--

Ongeza kisanduku kingine cha kuteua baada ya kile ulichoongeza hivi punde. Sifa ya `id` inapaswa kuwa `lazy` na thamani ya sifa ya `name` inapaswa kuwa sawa na kisanduku cah kuteua cha mwisho.

Pia ongeza kipengele cha `label` upande wa kulia wa kisanduku kipya cha kuteua chenye maandishi `Lazy`. Hakikisha kuwa umehusisha kipengele cha `label` na kisanduku cha kuteua kipya kwa kutumia sifa ya `for`.

# --hints--

Unahitaji kuongeza kisanduku cha kuteua kipya.

```js
assert($('input[type="checkbox"]').length === 2);
```

Kisanduku chako kipya cha kuteua kinapaswa kuwa na sifa ya `id` yenye thamani `lazy` na sifa ya `name` yenye thamani `personality`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
const checkboxes = [...$('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

Kisanduku chako kipya cha kuteua kinapaswa kuwa baada ya cha kwanza. Unazo katika mpangilio mbaya.

```js
const checkboxes = [...$('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

Katika upande wa kulia wa kisanduku chako kipya cha kuteua, kunapaswa kuwa na kipengele cha `label` chenye maandishi `Lazy`.

```js
const nextElementSibling = $('input[type="checkbox"]')[1].nextElementSibling;
assert(
  nextElementSibling.nodeName === 'LABEL' &&
    nextElementSibling.innerText.replace(/\s+/g, '').match(/^Lazy$/i)
);
```

`label` mpya inapaswa kuwa na sifa ya `for` yenye thamani sawa na sifa ya `id` ya kisanduku kipya cha kuteua. Aidha umesahau maandishi, au kuna makosa ya maandishi.

```js
assert(
  $('input[type="checkbox"]')[1].nextElementSibling.getAttribute('for') ===
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

