---
id: 5ef9b03c81a63668521804d4
title: Hatua ya 32
challengeType: 0
dashedName: step-32
---

# --description--

Kipengele cha `strong` kinatumika kuonyesha kwamba maandishi fulani ni ya umuhimu mkubwa au ya dharura.

Katika `figcaption` uliyoongeza hivi punde, onyesha kuwa `hate` ni ya umuhimu mkubwa kwa kuifunga katika kipengele cha `strong`.

# --hints--

Kipengee chako cha `strong` kinapaswa kuwa na tagi ya ufunguzi. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('strong'));
```

Kipengee chako cha `strong` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/strong\>/));
```

Kipengele chako cha `strong` kinapaswa kuzunguka neno `hate` katika maandishi `Cats hate other cats.` Aidha umesahau maandishi, au kuna makosa ya herufi.

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.toLowerCase() === 'hate'
);
```

Maandishi ya kipengele cha `figcaption` yanapaswa kuwa `Cats hate other cats.` Angalia kama kuna makosa ya kuandika na kwamba nafasi zinazohitajika zipo karibu na kipengele cha `strong` katika tagi za kufungua na kufunga.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

