---
id: 5ef9b03c81a63668521804d8
title: Hatua ya 37
challengeType: 0
dashedName: step-37
---

# --description--

Kipengele cha `input` hukuruhusu njia kadhaa za kukusanya data kutoka kwa fomu ya wavuti. Kama vile vipengele vya `img`, vipengee vya `input` ni <dfn>self-closing</dfn> na havihitaji tagi za kufunga.

Ongeza kipengele cha `input` ndani ya kipengele chako cha `form`.

# --hints--

Kipengele chako cha `form` kinapaswa kuwa na tagi ya ufunguzi na tagi ya kufunga kwa mpangilio sahihi. Huenda unakosa tagi moja au zote mbili zinazohitajika, au kuwa nazo katika mpangilio usio sahihi.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Tagi ya ufunguzi ya kipengele chako cha `form` inapaswa kuwa na sifa ya `action` pekee. Ondoa kitu kingine chochote ambacho unaweza kuwa umeandika ndani yake.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Unapaswa kuunda kipengele cha `input`. Angalia sintaksia.

```js
assert(document.querySelector('input'));
```

Kipengele chako cha `input` kinafaa kuwa na tagi ya kufungua na tagi ya kufunga.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Kipengele cha `input` kinafaa kuwa ndani ya kipengele cha `form`.

```js
assert(document.querySelector('form > input'));
```

`form` yako inapaswa kuwa na kipengele cha `input` pekee. Ondoa vipengele au maandishi yoyote ya HTML kati ya tagi za kipengele cha `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">

        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

