---
id: 5ef9b03c81a63668521804db
title: Hatua ya 41
challengeType: 0
dashedName: step-41
---

# --description--

Ili kuzuia mtumiaji kuwasilisha fomu yako wakati maelezo yanayohitajika hayapo, unahitaji kuongeza sifa ya `required` kwa kipengele cha `input`. Hakuna haja ya kuweka thamani kwa sifa ya `required`. Badala yake, ongeza tu neno `required` kwa kipengele cha `input`, kuhakikisha kuna nafasi kati yake na sifa nyingine.

# --hints--

Umefuta kipengele chako cha `input` au kina sintaksia batili. Thamani zote za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert($('input').length);
```

`form` yako inapaswa kuwa na kipengele cha `input` pekee. Ondoa vipengele au maandishi yoyote ya HTML kati ya tagi za kipengele cha `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Kipengele chako cha `input` kinafaa kuwa na sifa ya `required`. Kumbuka, unaongeza tu neno `required` ndani ya tagi ya kipengele cha `input`.

```js
assert($('input')[0].hasAttribute('required'));
```

Thamani haipaswi kutolewa kwa sifa ya `required`.

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

