---
id: 5ef9b03c81a63668521804d9
title: Hatua ya 40
challengeType: 0
dashedName: step-40
---

# --description--

Placeholder text is used to give people a hint about what kind of information to enter into an input.

Here is an example of an `input` element with a placeholder set to `Ex. Jane Doe`:

```html
<input type="text" placeholder="Ex. Jane Doe">
```

Ongeza maandishi ya placeholder `cat photo URL` kwenye kipengele chako cha `input`.

# --hints--

Umefuta kipengele chako cha `input` au kina sintaksia batili. Thamani zote za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert($('input').length);
```

`form` yako inapaswa kuwa na kipengele cha `input` pekee. Ondoa vipengele au maandishi yoyote ya HTML kati ya tagi za kipengele cha `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Kipengele chako cha `input` hakina sifa ya `placeholder`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
assert($('input')[0].hasAttribute('placeholder'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya `placeholder` yenye thamani ya `cat photo URL`.  Umeacha thamani au umeandika makosa. Kumbuka kwamba thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert(
  $('input')[0]
    .getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

Ingawa umeweka kipengele cha `input` cha sifa ya `placeholder` kuwa `cat photo URL`, inashauriwa kuzunguka thamani ya sifa kwa alama za nukuu kila wakati.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

