---
id: 5ef9b03c81a63668521804d6
title: Hatua ya 35
challengeType: 0
dashedName: step-35
---

# --description--

Sasa utaongeza fomu ya wavuti ili kukusanya taarifa kutoka kwa watumiaji.

The `form` element is used to get information from a user like their name, email, and other details.

Baada ya kichwa cha `Cat Form`, ongeza kipengele cha `form`.

# --hints--

Kipengele chako cha `form` kinafaa kuwa na tagi ya kufungua na tagi ya kufunga. Huenda unakosa tagi moja au zote mbili zinazohitajika, au kuwa nazo katika mpangilio usio sahihi.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Tagi zako za vipengee vya `form` haziko katika mpangilio sahihi.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Kipengele cha `form` kilichowekwa katika `section` ya mwisho lazima kiwe chini ya kipengele cha `h2`. Una kipengele cha `h2` na kipengele cha `form` katika mpangilio usio sahihi.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Kipengele cha `form` hakipaswi kuwa na maudhui. Ondoa vipengele au maandishi yoyote ya HTML kati ya tagi za kipengele cha `form`.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

