---
id: 5ef9b03c81a63668521804d2
title: Hatua ya 27
challengeType: 0
dashedName: step-27
---

# --description--

Msimbo wa orodha iliyoapangwa (`ol`) ni sawa na orodha isiyopangwa, lakini orodha ya vitu katika orodha iliyopangwa huwekwa nambari inapoonyeshwa.

Baada ya kipengele cha pili cha `section` cha `h3` cha mwisho, ongeza orodha iliyopangwa na orodha hizi tatu:

`flea treatment` `thunder` `other cats`

# --hints--

Kipengele chako cha `ol` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia ifuatayo: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

Kipengele chako cha `ol` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/ol>/));
```

Kipengele cha `ol` kinafaa kuwa juu ya tagi ya kufunga ya kipengele cha pili cha `section`. Unazo katika mpangilio mbaya.

```js
assert($('main > section')[1].lastElementChild.nodeName === 'OL');
```

Vipengee vitatu vya `li` lazima viwekwe ndani ya kipengele cha `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Unapaswa kuwa na vipengele vitatu vya `li` vyenye maandishi `flea treatment`, `thunder` na `other cats` kwa mpangilio wowote.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Unapaswa kuwa na kipengele kimoja cha `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

