---
id: 5ef9b03c81a63668521804d5
title: Hatua ya 34
challengeType: 0
dashedName: step-34
---

# --description--

Ndani ya kipengele cha tatu cha `section`, ongeza kipengele cha `h2` chenye maandishi haya:

`Cat Form`

# --hints--

Haujaweza kupata kipengele cha tatu cha `section`. Huenda umeifuta kimakosa au tagi ya ufunguzi au tagi ya kufunga.

```js
assert(
  document.querySelectorAll('section').length === 3 &&
    code.match(/<\/section>/g).length === 3
);
```

Kipengele chako cha `h2` kinafaa kuwa na tagi ya kufungua na tagi ya kufunga. Huenda unakosa tagi moja au zote mbili zinazohitajika.

```js
assert(
  document.querySelectorAll('h2').length >= 3 &&
    code.match(/<\/h2>/g).length >= 3
);
```

Unapaswa kuongeza kipengele kimoja tu cha `h2`. Tafadhali ondoa ziada yoyote.

```js
assert(document.querySelectorAll('h2').length === 3);
```

Kipengele kipya cha `h2` kinapaswa kupatikana juu kabisa ya tagi ya kufunga ya kipengele cha tatu cha `section`.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(thirdSection.lastElementChild.nodeName === 'H2');
```

Maandishi ya kipengele cha `h2` yanapaswa kuwa `Cat Form`.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(
  thirdSection
    .querySelector('h2')
    .innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'cat form'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>
        </figure>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

