---
id: 5ef9b03c81a63668521804dc
title: Hatua ya 44
challengeType: 0
dashedName: step-44
---

# --description--

Unaweza kutumia vitufe vya redio (radio buttons) kwa maswali ambapo unataka jibu moja tu kati ya chaguo nyingi.

Here is an example of a radio button with the option of `cat`:

```html
<input type="radio"> cat
```

Kumbuka kwamba vipengele vya `input` vinajifunga.

Kabla ya kuingiza maandishi, ongeza kitufe cha redio na chaguo limewekwa kama:

`Indoor`

# --hints--

Unapaswa kuunda kipengele cha `input` kwa ajili ya kitufe chako cha redio. Angalia sintaksia.

```js
assert($('form > input').length >= 2);
```

Kipengele chako cha `input` kinafaa kuwa na tagi ya kufungua na tagi ya kufunga.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

Unapaswa kuwa umeongeza kipengele kimoja tu cha `input` kwa kitufe chako cha redio. Ondoa ziada yoyote.

```js
assert($('form > input').length === 2);
```

Kipengele chako kipya cha `input` lazima kiwe juu ya `input` iliyopo na sifa ya `type` iliyowekwa kuwa `text`. Unazo katika mpangilio mbaya.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Kipengele chako cha kipya cha `input` hakina sifa ya `type`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua.

```js
assert($('input')[0].hasAttribute('type'));
```

Kipengele chako kipya cha `input` kinapaswa kua na sifa moja pekee ya `type`. Ondoa ziada yoyote.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Kipengele chako kipya cha `input` kinafaa kuwa na sifa ya `type` yenye thamani ya`radio`. Aidha umesahau maandishi, au kuna makosa ya maandishi. Kumbuka kwamba thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Ingawa umeweka sifa mpya ya kipengele cha `input` ya `type` kuwa `radio`, inashauriwa kuzunguka thamani ya sifa kwa alama za nukuu kila wakati.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Maandishi ya kitufe cha `radio` ya `Indoor` yanapaswa kupatikana baada yake badala ya kabla yake.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

Maandishi `Indoor` yanapaswa kupatikana moja kwa moja upande wa kulia wa kitufe chako cha `radio`. Aidha umesahau maandishi, au kuna makosa ya maandishi.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

