---
id: 5ef9b03c81a63668521804d0
title: Hatua ya 25
challengeType: 0
dashedName: step-25
---

# --description--

To place emphasis on a specific word or phrase, you can use the `em` element.

Sisitiza neno `love` katika kipengele cha `figcaption` kwa kukifunga katika kipengele cha msisitizo cha `em`.

# --hints--

Kipengele cha msisitizo chako cha `em` kinapaswa kuwa na tagi ya ufunguzi. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Kipengele cha msisitizo chako cha `em` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/em\>/));
```

Aidha umefuta kipengele cha `figcaption` au kinakosa tagi ya ufunguzi au ya kufunga.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Msisitizo wako wa kipengele cha `em` unapaswa kuzunguka maandishi `love`. Aidha umesahau maandishi, au kuna makosa ya herufi.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Maandishi ya kipengele cha `figcaption` yanapaswa kuwa `Cats love lasagna`. Angalia kama kuna makosa ya kuandika na kwamba nafasi zinazohitajika zipo karibu na tagi za `em` za kufungua na kufunga za kipengele.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

