---
id: 5efb2c990dc218d6c85f89b2
title: Hatua ya 43
challengeType: 0
dashedName: step-43
---

# --description--

Ingawa uliongeza kitufe chako chini ya maandishi, yanaonekana karibu na kila moja kwenye ukurasa. Hiyo ni kwa sababu vipengele vyote viwili vya `input` na `button` ni <dfn>inline elements</dfn>, ambavyo havionekani kwenye laini mpya.

Kitufe ulichoongeza kitawasilisha fomu kwa chaguomsingi. Walakini, kutegemea tabia chaguo-msingi kunaweza kusababisha mkanganyiko. Ongeza sifa ya `type` yenye thamani `submit` kwenye `button` ili kuweka wazi kuwa ni kitufe cha kuwasilisha.

# --hints--

Kipengee chako cha `button` kinapaswa kuwa na tagi ya ufunguzi. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Kipengee chako cha `button` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/button\>/));
```

Kipengele chako cha `button` hakina sifa ya `type`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua.

```js
assert($('button')[0].hasAttribute('type'));
```

Kipengele chako cha `button` kinafaa kuwa na sifa ya `type` yenye thamani ya`submit`. Aidha umesahau maandishi, au kuna makosa ya maandishi. Kumbuka kwamba thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

Ingawa umeweka sifa mpya ya kipengele cha `button` ya `type` kuwa `submit`, inashauriwa kuzunguka thamani ya sifa kwa alama za nukuu kila wakati.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

