---
id: 5efb23e70dc218d6c85f89b1
title: Hatua ya 38
challengeType: 0
dashedName: step-38
---

# --description--

Kuna aina nyingi za ingizo unazoweza kuunda kwa kutumia sifa ya `type`. Unaweza kuunda sehemu ya nenosiri kwa urahisi, kitufe cha kuweka upya au kidhibiti ili kuwaruhusu watumiaji kuchagua faili kutoka kwa kompyuta zao.

Unda sehemu ya maandishi ili kupata maandishi kutoka kwa mtumiaji kwa kuongeza sifa ya `type` yenye thamani `text` kwenye kipengele cha `input`.

# --hints--

Umefuta kipengele chako cha `input` au kina sintaksia batili. Ikiwa umeongeza sifa, hakikisha thamani zao zimezungukwa na alama za nukuu.

```js
assert($('input').length);
```

`form` yako inapaswa kuwa na kipengele cha `input` pekee. Ondoa vipengele au maandishi yoyote ya HTML kati ya tagi za kipengele cha `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Kipengele chako cha `input` hakina sifa ya `type` yenye thamani ya`text`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
assert($('input')[0].hasAttribute('type'));
```

Kipengele chako kipya cha `input` kinafaa kuwa na sifa ya `type` yenye thamani ya`text`. Aidha umesahau maandishi, au kuna makosa ya maandishi. Kumbuka kwamba thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

Ingawa umeweka sifa mpya ya kipengele cha `input` ya `type` kuwa `text`, inashauriwa kuzunguka thamani ya sifa kwa alama za nukuu kila wakati.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

