---
id: 5efc4f528d6a74d05e68af74
title: Hatua ya 56
challengeType: 0
dashedName: step-56
---

# --description--

Kuna njia nyingine ya kuhusisha maandishi ya kipengele cha `input` na kipengele chenyewe. Unaweza kuweka maandishi ndani ya `label` na kuongeza sifa ya `for` yenye thamani sawa na kipengele cha `input` cha sifa ya `id`.

Here is an example of using the `for` attribute to associate a `label` with an `input` element:

```html
<label for="breakfast"> Breakfast </label>
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

Husianisha maandishi `Loving` na kisanduku cha kuteua kwa kuweka tu maandishi `Loving` katika `label` na kukipa sifa sahihi ya `for`.

# --hints--

Unapaswa kuhakikisha kuwa kitufe cha kuteua bado kipo.

```js
assert($('input[type="checkbox"]')[0]);
```

Kitufe chako cha kuteua kinafaa kuwa na sifa ya `id` yenye thamani ya `loving`.  `--fcc-expected--` inatarajiwalakini inayopatikana ni `--fcc-actual--`.

```js
assert.equal($('input[type="checkbox"]')[0].id, 'loving');
```

Maandishi `Loving` yanapaswa kufungwa katika kipengele cha `label`.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  /Loving/i.test(checkboxInputElem.nextElementSibling.innerText.replace(/\s+/g, ' '))
);
```

Utahitaji kuongeza kipengele kipya cha `label` ili kuweka maandishi `Loving`. Hakikisha kuwa ina tagi ya kufungua na kufunga.

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

Kipengele kipya cha `label` kinapaswa kupatikana moja kwa moja upande wa kulia wa kisanduku chako cha kuteua. Hakikisha kuna nafasi kati ya vipengele viwili.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

Kipengele kipya cha `label` hakina sifa ya `for`. Hakikisha kuwa kuna nafasi baada ya jina la tagi inayofungua.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

Kipengele kipya cha `label` kinafaa kuwa na sifa ya `for` yenye thamani ya`loving`.  `--fcc-expected--` inatarajiwa lakini inayopatikana ni `--fcc-actual--`.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert.equal(labelElem.getAttribute('for'), 'loving');
```

Maandishi `Loving` yanapaswa kufungwa katika kipengele cha `label`. Aidha umesahau maandishi, au kuna makosa ya maandishi.

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

Kunapaswa kuwa na nafasi kati ya kisanduku chako cha kuteua na kipengele chako kipya cha `label`.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

