---
id: 5f05a1d8e233dff4a68508d8
title: Hatua ya 47
challengeType: 0
dashedName: step-47
---

# --description--

Unda kitufe kingine cha redio chini ya cha kwanza. Iweke ndani ya kipengele cha `label` chenye `Outdoor` kama maandishi ya `label`. Ipe kitufe cha redio sifa ya `id` iliyo na `outdoor` kama thamani.

# --hints--

Kitufe asili cha redio cha`input` pekee `Indoor` ndio kinapaswa kuwa na `id` iliyowekwa `indoor`. Unaweza kuanzisha upya hatua ili kurudisha HTML asili ikihitajika.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Hupaswi kufanya mabadiliko yoyote kwenye kitufe cha redio cha `Indoor`. Unaweza kuanzisha upya hatua ili kurudisha HTML asili ikihitajika.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Unapaswa kuongeza kipengele kimoja kipya cha `input` kilichowekwa katika kipengele kipya cha `label`. Hakikisha `label` yako mpya ina tagi ya kufungua na kufunga.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Hupaswi kuongeza sifa zozote kwenye tagi ya kufungua ya `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

`input` yako mpya inapaswa kuwa na sifa ya `id` yenye thamani `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

`input` yako mpya inapaswa kuwa na sifa ya `type` yenye thamani `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Hupaswi kuongeza vipengele vyovyote vipya isipokuwa `input` iliyowekwa katika `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Kusiwe na maandishi upande wa kushoto wa `input` yako mpya.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Maandishi ya lebo ya kitufe chako kipya cha redio lazima yawe `Outdoor`. Hii ni pamoja na herufi kubwa.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Kitufe na lebo yako mpya ya redio inapaswa kuwa mara moja chini/baada ya kitufe cha redio na lebo ya `Indoor`. Kusiwe na tagi zingine kati yao.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

