---
id: 5f07fb1579dc934717801375
title: Hatua ya 33
challengeType: 0
dashedName: step-33
---

# --description--

Ni wakati wa kuongeza sehemu mpya. Ongeza kipengele cha tatu cha `section` chini ya kipengele cha pili cha `section`.

# --hints--

Kipengee chako cha `section` kinapaswa kuwa na tagi ya ufunguzi. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

Unapaswa kuongeza tu tagi moja ya ufunguzi ya `section`. Tafadhali ondoa ziada yoyote.

```js
assert(document.querySelectorAll('section').length === 3);
```

Kipengee chako cha `section` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

Unapaswa kuongeza tu tagi moja ya kufunga ya `section`. Tafadhali ondoa ziada yoyote.

```js
assert(code.match(/<\/section>/g).length === 3);
```

Vipengele vyote vya `section` vinapaswa kuwa kati ya tagi za kufungua na kufunga za kipengele cha `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

Kipengele cha mwisho cha `section` hakipaswi kuwa na maudhui. Ondoa vipengele au maandishi yoyote ya HTML kati kipengele cha `section`.

```js
assert($('main > section')[2].children.length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

