---
id: 62dabe2ef403a12d5d295273
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Kwa kuwa sasa umegeuza maandishi `cat photos` ndani ya kipengele cha `p` kuwa kiungo, huhitaji kiungo cha pili chini ya kipengele cha `p`. Futa kipengele kizima cha nanga chini ya kipengele cha `p`.

# --hints--

Msimbo wako unapaswa kuwa na kipengee kimoja tu cha nanga (`a`).

```js
assert(document.querySelectorAll('a').length === 1);
```

Kipengele cha `p` bado kinapaswa kuwa na kipengele cha nanga (`a`) ndani yake.

```js
assert($('p > a').length);
```

Thamani ya `href` ya kiungo inapaswa kuwa `https://freecatphotoapp.com`. Umeacha thamani ya `href` au umeandika makosa.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Maandishi ya kiungo yanapaswa kuwa `cat photos`. Aidha umesahau maandishi, au kuna makosa ya maandishi.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Maudhui ya kipengele cha `p` yanayoonekana kwenye kivinjari yanapaswa kuwa `See more cat photos in our gallery.` Angalia tena maandishi, nafasi, au uakifishaji wa kipengee cha `p` na nanga iliyowekwa.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

Haipaswi kuwa na maandishi ya `cat photos` chini ya kipengele cha `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
