---
id: 62dabe2ef403a12d5d295273
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Kwa kuwa sasa umegeuza maandishi `cat photos` ndani ya kipengele cha `p` kuwa kiungo, huhitaji kiungo cha pili chini ya kipengele cha `p`. Futa kipengele kizima cha nanga chini ya kipengele cha `p`.

# --hints--

Your code should only have one paragraph (`p`) element.

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

Msimbo wako unapaswa kuwa na kipengee kimoja tu cha nanga (`a`).

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

Do not make any changes to the elements you created in the previous step.

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

Haipaswi kuwa na maandishi ya `cat photos` chini ya kipengele cha `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
