---
id: 66faa20000000000000000ab
title: Build a Space Shooter with Phaser
challengeType: 26
dashedName: lab-space-shooter-with-phaser
---

# --description--

**Objective:** Build a top-down space shooter using Phaser 3 that fulfills the user stories below.

The Phaser library is loaded for you on the global `Phaser` object.

**User Stories:**

1. You should declare a class named `MainScene` that extends `Phaser.Scene`.
1. `MainScene` should have a constructor that calls `super("MainScene")`.
1. `MainScene` should have `preload`, `create`, and `update` methods.
1. You should declare a class named `GameOverScene` that extends `Phaser.Scene`.
1. `GameOverScene` should have a constructor that calls `super("GameOverScene")`.
1. You should declare a `config` object with a `type` of `Phaser.AUTO`, a `width` of `480`, a `height` of `640`, and a `parent` of `"game-container"`.
1. `config` should reference both scenes via `scene: [MainScene, GameOverScene]`.
1. `config` should configure `physics` with `default: "arcade"` and arcade gravity `{ y: 0 }`.
1. You should declare a `game` constant assigned to `new Phaser.Game(config)`.
1. `MainScene`'s `create` should set up `this.player`, `this.bullets`, `this.enemies`, `this.cursors`, `this.fireKey`, `this.score`, and `this.lives`.

# --hints--

You should declare a class named `MainScene` that extends `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{/);
```

`MainScene` should call `super("MainScene")` in its constructor.

```js
assert.match(code, /super\(\s*(['"`])MainScene\1\s*\)/);
```

`MainScene` should define `preload`, `create`, and `update` methods.

```js
assert.match(code, /preload\s*\(\s*\)\s*\{/);
assert.match(code, /create\s*\(\s*\)\s*\{/);
assert.match(code, /update\s*\(\s*[a-zA-Z_$][\w$]*\s*\)\s*\{/);
```

You should declare a class named `GameOverScene` that extends `Phaser.Scene`.

```js
assert.match(code, /class\s+GameOverScene\s+extends\s+Phaser\.Scene\s*\{/);
```

`GameOverScene` should call `super("GameOverScene")` in its constructor.

```js
assert.match(code, /super\(\s*(['"`])GameOverScene\1\s*\)/);
```

You should declare a `config` object with the required Phaser settings.

```js
assert.match(code, /const\s+config\s*=\s*\{/);
assert.match(code, /type\s*:\s*Phaser\.AUTO/);
assert.match(code, /width\s*:\s*480\b/);
assert.match(code, /height\s*:\s*640\b/);
assert.match(code, /parent\s*:\s*(['"`])game-container\1/);
```

`config` should register both scenes as an array.

```js
assert.match(code, /scene\s*:\s*\[\s*MainScene\s*,\s*GameOverScene\s*\]/);
```

`config` should configure arcade physics with zero gravity.

```js
assert.match(code, /physics\s*:\s*\{[\s\S]*?default\s*:\s*(['"`])arcade\1[\s\S]*?arcade\s*:\s*\{\s*gravity\s*:\s*\{\s*y\s*:\s*0\s*\}\s*\}/);
```

You should declare a `game` constant assigned to `new Phaser.Game(config)`.

```js
assert.match(code, /const\s+game\s*=\s*new\s+Phaser\.Game\(\s*config\s*\)/);
```

`MainScene`'s `create` should set up the player, groups, input, and score state.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\(/);
assert.match(code, /this\.bullets\s*=\s*this\.physics\.add\.group\(/);
assert.match(code, /this\.enemies\s*=\s*this\.physics\.add\.group\(/);
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(/);
assert.match(code, /this\.fireKey\s*=\s*this\.input\.keyboard\.addKey\(\s*(['"`])SPACE\1/);
assert.match(code, /this\.score\s*=\s*0\b/);
assert.match(code, /this\.lives\s*=\s*3\b/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js

const SCORE_TO_WIN = 100;
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image("player", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><polygon points='16,2 30,30 2,30' fill='%2300ff88'/></svg>");
    this.load.image("enemy", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><rect x='2' y='2' width='28' height='28' fill='%23ff3366'/></svg>");
    this.load.image("bullet", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='6' height='16'><rect width='6' height='16' fill='%23ffff00'/></svg>");
    this.load.image("star", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='40' height='40'><circle cx='10' cy='10' r='1' fill='white'/><circle cx='30' cy='25' r='1' fill='white'/><circle cx='20' cy='35' r='1' fill='white'/></svg>");
  }

  create() {
    this.background = this.add.tileSprite(240, 320, 480, 640, "star");
    this.player = this.physics.add.sprite(240, 580, "player");
    this.player.setCollideWorldBounds(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.fireKey = this.input.keyboard.addKey("SPACE");
    this.lastFired = 0;
    this.bullets = this.physics.add.group();
    this.enemies = this.physics.add.group();
    this.time.addEvent({
      delay: 1000,
      callback: this.spawnEnemy,
      callbackScope: this,
      loop: true
    });
    this.physics.add.overlap(this.bullets, this.enemies, this.hitEnemy, null, this);
    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Score: 0", { fontSize: "20px", fill: "#fff" });
    this.lives = 3;
    this.livesText = this.add.text(380, 10, "Lives: 3", { fontSize: "20px", fill: "#fff" });
    this.physics.add.overlap(this.player, this.enemies, this.hitPlayer, null, this);
  }

  update(time) {
    this.background.tilePositionY -= 2;
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-300);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(300);
    } else {
      this.player.setVelocityX(0);
    }
    if (this.cursors.up.isDown) {
      this.player.setVelocityY(-300);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(300);
    } else {
      this.player.setVelocityY(0);
    }
    if (this.fireKey.isDown && time > this.lastFired) {
      this.fireBullet();
      this.lastFired = time + 250;
    }
    this.bullets.children.each((bullet) => {
      if (bullet.y < 0) bullet.destroy();
    });
    this.enemies.children.each((enemy) => {
      if (enemy.y > 640) enemy.destroy();
    });
  }

  fireBullet() {
    const bullet = this.bullets.create(this.player.x, this.player.y - 30, "bullet");
    bullet.setVelocityY(-500);
  }

  spawnEnemy() {
    const x = Phaser.Math.Between(30, 450);
    const enemy = this.enemies.create(x, -30, "enemy");
    enemy.setVelocityY(150);
  }

  hitEnemy(bullet, enemy) {
    bullet.destroy();
    enemy.destroy();
    this.score += 10;
    this.scoreText.setText("Score: " + this.score);
    if (this.score >= SCORE_TO_WIN) {
      this.scene.start("GameOverScene", { score: this.score, win: true });
    }
  }

  hitPlayer(player, enemy) {
    enemy.destroy();
    this.lives -= 1;
    this.livesText.setText("Lives: " + this.lives);
    this.cameras.main.flash(150, 255, 0, 0);
    if (this.lives <= 0) {
      this.scene.start("GameOverScene", { score: this.score });
    }
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create(data) {
    this.add.text(240, 280, "GAME OVER", { fontSize: "48px", fill: "#ff0000" }).setOrigin(0.5);
    this.add.text(240, 340, "Score: " + data.score, { fontSize: "24px", fill: "#fff" }).setOrigin(0.5);
    if (data.win) {
      this.add.text(240, 220, "YOU WIN!", { fontSize: "36px", fill: "#00ff88" }).setOrigin(0.5);
    }
    this.add.text(240, 400, "Press R to restart", { fontSize: "18px", fill: "#fff" }).setOrigin(0.5);
    this.input.keyboard.on("keydown-R", () => {
      this.scene.start("MainScene");
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene, GameOverScene]
};

const game = new Phaser.Game(config);
```
