---
id: 612e813b3ba67633222cbe54
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Denke daran, dass ein `class`-Attribut mehrere Werte enthalten kann. Um deine weißen Schlüssel von deinen schwarzen Schlüsseln zu trennen, wirst du einen zweiten `class`-Wert von `black--key` hinzufügen. Füge dies zu deinem zweiten, dritten, fünften, sechsten und siebten `.key`-Element hinzu.

# --hints--

Dein zweites `.key`-Element sollte auch eine `class` von `black--key` enthalten.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

Dein dritter `.key` sollte `black--key` in der `class` haben.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

Dein fünfter `.key` sollte `black--key` in der `class` haben.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

Dein sechster `.key` sollte `black--key` in der `class` haben.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

Dein siebter `.key` sollte `black--key` in der `class` haben.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

Du solltest fünf `.black--key`-Elemente haben.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

Du solltest sieben `.key`-Elemente haben.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
