---
id: 612e96fc87fe8e44f69f7ec5
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Es ist Zeit, die Tasten zu gestalten. Visiere unter der `#piano`-Regel das `.keys`-Element mit einem `class`-Selektor an. Weise der neuen Regel eine `background-color`-Eigenschaft von `#040404`, eine `width`-Eigenschaft von `949px` und eine `height`-Eigenschaft von `180px` zu.

# --hints--

Du solltest einen `.keys`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

Dein `.keys`-Selektor sollte eine `background-color`-Eigenschaft von `#040404` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

Dein `.keys`-Selektor sollte die `width`-Eigenschaft auf `949px` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

Dein `.keys`-Selektor sollte eine `height`-Eigenschaft von `180px` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
