---
id: 612e8279827a28352ce83a72
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Kopiere nun die Menge von sieben `.key`-Elementen und füge zwei weitere Mengen in das `.keys`-div ein.

# --hints--

Du solltest 21 `.key`-Elemente haben.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Du solltest 15 `.black--key`-Elemente haben.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Alle `.key`-Elemente sollten sich innerhalb deines `.keys`-Elements befinden.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Alle `.black--key`-Elemente sollten sich innerhalb deines `.keys`-Elements befinden.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
