---
id: 612e89d254fe5d3df7d6693d
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Jetzt, da du das `html`-Box-Modell zurückgesetzt hast, musst du das auch an die darin liegenden Elemente weitergeben. Um dies zu tun, kannst du die `box-sizing`-Eigenschaft auf `inherit` setzen, wodurch die ausgewählten Elemente angewiesen werden, den gleichen Wert wie das Elternelement zu verwenden.

Du musst auch die Pseudo-Elemente auswählen, die spezielle Schlüsselwörter sind, die einem Selektor folgen. Die beiden Pseudo-Elemente, die du verwenden wirst, sind die `::before`- und `::after`-Pseudo-Elemente.

The `::before` selector creates a pseudo-element which is the first child of the selected element, while the `::after` selector creates a pseudo-element which is the last child of the selected element. Diese Pseudo-Elemente werden oft zur Erstellung von kosmetischem Inhalt verwendet, den du später in diesem Projekt sehen wirst.

For now, create a CSS selector to target all elements with `*`, and include the pseudo-elements with `::before` and `::after`. Setze die `box-sizing`-Eigenschaft auf `inherit`.

# --hints--

Du solltest einen `*, ::before, ::after`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Dein `*, ::before, ::after`-Selektor sollte die `box-sizing`-Eigenschaft auf `inherit` eingestellt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
