---
id: 612e6afc009b450a437940a1
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Beginne mit der grundlegenden HTML-Struktur. Füge eine `DOCTYPE`-Deklaration sowie `html`-, `head`-, `body`- und `title`-Elemente hinzu.

Stelle Englisch als Sprache für diese Seite ein. Setze den `title` auf `Piano`.

# --hints--

Dein Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut mit dem Wert `en` enthalten.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML-Codes stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Deine `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein Projekt sollte den Titel `Piano` tragen.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
