---
id: 612e77aba7ca691f598feb02
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge zwei `meta`-Tags hinzu, eins, um deine Seite für mobile Geräte zu optimieren und eins, um ein akzeptiertes `charset` für die Seite zu bestimmen.

# --hints--

Du solltest deiner Seite zwei `meta`-Elemente hinzufügen.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Du solltest zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ein `meta`-Element sollte einen `name` mit dem Wert `viewport` und `content` mit dem Wert `width=device-width, initial-scale=1.0` enthalten.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset` auf `UTF-8` eingestellt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
