---
id: 612ea4c4993aba52ab4aa32e
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Now it is time to use the pseudo-selectors you prepared for earlier. Füge einen neuen `.key.black--key::after`-Selektor hinzu, um die schwarzen Tasten zu erstellen. This will target the elements with the class `key black--key`, and select the pseudo-element after these elements in the HTML.

In the new selector, set the `background-color` to `#1d1e22`. Setze ebenfalls die `content`-Eigenschaft auf `""`. This will make the pseudo-elements empty.

Die `content`-Eigenschaft wird verwendet, um den Inhalt des Elements festzulegen oder zu überschreiben. Die Pseudo-Elemente, die von `::before`- und `::after`-Pseudo-Selektoren erstellt wurden, sind standardmäßig leer und die Elemente werden nicht auf der Seite gerendert. Wenn du die `content`-Eigenschaft auf einen leeren String `""` festlegst, wird sichergestellt, dass das Element auf der Seite gerendert wird, während es noch immer leer ist.

Wenn du experimentieren möchtest, versuche die `background-color`-Eigenschaft zu entfernen und verschiedene Werte für die `content`-Eigenschaft, wie z. B. `"♥"`, festzulegen. Denke daran, diese Änderungen rückgängig zu machen, wenn du damit fertig bist, damit die Tests erfolgreich durchgeführt werden.

# --hints--

Du solltest einen `.key.black--key::after`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Dein `.key.black--key::after`-Selektor sollte eine `background-color`-Eigenschaft auf `#1d1e22` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

Dein `.key.black--key::after`-Selektor sollte eine `content`-Eigenschaft auf `""` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
