---
id: 612e804c54d5e7308d7ebe56
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Füge innerhalb deines `.keys`-Elements sieben `div`-Elemente hinzu. Give them all the class `key`.

# --hints--

Du solltest sieben neue `div`-Elemente erstellen.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

Deine sieben neuen `div`-Elemente sollten sich innerhalb deines `.keys`-Elements befinden.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

Deine sieben neuen `div`-Elemente sollten die `class` auf `key` gesetzt haben.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
