---
id: 5d822fd413a79914d39e98c9
title: Part 1
challengeType: 0
dashedName: part-1
---

# --description--

Welcome to the CSS Variables skyline project! Start by adding the `!DOCTYPE html` declaration at the top of the document so the browser knows what type of document it's reading.

# --hints--

test-text

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

# --seed--

## --seed-contents--

```html
```

# --solutions--

```html
<!DOCTYPE html>
```
