---
id: 5d822fd413a79914d39e98ca
title: Part 2
challengeType: 0
dashedName: part-2
---

# --description--

Add opening and closing `html` tags below the doctype so you have a place to start putting some code.

# --hints--

test-text

```js
assert(code.match(/<!DOCTYPE\s+html\s*>\s*<html\s*>\s*<\/html\s*>/gi));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  
</html>
```
