---
id: 5d822fd413a79914d39e98cb
title: Part 3
challengeType: 0
dashedName: part-3
---

# --description--

Next, add opening and closing `head` and `body` tags within the `html` element.

# --hints--

You should have an opening `head` tag.

```js
```

You should have a closing `head` tag.

```js
```

You should have an opening `body` tag.

```js
```

You should have a closing `body` tag.

```js
```

The `head` and `body` elements should be siblings.

```js
```

The `head` element should be within the `html` element.

```js
```

The `body` element should be within the `html` element.

```js
assert(
  code.match(
    /<html\s*>\s*<head\s*>\s*<\/head\s*>\s*<body\s*>\s*<\/body\s*>\s*<\/html\s*>/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  
</html>
--fcc-editable-region--

```

