---
id: 660eebd83100d37862268781
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Deine `character` Variable hat aktuell noch keinen Wert. Du kannst der Variable einen Wert zuweisen indem du den <dfn>assignment</dfn>-Operator `=` benutzt. Hier ein Beispiel:

```js
let hello = "Hello";
```

Wenn man der Variable einen Wert zuweist, während sie gerade deklariert wird, nennt man das initialisieren (eng: <dfn>initialization</dfn>).

Initialisiere deine `character` variable, indem du ihr den Wert `"Hello"` zuweist, wenn sie deklariert wird.

# --hints--

Du solltest den assignment-Operator `=` benutzen.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Du solltest den Text `"Hello"` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Du solltest doppelte Anführungszeichen für den `"Hello"`-Text verwenden.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Du solltest deiner `character` Variable den Wert `"Hello"` zuweisen.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
