---
id: 660ef5c1904955978a986a5c
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Du solltest `let` verwenden, um eine `count` -Variable zu deklarieren. Weise ihr die <dfn>Zahl</dfn> `8` zu. Bei Zahlenwerten brauchst du keine Anführungszeichen. Hier ein Beispiel:

```js
let money = 100;
```

# --hints--

Du benötigst ein zweites `let` Schlüsselwort in deinem Programm.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Du solltest `let` verwenden, um eine `count`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Du solltest der `count`-Variablen den Wert `8` zuweisen.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
