---
id: 660ee6e3a242da6bd579de69
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Eins der wichtigsten Konzepte in der Programmierung sind Variablen. Eine <dfn>Variable</dfn> weist auf eine bestimmte Speicheradresse hin, die einen Wert speichert. Variablen bekommen einen Namen, der im gesamten Code verwendet werden kann, um auf den Wert der Variable zuzugreifen.

Eine Variable zu deklarieren bedeutet, ihr einen Namen zu geben. In JavaScript erstellt man eine Variable mit dem `let` Keyword. Hier ist ein Beispiel, wie man eine `hello` Variable deklariert:

```js
let hello;
```

Namen für Variablen haben bestimmte Regeln: Namen können Buchstaben, Zahlen, Dollarzeichen und Unterstriche enthalten, allerdings keine Leerzeichen und sie dürfen nicht mit einer Zahl beginnen.

Benutze das `let` Keyword um eine Variable mit dem Namen `character` zu erstellen.

_Hinweis_: Anweisungen in JavaScript werden normalerweise mit einem Semikolon beendet. ( ; )

# --hints--

Du solltest `let` in deinem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Du solltest `character` in deinem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Du solltest `let` verwenden, um eine Variable mit dem Namen `character` zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Your declaration should end with a semicolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
