---
id: 660ef31a5be625914a0102cd
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Wenn eine Variable mit dem `let`-Schlüsselwort zugewiesen wird, kann man später den Wert <dfn>zuordnen</dfn> (oder den Wert ändern). In diesem Beispiel wird der Wert von `programmer` von `"Naomi"` zu `"CamperChan"` geändert.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Beachte: Bei der Neuzuweisung wird **NICHT** mehr das Schlüsselwort `let` verwendet.

After your `console.log`, assign the value `"World"` to your `character` variable.

# --hints--

You should use `character` twice in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

You should use the assignment operator to reassign `character`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

You should use the string `"World"` in your code.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Your `character` variable should have the value `"World"`.

```js
assert.equal(character, "World");
```

Your reassignment should not use `let`.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
