---
id: 660ef19b95d3308e7dd31bb6
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

In der <dfn>Konsole</dfn> kann die JavaScript-Ausgabe ausgegeben und angezeigt werden. Mit `console.log()` kannst du Informationen an die Konsole senden. Beispielsweise wird dieser Code `"Naomi"` in der Konsole ausgeben:

```js
let developer = "Naomi";
console.log(developer);
```

Der obige Code greift auf die `developer`-Variable mit dem Namen im `consloe.log()` zu. Beachte, dass der Wert zwischen den Klammern der Wert ist, der ausgegeben wird.

Gibt den Wert der `character`-Variable in der Konsole aus. Klicke dann auf "Console", um die JavaScript Konsole zu sehen.

# --hints--

Du solltest im Code auf die `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Du solltest auf die `log` Eigenschaft der `Konsole` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Du solltest Klammern verwenden, um die `log()`-Methode zu verwenden.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Du solltest die `character`-Variable in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
