---
id: 660ee6e3a242da6bd579de69
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

JavaScript is the programming language that powers the web. Unlike the HTML and CSS you have learned previously, JavaScript is most commonly used to write logic instead of markup.

One of the most important concepts in programming is variables. A <dfn>variable</dfn> points to a specific memory address that stores a value. Variables are given a name which can be used throughout your code to access that value.

Declaring a variable means giving it a name. In JavaScript, this is often done with the `let` keyword. For example, here is how you would declare a `hello` variable:

```js
let hello;
```

Variable naming follows specific rules: names can include letters, numbers, dollar signs, and underscores, but cannot contain spaces and must not begin with a number.

Use the `let` keyword to declare a variable called `character`.

_Note_: It is common practice to end statements in JavaScript with a semicolon. `;`

# --hints--

You should use `let` in your code.

```js
assert.match(code, /let/);
```

You should use `character` in your code.

```js
assert.match(code, /character/);
```

You should use `let` to declare a `character` variable.

```js
assert.match(code, /let\s+character/);
```

Your declaration should end with a semi-colon.

```js
assert.match(code, /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
