---
id: 660ef55dd468079679ee0092
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

You are now ready to declare your next variable. Remove both `console.log` statements, and the `character` reassignment.

Also remove your `secondCharacter` variable, but leave the `character` initialization unchanged.

# --hints--

You should not have any `console.log()` statements in your code.

```js
assert.notMatch(code, /console/);
```

You should not reassign the value of `character`.

```js
assert.isAtMost(code.match(/character/).length, 1);
```

Your `character` variable should have the value `"Hello"`.

```js
assert.equal(character, "Hello");
```

You should not have a `secondCharacter` variable.

```js
assert.notMatch(code, /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
