---
id: 660ef857f2806aa626d29d17
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

You can also perform subtraction (`-`), multiplication (`*`), and division (`/`). Feel free to experiment with the operators and numbers in your `console.log`. When you are ready to move on, remove the `console.log`.

# --hints--

You should not have a `console.log()` in your code.

```js
assert.notMatch(code, /console/);
```

Your `character` variable should still have the value `"Hello"`.

```js
assert.equal(character, "Hello");
```

Your `count` variable should still have the value `8`.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
