---
id: 66c06d618d075c7f7f1b890a
title: Build a Fortune Teller App
challengeType: 14
dashedName: build-a-fortune-teller-app
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should declare five variables, each with a string value that describes a fortune.
1. You should select a random number between 1 and 5, inclusive, and assign it to the variable `randomNumber`.
You should use an `if/else if` statement to print a different fortune based on the value of `randomNumber`:

   - If `randomNumber` is 1, print the first fortune.
   - If `randomNumber` is 2, print the second fortune.
   - If `randomNumber` is 3, print the third fortune.
   - If `randomNumber` is 4, print the fourth fortune.
   - If `randomNumber` is 5, print the fifth fortune.

1. You should log the fortune to the console.
   
# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbours.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the colour red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
