---
id: 66c06d618d075c7f7f1b890a
title: Build a Fortune Teller App
challengeType: 14
dashedName: build-a-fortune-teller-app
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should initialize the following five variables, each with a string value that describes a fortune:
   
- `fortune1` should be set to `"Your cat will look very cuddly today."`.
- `fortune2` should be set to `"The weather will be nice tomorrow."`.
- `fortune3` should be set to `"Be cautious of your new neighbours."`.
- `fortune4` should be set to `"You will find a new hobby soon."`.
- `fortune5` should be set to `"It would be wise to avoid the colour red today."`.
  
1. You should select a random number between 1 and 5, inclusive, and assign it to the variable `randomNumber`.
You should use an `if/else if` statement to print a different fortune based on the value of `randomNumber`:

- If `randomNumber` is 1, print the first fortune.
- If `randomNumber` is 2, print the second fortune.
- If `randomNumber` is 3, print the third fortune.
- If `randomNumber` is 4, print the fourth fortune.
- If `randomNumber` is 5, print the fifth fortune.

1. You should log the fortune to the console.
   
# --hints--

You should initialize `fortune1` with the value `"Your cat will look very cuddly today."`.

```js
assert.equal(fortune1, "Your cat will look very cuddly today.");
```

You should initialize `fortune2` with the value `"The weather will be nice tomorrow."`.

```js
assert.equal(fortune2, "The weather will be nice tomorrow.");
```

You should initialize `fortune3` with the value `"Be cautious of your new neighbours."`.

```js
assert.equal(fortune3, "Be cautious of your new neighbours.");
```

You should initialize `fortune4` with the value `"You will find a new hobby soon."`.

```js
assert.equal(fortune4, "You will find a new hobby soon.");
```

You should initialize `fortune5` with the value `"It would be wise to avoid the colour red today."`.

```js
assert.equal(fortune5, "It would be wise to avoid the colour red today.");
```

You should generate a random number between 1 and 5, inclusive, and assign it to the variable `randomNumber`.

```js
assert.isAtLeast(randomNumber, 1);
assert.isAtMost(randomNumber, 5);

// assert.match(__helpers.removeJSComments(code), /let randomNumber\s*=\s*Math.floor\(Math.random\(\)\s*\*\s*5\)\s*\+\s*1;?/);
```

You should have a `selectedFortune` variable that is assigned a value based on the value of `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

You should use an `if` statement to check if `randomNumber` is equal to `1` and assign `selectedFortune` to `fortune1`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(randomNumber\s*===\s*1\)\s*{\s*selectedFortune\s*=\s*fortune1;\s*}/);
```

You should create another condition using `else if` to check if `randomNumber` is equal to `2` and assign `selectedFortune` to `fortune2`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(randomNumber\s*===\s*2\)\s*{\s*selectedFortune\s*=\s*fortune2;\s*}/);
```

You should create another condition using `else if` to check if `randomNumber` is equal to `3` and assign `selectedFortune` to `fortune3`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(randomNumber\s*===\s*3\)\s*{\s*selectedFortune\s*=\s*fortune3;\s*}/);
```

You should create another condition using `else if` to check if `randomNumber` is equal to `4` and assign `selectedFortune` to `fortune4`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(randomNumber\s*===\s*4\)\s*{\s*selectedFortune\s*=\s*fortune4;\s*}/);
```

You should create a final condition using `else` to assign `selectedFortune` to `fortune5`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*{\s*selectedFortune\s*=\s*fortune5;\s*}/);
```

You should output the `selectedFortune` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(selectedFortune\);?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbours.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the colour red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
