---
id: 60b69a66b6ddb80858c5157d
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Usa un selector de id para darle al elemento con el id `back-wall` un `background-color` de `#8B4513`.

# --hints--

Debe usar un selector `#back-wall`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall'));
```

Su selector `#back-wall` debe tener un `background-color` de `#8B4513`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.backgroundColor === 'rgb(139, 69, 19)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

--fcc-editable-region--

--fcc-editable-region--

```
