---
id: 60b69a66b6ddb80858c5157b
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Para comenzar a pintar, asigne a su elemento `body` un `background-color` de `rgb(184, 132, 46)`.

# --hints--

Debes usar el selector `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Su elemento `body` debe tener la propiedad `background-color` establecida en `rgb (184, 132, 46)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(184, 132, 46)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  </body>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

```
