---
id: 60b69a66b6ddb80858c51579
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de su elemento `head`, agregue una etiqueta `meta` con el atributo `charset` establecido en `utf-8`. Añade también un elemento `title` con el texto `Picasso Painting`.

# --hints--

Debe agregar exactamente un elemento `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Su elemento `meta` debe tener un atributo `charset`.

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

Su atributo `charset` debe establecerse en `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Debe agregar exactamente un elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Tu elemento `title` debe tener el texto `Picasso Painting`. Tenga en cuenta que la ortografía y las mayúsculas son importantes.

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
