---
id: 60b69a66b6ddb80858c51583
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Cree cuatro elementos `div` dentro del elemento `offwhite-character`. Asigne a esos elementos `div` los siguientes valores `id`, en orden: `white-hat`, `black-mask`, `gray-instrument`, `tan-table`.

# --hints--

Debe agregar cuatro elementos `div` dentro de su elemento `.offwhite-character`.

```js
assert(document.querySelectorAll('#offwhite-character div').length === 4);
```

El primer elemento nuevo `div` debe tener el `id` de `white-hat`.

```js
assert(document.querySelectorAll('#offwhite-character div')[0]?.getAttribute('id') === 'white-hat');
```

El segundo elemento nuevo `div` debe tener el `id` de `black-mask`.

```js
assert(document.querySelectorAll('#offwhite-character div')[1]?.getAttribute('id') === 'black-mask');
```

El tercer elemento nuevo `div` debe tener el `id` de `gray-instrument`.

```js
assert(document.querySelectorAll('#offwhite-character div')[2]?.getAttribute('id') === 'gray-instrument');
```

El cuarto elemento nuevo `div` debe tener el `id` de `tan-table`.

```js
assert(document.querySelectorAll('#offwhite-character div')[3]?.getAttribute('id') === 'tan-table');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
