---
id: 60b69a66b6ddb80858c5157c
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro de su etiqueta de cuerpo, agregue un elemento `div`. Dale un `id` de `back-wall`.

# --hints--

Debe agregar exactamente 1 elemento `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

Su elemento `div` debe tener el valor `id` de `back-wall`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'back-wall');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}
```
