---
id: 60b69a66b6ddb80858c5157a
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

FontAwesome es una biblioteca de iconos basados ​​en SVG, muchos de los cuales están disponibles para su uso de forma gratuita. Utilizará algunos de estos íconos en este proyecto, por lo que deberá vincular la hoja de estilo externa a su HTML.

Agregue un elemento `link` con un `rel` de `stylesheet` y un `href` de `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

# --hints--

Debes tener dos elementos `link`.

```js
assert(document.querySelectorAll('link').length === 2);
```

Su elemento `link` debe tener un `rel` de `stylesheet`.

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('rel') === 'stylesheet');
```

Su elemento `link` debe tener un `href` de `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
