---
id: 60b69a66b6ddb80858c51578
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Comience configurando su estructura HTML. Agregue una declaración`<!DOCTYPE>` y un elemento `html` con un atributo `lang` establecido en `en`. Dentro del elemento `html`, agregue un elemento `head` y un elemento `body`.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debe incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/html/gi));
```

Debe cerrar la declaración `DOCTYPE` con `>` después del tipo.

```js
assert(code.match(/html\s*>/gi));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Su elemento `html` debe tener una etiqueta de apertura con un atributo `lang` de `en`

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Debe tener una etiqueta de apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debe tener una etiqueta de cierre `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Debe tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debe tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

El elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

--fcc-editable-region--

```

```css

```
