---
id: 60b80da8676fb3227967a731
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Continúe y vincule su archivo CSS ahora, aunque aún no haya escrito ningún CSS.

Agregue un elemento `link` con un `rel` de `stylesheet` y un `href` de `styles.css`.

# --hints--


El código debe tener un elemento `link`.

```js
assert.match(code, /<link/)
```

Debe tener un elemento `link` de cierre automático.

```js
assert(document.querySelectorAll('link').length === 1);
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

El elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

El elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
