---
id: 60b69a66b6ddb80858c51579
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele chako cha `head`, ongeza tagi ya `meta` yenye sifa ya `charset` iliyowekwa kuwa `utf-8`. Pia, ongeza kipengele cha `title` chenye maandishi `Picasso Painting`.

# --hints--

Unapaswa kuwa na kipengele kimoja cha `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Kipengele chako cha `meta` kinafaa kuwa na sifa ya `charset`.

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

Sifa yako ya `charset` inapaswa kuwa na thamani ya `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Unapaswa kuongeza kipengele kimoja cha `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Kipengele chako cha `title` kinapaswa kuwa na maandishi `Picasso Painting`. Kumbuka kwamba tahajia na herufi ni muhimu.

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
