---
id: 60b69a66b6ddb80858c51582
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Ndani ya kipengele hicho cha `.characters`, unda `div` nyingine yenye `id` ya `offwhite-character`.

# --hints--

Unapaswa kuunda kipengele 1 pekee cha `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Kipengee chako kipya cha `div` kinapaswa kuwekwa ndani ya kipengele chako cha `.characters`.

```js
assert(document.querySelector('.characters div'));
```

Kipengee chako kipya cha `div` kinapaswa kuwa na `id` iliyowekwa kuwa `offwhite-character`.

```js
assert(document.querySelector('.characters div')?.getAttribute('id') === 'offwhite-character');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
