---
id: 60b69a66b6ddb80858c5157b
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Ili uanze uchoraji wako, kipe kipengele cha `body` `background-color` ya `rgb(184, 132, 46)`.

# --hints--

Unapaswa kutumia kichaguzi cha `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Kichaguzi chako cha `body` kinapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb (184, 132, 46)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(184, 132, 46)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  </body>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

```
