---
id: 60b69a66b6ddb80858c5157a
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

FontAwesome ni maktaba ya ikoni zinazoendeshwa na SVG, nyingi zinapatikana kwa matumizi bila malipo. Utakuwa ukitumia baadhi ya aikoni hizi katika mradi huu, kwa hivyo utahitaji kuunganisha laha ya nje ya mtindo kwenye HTML yako.

Ongeza kipengele cha `link` kilicho na `rel` ya `stylesheet` na `href` ya `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

# --hints--

Unapaswa kuwa na vipengele viwili vya `link`.

```js
assert(document.querySelectorAll('link').length === 2);
```

Kipengele chako cha `link` kinafaa kuwa na `rel` ya `stylesheet`.

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('rel') === 'stylesheet');
```

Kipengele chako kipya cha `link` kinafaa kuwa na `href` ya `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
