---
id: 60b69a66b6ddb80858c5157f
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Kwa kawaida, HTML hutolewa kwa namna ya juu-chini. Vipengele vilivyo juu ya msimbo vimewekwa juu ya ukurasa. Hata hivyo, mara nyingi unaweza kutaka kuhamisha vipengele kwenye nafasi tofauti. Unaweza kufanya hivi kwa sifa ya `position`.

Weka `position` ya kipengele cha `#back-wall` kuwa `absolute`. Nafasi ya `absolute` huchukua kipengele kutoka kwenye mtiririko huo wa hati unaotoka juu chini na hukuruhusu kukirekebisha kulingana na chombo chake.

Wakati kipengele kinapowekwa nafasi kwa mkono, unaweza kubadilisha mpangilio wake na `top`, `left`, `right` na `bottom`. Weka kipengele cha `#back-wall` kuwa na thamani ya `top` ya `0`, na thamani ya `left` ya `0`.


# --hints--

Kichaguzi chako cha `#back-wall` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.position === 'absolute');
```

Kichaguzi chako cha `#back-wall` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.top === '0px');
```

Kichaguzi chako cha `#back-wall` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.left === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
