---
id: 60b80da8676fb3227967a731
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Endelea na uunganishe faili yako ya CSS sasa, ingawa bado hujaandika CSS yoyote.

Ongeza kipengele cha `link` kilicho na `rel` ya `stylesheet` na `href` ya `styles.css`.

# --hints--


Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/)
```

Unapaswa kuwa na kipengele kimoja cha kujifunga cha `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `rel` yenye thamani `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
