---
id: 614385513d91ae5c251c2052
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi un elemento `title` con il testo `Magazine`, un elemento `link` per il foglio di stile dei caratteri `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, un `link` per il foglio di stile FontAwesome `https://use.fontawesome.com/releases/v5.8.2/css/all.css` e un `link` per il foglio di stile `./styles.css`.

Il foglio di stile dei caratteri caricherà tre caratteri separati: `Anton`, `Baskervville` e `Raleway`.

# --hints--

Il tuo codice dovrebbe avere tre elementi `link` auto-concludenti.

```js
assert(document.querySelectorAll('link').length === 3);
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert(document.querySelectorAll('head > link').length === 3);
```

I tre elementi `link` dovrebbero avere un attributo `rel` con il valore `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Uno degli elementi link dovrebbe avere un `href` impostato su `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Uno degli elementi link dovrebbe avere un `href` impostato su `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Uno degli elementi `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/)
```

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il progetto dovrebbe avere il titolo `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

Fai attenzione all'uso delle maiuscole/minuscole e all'ortografia nel titolo.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
