---
id: 61437d575fb98f57fa8f7f36
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Inizia con il testo HTML standard. Aggiungi la dichiarazione `DOCTYPE` e gli elementi `html` (specificando che la pagina è in inglese), `head` e `body`.

All'interno dell'elemento `head`, aggiungi un tag `<meta>` con il valore appropriato dell'attributo `charset` e un tag `<meta>` per la responsività mobile.

# --hints--

Il codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un carattere `>`.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dovresti avere un tag `<html>` di apertura con un attributo `lang` con il valore `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag di apertura `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag di chiusura `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag di apertura `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag di chiusura `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe trovarsi all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe trovarsi all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Un elemento `meta` dovrebbe avere un attributo `name` impostato su `viewport` e un attributo `content` impostato su `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

L'altro elemento `meta` dovrebbe avere l'attributo `charset` impostato su `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
