---
id: 614385513d91ae5c251c2052
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi un elemento `title` con il testo `Magazine`, un elemento `link` per il foglio di stile dei caratteri `https://fonts.googleapis.com/css?family=Anton|Baskervville|Raleway&display=swap`, un `link` per il foglio di stile FontAwesome `https://use.fontawesome.com/releases/v5.8.2/css/all.css` e un `link` per il foglio di stile `./styles.css`.

Il foglio di stile dei caratteri caricherà tre caratteri separati: `Anton`, `Baskervville` e `Raleway`.

# --hints--

Il codice dovrebbe avere tre elementi `link`.

```js
assert(code.match(/<link/g)?.length === 3);
```

Gli elementi `link` dovrebbero essere elementi autoconcludenti.

```js
assert(code.match(/<\/link>/i) === null);
```

Gli elementi `link` dovrebbero trovarsi all'interno dell'elemento `head`.

```js
const head = code.match(/<head>(.|\r|\n)*<\/head>/i)?.[0];
assert(head.match(/<link/g)?.length === 3)
```

I tre elementi `link` dovrebbero avere un attributo `rel` con il valore `stylesheet`.

```js
assert(code.match(/<link[\s\S]*?rel=('|"|`)stylesheet\1/gi)?.length === 3);
```

Uno dei tuoi elementi link dovrebbe avere un `href` impostato su `https://fonts.googleapis.com/css?family=Anton|Baskervville|Raleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton|Baskervville|Raleway&display=swap'));
```

Uno dei tuoi elementi link dovrebbe avere un `href` impostato su `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Uno dei tuoi elementi `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il progetto dovrebbe avere il titolo `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

Fai attenzione all'uso delle maiuscole/minuscole e all'ortografia nel titolo.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
