---
id: 6852ef95ccc27660275692cd
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Now you'll create your first class called `Email`. A *class* is like a blueprint for creating objects. It defines what properties (attributes) and behaviors (methods) objects of that type will have.

Recall that a class is defined as follows:

```py
class ClassName:
    # class body
```

In your email simulator, each email will be an object created from the `Email` class. This class will store information about who sent the email, who received it, the subject, and the message body.

Create a class named `Email` using the `class` keyword. Add `pass` keyword in the class body for now.

# --hints--

You should create a class named `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

You should add the `pass` statement inside the class body for now.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_pass()`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import datetime

--fcc-editable-region--
```
