---
id: 6852f06fdde1b4634a263f03
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Inside the `__init__` method, you need to assign the parameter values to the object's attributes. In Python, `self` refers to the current instance of the class, and you use it to access the object's attributes.

Inside your `__init__` method, remove the `pass` keyword and assign the parameters to `self.sender`, `self.receiver`, `self.subject`, and `self.body`. Then set `self.timestamp` to `datetime.datetime.now()` to automatically set a timestamp for when the email was created. Also, set `self.read` to `False` to initally mark the email as unread.

# --hints--

You should assign the `sender` parameter to `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

You should assign the `receiver` parameter to `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

You should assign the `subject` parameter to `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

You should assign the `body` parameter to `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})

```

You should set `self.timestamp` to `datetime.datetime.now()`.

```js
({
    test: () => 
    {
        assert(_Node(_code.find_class("Email").find_function("__init__").find_variable("self.timestamp").is_equivalent("self.timestamp = datetime.datetime.now()")))
    }

})

```

You should set `self.read` to `False`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")`))
    }

})

```

You should not use the `pass` statement anymore.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_pass()`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
