---
id: 6852f19cdc3a4a67ef06f590
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

The `strftime` method allows you to format datetime objects into strings. It stands for “string format time”, and it's available through the datetime module you imported earlier. You'll use it to format the timestamp of the email to show the date and time in a readable format.

`datetime.now()` returns the current date and time, and `strftime` lets you specify how you want that date and time to be displayed as a string.

- `%Y` is the year with century as a decimal number.
- `%m` is the month as a zero-padded decimal number (01 to 12).
- `%d` is the day of the month as a zero-padded decimal number (01 to 31).
- `%H` is the hour (00 to 23).
- `%M` is the minute (00 to 59).
- `%S` is the second (00 to 59).

Enter the following code to see how it works:

```py
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: Year-Month-Date
print(now.strftime("%H:%M:%S"))  # Output: Hours:Minutes:Seconds
```

# --hints--

You should have the statement `now = datetime.datetime.now()` in your code to get the current date and time.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("now = datetime.datetime.now()")`));
    }
})
```

You should print `now.strftime("%Y-%m-%d")` to the console to see the formatted date.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(now.strftime('%Y-%m-%d'))")`));
    }
})
```

You should print `now.strftime("%H:%M:%S")` to the console to see the formatted time.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(now.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--

--fcc-editable-region--
```
