---
id: 6852f11d412c29665985a915
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Now that you see that read status of the email has been updated correctly, remove the  call to `mark_as_read` and the `print` statement after it.

# --hints--

You should remove the existing `email_obj.mark_as_read()` and `print(email_obj.read)` from the code.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_stmt("email_obj.mark_as_read()")`));
        assert.isFalse(runPython(`_Node(_code).has_stmt("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
email_obj.mark_as_read()
print(email_obj.read)

--fcc-editable-region--
```
