---
id: 6852f0c3fcedb964ab959921
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Remove all the existing `print` calls. 

# --hints--

You should remove the existing `print` statements from the code.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.receiver)")`));
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.body)")`));
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.timestamp)")`));
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False


--fcc-editable-region--
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

print(email_obj.sender)
print(email_obj.receiver)
print(email_obj.subject)
print(email_obj.body)
print(email_obj.timestamp)
print(email_obj.read)

--fcc-editable-region--
```
