---
id: 6852f11d412c29665985a915
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

<!-- split step into two parts
code removal in one, adding the method in the other -->

Now that you see that read status of the email has been updated correctly, remove the  call to `mark_as_read` and the `print` statement after it.

Next, you'll create a method to display the full email content. This method will show all the email details in a nicely formatted way, similar to how you'd see an email in your email client.

Create a method called `display_full_email` that takes only `self` as a parameter. Inside this method, using the `mark_as_read` method, mark the email as read since the user is viewing it.

# --hints--

You should define a method named `display_full_email` in the `Email` class that takes `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`))
    }

})
```

Your `display_full_email` method should call `self.mark_as_read()`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`))
    }

})
```

You should remove the existing `email_obj.mark_as_read()` and `print(email_obj.read)` from the code.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")
        `))
    }

})
```

You should remove the existing `print(email_obj.read)` from the code.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).has_call("print(email_obj.read)")
        `))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
email_obj.mark_as_read()
print(email_obj.read)

--fcc-editable-region--
```
