---
id: 6852f0c3fcedb964ab959921
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Now you'll add a method to mark an email as read. Recall that a *method* is a function that belongs to a class and can be called on objects of that class.

When someone opens an email, you want to change its status from unread to read. This method will be simple - it just needs to set the `read` attribute to `True`.

First remove all the existing `print` calls. After that, add a method called `mark_as_read` to your `Email` class. This method should take only `self` as a parameter and set the `read` attribute to `True`.

# --hints--

You should remove the existing `print` statements from the code.

```js
({
    test: () => 
    {
        assertisFalse(runPython(`      
        _Node(_code).has_call("print(email_obj.sender)")
        _Node(_code).has_call("print(email_obj.receiver)")
        _Node(_code).has_call("print(email_obj.subject)")
        _Node(_code).has_call("print(email_obj.body)")
        _Node(_code).has_call("print(email_obj.timestamp)")
        _Node(_code).has_call("print(email_obj.read)")
`))
    }

})
```

You should define a method named `mark_as_read` in the `Email` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("mark_as_read")`))
    }

})
```

Your `mark_as_read` method should only take `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

The method should set `self.read` to `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
print(email_obj.sender)
print(email_obj.receiver)
print(email_obj.subject)
print(email_obj.body)
print(email_obj.timestamp)
print(email_obj.read)

--fcc-editable-region--
```
