---
id: 6852f19cdc3a4a67ef06f590
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

<!-- Note: would be changed to self.sender and self.receiver should be changed to self.sender.name and self.receiver.name after the User Class has been created. -->

Now you'll add the print statements to display the email information. You'll format the timestamp to show only the date and time in a readable format using the `strftime` method.

The `strftime` method allows you to format datetime objects into strings. The format `'%Y-%m-%d %H:%M'` will show the year, month, day, hour, and minute.

In your `display_full_email` method, after calling `mark_as_read()`, add print statements to display the email header, sender name, receiver name, subject, formatted timestamp, and body. Use f-strings to create a format like this:

```md
\n--- Email ---
From: [sender name]
To: [receiver name]
Subject: [subject]
Received: [date formatted as '%Y-%m-%d %H:%M']
Body: [email body]
------------\n
```

# --hints--

You should print the header like `'--- Email ---'` at the beginning with a newline character at the start.

```js
({ // escape string issue
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call('print("\\n--- Email ---")')`))
    }
})
```

You should print the sender's name using an f-string with `self.sender`. Add `'From: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender}')")`))
    }

})
```

You should print the receiver's name using an f-string with `self.receiver`. Add `'To: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver}')")`))
    }

})
```

You should print the subject using an f-string with `self.subject`. Add `'Subject: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Subject: {self.subject}')")`))
    }

})
```

You should print the formatted timestamp using `self.timestamp.strftime('%Y-%m-%d %H:%M')`. Add `'Received: '` before it.

```js
({// escape string issue
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call('print(f"Received: {self.timestamp.strftime(\'%Y-%m-%d %H:%M\')}")')`))
    }
})
```

You should print `'Body:'` followed by the email body using `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Body: {self.body}')")`))
    }

})
```

You should end the email display with a line of dashes as shown in the example.

```js
({// escape string issue
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call('print("------------\\n")')`))
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        
--fcc-editable-region--

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

```
