---
id: 6852f01116912861c7d1b0e6
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Your `Email` class will need to store information about the sender, receiver, subject, and body of the email. The `__init__` method will take these as parameters and assign them to the object.

Inside your `Email` class, remove the existing `pass` keyword and create the `__init__` method that takes `self`, `sender`, `receiver`, `subject`, and `body` as parameters. Add the `pass` keyword inside the method body for now.

# --hints--

You should define an `__init__` method inside the `Email` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Your `__init__` method should have parameters: `self`, `sender`, `receiver`, `subject`, and `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
