---
id: 6852ef95ccc27660275692cd
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

In your email simulator, each email will be an object created from the `Email` class. This class will store information about who sent the email, who received it, the subject, and the message body.

Create a class named `Email` using the `class` keyword. Add `pass` keyword in the class body for now.

# --hints--

You should create a class named `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import datetime

--fcc-editable-region--
```
