---
id: 6853f70470f45d9c280af659
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--

Delete the existing `my_inbox` and `email_obj` object definitons from the code. Also delete method calls to `my_inbox`

# --hints--

You should remove the existing `my_inbox` and `email_obj` object definitions from the code.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox = Inbox()")`))
        assert.isFalse(runPython(`_Node(_code).has_stmt("email_obj = Email(sender='jessica@example.com',receiver='ilenia@example.com', subject='Meeting Reminder', body='Hi Ilenia, just reminding you about our meeting at 3 PM today.')")`))
    }

})
```

You should remove the existing `my_inbox.read_email(0)` line from the code.


```js
// error in code
```

You should remove the existing `my_inbox.list_emails()` line from the code.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox.list_emails()")`))
    }

})
```

Yous should remove the existing `my_inbox.receive_email(email_obj)` line from the code.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox.receive_email(email_obj)")`))
    }

})
```


# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')


--fcc-editable-region--
my_inbox = Inbox()
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

my_inbox.receive_email(email_obj)
my_inbox.list_emails()
my_inbox.read_email(0)
   
--fcc-editable-region--

```
