---
id: 6853f553d90a2f9874ff9d0e
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

The methods in the previous steps work as expected. Now remove the `receive_emails`, `list_emails`, and `display_full_email` method calls from the code. Also, remove the `email_obj2` definition.

# --hints--

You should remove the `receive_emails`, `list_emails`, and `display_full_email` method calls from the code.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox.receive_email(email_obj)")`));
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox.receive_email(email_obj2)")`));
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox.list_emails()")`));
        assert.isFalse(runPython(`_Node(_code).has_stmt("my_inbox.emails[0].display_full_email()")`));

    }
})
```

You should remove the `email_obj2` definition from the code.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_stmt("email_obj2 = Email(sender='lisa@example.com',receiver='mona@example.com',subject='Lunch?',body='Are you free for lunch tomorrow?')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')

my_inbox = Inbox()

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
email_obj2 = Email(
    sender='lisa@example.com',
    receiver='mona@example.com',
    subject='Lunch?',
    body='Are you free for lunch tomorrow?'
)

my_inbox.receive_email(email_obj)
my_inbox.receive_email(email_obj2)

my_inbox.list_emails()

my_inbox.emails[0].display_full_email()
my_inbox.list_emails()

--fcc-editable-region--
```
