---
id: 6853f90fd03acc9fd75ff860
title: Step 28
challengeType: 20
dashedName: step-28
---

# --description--

The `User` class needs a method to send emails to other users. When sending an email, you'll create a new `Email` object and add it to the receiver's inbox.

Create a method called `send_email` in your `User` class. This method should take parameters for `receiver`, `subject`, and `body`. 

Inside the method, create a new `Email` object with the following parameter values and assign it to a variable named `email`:

- `sender`: `self` (the current user)
- `receiver`: `receiver` (the user receiving the email)
- `subject`: `subject` (the subject of the email)
- `body`: `body` (the body of the email)

# --hints--

You should define a method named `send_email` in the `User` class that takes `self`, `receiver`, `subject`, and `body` as parameters.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

You should create a new `Email` object with the correct parameters.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").find_body().is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            del self.emails[index]
            print('Email deleted.\n')
        except IndexError:
            print('Invalid email number.\n')

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()
    
--fcc-editable-region--
```
