---
id: 6853f4ded9f84896ff182dd8
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--

After the empty inbox check, print the message `'\nYour Emails:'`. After that, iterate over all emails in the inbox using a `for` loop. Use enumeration to number the emails starting from `1`. Use `i` (the index) and `email` (the email object) as the iteration variables.

Inside the loop, print an f-string made by the iteration index followed by a `.`, then a space and the string representation of the email object like this:

```md
[index]. [string representation of email object]
```

# --hints--

You should print the header `'\nYour Emails:'` outside the `if` statement.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

You should have a `for` loop that iterates over `enumerate(self.emails, 1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Your `for` loop should use `i` and `email` to iterate over `enumerate(self.emails, 1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Inside the loop, you should print each email summary with a numbered prefix followed by the formatted string from the `__str__` method. Don't forget to add a `.` after the number.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body().is_equivalent("print(f'{i}. {str(email)}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return

--fcc-editable-region--
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)
```
