---
id: 68540aa7dfd449ad074c285c
title: Step 47
challengeType: 20
dashedName: step-47
---

# --description--

Now you'll demonstrate the inbox functionality. Oliver will check his inbox, read the first email, delete the second email, and then check his inbox again to see the changes.

In your `main` function, after sending the emails, add code to:

- Have Oliver check his inbox  using the `check_inbox` method.
- Have Oliver read the first email (using index 0).
- Have Oliver delete the first email (index 0).
- Have Oliver check his inbox again.

With this, you have completed the email simulator!

# --hints--

You should call `oliver.check_inbox()` to show Oliver's emails.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().has_stmt("oliver.check_inbox()")`))
    }
})
```

You should call `oliver.read_email(0)` to read the first email.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().has_stmt("oliver.read_email(0)")`))
    }
})

```

You should call `oliver.delete_email(1)` to delete the second email.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().has_stmt("oliver.delete_email(1)")`))
    }
})

```

You should call `oliver.check_inbox()` again to show the updated inbox.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body().is_ordered("oliver.delete_email(1)", "oliver.check_inbox()")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            del self.emails[index]
            print('Email deleted.\n')
        except IndexError:
            print('Invalid email number.\n')
        
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

    def check_inbox(self):
        print(f"\n{self.name}'s Inbox:")
        self.inbox.list_emails()

    def read_email(self, index):
        self.inbox.read_email(index)

    def delete_email(self, index):
        self.inbox.delete_email(index)

--fcc-editable-region--
def main():
    naomi = User('Naomi')
    oliver = User('Oliver')        
    
    naomi.send_email(oliver, 'Hello', 'Hi Oliver, just saying hello!')
    oliver.send_email(naomi, 'Re: Hello', 'Hi Naomi, hope you are fine.')


--fcc-editable-region--
if __name__ == '__main__':
    main()
```

# --solutions--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"


class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')
    

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            del self.emails[index]
            print('Email deleted.\n')
        except IndexError:
            print('Invalid email number.\n')

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

    def check_inbox(self):
        print(f"\n{self.name}'s Inbox:")
        self.inbox.list_emails()
    
    def read_email(self, index):
        self.inbox.read_email(index)

    def delete_email(self, index):
        self.inbox.delete_email(index)

def main():
    # Create users
    naomi = User('Naomi')
    oliver = User('Oliver')

    # Send emails
    naomi.send_email(oliver, 'Hello', 'Hi Oliver, just saying hello!')
    oliver.send_email(naomi, 'Re: Hello', 'Hi Naomi, hope you are fine.')

    # Oliver checks inbox
    oliver.check_inbox()

    # Oliver reads first email
    oliver.read_email(0)

    # Oliver deletes second email
    oliver.delete_email(1)

    # Oliver checks inbox again
    oliver.check_inbox()

if __name__ == '__main__':
    main()
```
