---
id: 68540938f7f51fa80721d352
title: Step 31
challengeType: 20
dashedName: step-31
---

# --description--

Users also need methods to read and delete emails from their inbox.

Add two methods to your `User` class: `read_email` and `delete_email`. Both should take an `index` parameter and call the corresponding method on `self.inbox`.

# --hints--

You should define a method named `read_email` in the `User` class with the parameters `self` and `index`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("read_email").has_args("self, index")`))
    }
})

```

The `read_email` method should call `self.inbox.read_email(index)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("read_email").find_body().is_equivalent("self.inbox.read_email(index)")`))
    }
})
```

You should define a method named `delete_email` in the `User` class with the parameters `self` and `index`.

```js
({// try now
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("delete_email").has_args("self, index")`))
    }
})

```

The `delete_email` method should call `self.inbox.delete_email(index)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("delete_email").find_body().is_equivalent("self.inbox.delete_email(index)")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            del self.emails[index]
            print('Email deleted.\n')
        except IndexError:
            print('Invalid email number.\n')

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

    def check_inbox(self):
        print(f"\n{self.name}'s Inbox:")
        self.inbox.list_emails()

--fcc-editable-region--

--fcc-editable-region--

```
