---
id: 685941422c30cd3dc41b1d0a
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Test the `display_full_email` on the `email_obj` and as a result you will see the full email details. Also, the email will be marked as read, confirm that by printing the `read` attribute of the `email_obj`.

# --hints--

You should call the `display_full_email` method on the `email_obj`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_stmt("email_obj.display_full_email()")`))
    }

})
```

You should print the `read` attribute of the `email_obj` to confirm that the email has been marked as read.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True
    
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
```
