---
id: 6859523c793377706e19e9a0
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

Now, test the `read_email` method you created in the `Inbox` class. First, send an email using the `receive_email` method and pass in the existing `email_obj` object. Then, call the `list_emails` method to display the emails in the inbox.

Once the emails are listed, call the `read_email` method with the index `0` to read the first email in the inbox. This should display the full content of the email and mark it as read.

# --hints--

You should send an email using the `receive_email` method on your `my_inbox` object, passing in the `email_obj` object.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_stmt("my_inbox.receive_email(email_obj)")`))
    }

})
```

You should call the `list_emails` method on your `my_inbox` object to display the emails in the inbox.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_stmt("my_inbox.list_emails()")`))
    }

})
```

You should call the `read_email` method on your `my_inbox` object with index `0` to read the first email in the inbox.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).has_stmt("my_inbox.read_email(0)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

--fcc-editable-region--
my_inbox = Inbox()

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
```
