---
id: 68594df7779a164e6a61e8ae
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

Now, it is time to test the `receive_email` and `list_emails` methods you created in the `Inbox` class. 

Create an `Inbox` class object `my_inbox`.

Create another `Email` class object `email_obj2` with the following values:

- `sender`=`'lisa@example.com'`,
- `receiver`=`'mona@example.com'`,
- `subject`=`'Lunch?'`,
- `body`=`'Are you free for lunch tomorrow?'`

# --hints--

You should create an `Inbox` class object called `inbox`.

```js
({
    test: () => runPython(`
    assert isinstance(my_inbox, Inbox)
    `)
})
```

You should create another `Email` class object called `email_obj2` with the correct arguments.

```js
({ 
    test: () => runPython(`
    assert email_obj2.sender == 'lisa@example.com'
    assert email_obj2.receiver == 'mona@example.com'
    assert email_obj2.subject == 'Lunch?'
    assert email_obj2.body =='Are you free for lunch tomorrow?'
    `)
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
```
