---
id: 685a785d6f1b5f2dcd164533
title: Step 32
challengeType: 20
dashedName: step-32
---

# --description--

Previously, we have been passing strings as the `sender` and `receiver` parameters in the `Email` class. Now, we will pass the `User` objects instead.

For that, in the `display_full_email` method, change `self.sender` to `self.sender.name` and `self.receiver` to `self.receiver.name`. 

# --hints--

You should change the `display_full_email` method to use `self.sender.name` instead of `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_stmt("print(f'From: {self.sender.name}')")`))
    }

})
```

You should change the `display_full_email` method to use `self.receiver.name` instead of `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_stmt("print(f'To: {self.receiver.name}')")`))
    }

})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')
--fcc-editable-region--

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            del self.emails[index]
            print('Email deleted.\n')
        except IndexError:
            print('Invalid email number.\n')

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

    def check_inbox(self):
        print(f"\n{self.name}'s Inbox:")
        self.inbox.list_emails()

    def read_email(self, index):
        self.inbox.read_email(index)

    def delete_email(self, index):
        self.inbox.delete_email(index)

```
