---
id: 685c1eba966c6396850736ea
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Next, you'll create a method to display the full email content. This method will show all the email details in a nicely formatted way, similar to how you'd see an email in your email client.

Create a method called `display_full_email` that takes only `self` as a parameter. Inside this method, using the `mark_as_read` method, mark the email as read since the user is viewing it.

# --hints--

You should define a method named `display_full_email` in the `Email` class that takes `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`))
    }

})
```

Your `display_full_email` method should call `self.mark_as_read()`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

```
