---
id: 685cdfd1932d8ac6ad986813
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Now you'll create the `Inbox` class. An inbox is a container that holds multiple emails. Think of it like a folder that stores all the emails a user receives.

The `Inbox` class will store a list of emails and provide methods to add new emails, list all emails, and manage them.

Create a new class called `Inbox` with an `__init__` method that initializes an empty list called `emails`.

# --hints--

You should create a class named `Inbox`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class('Inbox')`))
    }

})
```

The `Inbox` class should have an `__init__` method that takes only `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`))
    }

})

```

You should initialize `self.emails` as an empty list.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`))
    }

})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

--fcc-editable-region--

--fcc-editable-region--

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)
```
