---
id: 685cdabd4cded3b340193f72
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

Next, you'll use `__str__` to create a brief summary of each email that shows whether it's read or unread, who sent it, the subject, and when it was received.

Add a `__str__` method to your `Email` class. Add a `pass` statement in the body for now.


# --hints--

You should define an `__str__` method in the `Email` class that takes `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_args("self")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    
--fcc-editable-region--
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

```
