---
id: 685cdabd4cded3b340193f72
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Python has special methods that start and end with double underscores (called 'dunder methods'). One useful dunder method is `__str__`, which defines how an object should be displayed when converted to a string or printed.

For your email simulator, you'll use `__str__` to create a brief summary of each email that shows whether it's read or unread, who sent it, the subject, and when it was received.

Add a `__str__` method to your `Email` class. Within the method, use conditional expression to assign `Read` to a variable `status` if the email is read and `Unread` if it is not. Then, return a formatted string that includes the read status, sender name, subject, and formatted timestamp as follows:

```py
[status] From: [sender name] | Subject: [subject] | Time: [Time in the format '%Y-%m-%d %H:%M']
```

Test it by printing the `email_obj`.

# --hints--

You should define an `__str__` method in the `Email` class that takes `self` as a parameter.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").has_args("self")`))
    }

})
```

You should use conditional expression to assign `status` based on `self.read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").find_body()[0].is_equivalent("status = 'Read' if self.read else 'Unread'")`))
    }
})
```

You should return the email summary string formatted correctly.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").find_body()[1].is_equivalent("return f'[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime(\\"%Y-%m-%d %H:%M\\")}'")`))
    }
})
```

You should print the `email_obj` to see the output of the `__str__` method.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_stmt('print(email_obj)')`))
    }
}
)
```


# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--


email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)


--fcc-editable-region--

```
