---
id: 685d53ec88a16a4643c2437b
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

Your inbox also needs a way to delete emails.

Python's `del` keyword can be used to delete items from a list by their index.

Add a method called `delete_email` to your `Inbox` class. Like `read_email`, it should:

- Take an `index` parameter, check for an empty inbox, and 
print the message `'Inbox is empty.\n'` if it is. Use `return` to exit the method.
- Use try-except to handle invalid indices. 
- In the try block, use the `del` keyword to delete the email at the mentioned index and print a confirmation message `'Email deleted.\n'`.
- In the except block, catch `IndexError` and print the message `'Invalid email number.\n'`.

# --hints--

You should define a method named `delete_email` in the `Inbox` class that takes `self` and `index` as parameters.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").has_args("self,index")`))
    }
})

```

You should check if the inbox is empty first using `if not self.emails:`. 

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})

```

If inbox is empty, you should print the message `'Inbox is empty.\n'` and use `return` to exit the method.

```js
({
    test: () => 
    { // test now
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_body().is_equivalent("print('Inbox is empty.\\\\n')\nreturn")`))
    }
})
```

Within the `try` block, you should use `del self.emails[index]` to delete the email.
You should also print a confirmation message `'Email deleted.\n'` when the email is deleted.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_trys()[0].find_body().is_equivalent("del self.emails[index]\nprint('Email deleted.\\\\n')")`))
    }
})
// test now
```

You should catch `IndexError` for invalid email numbers and print the message `'Invalid email number.\n'`.

```js
({ //test now
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_trys()[0].find_except("IndexError").find_body().is_equivalent("print('Invalid email number.\\\\n')")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')

--fcc-editable-region--
   
--fcc-editable-region--
```
