---
id: 685d2fd1f309b9191f257c55
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

<!-- see if email_obj and Inbox() will be used again, if not then remove the declaration here -->

The inbox needs a method to read a specific email. When a user wants to read an email, they'll specify which email number(starting from index 0) they want to see, and the method will display the full email content.

Add a method called `read_email` to your `Inbox` class that takes an `index` parameter. First, check if the inbox is empty and print the message `'Inbox is empty.\n'` if it is. Add a `return` statement after that to exit the method.

# --hints--

You should define a method named `read_email` in the `Inbox` class that takes `self` and `index` as parameters.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index"`))
    }

})
```

You should check if the inbox is empty using `if not self.emails:`

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }

})
```

Within the `if` statement, you should print the message `'Inbox is empty.\n'`.


```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }

})
```

You should add a `return` statement after printing the message to exit the method.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\n')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {str(email)}')


--fcc-editable-region--

--fcc-editable-region--

my_inbox = Inbox()

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

```
