---
id: 688798fccea91f522d1f245e
title: Step 29
challengeType: 20
dashedName: step-29
---

# --description--

Use the `receive_email` method to add both `email_obj` and `email_obj2` to `my_inbox`. Then, call the `list_emails` method to display the emails in the inbox.

Display the full content of the first email in the inbox using the `display_full_email` method. Finally, call the `list_emails` method again to see that the first email is now marked as read.

# --hints--

You should use `receive_email` to add both `email_obj` and `email_obj2` to `inbox`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("my_inbox = Inbox()", "my_inbox.receive_email(email_obj)", "my_inbox.receive_email(email_obj2)")`))
    }

})
```

You should call the `list_emails` method to display the emails in the inbox.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("my_inbox.receive_email(email_obj)", "my_inbox.receive_email(email_obj2)", "my_inbox.list_emails()")`))
    }

})
```

You should display the content of the first email using `display_full_email` after adding emails to the inbox and listing them.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).has_stmt("my_inbox.emails[0].display_full_email()")`))
    }

})
```

After displaying the full email you should list the emails again to see that the first email is now marked as read.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("my_inbox.emails[0].display_full_email()", "my_inbox.list_emails()")`))
    }

})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
my_inbox = Inbox()

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

email_obj2 = Email(
    sender='lisa@example.com',
    receiver='mona@example.com',
    subject='Lunch?',
    body='Are you free for lunch tomorrow?'
)

--fcc-editable-region--
```
