---
id: 68875521efc1a4e3f0c570e3
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Next, print the received timestamp using `strftime` to format it as `'%Y-%m-%d %H:%M'`. Use the following format:

```py
Received: [date formatted as '%Y-%m-%d %H:%M']
```

# --hints--

You should print the formatted timestamp using `self.timestamp.strftime('%Y-%m-%d %H:%M')`. Add `'Received: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call('print(f"Received: {self.timestamp.strftime(\\'%Y-%m-%d %H:%M\\')}")')`))
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')

--fcc-editable-region--

email_obj = Email(
    sender='quinn@example.com',
    receiver='ash@example.com',
    subject='Meeting Reminder',
    body='Hi Ash, just reminding you about our meeting at 3 PM today.'
)

```
