---
id: 68879d6de6398893be8c4b7b
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Conditional expressions allow you to assign a value based on a condition in a single line. Here is an example:

```py
x = 10
y = 'Even' if x % 2 == 0 else 'Odd' # y will be even
```

Within the method, use conditional expression to assign the string `Read` to a variable `status` if the email is read and `Unread` if it is not. 

# --hints--

You should use conditional expression to assign `status` based on `self.read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").find_body()[0].is_equivalent("status = 'Read' if self.read else 'Unread'")`))
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    def __str__(self):
        pass    
--fcc-editable-region--

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

```
