---
id: 688754d6d59578df1313cb3c
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Remove the existing datetime and print calls.

# --hints--


You should not have the statement `now = datetime.datetime.now()` in your code.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_stmt("now = datetime.datetime.now()")`));
    }
})
```

You should not have the statement `print(now.strftime("%Y-%m-%d"))` in your code.


```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_call("print(now.strftime('%Y-%m-%d'))")`));
    }
})
```


You should not have `print(now.strftime("%H:%M:%S"))` in your code.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).has_call("print(now.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        
email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)

--fcc-editable-region--
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))
print(now.strftime("%H:%M:%S"))
--fcc-editable-region--
```
