---
id: 68879d0a2aeb368e26559635
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

Return a formatted string that includes the read status, sender name, subject, and formatted timestamp as follows:

```py
[status] From: [sender name] | Subject: [subject] | Time: [Time in the format '%Y-%m-%d %H:%M']
```

Test it by printing the `email_obj`.

# --hints--


You should return the email summary string formatted correctly.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Email").find_function("__str__").find_body()[1].is_equivalent("return f'[{status}] From: {self.sender} | Subject: {self.subject} | Time: {self.timestamp.strftime(\\"%Y-%m-%d %H:%M\\")}'")`))
    }
})
```

You should print the `email_obj` to see the output of the `__str__` method.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_stmt('print(email_obj)')`))
    }
}
)
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')
--fcc-editable-region--
    def __str__(self):
        status = 'Read' if self.read else 'Unread'

email_obj = Email(
    sender='jessica@example.com',
    receiver='ilenia@example.com',
    subject='Meeting Reminder',
    body='Hi Ilenia, just reminding you about our meeting at 3 PM today.'
)
--fcc-editable-region--

```
