---
id: 6551ecf597038c36637f3759
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

Variables are an essential part of Python and any programming language. A variable is a name that references or points to an object. You can declare a variable by writing the variable name on the left side of the assignment operator `=` and specifying the value to assign to that variable on the right side of the assignment operator:

```py
variable_name = value
```

Create a variable called `number` and assign the value `5` to your new variable.

# --hints--

You should declare a variable called `number`. Pay attention to place the variable name at the beginning of the line.

```js
assert.match(code, /^number\s*=/m)
```

You should assign the value `5` to your `number` variable.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
