---
id: 6552008c0d9d9075cbec9772
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

The first kind of cipher you are going to build is called a *Caesar* cipher. Specifically, you will take each letter in your message, find its position in the alphabet, take the letter located after 3 positions, and replace the original letter with the new letter.

Start by finding the position of the first letter in the string. One way is to use the built-in `find()` function:

```py
a_string.find(char)
```

Above, `char` is the character you want to locate, and `a_string` is the string you want to parse.

At the end of your code, call `find()` on your `alphabet` string and pass `text[0]` to the function.

# --hints--

You should call the `find()` method.

```js
({ test: () => assert.match(code, /find\s*\(.*\)/) })

```

You should call the `find()` method on the `alphabet` variable.

```js
({ test: () => assert.match(code, /alphabet\.find\s*\(.*\)/) })

```

You should call the `find()` method on the `alphabet` variable and pass `text[0]` to the function.

```js
({ test: () => assert.match(code, /alphabet\.find\s*\(\s*text\[\s*\s*0\s*\]\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
