---
id: 6553995f412dd8122ed38e4a
title: الخطوة 40
challengeType: 20
dashedName: step-40
---

# --description--

A conditional statement can also have an `else` clause. This clause can be added to the end of an `if` statement to execute alternative code if the condition is false:

```py
if condition:
    <code>
else:
    <code>
```

As you can see in your output, when the loop iterations reach the space, a space is added to the encrypted string. Then the code outside the `if` block executes and a `c` is added to the encrypted string.

To fix it, add an `else` clause after `encrypted_text += char` and indent all the subsequent lines of code.

# --hints--

You should create an `else` clause.

```js
({ test: () => assert.match(code, /else\s*:/) })
```

You should indent the lines of code after your `else` clause.

```js
({ test: () => assert.match(code, /(\s*)else\s*:\s*?^\1\1\w.*?\n\s*?^\1\1\w.*?\n\s*?^\1\1\w.*?\n\s*?^\1\1\w.*?\n*/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
