---
id: 6568cc9301a0c41058ed95c5
title: الخطوة 34
challengeType: 20
dashedName: step-34
---

# --description--

Now, replace `new_char` with `encrypted_text`. Also, modify the `print()` call to reflect this change.

# --hints--

You should replace `new_char` with `encrypted_text`.

```js
({ test: () => assert.match(code, /encrypted_text\s*=\s*alphabet\s*\[\s*new_index\s*\]/) })
```

You should turn your `print()` call into `print('char:', char, 'encrypted text:', encrypted_text)`.

```js
({
  test: () => {
    assert.isFalse(/print\s*\(\s*new_char\s*\)/.test(code));
    assert.match(code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)    
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
