---
id: 66c8ba975ee7230e29f6c4ad
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

To check which loan a user is qualified for based on the `annualIncome` and `creditScore` they entered, you have to use `if/else if` statement, AKA conditonals.

Starting with the duplex loan, create an `if` statement that checks if `annualIncome` is greater than or equal to `minIncomeForDuplex` AND if `creditScore` is greater than `minCreditScoreForDuplex`.

Inside the `if` statement, reassign the `loanMessage` to `"You're qualified for a duplex loan, car loan, and condo loan"`, then log the `loanMessage` variable to the console.

# --hints--

You should create an `if` statement.

```js
assert.match(code, /if\s*\(\s*/)
```

Your `if` statement should use the `&&` operator to check if `annualIncome` is greater than or equal to `minIncomeForDuplex`, and if `creditScore` is greater than or equal to `minCreditScoreForDuplex`.

```js
assert.match(code, /if\s*\(\s*annualIncome\s*>=\s*minIncomeForDuplex\s*&&\s*creditScore\s*>=\s*minCreditScoreForDuplex\s*\)/)
```

You should reassign the `loanMessage` varible to `"You're qualified for a duplex loan, car loan, and condo loan"`.

```js
assert.match(code, /if\s*\(\s*annualIncome\s*>=\s*minIncomeForDuplex\s*&&\s*creditScore\s*>=\s*minCreditScoreForDuplex\s*\)\s\{\s*loanMessage\s*=\s*("|'|`)You're\s*qualified\s*for\s*a\s*duplex\s*loan,\s*car\s*loan,\s*and\s*condo\s*loan("|'|`);?\s*/)
```

You should log the `loanMessage` to the console.

```js
assert.match(code, /if\s*\(\s*annualIncome\s*>=\s*minIncomeForDuplex\s*&&\s*creditScore\s*>=\s*minCreditScoreForDuplex\s*\)\s\{\s*loanMessage\s*=\s*("|'|`)You're\s*qualified\s*for\s*a\s*duplex\s*loan,\s*car\s*loan,\s*and\s*condo\s*loan("|'|`);?\s*console\.log\(loanMessage\);?\s*}/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';
--fcc-editable-region--

--fcc-editable-region--
```
