---
id: 66c8ba975ee7230e29f6c4b6
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

You don't need the console log statements inside the `if` and `else if` blocks anymore. So,  remove all of them.

# --hints--

You should not modify the condition and the existing code in the `if` block.

```js
assert.match(code, /if\s*\(\s*annualIncome\s*>=\s*minIncomeForDuplex\s*&&\s*creditScore\s*>=\s*minCreditScoreForDuplex\s*\)\s*\{\s*loanMessage\s*=\s*"-\s*Duplex\s*loan\\n-\s*Car\s*loan\\n-\s*Condo\s*loan\\n"\;/)
```

You should not modify the existing condition and code in the first `else if` block.

```js
assert.match(code, /else\s*if\s*\(\s*annualIncome\s*>=\s*minIncomeForCondo\s*&&\s*creditScore\s*>=\s*minCreditScoreForCondo\s*\)\s*\{\s*loanMessage\s*=\s*"-\s*Car\s*loan\\n-\s*Condo\s*loan\\n";?/)
```

You should not modify the condition and the code inside the second `else if` bkock.

```js
assert.match(code, /else\s*if\s*\(\s*annualIncome\s*>=\s*minIncomeForCar\s*&&\s*creditScore\s*>=\s*minCreditScoreForCar\s*\)\s*\{\s*loanMessage\s*=\s*"-\s*Car\s*loan\\n";?/)
```

You should not modify the else block and its code.

```js
assert.match(code, /else\s*\{\s*loanMessage\s*=\s*"";?/)
```

You should remove all the console log statements.

```js
assert.notMatch(code, /console\.log\(loanMessage\);?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

--fcc-editable-region--
if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "- Duplex loan\n- Car loan\n- Condo loan\n";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
  loanMessage = "- Car loan\n- Condo loan\n";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCar &&
  creditScore >= minCreditScoreForCar
) {
  loanMessage = "- Car loan\n";
  console.log(loanMessage);
} else {
  loanMessage = "";
}
--fcc-editable-region--
```
