---
id: 66c8ba975ee7230e29f6c4af
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Create an `else if` block that checks if `annualIncome` is greater than or equal to `minIncomeForCondo`, AND if `creditScore` is greater than or equal to `minCreditScoreForCondo`.

# --hints--

You should create an `else if` block.

```js
assert.match(code, /\}\s*else\s*if\s*\(\s*/)
```

Your `else if` block should use the `&&` operator to check if `annualIncome` is greater than or equal to `minIncomeForCondo`, and if `creditScore` is greater than or equal to `minCreditScoreForCondo`.

```js
assert.match(code, /\}\s*else\s*if\s*\(\s*annualIncome\s*>=\s*minIncomeForCondo\s*&&\s*creditScore\s*>=\s*minCreditScoreForCondo\s*\)\s*\{\s*/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "You're qualified for a duplex loan, car loan, and condo loan";
  console.log(loanMessage);
--fcc-editable-region--
} 

--fcc-editable-region--
```
