---
id: 66c8ba975ee7230e29f6c4a9
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

The next thing to do is to create a variable for the minimum amount of `annualIncome` and `creditScore` required to be qualified for any of duplex, car, and condo loans.

Create a `minIncomeForDuplex` variable set to `60000`, and a `minCreditScoreForDuplex` variable set to `700`

# --hints--

You should create a `minIncomeForDuplex` variable set to `60000`.

```js
assert.match(code, /(const|let|var)\s*minIncomeForDuplex\s*=\s*6000;?/)
```

You should create a `minCreditScoreForDuplex` set to 700.

```js
assert.match(code, /(const|let|var)\s*minCreditScoreForDuplex\s*=\s*700;?/)
```

# --seed--

## --seed-contents--

```js
let annualIncome = 55000;
let creditScore = 650;
--fcc-editable-region--

--fcc-editable-region--
```
