---
id: 66c8ba975ee7230e29f6c4ac
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

You need to log a message that display loan eligibility to the console.

For that, use `let` to create a `loanMessage` variable and set it to an empty string.

You'll reassign this `loanMessage` a couple of time soon to correspond with the loan the applicant is eligible for.

# --hints--

You should use `let` to create a `loanMessage` variable.

```js
assert.match(code, /let\s*loanMessage\s*/)
```

Your `loanMessage` variable should be set to an empty string.

```js
assert.match(code, /let\s*loanMessage\s*=\s*('|"|`)('|"|`)/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;
--fcc-editable-region--

--fcc-editable-region--
```
