---
id: 66c8ba41a77db20a93f9d7a1
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you'll review JavaScript conditionals by building a loan qualification checker app.

The app will check whether the user is eligible for a duplex, car, and condo loan based on their annual income and credit score.

Use `let` to create an `annualIncome` and `creditScore` variables. Set them to `55000` and `650` respectively.

You will reassign these two variables later to test eligibility for loans based on different annual income and credit score values. That's why you're using `let` to declare both variables.

# --hints--

You should use `let` to create an `annualIncome` variable.

```js
assert.match(code, /let\s*annualIncome\s*/)
```

Your `annualIncome` variable should be set to `55000`.

```js
assert.match(code, /let\s*annualIncome\s*=\s*55000;?/)
```

You should use `let` to create a `creditScore` variable.

```js
assert.match(code, /let\s*creditScore\s*/)
```

Your `creditScore` variable should be set to `650`.

```js
assert.match(code, /let\s*creditScore\s*=\s*650;?/)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
