---
id: 66c8ba975ee7230e29f6c4b8
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

Create a `finalMessage` variable and set it to \` `Your credit score is ${creditScore} and your annual income is ${annualIncome}. So, ${qualificationMessage}` \`.

Feel free to use the plus operator for the concatenation if you want to.

After that, log the `finalMessage` variable to the console to see what it looks like.

# --hints--

You should create a `finalMessage` variable and set it to \` `Your credit score is ${creditScore} and your annual income is ${annualIncome}. So, ${qualificationMessage}` \`. You're free to use the plus operator to concatenate.

```js
assert.match(code, /(const|let|var)\s+finalMessage\s*=\s*`Your\s+credit\s+score\s+is\s+\$\{creditScore\}\s+and\s+your\s+annual\s+income\s+is\s+\$\{annualIncome\}\.\s+So\,\s+\$\{qualificationMessage\}`;?|(const|let|var)\s+finalMessage\s*=\s*("|'|`)Your\s+credit\s+score\s+is\s+("|'|`)\s+\+\s+creditScore\s+\+\s+("|'|`)\s+and\s+your\s+annual\s+income\s+is\s+("|'|`)\s+\+\s+annualIncome\s+\+\s+("|'|`)\.\s+("|'|`)\s+\+\s+("|'|`)So\,\s+("|'|`)\s+\+\s+qualificationMessage;?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "- Duplex loan\n- Car loan\n- Condo loan\n";
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
  loanMessage = "- Car loan\n- Condo loan\n"; 
} else if (
  annualIncome >= minIncomeForCar &&
  creditScore >= minCreditScoreForCar
) {
  loanMessage = "- Car loan\n";
} else {
  loanMessage = "";
}

let qualificationMessage =
  'You qualify for the following loan(s):\n \n' + loanMessage;

--fcc-editable-region--

--fcc-editable-region--
```
