---
id: 66c8ba975ee7230e29f6c4aa
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Create a `minIcomeForCar` and `minCreditScoreForCar` variables. Set them to `30000` and `650` respectively.

# --hints--

You should create a `minIcomeForCar` variable set to `30000`.

```js
assert.match(code, /(const|let|var)\s*minIcomeForCar\s*=\s*30000;?/)
```

You should create a `minCreditScoreForCar` variable set to `650`.

```js
assert.match(code, /(const|let|var)\s*minCreditScoreForCar\s*=\s*650;?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;
--fcc-editable-region--

--fcc-editable-region--
```
