---
id: 66c8ba975ee7230e29f6c4ab
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Lastly, create a `minIncomeForCondo` and `minCreditScoreForCondo` variables. Set them both to `45000` and `680` respectively.

# --hints--

You should create a `minIncomeForCondo` variable set to `45000`.

```js
assert.match(code, /(const|let|var)\s*minIncomeForCondo\s*=\s*45000;?/)
```

You should create a `minCreditScoreForCondo` variable set to `680`.

```js
assert.match(code, /(const|let|var)\s*minCreditScoreForCondo\s*=\s*680;?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;
--fcc-editable-region--

--fcc-editable-region--
```
