---
id: 66c8ba975ee7230e29f6c4b5
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Change the `loanMessage` inside the first `else if` block to `"- Car loan\n- Condo loan\n"`, and that of the second `else if` block to `"- Car loan\n"`.

# --hints--

You should not modify the existing code inside the first `else if block`.

```js
assert.match(code, /\}\s*else\s*if\s*\(\s*annualIncome\s*>=\s*minIncomeForCondo\s*&&\s*creditScore\s*>=\s*minCreditScoreForCondo\s*\)\s*\{/)
```

You should change the loanMessage inside the first `else if` block to `"- Car loan\n- Condo loan\n"`.

```js
assert.match(code, /("|'|`)-\s*Car\s*loan\\n-\s*Condo\s*loan\\n("|'|`);?/)
```

You should not modify the existing code inside the second `else if` block.

```js
assert.match(code, /else\s*if\s*\(\s*annualIncome\s*>=\s*minIncomeForCar\s*&&\s*creditScore\s*>=\s*minCreditScoreForCar/)
```

You should change the `loanMessage` inside the second `else if` block to `"- Car loan\n"`.

```js
assert.match(code, /("|'|`)-\s*Car\s*loan\\n("|'|`);?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "- Duplex loan\n- Car loan\n- Condo loan\n";
  console.log(loanMessage)
--fcc-editable-region--
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
  loanMessage = "You're qualified for a car loan and condo loan";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCar &&
  creditScore >= minCreditScoreForCar
) {
  loanMessage = "You're qualified for a car loan";
  console.log(loanMessage);
--fcc-editable-region--
} else {
  loanMessage = "";
}
```
