---
id: 66c8ba975ee7230e29f6c4b1
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Create another `else if` block that checks if `annualIncome` is greater than `minIncomeForCar` AND if `creditScore` is greater than `minCreditScoreForCar`.

# --hints--

You should create another `else if` block.

```js
const splitter = code.split("console.log(loanMessage);")
assert.match(splitter[2], /\}\s*else\s*if\s*\(\s*/)
```

The `else if` block should use the `&&` operator to check if `annualIncome` is greater than or equal to `minIncomeForCar`, and if `creditScore` is greater than or equal to `minCreditScoreForCar`.

```js
const splitter = code.split("console.log(loanMessage);")
assert.match(splitter[2], /\}\s*else\s*if\s*\(\s*annualIncome\s*>=\s*minIncomeForCar\s*&&\s*creditScore\s*>=\s*minCreditScoreForCar\s*\)\s*\{\s*\}/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "You're qualified for a duplex loan, car loan, and condo loan";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
  loanMessage = "You're qualified for a car loan and condo loan";
  console.log(loanMessage);
--fcc-editable-region--
} 

--fcc-editable-region--
```
