---
id: 66c8ba975ee7230e29f6c4b9
title: Step 17
challengeType: 1
dashedName: step-17
---

# --description--

Now it is time to test out your loan qualification checker.

Update the values for the `annualIncome` and `creditScore` variables to see the different loan messages. You can update the values to whatever positive integers you like.

# --hints--

You should update the `annualIncome` to any positive integer.

```js
assert.isTrue(Number.isInteger(annualIncome) && annualIncome > 0);
```

You should update the `creditScore` to any positive integer.

```js
assert.isTrue(Number.isInteger(creditScore) && creditScore > 0);
```

# --seed--

## --seed-contents--


```js
--fcc-editable-region--
const annualIncome = 55000;
const creditScore = 650;
--fcc-editable-region--

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "- Duplex loan\n- Car loan\n- Condo loan\n";
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
  loanMessage = "- Car loan\n- Condo loan\n"; 
} else if (
  annualIncome >= minIncomeForCar &&
  creditScore >= minCreditScoreForCar
) {
  loanMessage = "- Car loan\n";
} else {
  loanMessage = "";
}

let qualificationMessage =
  'You qualify for the following loan(s):\n \n' + loanMessage;

const finalMessage = `Your credit score is ${creditScore} and your annual income is ${annualIncome}. So, ${qualificationMessage}`;
console.log(finalMessage);
```
