---
id: 66c8ba975ee7230e29f6c4b0
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

If the last condition is `true`, then the applicant is qualified for a car and condo loans.

That means you should reassign the `loanMessage` to `"You're qualified for a car loan and condo loan"` and log the `loanMessage` variable to the console.

# --hints--

You should reassign the `loanMessage` variable to `"You're qualified for car loan and condo loan"`.

```js
assert.match(code, /\s*loanMessage\s*=\s*("|'|`)You're\s*qualified\s*for\s*a\s*car\s*loan\s*and\s*condo\s*loan("|'|`);?/)
```

You should log the `loanMessage` variable to the console.

```js
assert.match(code, /\s*loanMessage\s*=\s*("|'|`)You're\s*qualified\s*for\s*a\s*car\s*loan\s*and\s*condo\s*loan("|'|`);?\s*console\.log\(loanMessage\);?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "You're qualified for a duplex loan, car loan, and condo loan";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
--fcc-editable-region--


--fcc-editable-region--
}
```
