---
id: 66c8ba975ee7230e29f6c4b2
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

If the last condition is true, then the applicant is qualified for a car loan only.

Reassign the `loanMessage` variable to `"You're qualified for a car loan"` and log it to the console.

# --hints--

You should reassign the `loanMessage` variable to `"You're qualified for a car loan"`.

```js
assert.match(code, /loanMessage\s*=\s*("|'|`)You're\s*qualified\s*for\s*a\s*car\s*loan("|'|`);?/)
```

You should log the `loanMessage` variable to the console.

```js
assert.match(code, /loanMessage\s*=\s*("|'|`)You're\s*qualified\s*for\s*a\s*car\s*loan("|'|`);?\s*console\.log\(loanMessage\);?/)
```

# --seed--

## --seed-contents--

```js
const annualIncome = 55000;
const creditScore = 650;

const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

let loanMessage = '';

if (
  annualIncome >= minIncomeForDuplex &&
  creditScore >= minCreditScoreForDuplex
) {
  loanMessage = "You're qualified for a duplex loan, car loan, and condo loan";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCondo &&
  creditScore >= minCreditScoreForCondo
) {
  loanMessage = "You're qualified for a car loan and condo loan";
  console.log(loanMessage);
} else if (
  annualIncome >= minIncomeForCar &&
  creditScore >= minCreditScoreForCar
) {
--fcc-editable-region--


--fcc-editable-region--
}
```
