---
id: 69de4308e652aaeb5ccdb03e
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will practice working with Python's string methods.

Start by creating a variable named `email` and assign it the string `alice.johnson@company.com`.

# --hints--

You should create a variable named `email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('email')`)) })
```

You should assign the string `alice.johnson@company.com` to `email`. Remember to enclose the text between either single or double quotes.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
