---
id: 69de4308e652aaeb5ccdb03e
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will practice working with Python's string methods.

Start by creating a variable named `email` and assign it an `empty string`.

# --hints--

You should define a variable called `email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("email")`)) })
```

You should assign a empty string value to `email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("email")
not node.is_empty() and (
  node.is_equivalent("email = ''")
)
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
