---
id: 69e20a0451c7a3ef189514d3
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now that you know the position of `@` in the email, you can extract the username, which is placed before the `@` sign.

In the previous workshop, you extracted a portion of a string using slicing. As a reminder, the syntax is:

```python
text = 'hello world'
print(text[:5])  # hello
```

Use slicing on `email` to extract everything before `at_position`, and store the result in a variable named `username`.

Then print `username` so that the output reads `Username: alice.johnson`.

# --hints--

You should create a variable named `username`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('username')`)) })
```

You should use slicing with `at_position` on `email` to assign `username` a value.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('username')
not node.is_empty() and node.is_equivalent("username = email[:at_position]")
`)) })
```

You should print `username` so that the output reads `Username: alice.johnson`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Username: {username}')") or
_Node(_code).has_call("print('Username: ' + username)")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
