---
id: 69e20876baed33de636b989a
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Sometimes, you may want to find the position of a specific character inside a string. Python provides the built-in `.find()` method for this purpose. It returns an integer representing the index (position) of the first occurrence of the given substring. For example:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Now, call the `.find()` method on `email` and pass it the string `@` to find the position of the at symbol in the `email` string. Assign the result to a variable named `at_position`.

Then, print the result of concatenating the string `Position of @: ` with the value of `at_position`. Remember to use `str()` to convert `at_position` before concatenating it.

# --hints--

You should create a variable named `at_position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('at_position')`)) })
```

You should assign `at_position` using `email.find('@')`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('at_position')
not node.is_empty() and node.is_equivalent("at_position = email.find('@')")
`)) })
```

You should print the `email` with the label `'Email: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Email: ' + email)")`)) })
```

You should print the `position` with the label `'Position of @: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Position of @: ' + str(at_position))")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

--fcc-editable-region--

--fcc-editable-region--
```
