---
id: 69e6284f3b7659c5a05c2a79
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

You can check whether a substring exists inside a string using the `in` keyword. It returns `True` if the substring exists and `False` otherwise. Here's an example:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Create a variable named `has_at` that checks if the character `@` is in `email`.

Then print `has_at` so that the output reads: `Has @ symbol: True`.

# --hints--

You should create a variable named `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

Your `has_at` variable should have the value `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

You should print `has_at` so that the output reads `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
