---
id: 69e62ce96df237884da25282
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

The `.lower()` method converts all characters in a string to lowercase. This is useful for normalizing email addresses:


Create a variable `messy_email` and assign it the string `  SARAH.DAVIS@COMPANY.COM  `. Pay attention to the two leading spaces and the two trailing spaces.

Also, call `.strip().lower()` on `messy_email` and assign the result to a variable `clean_email`.

Then print both like this:

```python
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")
```

# --hints--

You should create a variable named `messy_email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('messy_email')`)) })
```

You should assign <code>'&nbsp;&nbsp;SARAH.DAVIS@COMPANY.COM&nbsp;&nbsp;'</code> to `messy_email`.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`)) })
```

You should create a variable named `clean_email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('clean_email')`)) })
```

You should call `.strip().lower()` on `messy_email` and assign the result to a `clean_email` variable.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('clean_email')
not node.is_empty() and node.is_equivalent("clean_email = messy_email.strip().lower()")
`)) })
```

You should have `print("Original email: '" + messy_email + "'")` in your code.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(\\"Original email: '\\" + messy_email + \\"'\\")")`)) })
```

You should have `print("Cleaned email: '" + clean_email + "'")` in your code.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(\\"Cleaned email: '\\" + clean_email + \\"'\\")")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

--fcc-editable-region--

--fcc-editable-region--
```
