---
id: 69e6284f3b7659c5a05c2a79
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

You can check whether a substring exists inside a string using the `in` keyword. It returns `True` if the substring exists and `False` otherwise. Here's an example:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Create a variable named `has_at` that checks if the character `@` is in `email`.

Then print `has_at` so that the output reads: `Has @ symbol: True`

# --hints--

You should create a variable named `has_at`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('has_at')`)) })
```

Your `has_at` variable should have the value `'@' in email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('has_at')
not node.is_empty() and (
  node.is_equivalent("has_at = '@' in email")
)
`)) })
```

You should print `has_at` so that the output reads `Has @ symbol: True`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Has @ symbol: {has_at}')") or
_Node(_code).has_call("print('Has @ symbol: ' + str(has_at))")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
