---
id: 69e6299710319e7cb8ebdc37
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Real-world data is often messy with extra spaces and inconsistent casing.

Create a variable `messy_name` and assign it `'  sARaH dAVis  '`. Then print it like this:

```python
print("Original name: '" + messy_name + "'")
```

You will clean this up in the coming steps.

# --hints--

You should create a variable named `messy_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('messy_name')`)) })
```

You should assign <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> to `messy_name`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('messy_name')
not node.is_empty() and node.is_equivalent("messy_name = '  sARaH dAVis  '")
`)) })
```


You should print the original name with the prefix `'Original name: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(\\"Original name: '\\" + messy_name + \\"'\\")")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

--fcc-editable-region--

--fcc-editable-region--
```
