---
id: 69e6299710319e7cb8ebdc37
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Real-world data is often messy with extra spaces and inconsistent casing.

Create a variable `messy_name` and assign it the string `  sARaH dAVis  `. Pay attention to the two leading spaces and the two trailing spaces. Then print it so that the output reads:

```md
Original name: '  sARaH dAVis  '
```

You will clean this up in the coming steps.

# --hints--

You should create a variable named `messy_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('messy_name')`)) })
```

You should assign <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> to `messy_name`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('messy_name')
not node.is_empty() and node.is_equivalent("messy_name = '  sARaH dAVis  '")
`)) })
```

You should print `messy_name` so that the output reads `Original name: '  sARaH dAVis  '`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f\\"Original name: '{messy_name}'\\")") or
_Node(_code).has_call("print(\\"Original name: '\\" + messy_name + \\"'\\")")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
