---
id: 69e62ce96df237884da25282
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

The `.lower()` method converts all characters in a string to lowercase. This is useful for normalizing email addresses:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

You can call both `.strip()` and `.lower()` on a string this way:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

Create a variable `messy_email` and assign it the string `  SARAH.DAVIS@COMPANY.COM  `. Pay attention to the two leading spaces and the two trailing spaces.

Also, call `.strip().lower()` on `messy_email` and assign the result to a variable `clean_email`.

Then print both so that the output reads:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

You should create a variable named `messy_email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('messy_email')`)) })
```

You should assign <code>'&nbsp;&nbsp;SARAH.DAVIS@COMPANY.COM&nbsp;&nbsp;'</code> to `messy_email`.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`)) })
```

You should create a variable named `clean_email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('clean_email')`)) })
```

You should call `.strip().lower()` on `messy_email` and assign the result to a `clean_email` variable.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('clean_email')
not node.is_empty() and node.is_equivalent("clean_email = messy_email.strip().lower()")
`)) })
```

You should print `messy_email` so that the output reads `Original email: '  SARAH.DAVIS@COMPANY.COM  '`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f\\"Original email: '{messy_email}'\\")") or
_Node(_code).has_call("print(\\"Original email: '\\" + messy_email + \\"'\\")")
)
`)) })
```

You should print `clean_email` so that the output reads `Cleaned email: 'sarah.davis@company.com'`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f\\"Cleaned email: '{clean_email}'\\")") or
_Node(_code).has_call("print(\\"Cleaned email: '\\" + clean_email + \\"'\\")")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
