---
id: 69e73a141addc6f442924b0d
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

You already used `.find()` and slicing on the `email` variable. Now you will use the same technique on `clean_email`.

Find the position of the `@` character in `clean_email` and store it in `at_pos`. Then use slicing to extract everything before `@` and store it in `email_user`.

Then print `clean_email` and `email_user` so that the output reads:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

You should create a variable named `at_pos`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('at_pos')`)) })
```

You should call `.find('@')` on `clean_email` and assign the result to the `at_pos` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`)) })
```

You should create a variable named `email_user`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('email_user')`)) })
```

You should use slicing with `at_pos` on `clean_email` and assign the result to the `email_user` variable.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('email_user')
node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`)) })
```

You should print `clean_email` so that the output reads `Email: sarah.davis@company.com`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Email: {clean_email}')") or
_Node(_code).has_call("print('Email: ' + clean_email)")
)
`)) })
```

You should print `email_user` so that the output reads `Username: sarah.davis`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Username: {email_user}')") or
_Node(_code).has_call("print('Username: ' + email_user)")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
