---
id: 69e73a141addc6f442924b0d
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

You already used `.find()` and slicing on the `email` variable. Now you will use the same technique on `clean_email`.

Find the position of the `@` character in `clean_email` and store it in `at_pos`. Then use slicing to extract everything before `@` and store it in `email_user`.

Then print the result of concatenating the string `Email: ` and the `clean_email` variables. Also concatenate the string `Username: ` and the `email_user` variable and print the result.

# --hints--

You should create a variable named `at_pos`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('at_pos')`)) })
```

You should call `.find('@')` on `clean_email` and assign the result to the `at_pos` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`)) })
```

You should create a variable named `email_user`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('email_user')`)) })
```

You should use slicing with `at_pos` on `clean_email` and assign the result to the `email_user` variable.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('email_user')
node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`)) })
```

You should print the result of concatenating the string `Email: ` and the `clean_email` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Email: ' + clean_email)")`)) })
```

You should print the result of concatenating the string `Username: ` and the `email_user` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Username: ' + email_user)")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

--fcc-editable-region--

--fcc-editable-region--
```
