---
id: 69e75dc335b6ec9903a87d3e
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

You can also use `.replace()` to remove unwanted characters. For example, removing dashes from a phone number:

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

Create a variable `phone` and assign it `123-456-7890`. Then replace `-` with an empty string and assign the result to a variable named `phone_clean`.

Then print both so that the output reads:

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

You should create a variable named `phone`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('phone')`)) })
```

You should assign `123-456-7890` to `phone`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('phone')
not node.is_empty() and node.is_equivalent("phone = '123-456-7890'")
`)) })
```

You should create a variable named `phone_clean`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('phone_clean')`)) })
```

You should use `.replace('-', '')` on `phone` and assign the result to the `phone_clean` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`)) })
```

You should print `phone` so that the output reads `Phone with dashes: 123-456-7890`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Phone with dashes: {phone}')") or
_Node(_code).has_call("print('Phone with dashes: ' + phone)")
)
`)) })
```

You should print `phone_clean` so that the output reads `Phone without dashes: 1234567890`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Phone without dashes: {phone_clean}')") or
_Node(_code).has_call("print('Phone without dashes: ' + phone_clean)")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
