---
id: 69e774b82fe0bfdb684f95b4
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Sometimes you need to check how a string starts. Python provides a method `.startswith()` for this. It returns `True` if the string begins with a given prefix, otherwise it returns `False`.

Here's an example:

```python
code = 'EMP-001'
print(code.startswith('EMP')) # True
print(code.startswith('001')) # False
```

Create a variable `employee_code` and assign it the string `EMP-2024-SD`. Then create a `starts_with_emp` variable that checks if `employee_code` starts with `EMP`.

# --hints--

You should create a variable named `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('employee_code')`)) })
```

You should assign `'EMP-2024-SD'` to `employee_code`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('employee_code')
not node.is_empty() and node.is_equivalent("employee_code = 'EMP-2024-SD'")
`)) })
```

You should create a variable named `starts_with_emp`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('starts_with_emp')`)) })
```

You should use `.startswith('EMP')` on `employee_code` and assign the result to the `starts_with_emp` variable.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('starts_with_emp')
not node.is_empty() and node.is_equivalent("starts_with_emp = employee_code.startswith('EMP')")
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print('Email: ' + clean_email)
print('Username: ' + email_user)

display_name = email_user.replace('.', ' ').title()
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

rejoined = ' | '.join(address_parts)
print('Rejoined: ' + rejoined)

--fcc-editable-region--

--fcc-editable-region--
```
