---
id: 69e7766b93d27261775d1cf1
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

Now print `employee_code`, `starts_with_emp`, `clean_email`, and `ends_with_com` so that the output reads:

```md
Code: EMP-2024-SD
Starts with EMP: True
Email: sarah.davis@company.com
Ends with .com: True
```

# --hints--

You should print `employee_code` so that the output reads `Code: EMP-2024-SD`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Code: {employee_code}')") or
_Node(_code).has_call("print('Code: ' + employee_code)")
)
`)) })
```

You should print `starts_with_emp` so that the output reads `Starts with EMP: True`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Starts with EMP: {starts_with_emp}')") or
_Node(_code).has_call("print('Starts with EMP: ' + str(starts_with_emp))")
)
`)) })
```

You should print `clean_email` so that the output reads `Email: sarah.davis@company.com`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Email: {clean_email}')") or
_Node(_code).has_call("print('Email: ' + clean_email)")
)
`)) })
```

You should print `ends_with_com` so that the output reads `Ends with .com: True`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Ends with .com: {ends_with_com}')") or
_Node(_code).has_call("print('Ends with .com: ' + str(ends_with_com))")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
has_com = '.com' in email
print(f'Has @ symbol: {has_at}')
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

--fcc-editable-region--

--fcc-editable-region--
```
