---
id: 69e736c39dfa403496e8071e
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

The `.upper()` method converts all characters to uppercase. This is useful for name badges or headings:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Create a variable named `name_badge` and assign it the result of calling the `.upper()` method on `clean_name`.

Then print the result of concatenating the string `Name badge: ` and the `name_badge` variable.

# --hints--

You should create a variable named `name_badge`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('name_badge')`)) })
```

You should call `.upper()` on `clean_name` and assign the result to the `name_badge` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`)) })
```

You should print the result of concatenating the string `Name badge: ` and the `name_badge` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Name badge: ' + name_badge)")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

--fcc-editable-region--

--fcc-editable-region--
```
