---
id: 69e736c39dfa403496e8071e
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

The `.upper()` method converts all characters to uppercase. This is useful for name badges or headings:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Create a variable named `name_badge` and assign it the result of calling the `.upper()` method on `clean_name`.

Then print `name_badge` so that the output reads: `Name badge: SARAH DAVIS`.

# --hints--

You should create a variable named `name_badge`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('name_badge')`)) })
```

You should call `.upper()` on `clean_name` and assign the result to the `name_badge` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`)) })
```

You should print `name_badge` so that the output reads `Name badge: SARAH DAVIS`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Name badge: {name_badge}')") or
_Node(_code).has_call("print('Name badge: ' + name_badge)")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
has_com = '.com' in email
print(f'Has @ symbol: {has_at}')
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
