---
id: 69e7766b93d27261775d1cf1
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

In previous steps, you used `+` to build print statements.

As you learned in the previous workshop, Python also has f-strings which make it easier to include variables inside strings.

Now use f-strings to print the following variables, each concatenated with the given text:

- `employee_code` concatenated with `'Code: '
- `starts_with_emp` concatenated with `'Starts with EMP: '`
- `clean_email` concatenated with `'Email: '`
- `ends_with_com` concatenated with `'Ends with .com: '`

# --hints--

You should print `employee_code` by concatenating `'Code: '` with the `employee_code` variable, using an f-string.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'Code: {employee_code}')")`)) })
```

You should print `starts_with_emp` by concatenating `'Starts with EMP: '` with the `starts_with_emp` using an f-string.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'Starts with EMP: {starts_with_emp}')")`)) })
```

You should print `clean_email` by concatenating`'Email: '` with the `clean_email` variable using an f-string.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'Email: {clean_email}')")`)) })
```

You should print `ends_with_com` by concatenating `'Ends with .com: '` with the `ends_with_com` variable using an f-string.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'Ends with .com: {ends_with_com}')")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print('Email: ' + clean_email)
print('Username: ' + email_user)

display_name = email_user.replace('.', ' ').title()
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

rejoined = ' | '.join(address_parts)
print('Rejoined: ' + rejoined)

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

--fcc-editable-region--

--fcc-editable-region--
```
