---
id: 69e73f4f5f986657647bbd63
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

The `.replace(old, new)` method replaces every occurrence of the substring `old` with another, `new`. You can also chain it with other methods. For example:

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

Replace `.` with a single space in `email_user`, and then call `.title()` on the result. Assign the result to a variable named `display_name`.

Then print `display_name` so that the output reads: `Display name: Sarah Davis`.

# --hints--

You should create a variable named `display_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('display_name')`)) })
```

You should use `.replace('.', ' ')` and `.title()` on `email_user` and assign the result to the `display_name` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`)) })
```

You should print `display_name` so that the output reads `Display name: Sarah Davis`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Display name: {display_name}')") or
_Node(_code).has_call("print('Display name: ' + display_name)")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
