---
id: 69e75dc335b6ec9903a87d3e
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

You can also use `.replace()` to remove unwanted characters. For example, removing dashes from a phone number:

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

Create a variable `phone` and assign it `'123-456-7890'`. Then replace `-` with an empty string and assign the result to a variable named `phone_clean`.

Then, print the result of concatenating the string `Phone with dashes: ` and the `phone` variable. Also, print the result of concatenating the string `Phone without dashes: ` and the `phone_clean` variable.

# --hints--

You should create a variable named `phone`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('phone')`)) })
```

You should assign `'123-456-7890'` to `phone`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('phone')
not node.is_empty() and node.is_equivalent("phone = '123-456-7890'")
`)) })
```

You should create a variable named `phone_clean`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('phone_clean')`)) })
```

You should use `.replace('-', '')` on `phone` and assign the result to the `phone_clean` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`)) })
```

You should print the result of concatenating the string `Phone with dashes: ` and the `phone` variable .

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Phone with dashes: ' + phone)")`)) })
```

You should print the result of concatenating the string `Phone without dashes: ` and the `phone_clean` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Phone without dashes: ' + phone_clean)")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print('Email: ' + clean_email)
print('Username: ' + email_user)

display_name = email_user.replace('.', ' ').title()
print('Display name: ' + display_name)

--fcc-editable-region--

--fcc-editable-region--
```
