---
id: 69e76e91402d7850510a5f18
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

The `.join(list)` method is the opposite of `.split()`. It joins a list of strings into one string using a separator. For example:

```py
parts = ['apple', 'banana', 'cherry']
joined = ' - '.join(parts)
print(joined)  # Output: apple - banana - cherry
```

Using `address_parts`, create a variable `rejoined` by joining with `' | '`.

Then print `rejoined` with the prefix `'Rejoined: '`.

# --hints--

You should create a variable named `rejoined`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('rejoined')`)) })
```

You should assign `rejoined` the value `' | '.join(address_parts)`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('rejoined')
not node.is_empty() and node.is_equivalent("rejoined = ' | '.join(address_parts)")
`)) })
```

You should print `rejoined` with the prefix `'Rejoined: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Rejoined: ' + rejoined)")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print('Email: ' + clean_email)
print('Username: ' + email_user)

display_name = email_user.replace('.', ' ').title()
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

--fcc-editable-region--

--fcc-editable-region--
```
