---
id: 69e76d1997661be7f0040e2d
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

The `.split(separator)` method splits a string into a list of substrings at each occurrence of `separator`. For example:

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

Create a variable `full_address` and assign it `123 Main Street, Springfield, IL`. Then create `address_parts` by splitting on a comma followed by a single space.

Then print `address_parts` so that the output reads `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

# --hints--

You should create a variable named `full_address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('full_address')`)) })
```

You should assign `123 Main Street, Springfield, IL` to `full_address`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('full_address')
not node.is_empty() and node.is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`)) })
```

You should create a variable named `address_parts`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('address_parts')`)) })
```

You should use `.split(', ')` on `full_address` and assign the result to the `address_parts` variable.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`)) })
```

You should print `address_parts` so that the output reads `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Address parts: {address_parts}')") or
_Node(_code).has_call("print('Address parts: ' + str(address_parts))")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
