---
id: 69eb3158c33d9a229a79951c
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

The `.count(substring)` method returns the number of times a substring appears in a string. For example:

```py
text = 'banana'
print(text.count('a'))  # Output: 3
```

Create two variables named `dot_count` and `dash_count`. `dot_count` should count how many `.` characters appear in `clean_email`, and `dash_count` should count how many `-` characters appear in `employee_code`.

# --hints--

You should create a variable named `dot_count`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('dot_count')`)) })
```

You should use `.count('.')` on `clean_email` to assign `dot_count` a value.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('dot_count')
not node.is_empty() and node.is_equivalent("dot_count = clean_email.count('.')")
`)) })
```

You should create a variable named `dash_count`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('dash_count')`)) })
```

You should use `.count('-')` on `employee_code` to assign `dash_count` a value.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('dash_count')
not node.is_empty() and node.is_equivalent("dash_count = employee_code.count('-')")
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print('Email: ' + clean_email)
print('Username: ' + email_user)

display_name = email_user.replace('.', ' ').title()
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

rejoined = ' | '.join(address_parts)
print('Rejoined: ' + rejoined)

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

print(f'Code: {employee_code}')
print(f'Starts with EMP: {starts_with_emp}')
print(f'Email: {clean_email}')
print(f'Ends with .com: {ends_with_com}')

--fcc-editable-region--

--fcc-editable-region--
```
