---
id: 6a71b3e49668c3be461257b5
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

The `.title()` method capitalizes the first letter of each word. Here is an example:

```py
lowercase = 'hello world'
uppercase = lowercase.title()
print(uppercase)  # Output: Hello World
```

Create a variable named `clean_name` and assign it the result of calling `.title()` on `messy_name`.

# --hints--

You should create a variable named `clean_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('clean_name')`)) })
```

You should use `.title()` on `messy_name` to assign `clean_name` a value.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('clean_name')
not node.is_empty() and node.is_equivalent("clean_name = messy_name.title()")
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
