---
id: 6a71a8df10e823e77358d896
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Create another variable named `has_com` that checks if `.com` is in `email`.

Then print the result of concatenating the string`Has .com: ` and the `has_com` variable.

Don't forget to use the `str()` function to properly convert the numerical variable.

# --hints--

You should create a variable named `has_com`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('has_com')`)) })
```

Your `has_com` variable should have the value `'.com' in email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('has_com')
not node.is_empty() and (
  node.is_equivalent("has_com = '.com' in email")
)
`)) })
```

You should print the result of concatenating the string `Has .com: ` and the `has_com` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Has .com: ' + str(has_com))")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
print('Has @ symbol: ' + str(has_at))
--fcc-editable-region--

--fcc-editable-region--
```
