---
id: 6a71a8df10e823e77358d896
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Create another variable named `has_com` that checks if `'.com'` is in `email`.

Then print both `has_at` and `has_com` with the prefixes `'Has @ symbol: '` and `'Has .com: '` respectively.

Don't forget to add `str()` to both variables to properly concatenate them with their prefixes.

# --hints--

You should create a variable named `has_com`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('has_com')`)) })
```

Your `has_com` variable should have the value `'.com' in email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('has_com')
not node.is_empty() and (
  node.is_equivalent("has_com = '.com' in email")
)
`)) })
```

You should print `has_at` with the prefix `'Has @ symbol: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Has @ symbol: ' + str(has_at))")`)) })
```

You should print `has_com` with the prefix `'Has .com: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Has .com: ' + str(has_com))")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
--fcc-editable-region--

--fcc-editable-region--
```
