---
id: 6a71a8df10e823e77358d896
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Create another variable named `has_com` that checks if `.com` is in `email`.

Then print `has_com` so that the output reads: `Has .com: True`.

# --hints--

You should create a variable named `has_com`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('has_com')`)) })
```

Your `has_com` variable should have the value `'.com' in email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('has_com')
not node.is_empty() and (
  node.is_equivalent("has_com = '.com' in email")
)
`)) })
```

You should print `has_com` so that the output reads `Has .com: True`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Has .com: {has_com}')") or
_Node(_code).has_call("print('Has .com: ' + str(has_com))")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
