---
id: 6a71b3e49668c3be461257b5
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

The `.title()` method capitalizes the first letter of each word. Here is an example:

```py
lowercase = 'hello world'
uppercase = lowercase.title()
print(uppercase)  # Output: Hello World
```

Create a variable named `clean_name` and assign it the result of calling `title()` on `messy_name`.

# --hints--

You should create a variable named `clean_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('clean_name')`)) })
```

You should use `.title()` on `messy_name` to assign a value to `clean_name`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('clean_name')
not node.is_empty() and node.is_equivalent("clean_name = messy_name.title()")
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

--fcc-editable-region--

--fcc-editable-region--
```
