---
id: 6a71a8df10e823e77358d896
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Create another variable named `has_com` that checks if `.com` is in `email`.

Then print `has_com` so that the output reads: `Has .com: True`.

# --hints--

You should create a variable named `has_com`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

Your `has_com` variable should have the value `'.com' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

You should print `has_com` so that the output reads `Has .com: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
