---
id: 6a73427af41619a909fb21b4
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

Now print both variables using f-strings with the labels `"Dots in email: "` and `"Dashes in code: "`.

With that, your employee profile validator workshop is complete!

# --hints--


You should print `dot_count` with the label `'Dots in email: '` using an f-string.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'Dots in email: {dot_count}')")`)) })
```

You should print `dash_count` with the label `'Dashes in code: '` using an f-string.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'Dashes in code: {dash_count}')")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '

messy_email = '  SARAH.DAVIS@COMPANY.COM  '

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
display_name = email_user.replace('.', ' ').title()

print('Email: ' + clean_email)
print('Username: ' + email_user)
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

rejoined = ' | '.join(address_parts)
print('Rejoined: ' + rejoined)

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

print(f'Code: {employee_code}')
print(f'Starts with EMP: {starts_with_emp}')
print(f'Email: {clean_email}')
print(f'Ends with .com: {ends_with_com}')

dot_count = clean_email.count('.')
dash_count = employee_code.count('-')

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '

messy_email = '  SARAH.DAVIS@COMPANY.COM  '

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
display_name = email_user.replace('.', ' ').title()

print('Email: ' + clean_email)
print('Username: ' + email_user)
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

rejoined = ' | '.join(address_parts)
print('Rejoined: ' + rejoined)

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

print(f'Code: {employee_code}')
print(f'Starts with EMP: {starts_with_emp}')
print(f'Email: {clean_email}')
print(f'Ends with .com: {ends_with_com}')

dot_count = clean_email.count('.')
dash_count = employee_code.count('-')
print(f'Dots in email: {dot_count}')
print(f'Dashes in code: {dash_count}')
```
