---
id: 6a86b9e64da746838e67fb13
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

The `.strip()` method removes whitespace from both ends of a string.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

You can chain `.strip()` and `.title()` to remove whitespace from both ends of a string and capitalize the first letter of each word:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Update the value of `clean_name` to call `.strip()` on `messy_name` before `.title()`.

After that, print `clean_name` so that the output reads `Cleaned name: 'Sarah Davis'`.

# --hints--

You should update `clean_name` to call `.strip()` on `messy_name` before `.title()`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('clean_name')
not node.is_empty() and node.is_equivalent("clean_name = messy_name.strip().title()")
`)) })
```

You should print `clean_name` so that the output reads `Cleaned name: 'Sarah Davis'`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f\\"Cleaned name: '{clean_name}'\\")") or
_Node(_code).has_call("print(\\"Cleaned name: '\\" + clean_name + \\"'\\")")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
