---
id: 6a86be3e8cc29e9ae239e70a
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

You might also need to check how a string ends. Python provides the `.endswith()` method for this. It returns `True` if the string ends with the given suffix, otherwise it returns `False`.

Here's an example:

```python
code = 'EMP-001'
print(code.endswith('001')) # True
print(code.endswith('EMP')) # True
```

Create a variable `ends_with_com` that checks if `clean_email` ends with `'.com'`

# --hints--

You should create a variable named `ends_with_com`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('ends_with_com')`)) })
```

You should use `.endswith('.com')` on `clean_email` to assign `ends_with_com` a value.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('ends_with_com')
not node.is_empty() and node.is_equivalent("ends_with_com = clean_email.endswith('.com')")
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

has_at = '@' in email
has_com = '.com' in email
print('Has @ symbol: ' + str(has_at))
print('Has .com: ' + str(has_com))

messy_name = '  sARaH dAVis  '
print("Original name: '" + messy_name + "'")

clean_name = messy_name.strip().title()
print("Cleaned name: '" + clean_name + "'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print('Name badge: ' + name_badge)

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print('Email: ' + clean_email)
print('Username: ' + email_user)

display_name = email_user.replace('.', ' ').title()
print('Display name: ' + display_name)

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print('Phone with dashes: ' + phone)
print('Phone without dashes: ' + phone_clean)

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print('Address parts: ' + str(address_parts))

rejoined = ' | '.join(address_parts)
print('Rejoined: ' + rejoined)

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')

--fcc-editable-region--

--fcc-editable-region--
```
