---
id: 6a86be3e8cc29e9ae239e70a
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

You might also need to check how a string ends. Python provides the `.endswith()` method for this. It returns `True` if the string ends with the given suffix, otherwise it returns `False`.

Here's an example:

```python
code = 'EMP-001'
print(code.endswith('EMP')) # False
print(code.endswith('001')) # True
```

Create a variable `ends_with_com` that checks if `clean_email` ends with `.com`.

# --hints--

You should create a variable named `ends_with_com`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('ends_with_com')`)) })
```

You should use `.endswith('.com')` on `clean_email` and assign the result to the `ends_with_com` variable.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('ends_with_com')
not node.is_empty() and node.is_equivalent("ends_with_com = clean_email.endswith('.com')")
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print(f'Email: {email}')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]
print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')

--fcc-editable-region--

--fcc-editable-region--
```
