---
id: 6a870cc60c23f0332a2fe788
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Remeber from a previous workshop that you can join strings together using the `+` operator. This is called concatenation.

Here's a reminder of how to do it:

```py
name = 'Alice'
print('Hello, ' + name)  # Hello, Alice
```

Now, print `email` and concatenate the prefix `'Email: '` with it.

# --hints--

You should print the result of concatenating the string `Email: ` with the `email` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Email: ' + email)")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
