---
id: 6a870cc60c23f0332a2fe788
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Remember from a previous workshop that you can use an f-string to include a variable inside a string, by writing an `f` before the opening quote and wrapping the variable in `{}`.

Here's a reminder of how it works:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Now, print `email` so that the output reads: `Email: alice.johnson@company.com`

# --hints--

You should print `email` so that the output reads `Email: alice.johnson@company.com`.

```js
({ test: () => assert(runPython(`
(
_Node(_code).has_call("print(f'Email: {email}')") or
_Node(_code).has_call("print('Email: ' + email)")
)
`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
