---
id: 612e96fc87fe8e44f69f7ec5
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

では、鍵盤のスタイルを設定しましょう。 `#piano` ルールの下で、`class` セレクターを使用して `.keys` 要素を選択してください。 新しいルールには、`background-color` プロパティの値を `#040404` に、`width` プロパティの値を `949px` に、`height` プロパティの値を `180px` に設定してください。

# --hints--

`.keys` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

`.keys` セレクターには値が `#040404` に設定されている `background-color` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

`.keys` セレクターには値が `949px` に設定されている `width` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

`.keys` セレクターには値が `180px` に設定されている `height` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
