---
id: 612e9d142affc44a453655db
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

これから鍵盤の鍵一つ一つのスタイルを設定しましょう。 `.key` 要素用の `class` セレクターを作成してください。 `background-color` の値を `#ffffff` に、`position` プロパティの値を `relative` に、`width` プロパティの値を `41px` に、`height` プロパティの値を `175px` に設定してください。

# --hints--

`.key` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

`.key` セレクターには値が `#ffffff` に設定されている `background-color` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

`.key` セレクターには値が `relative` に設定されている `position` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

`.key` セレクターには値が `41px` に設定されている `width` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

`.key` セレクターには値が `175px` に設定されている `height` プロパティが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
