---
id: 612e83ec2eca1e370f830511
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`link` 要素を 1 つ `head` 要素内に追加してください。 その `link` 要素の、`rel` 属性の値を `stylesheet` に、`href` 属性の値を `./styles.css` に設定してください。

# --hints--

コードには `link` 要素が 1 つ必要です。

```js
assert.match(code, /<link/)
```

既存の `head` タグを変更しないようにしてください。 終了タグを削除していないか確認してください。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

`link` 要素は自己閉じ要素である必要があります。

```js
assert(code.match(/<link[\w\W\s]+\/>/i));
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

`link` 要素は値が `stylesheet` に設定されている `rel` 属性を持つ必要があります。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
