---
id: 612e77aba7ca691f598feb02
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`meta` タグを 2 つ追加してください。ひとつはモバイル機器用にページを最適化し、もうひとつはページで使用する `charset` を指定するようにしてください。

# --hints--

`meta` 要素をページに 2 つ追加する必要があります。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

`meta` 要素のひとつは `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もうひとつの `meta` 要素は `charset` 属性を `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
