---
id: 612e804c54d5e7308d7ebe56
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`.keys` 要素の中に、`div` 要素を 7 つ追加してください。 それらすべての class の値を `key` に設定してください。

# --hints--

新しい `div` 要素を 7 つ作成する必要があります。

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

新しい 7 つの `div` 要素は `.keys` 要素の中にある必要があります。

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

新しい 7 つの `div` 要素はすべて、`class` の値が `key` に設定される必要があります。

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
