---
id: 612e6afc009b450a437940a1
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML の基本的な構造から始めましょう。 `DOCTYPE` 宣言と、`html`、`head`、`body`、`title` 要素を追加してください。 `title` を `Piano` に設定してください。

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントのタイプが `html` であると定義する必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

そのタイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素には開始タグが必要です。

```js
assert(code.match(/<html\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトには `Piano` というタイトルが必要です。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
