---
id: 6551fb531ce7f773071a0b84
title: 步驟 8
challengeType: 20
dashedName: step-8
---

# --description--

As you can see, the output of printing `type(message)` is `<class 'str'>`, which means that your variable is a string, indicated as `str`.

Now go to a new line and create another variable called `shift` and assign the value `3` to this variable.

# --hints--

You should declare a variable called `shift`.

```js
({ test: () => assert(__userGlobals.has("shift")) })
```

You should assign the value `3` to your `shift` variable.

```js
({ test: () => assert.equal(__userGlobals.get("shift"), 3) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
