---
id: 65522255d5b9cd80f335c6f2
title: 步驟 26
challengeType: 20
dashedName: step-26
---

# --description--

At the end of your loop body, declare a variable called `new_index` and assign the value of `index + shift` to this variable.

# --hints--

You should declare a variable called `new_index`.

```js
({ test: () => assert(__userGlobals.has("new_index")) })
```

You should assign `index + shift` to your new variable.

```js
({ test: () => assert.match(code, /new_index\s*=\s*index\s*\+\s*shift/) })
```

Your new variable should be inside your `for` loop body.

```js
({ test: () => assert.match(code, /:(.*?)index.*\1new_index\s*=\s*index\s*\+\s*shift/s) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
