---
id: 6553a572f7a65718f1e42e18
title: 步驟 46
challengeType: 20
dashedName: step-46
---

# --description--

A function is essentially a reusable block of code. You have already met some built-in functions, like `print()`, `find()` and `len()`. But you can also define custom functions like this:

```py
def function_name():
    <code>
```

A function declaration starts with the `def` keyword followed by the function name — a valid variable name — and a pair of parentheses. The declaration ends with a colon.

Right after your `shift` variable, declare a function called `caesar` and indent the following lines, so they become the function body.

# --hints--

You should have a `caesar` function.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    a = __locals.get("caesar")
    inspect.isfunction(a)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
