---
id: 65688b5a1655a7a6caede847
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

If you try to change the individual characters of a string, you will get a `TypeError`. However, a variable can be reassigned another string:

```py
message = 'Hello World'
message = 'Hello there!'
```

Delete the `text[0]` line and reassign `text` the string `Albatross`.

# --hints--

You should delete `text[0] = 'C'`.

```js
({ test: () => assert.isFalse( /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m.test(code)) })
```

You should reassign `text` the string `Albatross`. Do not modify `text = 'Hello World'`.

```js
({ test: () => assert.match(code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s*\n^text\s*=\s*("|')Albatross\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
