---
id: 6568b7a40c7ba0ccbc4b4425
title: 步驟 29
challengeType: 20
dashedName: step-29
---

# --description--

As you can see, each character of the reassigned string gets printed inside the loop.

Go back to the original string by deleting the `text` reassignment.

# --hints--

You should remove your `text = 'Albatross'` line.

```js
({ test: () => assert.isFalse( /^text\s*=\s*("|')Albatross\1/m.test(code)) })
```

You should still have a `text` variable with the value `Hello World`.

```js
({ test: () => assert.equal(__userGlobals.get("text"), "Hello World") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
