---
id: 60b69a66b6ddb80858c5157b
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

لبدء لوحتك، أعط عنصر `body` الـ `background-color` من `rgb(184, 132, 46)`.

# --hints--

يجب عليك استخدام منتقي `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب أن يكون لعنصر `body` الخاص بك خاصية `background-color` بقيمة `rgb (184, 132, 46)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(184, 132, 46)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  </body>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

```
