---
id: 60b69a66b6ddb80858c5157c
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

ضمن body tag الخاص بك، أضف عنصر `div`. أعطيه `id` من `back-wall`.

# --hints--

يجب عليك إضافة عنصر واحد `div` بالضبط.

```js
assert(document.querySelectorAll('div').length === 1);
```

يجب أن يكون لعنصر `div` الخاص بك `id` بقيمة `back-wall`.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'back-wall');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}
```
