---
id: 60b69a66b6ddb80858c51578
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

ابدأ بإعداد بنية HTML الخاصة بك. قم بإضافة تعريف `<!DOCTYPE>` وعنصر `html` مع سمة `lang` تم تعيينها إلى `en`. ضمن عنصر `html` أضف عنصر `head` وعنصر `body`.

# --hints--

يجب أن يحتوي الكود الخاص بك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` باستخدام `>` بعد تحديد النوع.

```js
assert(code.match(/html\s*>/gi));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة افتتاحية مع سمة `lang` بقيمة `en`

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على closing tag.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يكون لديك الـ opening tag للـ `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك الـ opening tag للـ `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

عناصر `head` و `body` يجب أن يكونا أشقاء (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

عنصر `head` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

عنصر `body` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

--fcc-editable-region--

```

```css

```
