---
id: 60b69a66b6ddb80858c51584
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

هذه الشخصية تحتاج إلى عينين. أنشئ عنصرين `div` في عنصر `#black-mask`. اعطيهم الـ classes الآتية `eyes left` و `eyes right`، بذلك الترتيب.

# --hints--

يجب عليك إنشاء عنصرين `div` ضمن عنصر `#black-mask` الخاص بك.

```js
assert(document.querySelectorAll('#black-mask div').length === 2);
```

يجب أن تحتوي أول `div` جديدة على `eyes` و `left`.

```js
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('left'));
```

يجب أن تحتوي ثاني `div` جديدة على classes تدعي `eyes` و `right`.

```js
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
        <div id="gray-instrument"></div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
