---
id: 60b69a66b6ddb80858c51579
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

ضمن عنصر `head` الخاص بك، أضف وسم `meta` مع سمة `charset` معينة إلى `utf-8`. ايضا قم بإضافة عنصر `title` مع النص `Picasso Painting`.

# --hints--

يجب عليك إضافة عنصر `meta` واحد بالضبط.

```js
assert(document.querySelectorAll('meta').length === 1);
```

يجب أن يحتوي وسم `meta` الخاص بك علي سمة `charset`.

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

يجب تعيين سمة `charset` الخاصة بك إلى `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

يجب عليك إضافة عنصر `title` واحد بالضبط.

```js
assert(document.querySelectorAll('title').length === 1);
```

عنصر `title` الخاص بك يجب أن يحتوي على النص `Picasso Painting`. لاحظ أن التهجئة والـ casing مهمان.

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
