---
id: 60b80da8676fb3227967a731
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

امضي قدما واربط مِلَفّ CSS الخاص بك الآن، حتى وإن لم تكتب أي CSS حتى الآن.

قم بإضافة عنصر `link` مع `rel` من `stylesheet` و `href` من `styles.css`.

# --hints--


الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
assert.match(code, /<link/)
```

الكود الخاص بك يجب أن يحتوي على عنصر `link` مغلق ذاتيا.

```js
assert(document.querySelectorAll('link').length === 1);
```

عنصر `link` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert.exists(document.querySelector('head > link'));
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
