---
id: 61967e74a8e3690ab6292daa
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Lenga kipengele cha `body` ili kuweka `background` kuwa linear gradient yenye pembe ya digrii 45, kuanzia `rgb(118, 201, 255)` na kuishia katika `rgb(247, 255, 222)`.

# --hints--

Unapaswa kutumia kipengele kichaguzi cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Unapaswa kutumia sifa ya `background` katika kichaguzi cha `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Unapaswa kuweka `background` kuwa `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
