---
id: 619691693bc14b0e528f5a20
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Lenga kipengele cha `.ground`, na uweke `width` ili kuchukua upana kamili wa viewport. Kisha, weka `height` kuwa `400px`.

# --hints--

Unapaswa kutumia kichaguzi cha `.ground`.

```js
assert.match(code, /\.ground\s*\{/);
```

Unapaswa kukipa kipengele cha `.ground` `width` ya `100vw`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

Unapaswa kukipa kipengele cha `.ground` `height` ya `400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
