---
id: 5daa813381b9e3db6c126b43
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

PYTHON

Setze die `hello`-Variable auf "world". Gib dann den Wert aus.

# --hints--

The source code should include `one = 2`

```js
assert(code.match(/one\s*=\s*2/))
```

Die `hello`-Variable sollte "world" entsprechen.

```js
({ test: () => assert.equal(__userGlobals.get("hello"), "world") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
one = 1
hello = "goodbye"
print(hello)
--fcc-editable-region--
```

# --solutions--

```py
one = 2
hello = "world"
print(hello)
```
