---
id: 64b163c20e59cbd4a64940b0
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Erstelle eine Funktion, um zwei Zahlen zu addieren.

# --hints--

Die Addition von 3 und 5 sollte 8 ergeben.

```js
({ test: () => assert.equal(__userGlobals.get("add")(3,5), 8) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def add(a, b):
    return a + b
```
