---
id: 64b171849f925b0773aa434c
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Erstelle eine Liste mit Zahlen von 1 bis 3 und weise sie der Variable `xs` zu.

# --hints--

`xs` sollte die Liste `[1,2,3]` darstellen.

```js
({ test: () => assert.deepEqual(Array.from(__userGlobals.get("xs")), [1,2,3])})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
xs = [1,2,3]
```
