---
id: 655cd899f8de09431eabb40c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви дізнаєтесь про структури даних, створивши алгоритм сортування злиттям.

Це алгоритм сортування, який використовує принцип «розділяй та володарюй» для сортування колекцій даних. Тобто він «розділяє» колекцію на менші частини та «володарює» ними, незалежно сортуючи, а потім об’єднує відсортовані частини.

Створіть функцію під назвою `merge_sort`. Ця функція буде обробляти завдання сортування списку чисел.

Використайте ключове слово `pass` у тілі функції.

# --hints--

Оголосіть функцію під назвою `merge_sort`. Не забудьте використати ключове слово `pass` у тілі функції.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
