---
id: 6564a9fe51964c229d5b7f4c
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Вам знадобиться параметр, який позначає колекцію даних, що потрібно відсортувати. Створіть параметр під назвою `array` у функції `merge_sort`.

# --hints--

Функція `merge_sort` має приймати параметр `array`.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
