---
id: 656627b47bd2d2a4afbc945d
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

У попередньому кроці ви отримали середню точку. Використайте її, щоб поділити `array` на дві частини та призначте ці частини до нових змінних.

Використайте синтаксис зрізів, щоб отримати ліву половину `array` та призначте її до змінної під назвою `left_part`.


# --hints--

Функція `merge_sort` повинна містити змінну під назвою `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

Призначте `array[:middle_point]` до змінної `left_part`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
