---
id: 656639771fed09b5c0e8fe71
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли ви розділили список `array` на два окремі списки, ви продовжуватимете ділити кожен список, доки кожен елемент не залишиться окремо у власному списку. Список з одним числом завжди відсортований.

Для цього рекурсивно викличте `merge_sort` у функції.

# --hints--

Викличте `merge_sort` у нижній частині тіла функції.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
