---
id: 656709e50ed928da35260276
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

В попередньому кроці ви призначили елемент з `left_part` до списку `array`, оскільки він був меншим. Але так буде не завжди. У деяких випадках порівняння елемент праворуч може бути меншим.

Створіть умову `else`, яка виконуватиметься, якщо індекс `left_part` не менший за індекс `right_part`.

Призначте `right_part[right_array_index]` до `array[sorted_index]` в межах блоку `else`.

# --hints--

Тіло умови `else` повинне містити `array[sorted_index] = right_part[right_array_index]`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\s*\[\s*sorted_index\s*\]\s*=\s*right_part\s*\[\s*right_array_index\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
