---
id: 65670d1ef177e7e2b76d9528
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Інструкції `if` та `else`, які ви створили в попередніх кроках, призначатимуть елементи до відсортованого масиву.

Кожен елемент, призначений до відсортованого масиву, займає індекс в списку. Тож вам потрібно перейти до наступного індексу без елемента.

Під блоком `if`/`else`, але все ще в межах циклу `while`, збільште `sorted_index` на 1. Це не повинно бути в тілі інструкції `if` або `else`.

# --hints--

Збільште поточне значення `sorted_index` на 1.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{8}sorted_index\s*(\+=\s*1|=\s*sorted_index\s*\+\s*1)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
