---
id: 656707e11e671ed54c96f7ec
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Якщо умова `if` оцінюється як `True`, це означає, що елемент в списку `left_part` менший за елемент в списку `right_part`, з яким він порівнюється.

У такому випадку можна призначити індекс `left_part` до відсортованого масиву.

Видаліть `pass` в межах блоку `if` та призначте `left_part[left_array_index]` до `array[sorted_index]`.

# --hints--

Замініть `pass` на `array[sorted_index] = left_part[left_array_index]` у тілі інструкції `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
