---
id: 6569b4e0bd29d17d4b53b16c
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Зараз ви повторите ту саму логіку циклу `while` для `right_part`.

Створіть цикл `while`, який запускається, якщо `right_array_index` менший за `len(right_part)`, та використайте ключове слово `pass` в тілі циклу.

# --hints--

Цикл `while` повинен містити умову `right_array_index < len(right_part)`. Не забудьте використати ключове слово `pass` у тілі циклу.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/m);
  }
})
```

Тіло циклу `while` повинне містити ключове слово `pass`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
