---
id: 6577254891048c8f2c19e961
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Знову викличте функцію `merge_sort()`. Наразі не передавайте жодних аргументів.

# --hints--

Знову викличте `merge_sort()` після попереднього виклику функції.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py

--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)

--fcc-editable-region--
```
