---
id: 65772ef923f922cd720e5008
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Pass `left_part` as the argument to the `merge_sort()` call you added in the last step.

# --hints--

Передайте `left_part` як аргумент до виклику `merge_sort()`.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*left_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort()
--fcc-editable-region--
```
