---
id: 680a5e7f3d47715196b12ca0
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Now that you have your entire shifted alphabet, remove the `print(shifted_alphabet)` call.

# --hints--

You should remove `print(shifted_alphabet)` from your code.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
print(shifted_alphabet)
--fcc-editable-region--
```
