---
id: 680a549f60f5a3fedfc0edd9
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

As you can see from the output, the shifted alphabet starts at the letter `f` because `shift` has the value `5`. But now the first five letters of the alphabet – `a`, `b`, `c`, `d` and `e` – are missing from the shifted alphabet, so you'll need to add them at the end of the shifted alphabet.

The `+` operator is used to combine two or more strings together in a process called concatenation like this:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

Use the slicing syntax to extract the missing first portion of `alphabet` and concatenate it to `alphabet[shift:]`.

As a reminder, `sentence[start:stop]` returns the characters of `sentence` from position `start` (included) to `stop` (excluded).

# --hints--

You should not directly modify the `alphabet` variable.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

You should have extracted the missing first portion of `alphabet` and concatenate it to `alphabet[shift:]`. 

```js
({ test: () => runPython(`
  shift = 5
  assert shifted_alphabet == "fghijklmnopqrstuvwxyzabcde"
`) })
```

You should assign the concatenation of `alphabet[shift:]` and the missing first portion of `alphabet` to the `shifted_alphabet` variable. Use `shift` to specify where to stop the slicing.

```js
({ test: () => runPython(`
  import re
  usesForwardShift = re.findall('shifted_alphabet\s*=\s*alphabet\[shift:\]',_code)
  usesBackwardShift = re.findall('shifted_alphabet.*\s*alphabet\[0?:shift\]',_code)
  assert usesForwardShift is not None and usesBackwardShift is not None
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:]
print(shifted_alphabet)
--fcc-editable-region--
```
