---
id: 6ac277dce914e9fd6a0d920f
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

In Python, `bool` is a subclass of `int`, where `True` behaves like `1` and `False` behaves like `0`. This means that `isinstance()` treats Boolean values as integers:

```py
print(isinstance(True, int)) # True
print(isinstance(False, int)) # True
```

Because of this, your function currently accepts `True` and `False` as shift values.

Use the `or` operator to add a second condition to your `if` statement that checks if `shift` is a Boolean with `isinstance(shift, bool)`. This way, your function also returns `Shift must be an integer value.` when `shift` is `True` or `False`.

# --hints--

You should use the `or` operator to add `isinstance(shift, bool)` as a second condition to your `if` statement.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0]
_options = ["not isinstance(shift, int) or isinstance(shift, bool)", "isinstance(shift, bool) or not isinstance(shift, int)"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

Your function should return `Shift must be an integer value.` when `shift` is `True` or `False`.

```js
({ test: () => runPython(`
assert caesar('freeCodeCamp', True) == 'Shift must be an integer value.'
assert caesar('freeCodeCamp', False) == 'Shift must be an integer value.'
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if not isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
