---
id: 60b69a66b6ddb80858c5157b
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

では絵画を作り始めます。まず `body` 要素の `background-color` を `rgb(184, 132, 46)` に設定してください。

# --hints--

`body` セレクターを使用してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 要素の `background-color` プロパティが `rgb (184, 132, 46)` に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(184, 132, 46)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  </body>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

```
