---
id: 60b69a66b6ddb80858c51588
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

ID セレクターを使用して、`white-hat` という ID を持つ要素のスタイルを設定しましょう。 `width` と `height` を `0` に、`border-style` を `solid` に設定してください。

# --hints--

`#white-hat` セレクターを使用する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat'));
```

`#white-hat` セレクターの `width` プロパティを `0` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.width === '0px');
```

`#white-hat` セレクターの `height` プロパティを `0` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.height === '0px');
```

`#white-hat` セレクターの `border-style` プロパティを `solid` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.borderStyle === 'solid');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

--fcc-editable-region--

--fcc-editable-region--

```
