---
id: 60b69a66b6ddb80858c51581
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

`#back-wall` の要素の下に、`div` を作成して `class` を `characters` に設定してください。 ここに絵画のキャラクターを作成していきます。

# --hints--

新しい `div` 要素を 1 つだけ追加してください。

```js
assert(document.querySelectorAll('div').length === 2);
```

新しい `div` 要素は `#back-wall` 要素の後に置く必要があります。

```js
assert(document.querySelector('#back-wall')?.nextElementSibling?.localName === 'div');
```

新しい `div` 要素の `class` を `characters` に設定する必要があります。

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('characters'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
