---
id: 60b69a66b6ddb80858c5157f
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

通常、HTML は上から下にレンダリングされます。 コードの上部にある要素は、ページの上部に表示されます。 しかし、要素を異なる位置に移動させたい場合も多々あります。 それには `position` プロパティを使います。

`#back-wall` の要素の `position` プロパティを `absolute` に設定してください。 `absolute` 位置指定は、上から下に向かうドキュメントフローから要素を除外し、その要素のコンテナーからの相対位置での位置調整を可能にします。

要素を手動で位置指定する場合、`top`、`left`、`right`、`bottom` を使用してレイアウトを調整できます。 `#back-wall` 要素の `top` の値が `0`、`left` の値が `0` になるよう設定してください。


# --hints--

`#back-wall` セレクターの `position` プロパティを `absolute` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.position === 'absolute');
```

`#back-wall` セレクターの `top` プロパティを `0` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.top === '0px');
```

`#back-wall` セレクターの `left` プロパティを `0` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.left === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
