---
id: 60b69a66b6ddb80858c51583
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

ID が `offwhite-character` の要素内に `div` 要素を 4 つ作成してください。 その `div` 要素に、以下の `id` の値をこの順で設定してください: `white-hat`、`black-mask`、`gray-instrument`、`tan-table`

# --hints--

`.offwhite-character` 要素内に `div` 要素を 4 つ追加する必要があります。

```js
assert(document.querySelectorAll('#offwhite-character div').length === 4);
```

1 つ目の新しい `div` 要素の `id` は `white-hat` に設定する必要があります。

```js
assert(document.querySelectorAll('#offwhite-character div')[0]?.getAttribute('id') === 'white-hat');
```

2 つ目の新しい `div` 要素の `id` は `black-mask` に設定する必要があります。

```js
assert(document.querySelectorAll('#offwhite-character div')[1]?.getAttribute('id') === 'black-mask');
```

3 つ目の新しい `div` 要素の `id` は `gray-instrument` に設定する必要があります。

```js
assert(document.querySelectorAll('#offwhite-character div')[2]?.getAttribute('id') === 'gray-instrument');
```

4 つ目の新しい `div` 要素の `id` は `tan-table` に設定する必要があります。

```js
assert(document.querySelectorAll('#offwhite-character div')[3]?.getAttribute('id') === 'tan-table');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
