---
id: 60b69a66b6ddb80858c51585
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

楽器に付ける「点」を作成します。 `#gray-instrument` の要素内に `div` 要素を 5 つ追加してください。 それぞれの `class` を `black-dot` に設定してください。

# --hints--

`#gray-instrument` 要素内に新しい `div` 要素が 5 つ必要です。

```js
assert(document.querySelectorAll('#gray-instrument div').length === 5);
```

5 つの新しい `div` 要素すべてにクラス `black-dot` が必要です。

```js
assert(document.querySelectorAll('#gray-instrument .black-dot').length === 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
