---
id: 60b69a66b6ddb80858c51582
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`.characters` の要素内に、もう 1 つ `div` を作成し、`id` を `offwhite-character` に設定してください。

# --hints--

`div` 要素を 1 つだけ追加してください。

```js
assert(document.querySelectorAll('div').length === 3);
```

新しい `div` 要素は `.characters` 要素内にネストされる必要があります。

```js
assert(document.querySelector('.characters div'));
```

新しい `div` 要素の `id` を `offwhite-character` に設定する必要があります。

```js
assert(document.querySelector('.characters div')?.getAttribute('id') === 'offwhite-character');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
