---
id: 60b69a66b6ddb80858c5158c
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

ID セレクターを使用して、`black-mask` という ID を持つ要素を対象とするルールを作成してください。 `width` を `100%`、`height` を `50px`、`background-color` を `rgb(45, 31, 19)` に設定してください。

# --hints--

`#black-mask` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask'));
```

`#black-mask` セレクターの `width` プロパティを `100%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask')?.width === '100%');
```

`#black-mask` セレクターの `height` プロパティを `50px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask')?.height === '50px');
```

`#black-mask` セレクターの `background-color` プロパティを `rgb(45, 31, 19)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask')?.backgroundColor === 'rgb(45, 31, 19)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

#white-hat {
  width: 0;
  height: 0;
  border-style: solid;
  border-width: 0 120px 140px 180px;
  border-top-color: transparent;
  border-right-color: transparent;
  border-bottom-color: GhostWhite;
  border-left-color: transparent;
  position: absolute;
  top: -140px;
  left: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
