---
id: 60b69a66b6ddb80858c51584
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

キャラクターの目を作りましょう。 `#black-mask` の要素内に `div` 要素を 2 つ作成してください。 作成した要素に、順番に `eyes left` と `eyes right` というクラスを設定してください。

# --hints--

`#black-mask` 要素内に `div` 要素を 2 つ作成する必要があります。

```js
assert(document.querySelectorAll('#black-mask div').length === 2);
```

1 つ目の新しい `div` には、`eyes` と `left` というクラスが必要です。

```js
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('left'));
```

2 つ目の新しい `div` には、`eyes` と `right` というクラスが必要です。

```js
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
        <div id="gray-instrument"></div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
