---
id: 60b69a66b6ddb80858c51578
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML 構造の準備から始めましょう。 `<!DOCTYPE>` 宣言と、`lang` 属性を `en` に設定した `html` 要素を追加してください。 `html` 要素の中に `head` 要素と `body` 要素を追加してください。

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert(code.match(/html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 要素には `lang` 属性 `en` を持つ開始タグが必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

--fcc-editable-region--

```

```css

```
