---
id: 60b69a66b6ddb80858c5157d
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

ID セレクターを使用して、ID が `back-wall` の要素の `background-color` を `#8B4513` に設定してください。

# --hints--

`#back-wall` セレクターを使用する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall'));
```

`#back-wall` セレクターの `background-color` プロパティを `#8B4513` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.backgroundColor === 'rgb(139, 69, 19)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

--fcc-editable-region--

--fcc-editable-region--

```
