---
id: 60b69a66b6ddb80858c51580
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`z-index` プロパティは、HTML 要素に「レイヤー」を作成するために使われます。 画像編集ツールを使い慣れている人であれば、レイヤーを使って作業したことがあるかもしれません。 それとよく似た概念です。

`z-index` の値が大きい要素は、`z-index` の値が小さい要素の上 (手前) に重なって表示されます。 これと 1 つ前のレッスンの位置指定 (position) を組み合わせて、ユニークな効果を生み出すことができます。

ID を `back-wall` にした要素を、これから作成する他の要素の「後ろ」に表示したいので、`back-wall` の要素の `z-index` を `-1` に設定してください。

# --hints--

`#back-wall` セレクターの `z-index` プロパティを `-1` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.zIndex === '-1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
