---
id: 60b69a66b6ddb80858c51579
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素内に `meta` タグを追加し、`charset` 属性を `utf-8` に設定してください。 また、`title` 要素を追加して、`Picasso Painting` というテキストを設定してください。

# --hints--

`meta` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 要素には `charset` 属性が必要です。

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

`charset` 属性を `utf-8` に設定する必要があります。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

`title` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 要素は `Picasso Painting` というテキストをもつ必要があります。 スペルや大文字小文字の区別に気をつけましょう。

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
