---
id: 60b69a66b6ddb80858c5157c
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

本文 (body) タグ内に `div` 要素を追加してください。 その `id` を `back-wall` に設定してください。

# --hints--

`div` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('div').length === 1);
```

`div` 要素の `id` の値を `back-wall` に設定する必要があります。

```js
assert(document.querySelector('div')?.getAttribute('id') === 'back-wall');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}
```
