---
id: 60b69a66b6ddb80858c51586
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

ID セレクターを使用して、`offwhite-character` という ID を持つ要素を対象とするルールを作成してください。 その `width` を `300px`、`height` を `550px`、`background-color` を `GhostWhite` に設定してください。

# --hints--

`#offwhite-character` セレクターを使用する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character'));
```

`#offwhite-character` の `width` プロパティを `300px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.width === '300px');
```

`#offwhite-character` の `height` プロパティを `550px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.height === '550px');
```

`#offwhite-character` の `background-color` プロパティを `GhostWhite` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.backgroundColor === 'ghostwhite');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

--fcc-editable-region--

--fcc-editable-region--

```
