---
id: 60b80da8676fb3227967a731
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

まだ CSS は書いていませんが、先に CSS ファイルをリンクしておきましょう。

`link` 要素を追加して、`rel` を `stylesheet` に、 `href` を `styles.css` に設定してください。

# --hints--


コードには `link` 要素が 1 つ必要です。

```js
assert.match(code, /<link/)
```

自己終了要素 `link` が 1 つ必要です。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素は値が `stylesheet` に設定されている `rel` 属性を持つ必要があります。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
