---
id: 67116d7584d0b469b14579c1
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

The rest of the objects representing the books have been filled in for you. You can take a look at them.

Now, you should start working on getting several common parts of the books.

Start by displaying all the books. Create a `displayBooks` function with a `catalog` parameter. The function should return a string that contains the `title`, `author`, and `pages` of all the books in the `library` array.

# --hints--

You should create a `displayBooks` function.

```js
assert.isFunction(displayBooks)
```

Your `displayBooks` function should have a `catalog` parameter.

```js
assert.match(displayBooks.toString(), /catalog/)
```

Your `displayBooks` function should use a higher order function. Ex. (`filter`, `map`, `reduce`)

```js
assert.match(displayBooks.toString(), /filter|map|reduce|forEach/);
```

Your `displayBooks` function should return a string.

```js
assert.isString(displayBooks(library))
```

Your `displayBooks` function should have each book title, author and pages in it.

```js
library.forEach((book) => {
  assert.include(displayBooks(library), book.title);
  assert.include(displayBooks(library), book.author);
  assert.include(displayBooks(library), book.pages.toString());
});
```

Your `displayBooks` function should contain the total of `8` books avaialble in the library.

```js
function getNumOfBooks(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`);
}

assert.lengthOf(getNumOfBooks(library), 8);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

--fcc-editable-region--
  
--fcc-editable-region--
```
