---
id: 67116c175a77cb64c5c10d49
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this project, you will learn higher order array methods by building a library manager. The methods you'll learn include `.map()`, `.filter()`, `.reduce()` and `.forEach()`.

To get started, create an empty library array.

# --hints--

You should create a `library` array.

```js
assert.isArray(library)
```

Your `library` array should be empty.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
