---
id: 67116c175a77cb64c5c10d49
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In the previous lecture videos, you learned how to work with higher order functions like `map`, `filter` and `reduce`.

In this workshop, you will build a library manager app that will give you an opportunity to practice working with different higher order functions.

To get started, create a variable called `library` and assign it an empty array.

# --hints--

You should create a `library` array.

```js
assert.isArray(library)
```

Your `library` array should be empty.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
