---
id: 67190ce4fe18911edc2018d2
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Now, you should test out the functions by calling them with the appropriate arguments.

Create four new variables: `libraryBooks`, `bookSummaries`, `booksByArvidKahl`, and `totalPagesOfBooksInLibrary`. Set them all to the calling of `displayBooks()`, `getBookSummaries()`, `getBooksByAuthor()`, and `getTotalPages()` respectively with the appropriate parameters.

Log all the variables to the console.

With that, your library manager workshop is complete.

# --hints--

You should create a `libraryBooks` variable.

```js
assert.isNotNull(libraryBooks)
```

Your `libraryBooks` variable should be set to `displayBooks(library)`.

```js
assert.equal(libraryBooks, displayBooks(library));
```

You should log the `libraryBooks` variable to the console.

```js
assert.match(code, /console\.log\(libraryBooks\);?/)
```

You should create a `bookSummaries` variable.

```js
assert.isNotNull(bookSummaries)
```

Your `bookSummaries` variable should be set to `getBookSummaries(library)`.

```js
assert.deepEqual(bookSummaries, getBookSummaries(library));
```

You should log the `bookSummaries` variable to the console.

```js
assert.match(code, /console\.log\(bookSummaries\);?/)
```

You should create a `booksByArvidKahl` variable.

```js
assert.isNotNull(booksByArvidKahl)
```

Your `booksByArvidKahl` variable should be set to `getBooksByAuthor(library, 'Arvid Kahl')`.

```js
assert.deepEqual(booksByArvidKahl, getBooksByAuthor(library, 'Arvid Kahl'));
```

You should log the `booksByArvidKahl` variable to the console.

```js
assert.match(code, /console\.log\(booksByArvidKahl\);?/)
```

You should create a `totalPagesOfBooksInLibrary` variable.

```js
assert.isNotNull(totalPagesOfBooksInLibrary)
```

Your `totalPagesOfBooksInLibrary` variable should be set to `getTotalPages(library)`.

```js
assert.deepEqual(totalPagesOfBooksInLibrary, getTotalPages(library));
```

You should log the `totalPagesOfBooksInLibrary` variable to the console.

```js
assert.match(code, /console\.log\(totalPagesOfBooksInLibrary\);?/)
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

function displayBooks(catalog) {
  let output = 'Books in the Library:\n';

  catalog.forEach((book) => {
    output += `- ${book.title} by ${book.author} (${book.pages} pages)\n`;
  });

  return output;
}

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about);
}

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

function getTotalPages(catalog) {
  return catalog.reduce((acc, book) => acc + book.pages, 0);
}

--fcc-editable-region--
  
--fcc-editable-region--
```

# --solutions--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

function displayBooks(catalog) {
  let output = 'Books in the Library:\n';

  catalog.forEach((book) => {
    output += `- ${book.title} by ${book.author} (${book.pages} pages)\n`;
  });

  return output;
}

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about);
}

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

function getTotalPages(catalog) {
  return catalog.reduce((acc, book) => acc + book.pages, 0);
}

const bookSummaries = getBookSummaries(library);
console.log(bookSummaries);

const libraryBooks = displayBooks(library);
console.log(libraryBooks);

const booksByArvidKahl = getBooksByAuthor(library, 'Arvid Kahl');
console.log(booksByArvidKahl);

const totalPagesOfBooksInLibrary = getTotalPages(library);
console.log(totalPagesOfBooksInLibrary);
```
