---
id: 6140d10d50636e14695013b2
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

2 つ目の `.line` 要素を選択するセレクターを作成してください。 `transform` プロパティを `rotate(60deg)` に設定してください。

`transform` プロパティで要素の形を操作できることを覚えておきましょう。 ここでは、`rotate(60deg)` という値を使用することにより、要素を `transform-origin` の点を中心に時計回りに 60 度回転させます。

# --hints--

`.line:nth-of-type(2)` セレクターを作成する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

`.line:nth-of-type(2)` セレクターの `transform` プロパティを `rotate(60deg)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
