---
id: 6140c7e645d8e905819f1dd4
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

標準的なボイラープレートから始めましょう。 `DOCTYPE` 宣言、言語を英語に設定した `html` 要素、`head` 要素、および `body` 要素を追加してください。

適切な `charset` を指定する `meta` 要素、`title` 要素、および `./styles.css` ファイルを指定する `link` 要素を追加してください。

`title` を `Ferris Wheel` に設定してください。

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

コードには `meta` 要素が 1 つ必要です。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 要素には値が `UTF-8` に設定されている `charset` 属性が必要です。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトには `Ferris Wheel` というタイトルが必要です。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

コードには `link` 要素が 1 つ必要です。

```js
assert.match(code, /<link/)
```

自己終了要素 `link` が 1 つ必要です。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素には値が `stylesheet` に設定されている `rel` 属性が必要です。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 要素には値が `styles.css` に設定されている `href` 属性が必要です。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
