---
id: 6140c7e645d8e905819f1dd4
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

標準的なボイラープレートから始めましょう。 `DOCTYPE` 宣言、言語を英語に設定した `html` 要素、`head` 要素、および `body` 要素を追加してください。

適切な `charset` を指定する `meta` 要素、`title` 要素、および `./styles.css` ファイルを指定する `link` 要素を追加してください。

`title` を `Ferris Wheel` に設定してください。

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

`html` 要素には終了タグが必要です。

```js
assert.match(code, /<\/html\s*>/);
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

`head` の開始タグが 1 つ必要です。

```js
assert.match(code, /<head\s*>/i);
```

`head` の終了タグが 1 つ必要です。

```js
assert.match(code, /<\/head\s*>/i);
```

`body` の開始タグが 1 つ必要です。

```js
assert.match(code, /<body\s*>/i);
```

`body` の終了タグが 1 つ必要です。

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

コードには `meta` 要素が 1 つ必要です。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 要素には値が `UTF-8` に設定されている `charset` 属性が必要です。

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトには `Ferris Wheel` というタイトルが必要です。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

コードには `link` 要素が 1 つ必要です。

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素には値が `stylesheet` に設定されている `rel` 属性が必要です。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素には値が `styles.css` に設定されている `href` 属性が必要です。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
