---
id: 6140cbeec34e970dfe75e710
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`.wheel` 要素のセレクターを作成してください。 まずは `border` を `2px solid black` に、`border-radius` を `50%` に、`margin-left` を `50px` に設定してください。

# --hints--

`.wheel` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

`.wheel` セレクターの `border` プロパティを `2px solid black` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

`.wheel` セレクターの `border-radius` プロパティを `50%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

`.wheel` セレクターの `margin-left` プロパティを `50px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
