---
id: 6140d94b5fab7f1d73c9bedb
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

`@keyframes` アットルールは、CSS アニメーションの流れを定義するために使われます。 `@keyframes` ルールの中では、`0%` や `25%` など、アニメーションシーケンスの特定の時点に対するセレクターを作成できます。また、`from` と `to` を使用してシーケンスの始まりと終わりの時点について定義することもできます。

`@keyframes` ルールには名前をつける必要があります。この名前は別のルールから参照するために使用します。 たとえば、`@keyframes freeCodeCamp { }` というルールは `freeCodeCamp` という名前になります。

ではアニメーションを作り始めましょう。 `@keyframes` ルールを作成し、`wheel` という名前をつけてください。

# --hints--

`@keyframes` ルールが必要です。

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.length === 1);
```

新しい `@keyframes` ルールには `wheel` という名前をつける必要があります。

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name === 'wheel');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

--fcc-editable-region--

--fcc-editable-region--
```
