---
id: 6140d4bc9db3c81c51a09ab7
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

同じパターンで、下記の `.cabin` 要素を選択してそれぞれ個別のルールを適用してください:

- 2 つ目の `.cabin` の `right` プロパティを `17%` に、`top` プロパティを `93.5%` に設定してください。
- 3 つ目の `.cabin` の `right` プロパティを `67%` に、`top` プロパティを `93.5%` に設定してください。
- 4 つ目の `.cabin` の `left` プロパティを `-8.5%` に、`top` プロパティを `50%` に設定してください。
- 5 つ目の `.cabin` の `left` プロパティを `17%` に、`top` プロパティを `7%` に設定してください。
- 6 つ目の `.cabin` の `right` プロパティを `17%` に、`top` プロパティを `7%` に設定してください。

# --hints--

`.cabin:nth-of-type(2)` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

`.cabin:nth-of-type(2)` セレクターの `right` プロパティを `17%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right === '17%');
```

`.cabin:nth-of-type(2)` セレクターの `top` プロパティを `93.5%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top === '93.5%');
```

`.cabin:nth-of-type(3)` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

`.cabin:nth-of-type(3)` セレクターの `right` プロパティを `67%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right === '67%');
```

`.cabin:nth-of-type(3)` セレクターの `top` プロパティを `93.5%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top === '93.5%');
```

`.cabin:nth-of-type(4)` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

`.cabin:nth-of-type(4)` セレクターの `left` プロパティを `-8.5%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left === '-8.5%');
```

`.cabin:nth-of-type(4)` セレクターの `top` プロパティを `50%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top === '50%');
```

`.cabin:nth-of-type(5)` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

`.cabin:nth-of-type(5)` セレクターの `left` プロパティを `17%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left === '17%');
```

`.cabin:nth-of-type(5)` セレクターの `top` プロパティを `7%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top === '7%');
```

`.cabin:nth-of-type(6)` セレクターが必要です。

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

`.cabin:nth-of-type(6)` セレクターの `right` プロパティを `17%` に設定する必要があります。

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

`.cabin:nth-of-type(6)` セレクターの `top` プロパティを `7%` に設定する必要があります。

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
