---
id: 6140cd32d018ed0f600eefce
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`.wheel` セレクターの `position` プロパティを `absolute` に設定してください。 `height` と `width` をどちらも `55vw` に設定してください。

# --hints--

`.wheel` セレクターの `position` プロパティを `absolute` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

`.wheel` セレクターの `height` プロパティを `55vw` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

`.wheel` セレクターの `width` プロパティを `55vw` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
