---
id: 6140c9d35015ae0ba0c250e8
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`body` 要素内に `div` を 1 つ追加し 、その `class` を `wheel` に設定してください。

新しく追加した `div` の中に、`class` が `line` の `span` 要素を 6 つと、`class` が `cabin` の `div` 要素を 6 つ追加してください。

# --hints--

`body` 要素の中に新しい `div` 要素を作成する必要があります。

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

`body` 要素内の `div` には `wheel` という値の `class` が必要です。

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

`.wheel` 要素の中には `span` 要素が 6 つ必要です。

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

`.wheel` 要素内の 6 つの `span` 要素には `line` という値の `class` が必要です。

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

`.wheel` 要素の中には `div` 要素が 6 つ必要です。

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

`.wheel` 要素内の 6 つの `div` 要素には `cabin` という値の `class` が必要です。

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
