---
id: 65576ff7888f9e96f52a4be1
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

Graphs are data structures representing relations between pairs of elements.
These elements, called *nodes*, can be real-life objects, entities, points in space or others. The connections between the nodes are called the *edges*.
 
For example, a graph can be used to represent two points in the space, `A` and `B`, connected by a path. A graph like this will be made of two nodes connected by an edge.

Start by creating a variable called `my_graph` and assign it an empty dictionary.

# --hints--

You should have a variable called `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Your `my_graph` variable should have the value of an empty dictionary.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    graph == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
