---
id: 6557746aad2844a0cd864e12
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

While the algorithm explores the graph, it should keep track of the currently known shortest distance between the starting node and the other nodes.

Before your `for` loop, create a new variable named `distances` and assign it an empty dictionary.

# --hints--

You should have a `distances` variable.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(e.code.original["main.py"], "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Your `distances` variable should be an empty dictionary.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(e.code.original["main.py"], "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
