---
id: 65577236b056379d5dbc7000
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Your dictionary has been transformed for you into something slightly more complex. If you want you can try to figure out how this graph looks like (you can try to sketch it using circles connected by straight lines).

Now you are going to start developing the algorithm to calculate the shortest path between each node in your new graph.

Declare an empty function called `shortest_path` and don't forget the `pass` keyword. 

# --hints--

You should have a `shortest_path` function.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    f = __locals.get("shortest_path")
    inspect.isfunction(f)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}


--fcc-editable-region--
```
