---
id: 655777060d8ddea6741be1b1
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

All the distances in `distances` are set to infinite, except for the starting node. And the `unvisited` list contains all the nodes in your graph. But actually, you don't need that `for` loop to achieve this result.

Remove your `for` loop with its entire body.

# --hints--

You should remove your `for` loop and all the nested code.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(e.code.original["main.py"], "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(?<=(^\s*)distances\s*=\s*\{\s*\}\s*)\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
