---
id: 6557743527cb92a06417ea97
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Create a `for` loop to iterate over your graph, and append each node to the `unvisited` list.

# --hints--

You should create a `for` loop to iterate over `graph` inside the `shortest_path` function.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(e.code.original["main.py"], "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for\s+(\w+)\s+in\s+graph\s*:/m));
  }
})
```

You should append each node to `unvisited` inside your `for` loop.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(e.code.original["main.py"], "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for\s+(\w+)\s+in\s+graph\s*:\s*^\1\1unvisited\.append\s*\(\s*\2\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
--fcc-editable-region--
```
