---
id: 655773b0591c5f9f4045883e
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

The algorithm will start at a specified node. Then it will explore the graph to find the shortest path between the starting node, or *source*, and all the other nodes.

For that your function needs two parameters: `graph`, and `start`. Add them to your function declaration.

# --hints--

Your function should accept `graph` and `start` as the parameters, in this order.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    f = __locals.get("shortest_path")
    sig = str(inspect.signature(f))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
