---
id: 6557924d47c325bf27afbe51
title: Step 37
challengeType: 20
dashedName: step-37
---

# --description--

The algorithm is complete but you can improve the output. Also, you can provide the function with an additional argument to return only the path between two nodes.

Add `target` as the third parameter to your function declaration and give it the default value of an empty string.

# --hints--

Your function should take three parameters:`graph`, `start`, and `target`. The order matters.

```js
({ test: () => assert.match(e.code.original["main.py"], /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

The `target` parameter should have the default value of an empty string.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    foo = __locals.get("shortest_path")
    sig = str(inspect.signature(foo))
    sig == '(graph, start, target="")'
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
