---
id: 655775221059f5a20493d5d7
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

After your `for` loop, add a `print()` call and pass in the following string to see the values of the variables you have created: `f'Unvisited: {unvisited}\nDistances: {distances}'`.

# --hints--

You should print `f'Unvisited: {unvisited}\nDistances: {distances}'` after your `for` loop.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(e.code.original["main.py"], "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/(?<=(^\s*)for.*:.*)^\1print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\2\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 5), ('C', 3), ('E', 11)],
    'B': [('A', 5), ('C', 1), ('F', 2)],
    'C': [('A', 3), ('B', 1), ('D', 1), ('E', 5)],
    'D': [('C',1 ), ('E', 9), ('F', 3)],
    'E': [('A', 11), ('C', 5), ('D', 9)],
    'F': [('B', 2), ('D', 3)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
        else:
            distances[node] = float('inf')
--fcc-editable-region--
```
