---
id: 6557709b0aee699a6a00528c
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Now add `A` and `B` keys to your dictionary  — one for each node — and give each of them the connected node as the value.

# --hints--

Your dictionary should have an `A` key.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    "A" in graph
  `))
})
```

Your `A` key should have the connected node as the value.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    graph["A"] == "B"
  `))
})
```

Your dictionary should have an `B` key.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    "B" in graph
  `))
})
```

Your `B` key should have the connected node as the value.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    graph["B"] == "A"
  `))
})
```

Your dictionary should have two keys.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    len(graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {}
--fcc-editable-region--
```
