---
id: 687b3926419dec576850b814
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Now that you’ve calculated the new possible distance to the neighbor, check if it’s better than the one currently stored in the `distances` list.

Add a conditional that checks if the `new_distance` is less than `distances[node_no]`.

# --hints--

You should have a conditional statement.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).find_if()`))
  }
})
```

You should check if `new_distance` is less than `distances[node_no]`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).find_if().has_condition("new_distance < distances[node_no]")`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

--fcc-editable-region--
    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        new_distance = distances[current] + distance
        
--fcc-editable-region--
```
