---
id: 613297a923965e0703b64796
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

もう `meta` 要素についてはよく知っているかもしれませんが、タイトル、説明、キーワード、著者など、ページの情報を指定するために使用されます。

適切な `charset` の値を持つ `meta` 要素を 1 つ、ページに加えてください。

`charset` 属性はページの文字エンコーディングを指定します。今日では `UTF-8` が、大半のブラウザーでサポートされている唯一のエンコーディングとなっています。

# --hints--

`head` 要素の中に新しい `meta` 要素を作成する必要があります。

```js
assert.exists(document.querySelector('head > meta'));
```

`meta` タグに、値が `UTF-8` に設定されている `charset` 属性を加える必要があります。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
