---
id: 61329d52e5010e08d9b9d66b
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

アクセシビリティと SEO に重要な `meta` 要素としてもう一つ、`description` 定義があります。 `content` 属性の値は、検索エンジンがページの説明を表示するために使用します。

`meta` 要素を追加して、`name` 属性を `description` に設定し、`content` 属性には意味のある説明を追加しましょう。

# --hints--

`head` 要素の中に新しい `meta` 要素を加える必要があります。

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

`meta` タグに、値が `description` に設定されている `name` 属性を加える必要があります。

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

`meta` タグに `content` 属性を加える必要があります。

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

`content` 属性の値は 165 文字を超えないようにしてください。 _これは Google で表示される説明の最大の長さです。_

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
