---
id: 6133d11ef548f51f876149e3
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

ナビゲーションはアクセシビリティの中核をなす部分であり、スクリーンリーダーは、あなたが提供するページの構造情報を頼りにしています。 これはセマンティックな HTML 要素を用いて実現します。

ページに `header` 要素と `main` 要素を 1 つずつ追加してください。

`header` 要素はナビゲーションメニューの提供と、ページの紹介に使用されます。

`main` 要素にはページの中心的な内容を含めます。

# --hints--

`header` 要素を `body` 要素に追加する必要があります。

```js
assert.exists(document.querySelector('body > header'));
```

`main` 要素を `body` 要素に追加する必要があります。

```js
assert.exists(document.querySelector('body > main'));
```

`header` 要素は `main` 要素の前に来る必要があります。

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
