---
id: 613e2546d0594208229ada50
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`header` 要素内に、`img`、`h1`、`nav` の各要素を 1 つずつ入れ子にし、このページの文脈を提供しましょう。

`img` 要素は `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` を指し、また、`logo` という `id` を持つようにしてください。

`h1` 要素は `HTML/CSS Quiz` というテキストを含むようにしましょう。

# --hints--

`img` 要素を `header` 要素に 1 つ追加する必要があります。

```js
assert.exists(document.querySelector('header > img'));
```

`h1` 要素を `header` 要素に 1 つ追加する必要があります。

```js
assert.exists(document.querySelector('header > h1'));
```

`nav` 要素を `header` 要素に 1 つ追加する必要があります。

```js
assert.exists(document.querySelector('header > nav'));
```

`img`、`h1`、`nav` の各要素をこの順番で配置する必要があります。

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

`img` 要素に、値が `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` に設定されている `src` 属性を加える必要があります。

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

`img` 要素に、値が `logo` に設定されている `id` 属性を加える必要があります。

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

`h1` 要素に `HTML/CSS Quiz` というテキストを加える必要があります。

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
