---
id: 6140883f74224508174794c0
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`header` 要素について、幅が親コンテナーの幅全体に広がるよう設定し、`height` は `50px`、`background-color` は `#1b1b32` に設定してください。 次に、`display` の設定で_フレックスボックス_が使われるようにします。

# --hints--

`header` 要素セレクターを使用する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

`header` 要素の `width` の値を `100%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

`header` 要素の `height` の値を `50px` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

`header` 要素の `background-color` の値を `#1b1b32` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

`header` 要素の `display` の値を `flex` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
