---
id: 6141fabd6f75610664e908fd
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

クイズなので、ユーザーが解答を提出できるフォームが必要になります。 `section` 要素を使用すればフォーム内のコンテンツをセマンティックに分けることができます。

`main` 要素内に、3 つの `section` 要素がネストされたフォームを 1 つ作成してください。 次に、適切なメソッドを使用して、フォームを `https://freecodecamp.org/practice-project/accessibility-quiz` に送信できるようにしてください。

# --hints--

`main` 要素の中に `form` 要素を 1 つネストする必要があります。

```js
assert.exists(document.querySelector('main > form'));
```

`form` 要素の中に 3 つの `section` 要素をネストする必要があります。

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

`form` 要素には `action` 属性が必要です。

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

`action` 属性の値を `https://freecodecamp.org/practice-project/accessibility-quiz` に設定する必要があります。

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

`form` 要素には `method` 属性が必要です。

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

`form` 要素の `method` 属性の値を `post` に設定する必要があります。

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
