---
id: 614206033d366c090ca7dd42
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

書体はページのアクセシビリティにおいて重要な役割を果たします。 特に低解像度の画面では、一部のフォントは他のものよりも読みやすくなります。

`h1` と `h2` 両方の要素のフォントを `Verdana` に変更し、さらにサンセリフ体ファミリーの別の Web セーフフォントをフォールバックとして使用してください。

また、セクションの境目を分かりやすくするため、`h2` 要素に `4px solid #dfdfe2` の `border-bottom` を追加してください。

# --hints--

複数要素を選択するセレクターを使用して、`h1` と `h2` 両方の要素を選択する必要があります。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
assert.exists(gs('h1, h2') || gs('h2, h1'));
```

`font-family` プロパティの 1 つ目の値を `Verdana` に設定する必要があります。

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const style = gs('h1, h2') || gs('h2, h1');
assert.include(style?.fontFamily, 'Verdana');
```

`font-family` プロパティの 2 つ目の値を、1 つ目とは別のサンセリフ体の Web セーフフォントに設定する必要があります。 _ヒント: 例えば Tahoma などが良いでしょう_。

```js
// Acceptable fonts: Arial, sans-serif, Helvetica, Tahoma, Trebuchet MS.
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const style = gs('h1, h2') || gs('h2, h1');
assert.match(style?.fontFamily, /(Tahoma)|(Arial)|(sans-serif)|(Helvetica)|(Trebuchet MS)/);
```

`h2` 要素を選択する `h2` 要素セレクターが 1 つ必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h2'));
```

`h2` 要素の `border-bottom` プロパティの値を `4px solid #dfdfe2` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('h2')?.borderBottom, '4px solid rgb(223, 223, 226)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
        </section>
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

--fcc-editable-region--

--fcc-editable-region--

```
