---
id: 614ccc21ea91ef1736b9b578
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

アクセシビリティクイズへようこそ。 あなたは HTML と CSS の経験を積んだ開発者となりつつあるようなので、ここでは基本的なボイラープレートから始めることにします。

`html` 要素に `lang` 属性を加えるところから、このアクセシビリティの旅を始めましょう。 これはスクリーンリーダーがページの言語を特定するのに使用されます。

# --hints--

`html` 要素には `lang` 属性が必要です。 _ヒント: 英語のための値は `en` です。_

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
