---
id: 6140d263016325162fd076fe
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Crea un selector `.cabin`. Establece el `background-color` a `red`, el `width` a `20%`, y la `height` a `20%`.

# --hints--

Debes tener un selector de `.cabin`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin'));
```

Tu `.cabin` selector debería tener una propiedad `background-color` establecida a `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin')?.backgroundColor === 'red');
```

Tu selector `.cabin` debe tener una propiedad `width` establecida a `20%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin')?.width === '20%');
```

Tu selector `.cabin` debe tener una propiedad de `height` establecida a `20%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin')?.height === '20%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
