---
id: 6140c9d35015ae0ba0c250e8
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Agrega una etiqueta `div` dentro de tu elemento `body` y asignale una `class` titulada `wheel`.

Dentro de tu nueva etiqueta `div`, agrega seis elementos `span` con la `class` de `line`, ahora agrega otros seis elementos `div` con la `class` de `cabin`.

# --hints--

Debes crear una nueva etiqueta `div` dentro de tu elemento `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

La etiqueta `div` dentro de tu elemento `body` debe tener la `class` de `wheel`.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

El elemento `wheel` debe tener seis etiquetas `span` dentro de el.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

Los seis elementos `span` dentro del elemento `wheel` deben tener una `class` de nombre `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

El elemento `wheel` debe tener seis etiquetas `div` dentro de el.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

Las seis etiquetas `div` dentro del elemento `wheel` deben tener una `class` titulada `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
