---
id: 6140c7e645d8e905819f1dd4
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Comience con la codigo de repetición estándar. Añada su declaración de `DOCTYPE`, su elemento `html` con el idioma establecido en Inglés, con los elementos `head` y `body`.

Añade tu elemento `meta`, `charset` para el conjunto de caracteres correcto, tu elemento `title`, y un elemento `link` para el archivo `./styles.css`.

Establece el `title` a `Ferris Wheel`.

# --hints--

El código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debe incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debes definir que el tipo de documento es `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de apertura con un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu declaración `DOCTYPE` debe estar al inicio de tu código HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debes tener una etiqueta `head` de apertura.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta de cierre del `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta `body` de apertura.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta `body` de cierre.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

El elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Tu código debe tener un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Tu elemento `meta` debe tener un atributo `charset` con el valor`UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu proyecto debe tener un título de `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

Tu código debe tener un elemento `link`.

```js
assert.match(code, /<link/)
```

Debes tener un elemento `link` de autocierre.

```js
assert(document.querySelectorAll('link').length === 1);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
