---
id: 6140d1a351e88f159ed54fca
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Usando el mismo patrón, crea un nuevo selector para el tercer `.line`, el cuarto `.line`, el quinto `.line`, y el sexto `.line`.

Establece la propiedad `transform` para el tercer `.line` a `rotate(120deg)`, para el cuarto a `rotate(180deg)`, para el quinto a `rotate(240deg)`, y para el sexto a `rotate(300deg)`.

# --hints--

Debes crear un selector `.line:nth-of-type(3)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Tu selector `.line:nth-of-type(3)` debe tener una propiedad `transform` establecida en `rotate(120deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform === 'rotate(120deg)');
```

Debes crear un selector `.line:nth-of-type(4)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Tu selector `.line:nth-of-type(4)` debe tener una propiedad `transform` establecida en `rotate(180deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform === 'rotate(180deg)');
```

Debes crear un selector `.line:nth-of-type(5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Tu selector `.line:nth-of-type(5)` debe tener una propiedad `transform` establecida en `rotate(240deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform === 'rotate(240deg)');
```

Debes crear un selector `.line:nth-of-type(6)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Tu selector `.line:nth-of-type(6)` debe tener una propiedad `transform` establecida en `rotate(300deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform === 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
