---
id: 6140c7e645d8e905819f1dd4
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Comience con la codigo de repetición estándar. Añada su declaración de `DOCTYPE`, su elemento `html` con el idioma establecido en Inglés, con los elementos `head` y `body`.

Añade tu elemento `meta`, `charset` para el conjunto de caracteres correcto, tu elemento `title`, y un elemento `link` para el archivo `./styles.css`.

Establece el `title` a `Ferris Wheel`.

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Debes definir que el tipo de documento es `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Tu elemento `html` debe tener una etiqueta de apertura con un atributo `lang` con el valor `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Tu elemento `html` debe tener una etiqueta de cierre.

```js
assert.match(code, /<\/html\s*>/);
```

Tu declaración `DOCTYPE` debe estar al inicio de tu código HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Debes tener una etiqueta `head` de apertura.

```js
assert.match(code, /<head\s*>/i);
```

Debes tener una etiqueta de cierre del `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Debes tener una etiqueta `body` de apertura.

```js
assert.match(code, /<body\s*>/i);
```

Debes tener una etiqueta `body` de cierre.

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Tu código debe tener un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Tu elemento `meta` debe tener un atributo `charset` con el valor`UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu proyecto debe tener un título de `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Tu código debe tener un elemento `link`.

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
