---
id: 69431e585f7d14fbf7ed71ba
title: Step 0
challengeType: 0
dashedName: step-0
---

# --description--

Inside the `else` block, declare an empty array named `countArray`.  
This array will store the value of `n` at each function call.

# --hints--

declare an empty array named 'countArray'.

```js
assert(Array.isArray(countup(1)));

```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function countup(n) {
  if (n < 1) {
    return [];
  } else {
    
  }
}
--fcc-editable-region--
```
