---
id: 694470b9a65b278f981ec62e
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Finally, test your `countUp` function by calling it with the argument `5` inside a `console.log()` statement.  
The output should be `[1, 2, 3, 4, 5]`.


# --hints--

Call the function 'countUp' with argument '5' inside a 'console.log()' statement

```js
assert.deepStrictEqual(countUp(5), [1, 2, 3, 4, 5]);
assert(/console\.log\s*\(\s*countUp\s*\(\s*5\s*\)\s*\)/.test(code));


```

# --seed--

## --seed-contents--

```js

function countUp(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countUp(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```
