---
id: 69432376b7208ebe0d54331c
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Finally, test your `countup` function by calling it with the argument `5` inside a `console.log()` statement.  
The output should be `[1, 2, 3, 4, 5]`.


# --hints--

Call the function 'countup' with argument '5' inside a 'console.log()' statement

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);

```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function countup(n) {
  if (n < 1) {
    return [];
  } else {
    const countArray = countup(n - 1);
    countArray.push(n);
    return countArray;
  }
}

--fcc-editable-region--
```
