---
id: 69431aaae3a67b2a740243a3
title: Step 0
challengeType: 0
dashedName: step-0
---

# --description--

Add an 'if' statement inside your 'countup' function. 
Check if 'n' is less than 1 and return an empty array.
This will serve as the base case for your recursion.

# --hints--
Remember: recursion needs a base case to stop.
Use an 'if' statement to check when 'n' is less than 1

```js
assert.deepStrictEqual(countup(0), []);

```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function countup(n) {

}
--fcc-editable-region--
```
