---
id: 69431175143ddcc3360274c0
title: Step 0
challengeType: 0
dashedName: step-0
demoType: onClick
---

# --description--

In this workshop, you will be creating a countup function.
Start by creating an empty function named 'countup' that takes a single parameter 'n'.
This function will later be used to build an array that count up from 1 to 'n'.

# --hints--

"Your function should be named 'countup' and accept a single parameter"


```js
assert(typeof countup === 'function');
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
