---
id: 69447054012a22f38c944ff0
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Add an 'if' statement inside your 'countUp' function. 
Check if 'number' is less than 1 and return an empty array.
This will serve as the base case for your recursion.

# --hints--
Remember: recursion needs a base case to stop.
Use an 'if' statement to check when 'number' is less than 1

```js
assert.deepStrictEqual(countup(0), []);

```


# --seed--

## --seed-contents--

```js
function countUp(number) {
--fcc-editable-region--

--fcc-editable-region--
}
```
