---
id: 69446f18fdcb11aa516f7c5c
title: Step 1
challengeType: 1
dashedName: step-1
demoType: onClick
---

# --description--

In this workshop, you will be creating a `countUp` function.

Start by creating an empty function named `countUp` that takes a single parameter `number`.

This function will later be used to build an array that count up from 1 to `number`.

# --hints--

You should have a function named `countUp` that accepts a single parameter.


```js
assert.isFunction(countUp);
assert.strictEqual(countUp.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
