---
id: 6943211551c04963545251a7
title: Step 0
challengeType: 0
dashedName: step-0
---

# --description--

Inside the `else` block, call the `countup` function with `n - 1` and assign the result to `countArray`.  
Then use `.push()` to add `n` to `countArray`, and finally return `countArray`.  
This ensures the recursion eventually reaches the base case and terminates.


# --hints--


Call `countup(n - 1)` and assign it to `countArray`.  
Then push `n` into `countArray` and return it.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);

```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function countup(n) {
  if (n < 1) {
    return [];
  } else {

  }
}
--fcc-editable-region--
```
