---
id: 694870bcf3c6874112b6f951
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Inside the `else` block, call the `countUp` function with `number - 1` and assign the result to `countArray`.  
Then use `.push()` to add `number` to `countArray`, and finally return `countArray`.  
This ensures the recursion eventually reaches the base case and terminates.


# --hints--


Call `countUp(number - 1)` and assign it to `countArray`.  
Then push `number` into `countArray` and return it.

```js
assert.deepStrictEqual(countUp(3), [1, 2, 3]);

```

# --seed--

## --seed-contents--

```js

function countUp(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
--fcc-editable-region--

--fcc-editable-region--
  }
}

```
