---
id: 6948705160b66cf93a916e29
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Declare an empty array named `countArray`.  
This array will store the value of `number` at each function call.

# --hints--

declare an empty array named 'countArray'.

```js
assert(/(let|const)\s+countArray\s*=\s*\[\s*\]/.test(code));

```

# --seed--

## --seed-contents--

```js

function countUp(number) {
  --fcc-editable-region-- 

--fcc-editable-region--
  if (number < 1) {
    return [];
  } else {
  }
}

```
