---
id: 68c362b379059c388d3874f2
title: Build a Digital Pet Game
challengeType: 25
dashedName: lab-digital-pet-game
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

The application should have two views. The requirements for each view are encompassed by the rest of the user stories:

The first view:

- Should contain a form that accepts a single string input.
- Should ask the user to name their pet.
- Should contain a button for submission that will start the game.
- Should disappear on submission.

The second view:

- Should appear upon submission of the first view.
- Should contain three `<button>` elements.
- Each button element should pertain to one each of the three following actions: `[EAT, PLAY, SLEEP]`.

  - The `EAT` button should reduce the pets `Hunger` value, and increase the pets `Energy` value.
  - The `PLAY` button should reduce the pets `Energy` value, and increase the pets `Happiness` value.
  - The `SLEEP` button should increase the pets `Hunger` value, and increase the pets `Energy` value.

- Should visually track three properties of the pet. These are: `['Hunger', 'Happiness', 'Energy']`.
- Each element that tracks each of the three properties should:

  - Contain the text relevant to that property, case insensitive, e.g., `hunger`, `Hunger`, `hUnGeR`, are all valid.
  - NOT contain the text relevent to other properties, i.e., the `Hunger` tracking element must NOT contain the text `Happiness`.
  - Contain a clear numeric representation of that properties current value.

- Each `Action` available to the pet should be encompassed by an enumerator.
- Each `Action` `button` should contain one each of the following `id`'s relevant to that button's action: `['eat-action', 'play-action', 'sleep-action']`.
- The pet should have the following moods, encompassed by a `PetMood` enumerator: `['Happy', 'Excited', 'Content', 'Sad', 'Tired', 'Sick', 'Hungry']`.

  - The pet should be `Hungry` if `Hunger > 70`.
  - The pet should be `Tired` if `Energy < 30`.
  - The pet should be `Sad` if `Happiness < 30`.
  - The pet should be `Excited` if `Happiness > 80` and `Energy < 70`.
  - The pet should be `Happy` if `Happiness > 60`.
  - Otherwise, the pet should be `Content`.

- The pet should have a number of images or emoji tied to each `Mood`.

  - These emoji or images should be tied to the current `PetMood`, with a `Record`.
  - The key of each `Record` entry should be the relevant `PetMood` enum value.

# --before-all--

```js
globalThis._PET_NAME = 'Fitzgerald';
globalThis._clock = __FakeTimers.install();

// BFS for the innerText of the provided element, and all
//  children, until the regex is matched or the tree is
//  exhausted. Will return the matching element if found,
//  otherwise, null.
globalThis._deepMatchText = (elem, regex)  => {
  let nodes = [elem];
  while (nodes.length > 0) {
    const nextNodes = [];
    for (const node of nodes) {
      if (node.innerText.match(regex)) {
        return node;
      }
      for (const child of node.children) {
        nextNodes.push(child);
      }
    }
    nodes = nextNodes;
  }
  return null;
}

globalThis._findFirstStatFor = (targetStat) => {
  const statRegEx = {
    'hunger': /.*hunger.*/i,
    'energy': /.*energy.*/i,
    'happiness': /.*happiness.*/i,
  }

  for (const statEl of _getStatEls()) {
    const statTextEl = _deepMatchText(statEl, statRegEx[targetStat.toLowerCase()]);
    if (statTextEl) return statEl;
  }
}

globalThis._getStatEls = () => {
  return document.querySelectorAll('.stat');
}

// Real time wait for fn to evaluate to true
globalThis._waitFor = async (fn, { timeout = 100, interval = 50} = {}) => {
  const timeoutStamp = Date.now() + timeout;
  let err;

  while (true) {
    try {
      const res = fn();
      if (res) return res;
    } catch (e) {
      err = e;
    }
    if (Date.now() > timeoutStamp) break;
    _clock.tick(interval)
  }
  throw err || new Error('waitFor timeout');
}

globalThis._submitForm = async () => {
  const input = document.querySelector('input');
  input.value = _PET_NAME;
  input.dispatchEvent(new Event('change'));
  _clock.tick(50);

  const button = document.querySelector('button');
  button.click();
  _clock.tick(50);

  await _waitFor(() => !document.querySelector('form'));
}


globalThis._testEatButton = () => {
  const hungerStatEl = _findFirstStatFor('hunger');
  const hungerStatValueEl = hungerStatEl.querySelector('.stat-value');
  const originalHungerValue = parseFloat(hungerStatValueEl.textContent);

  const energyStatEl = _findFirstStatFor('energy');
  const energyStatValueEl = energyStatEl.querySelector('.stat-value');
  const originalEnergyValue = parseFloat(energyStatValueEl.textContent);

  const eatButton = document.querySelector('button#eat-action');
  eatButton.click();
  _clock.tick(50);

  const newHungerValue = parseFloat(hungerStatValueEl.textContent);

  if (originalHungerValue !== 0) {
    assert.isBelow(newHungerValue, originalHungerValue);
  } else assert.equal(newHungerValue, originalHungerValue);

  const newEnergyValue = parseFloat(energyStatValueEl.textContent);

  if (originalEnergyValue !== 100) {
    assert.isAbove(newEnergyValue, originalEnergyValue);
  } else assert.equal(newEnergyValue, originalEnergyValue);
}

globalThis._testPlayButton = () => {
  const energyStatEl = _findFirstStatFor('energy');
  const energyStatValueEl = energyStatEl.querySelector('.stat-value');
  const originalEnergyValue = parseFloat(energyStatValueEl.textContent);

  const happinessStatEl = _findFirstStatFor('happiness');
  const happinessStatValueEl = happinessStatEl.querySelector('.stat-value');
  const originalHappinessValue = parseFloat(happinessStatValueEl.textContent);

  const eatButton = document.querySelector('button#play-action');
  eatButton.click();
  _clock.tick(50);

  const newEnergyValue = parseFloat(energyStatValueEl.textContent);

  if (originalEnergyValue !== 0) {
    assert.isBelow(newEnergyValue, originalEnergyValue);
  } else assert.equal(newEnergyValue, originalEnergyValue);

  const newHappinessValue = parseFloat(happinessStatValueEl.textContent);

  if (originalHappinessValue !== 100) {
    assert.isAbove(newHappinessValue, originalHappinessValue);
  } else assert.equal(newHappinessValue, originalHappinessValue);
}

globalThis._testSleepButton = () => {
  const hungerStatEl = _findFirstStatFor('hunger');
  const hungerStatValueEl = hungerStatEl.querySelector('.stat-value');
  const originalHungerValue = parseFloat(hungerStatValueEl.textContent);

  const energyStatEl = _findFirstStatFor('energy');
  const energyStatValueEl = energyStatEl.querySelector('.stat-value');
  const originalEnergyValue = parseFloat(energyStatValueEl.textContent);

  const sleepButton = document.querySelector('button#sleep-action');
  sleepButton.click();
  _clock.tick(50);

  const newHungerValue = parseFloat(hungerStatValueEl.textContent);

  if (originalHungerValue !== 100) {
    assert.isAbove(newHungerValue, originalHungerValue);
  } else assert.equal(newHungerValue, originalHungerValue);

  const newEnergyValue = parseFloat(energyStatValueEl.textContent);

  if (originalEnergyValue !== 100) {
    assert.isAbove(newEnergyValue, originalEnergyValue);
  } else assert.equal(newEnergyValue, originalEnergyValue);
}
```

# --after-all--

```js
delete globalThis._PET_NAME;
globalThis._clock.uninstall();
delete globalThis._clock;

delete globalThis._deepMatchText;
delete globalThis._findFirstStatFor;
delete globalThis._getStatEls;
delete globalThis._waitFor;
delete globalThis._submitForm;
delete globalThis._testEatButton;
delete globalThis._testPlayButton;
delete globalThis._testSleepButton;
```

# --hints--

When the page loads, there should be a visible `form` element.

```js
assert.isNotNull(document.querySelector('form'), 'The form element does not exist at load.');
```

This form should contain an input field with an id of `pet-name`.

```js
assert.isNotNull(document.querySelector('#pet-name'));
```

This form should contain a `button`.

```js
assert.isNotNull(document.querySelector('button'));
```

When the input is given a value and the button is pressed, the form should no longer be visible.

```js
await _submitForm();

assert.isNull(document.querySelector('form'));
```

The second view should contain an element with class `pet-name`.

```js
const petNameEl = document.querySelector('.pet-name');
assert.exists(petNameEl);
```

The second view should contain the provided pet name.

```js
const petNameEl = document.querySelector('.pet-name');
assert.strictEqual(petNameEl.textContent, _PET_NAME);
```

The second view should contain three `button`s.

```js
const buttons = document.querySelectorAll('button');
assert.equal(buttons.length, 3);
```

Exactly one button should contain the id `#eat-action`.

```js
const eatButton = document.querySelectorAll('button#eat-action');
assert.equal(eatButton.length, 1);
```

Exactly one button should contain the id `#play-action`.

```js
const playButton = document.querySelectorAll('button#play-action');
assert.equal(playButton.length, 1);
```

Exactly one button should contain the id `#sleep-action`.

```js
const sleepButton = document.querySelectorAll('button#sleep-action');
assert.equal(sleepButton.length, 1);
```

There should be exactly three elements with the class `.stat`.

```js
const statEls = document.querySelectorAll('.stat');
assert.equal(statEls.length, 3);
```

Each pet stat should contain exactly one element with class `.stat-value`.

```js
const statEls = document.querySelectorAll('.stat');
for (const statEl of statEls) {
  assert.equal(statEl.querySelectorAll('.stat-value').length, 1);
}
```

Each `.stat-value` element should contain one number.

```js
const statValueEls = document.querySelectorAll('stat-value');
for (const statEl of statValueEls) {
  assert.match(statEl.textContent, /-?\d+\.*\d*/);
}
```

Each `.stat-value` number should be within the range `[0,100]`, inclusive.

```js
const statValueEls = document.querySelectorAll('stat-value');
for (const statEl of statValueEls) {
  const statValue = parseFloat(statEl.textContent);
  assert.isAtMost(statValue, 100);
  assert.isAtLeast(statValue, 0);
}
```

Exactly one `.stat` element should contain the text `Hunger`, case insensitive.

```js
const statEls = document.querySelectorAll('.stat');
let matches = 0;
for (const statEl of statEls) {
  if (_deepMatchText(statEl, /.*hunger.*/i)) matches++;
}
assert.equal(matches, 1);
```

One `.stat` element should contain the text `Energy`, case insensitive.

```js
const statEls = document.querySelectorAll('.stat');
let matches = 0;
for (const statEl of statEls) {
  if (_deepMatchText(statEl, /.*energy.*/i)) matches++;
}
assert.equal(matches, 1);
```

One `.stat` element should contain the text `Happiness`, case insensitive.

```js
const statEls = document.querySelectorAll('.stat');
let matches = 0;
for (const statEl of statEls) {
  if (_deepMatchText(statEl, /.*happiness.*/i)) matches++;
}
assert.equal(matches, 1);
```

The `Hunger` stat should start at 0.

```js
const hungerStatEl = _findFirstStatFor('hunger');
const hungerStatValueEl = hungerStatEl.querySelector('.stat-value');
assert.match(hungerStatValueEl.textContent, /-?\d+\.*\d*/);

const match = hungerStatValueEl.textContent.match(/-?\d+\.*\d*/)[0];
assert.equal(parseFloat(match), 0);
```

The `Energy` stat should start at 100.

```js
const energyStatEl = _findFirstStatFor('energy');
const energyStatValueEl = energyStatEl.querySelector('.stat-value');
assert.match(energyStatValueEl.textContent, /-?\d+\.*\d*/);

const match = energyStatValueEl.textContent.match(/-?\d+\.*\d*/)[0];
assert.equal(parseFloat(match), 100);
```

The `Happiness` stat should start at 100.

```js
const happinessStatEl = _findFirstStatFor('happiness');
const happinessStatValueEl = happinessStatEl.querySelector('.stat-value');
assert.match(happinessStatValueEl.textContent, /-?\d+\.*\d*/);

const match = happinessStatValueEl.textContent.match(/-?\d+\.*\d*/)[0];
assert.equal(parseFloat(match), 100);
```

Clicking the `Eat` button should reduce the `Hunger` value and increase the `Energy` value. If `Hunger` is 0 or `Energy` is 100, that value must not change.

```js
_testEatButton();
```

Clicking the `Play` button should reduce the `Energy` value and increase the `Happiness` value. If `Energy` is 0 or `Happiness` is 100, that value must not change.

```js
_testPlayButton();
```

Clicking the `Sleep` button should increase the `Hunger` value and increase the `Energy` value. If `Hunger` is 100 or `Energy` is 100, that value must not change.

```js
_testSleepButton();
```


The `Hunger` stat should be 100 after an excessive period of idle time.

```js
const hungerStatEl = _findFirstStatFor('hunger');
const hungerStatValueEl = hungerStatEl.querySelector('.stat-value');

// This will advance the time for subsequent tests as well.
for (let i = 0; i < 100; i++) {
  await _clock.tickAsync(10 * 1_1000); 
}

const hungerValue = parseFloat(hungerStatValueEl.textContent);
assert.equal(hungerValue, 100);
```

The `Energy` stat should be 100 after an excessive period of idle time.

```js
const energyStatEl = _findFirstStatFor('energy');
const energyStatValueEl = energyStatEl.querySelector('.stat-value');

const energyValue = parseFloat(energyStatValueEl.textContent);
assert.equal(energyValue, 100);
```

The `Happiness` stat should be 0 after an excessive period of idle time.

```js
const happinessStatEl = _findFirstStatFor('happiness');
const happinessStatValueEl = happinessStatEl.querySelector('.stat-value');

const happinessValue = parseFloat(happinessStatValueEl.textContent);
assert.equal(happinessValue, 0);
```

Clicking the `Eat` button should reduce the `Hunger` value and increase the `Energy` value. If `Hunger` is 0 or `Energy` is 100, that value must not change.

```js
// These tests are simply repeated to test with the decayed values as well.
_testEatButton();
```

Clicking the `Play` button should reduce the `Energy` value and increase the `Happiness` value. If `Energy` is 0 or `Happiness` is 100, that value must not change.

```js
_testPlayButton();
```

Clicking the `Sleep` button should increase the `Hunger` value and increase the `Energy` value. If `Hunger` is 100 or `Energy` is 100, that value must not change.

```js
_testSleepButton();
```

There should be a `PetMood` enumerator.

```js
assert.match(code, /\s*enum PetMood\s*/);
```

The `PetMood` enumerator should contain `HAPPY`.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HAPPY[\s*\w+,\s*]*}/);
```

The `PetMood` enumerator should contain `EXCITED`.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*EXCITED[\s*\w+,\s*]*}/);
```

The `PetMood` enumerator should contain `CONTENT`.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*CONTENT[\s*\w+,\s*]*}/);
```

The `PetMood` enumerator should contain `SAD`'.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SAD[\s*\w+,\s*]*}/);
```

The `PetMood` enumerator should contain `TIRED`.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*TIRED[\s*\w+,\s*]*}/);
```

The `PetMood` enumerator should contain `SICK`.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*SICK[\s*\w+,\s*]*}/);
```

The `PetMood` enumerator should contain `HUNGRY`.

```js
assert.match(code, /\s*enum\s+PetMood\s*{[\s*\w+,\s*]*HUNGRY[\s*\w+,\s*]*}/);
```

There should be a `Record<PetMood, string` to map a value to each `PetMood`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>/);
```

There should be a value mapped to `PetMood.HAPPY`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.HAPPY\]/);
```

There should be a value mapped to `PetMood.EXCITED`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.EXCITED\]/);

```

There should be a value mapped to `PetMood.CONTENT`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.CONTENT\]/);
```

There should be a value mapped to `PetMood.SAD`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.SAD\]/);
```

There should be a value mapped to `PetMood.TIRED`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.TIRED\]/);
```

There should be a value mapped to `PetMood.SICK`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.SICK\]/);
```

There should be a value mapped to `PetMood.HUNGRY`.

```js
assert.match(code, /\s*(?:const|let|var)\s+\w+:\s+Record<PetMood,\s*string>\s*=\s*{\s*(?:\[PetMood.\w+]:\s\\*".+"\\*,\s*(?:\/\/.*]*)*\s*)*\[PetMood\.HUNGRY\]/);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">

        <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState } = React;

export const PetGame = () => {};
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="utf-8"/>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Digital Pet Game</title>

    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Tektur:wght@400..900&display=swap" rel="stylesheet">
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.tsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { PetGame } from './index.tsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<PetGame />);
    </script>
  </body>
</html>
```

```css
:root {
  /* Backgrounds */
  --bg-gradient-start: #dbeafe;
  --bg-gradient-mid: #faf5ff;
  --bg-gradient-end: #fce7f3;

  /* Text */
  --text-primary: #1f2937;
  --text-muted: #6b7280;
  --text-muted-2: #9ca3af;

  /* Card / surface */
  --card-bg-translucent: rgba(255, 255, 255, 0.6);
  --card-bg-solid: rgba(255, 255, 255, 0.8);

  /* Accent / brand */
  --accent-purple: #8b5cf6;
  --accent-purple-dark: #7c3aed;

  /* Actions */
  --action-feed: #10b981;
  --action-feed-dark: #059669;
  --action-play: #3b82f6;
  --action-play-dark: #2563eb;
  --action-sleep: #8b5cf6;
  --action-sleep-dark: #7c3aed;
  --action-heal: #ec4899;
  --action-heal-dark: #db2777;

  /* Stats */
  --stat-high: #10b981;
  --stat-medium: #f59e0b;
  --stat-low: #ef4444;

  /* UI greys */
  --ui-bg-light: #e5e7eb;
  --ui-fill: #8b5cf6; /* used for exp fill as accent */

  /* Shadows */
  --shadow-strong: rgba(0, 0, 0, 0.25);
  --shadow-soft: rgba(0, 0, 0, 0.1);
}

html {
  /*color-scheme: light dark;*/
  font-family: system-ui;
  display: flex;
  justify-content: center;
  font-family:
    -apple-system, BlinkMacSystemFont, "Segoe UI", "Roboto", "Oxygen", "Ubuntu",
    "Cantarell", "Fira Sans", "Droid Sans", "Helvetica Neue", sans-serif;
  -webkit-font-smoothing: antialiased;
  -moz-osx-font-smoothing: grayscale;
  background: linear-gradient(135deg, #dbeafe 0%, #faf5ff 50%, #fce7f3 100%);
  min-height: 100vh;
  padding: 1rem;
}

header {
  text-align: center;
  margin-bottom: 2rem;
}

body {
  max-width: 400px;
}

.base-container {
  background: rgba(255, 255, 255, 0.6);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 1.5rem;
  padding: 2rem;
  box-shadow: 0 25px 50px -12px rgba(0, 0, 0, 0.25);
  .pet-name {
    text-align: center;
  }
}

.game-container {
  margin-bottom: 1.5rem;
  display: flex;
  flex-direction: column;
  justify-content: center;
  align-items: center;
}

/* Stats Grid */
.stats-grid {
  display: grid;
  grid-template-columns: repeat(2, 1fr);
  gap: 1rem;
  margin-bottom: 1.5rem;
}

.stat-bar {
  background: rgba(255, 255, 255, 0.8);
  -webkit-backdrop-filter: blur(10px);
  backdrop-filter: blur(10px);
  border-radius: 0.75rem;
  padding: 1rem;
  box-shadow: 0 1px 3px 0 rgba(0, 0, 0, 0.1);
  transition: all 0.2s ease;
}

.stat-bar:hover {
  box-shadow: 0 4px 6px -1px rgba(0, 0, 0, 0.1);
}

.stat-header {
  display: flex;
  align-items: center;
  justify-content: space-between;
  margin-bottom: 0.5rem;
}

.stat-label {
  display: flex;
  align-items: center;
  gap: 0.5rem;
}

.stat-icon {
  font-size: 1.25rem;
}

.stat-name {
  font-weight: 500;
  color: #374151;
}

.stat-value {
  font-weight: bold;
  color: #1f2937;
}

.stat-progress {
  width: 100%;
  height: 0.75rem;
  background: #e5e7eb;
  border-radius: 9999px;
  overflow: hidden;
}

.stat-fill {
  height: 100%;
  transition: width 0.5s ease-out;
  border-radius: 9999px;
}

.stat-fill.high {
  background: #10b981;
}

.stat-fill.medium {
  background: #f59e0b;
}

.stat-fill.low {
  background: #ef4444;
}

.pet-buttons {
  width: auto;
  display: flex;
  flex-direction: row;
  gap: 1rem;
}

.pet-buttons .pet-button {
  min-width: 80px;
  background: #c2214a;
  color: white;
  border-radius: 13%;
  border-style: solid;
  border-color: #870b04;
  box-shadow: rgba(120, 25, 23, 0.2) 2px 2px 4px 2px;
  &:active {
    transform: scale(0.95);
    box-shadow: 0 2px 4px rgba(0, 0, 0, 0.15) inset;
  }
  text-align: center;
  cursor: pointer;
  padding: 0.5rem 1rem;
}

.pet-sprite {
  font-size: xxx-large;
  text-align: center;
}

.data-management * {
  margin: 0 0.5rem 0 0;
}

.info-panel {
  text-align: center;
}
.info-panel #hud,
.info-panel .start-questions {
  display: flex;
  align-items: center;
  flex-direction: column;
  gap: 1rem;
}

.info-panel #hud p {
  margin: 0rem;
}

.info-panel .start-questions div * {
  margin: 0 0.5rem 0 0;
}

#pet-fact {
  color: var(--text-muted-2);
  margin: 0 0 0.5rem 0;
}
#set-name-btn {
  display: inline-block;
  padding: 14px 28px;
  font-size: 1.15rem;
  font-weight: 600;
  background: linear-gradient(135deg, #6b8cff, #5ec2ff);
  color: #fff;
  border: none;
  border-radius: 10px;
  box-shadow: 0 6px 18px rgba(92, 132, 255, 0.25);
  cursor: pointer;
  transition:
    transform 0.08s ease,
    box-shadow 0.12s ease,
    opacity 0.12s;
}

#set-name-btn:hover {
  opacity: 0.98;
}

#set-name-btn:active,
#set-name-btn.pressed {
  transform: translateY(2px) scale(0.995);
  box-shadow: 0 3px 8px rgba(92, 132, 255, 0.18);
}

#set-name-btn:focus {
  outline: 3px solid rgba(94, 190, 255, 0.28);
  outline-offset: 2px;
}
```

```tsx
const { useState, useEffect } = React;

interface StatBarProps {
  label: string;
  value: number;
  icon: string;
  reverse?: boolean;
}

function getStatColorClass(value: number): string {
  if (value >= 70) return "high";
  if (value >= 40) return "medium";
  return "low";
}

function StatBar({ label, value, icon, reverse = false }: StatBarProps) {
  const colorClass = getStatColorClass(!reverse ? value : 100 - value);


  return (
    <div className="stat-bar stat">
      <div className="stat-header">
        <div className="stat-label">
          <span className="stat-icon">{icon}</span>
          <span className="stat-name">{label}</span>
        </div>
        <span className="stat-value">{Math.round(value)}%</span>
      </div>

      <div className="stat-progress">
        <div
          className={`stat-fill ${colorClass}`}
          style={{ width: `${value}%` }}
        />
      </div>
    </div>
  );
}

enum PetAction {
  NONE,
  EAT,
  PLAY,
  SLEEP,
}

interface Pet {
  name: string;
  happiness: number;
  hunger: number;
  energy: number;
  species: string;
}

enum PetMood {
  HAPPY,
  EXCITED,
  CONTENT,
  SAD,
  TIRED,
  SICK,
  HUNGRY,
}

export const STAT_DECAY_RATES = {
  hunger: 10,
  happiness: 5,
  energy: 5,
};

const UPDATE_INTERVAL = 30 * 1000; // 30 seconds

export function calculatePetMood(pt: Pet): PetMood {
  const { hunger, happiness, energy } = pt;

  if (hunger > 70) return PetMood.HUNGRY;
  if (energy < 30) return PetMood.TIRED;
  if (happiness < 30) return PetMood.SAD;
  if (happiness > 80 && energy > 70) return PetMood.EXCITED;
  if (happiness > 60) return PetMood.HAPPY;

  return PetMood.CONTENT;
}

const moodEmojiMap: Record<PetMood, string> = {
  [PetMood.HAPPY]: "ðŸ˜º", // grinning cat
  [PetMood.EXCITED]: "ðŸ˜»", // heart-eyes cat
  [PetMood.CONTENT]: "ðŸ˜¸", // smiling cat
  [PetMood.SAD]: "ðŸ˜¿", // crying cat
  [PetMood.TIRED]: "ðŸ˜½", // kissing cat (closest to sleepy)
  [PetMood.SICK]: "ðŸ™€", // weary cat (looks unwell)
  [PetMood.HUNGRY]: "ðŸ˜¹", // joy cat (closest to hungry/funny)
};

interface usePetProps {
  isGameStarted: boolean;
  setGameStarted: React.Dispatch<React.SetStateAction<boolean>>;
}
// hook where all pet logic is handled
function usePet({ isGameStarted }: usePetProps) {
  let [pet, setPet] = useState<Pet>({
    name: "",
    happiness: 100,
    hunger: 0,
    energy: 100,
    species: "Cat",
  });
  // Auto-decay stats over time
  useEffect(() => {
    if (!isGameStarted) return;

    const interval = setInterval(() => {
      setPet((pet) => ({
        ...pet,
        happiness: Math.max(pet.happiness - 5, 0),
        hunger: Math.min(pet.hunger + 5, 100),
      }));
    }, UPDATE_INTERVAL);

    return () => clearInterval(interval);
  }, [isGameStarted]);

  function doAction(action: PetAction) {
    switch (action) {
      case PetAction.EAT:
        feedPet();
        break;
      case PetAction.SLEEP:
        restPet();
        break;
      case PetAction.PLAY:
        playWithPet();
        break;
    }
  }

  function feedPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.max(pet.hunger - STAT_DECAY_RATES.hunger, 0),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  function playWithPet() {
    setPet((pet) => ({
      ...pet,
      energy: Math.max(pet.energy - STAT_DECAY_RATES.energy, 0),
      happiness: Math.min(pet.happiness + STAT_DECAY_RATES.happiness, 100),
    }));
  }

  function restPet() {
    setPet((pet) => ({
      ...pet,
      hunger: Math.min(pet.hunger + STAT_DECAY_RATES.hunger, 100),
      energy: Math.min(pet.energy + STAT_DECAY_RATES.energy, 100),
    }));
  }

  return {
    pet,
    doAction,
    setName: (name: string) => setPet({ ...pet, name }),
  };
}

export function PetGame() {
  const [isGameStarted, setGameStarted] = useState(false);

  let { pet, doAction, setName } = usePet({
    isGameStarted,
    setGameStarted,
  });

  const [fact, setFact] = useState<string | null>("");


  useEffect(() => {
    fetch("https://cat-facts-api.freecodecamp.rocks/api/catfacts/random").then(
      res => res.json()
    ).then(
      fact => setFact(fact)
    ).catch(err => {
      setFact("Sorry, we're not able to retrieve your cat fact right now!");
    });
  }, []);

  function startGame() {
    const petName = (document.getElementById("pet-name") as HTMLInputElement)
      .value;
    if (!petName) return; //do not let user to start game if no name selected
    setName(petName);
    setGameStarted(true);
  }

  return (
    <main>
      <header>
        <h1>Digital Pet Game</h1>
        <p>Take care of your virtual companion!</p>
      </header>

      {isGameStarted && (
        <section className="base-container game-container">
          <div className="pet-screen">
            <div className="pet-sprite">
              {moodEmojiMap[calculatePetMood(pet)]}
            </div>
            <h2 className="pet-name">{pet.name}</h2>
          </div>
          <div className="pet-buttons">
            <button
              onClick={() => doAction(PetAction.EAT)}
              className="pet-button pet-buttons-left"
              id="eat-action"
            >
              EAT
            </button>
            <button
              onClick={() => doAction(PetAction.PLAY)}
              className="pet-button pet-buttons-center"
              id="play-action"
            >
              PLAY
            </button>
            <button
              onClick={() => doAction(PetAction.SLEEP)}
              className="pet-button pet-buttons-right"
              id="sleep-action"
            >
              SLEEP
            </button>
          </div>
        </section>
      )}

      {isGameStarted && (
        <section className="stats-grid">
          <StatBar label="Hunger" value={pet.hunger} icon="ðŸ½ï¸" reverse />
          <StatBar label="Happiness" value={pet.happiness} icon="ðŸ˜Š" />
          <StatBar label="Energy" value={pet.energy} icon="âš¡" />
        </section>
      )}

      <section className="base-container info-panel">
        {!isGameStarted ? (
          <form className="start-questions">
            <label htmlFor="pet-name">What is your pet's name?</label>
            <input
              id="pet-name"
              name="pet-name"
              required={true}
              pattern="[A-Za-z0-9]{1,20}"
            />

            <button id="set-name-btn" onClick={startGame}>
              Start Game
            </button>
          </form>
        ) : (
          <div id="hud">
            <p id="pet-species">Species: {pet.species}</p>
            <p id="pet-fact">
              <b>Pet Fact:</b> {fact}
            </p>
          </div>
        )}
      </section>
    </main>
  );
}
```

