---
id: 657effaa2a5e0277d71f9cbe
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

By this point, the `snake_cased_char_list` variable contains all the converted characters. To convert this list of characters into a string, you'll use the `.join()` method. For example:

```py
result_string = ''.join(characters)
```

Right after the `snake_cased_char_list` variable, add a new variable named `snake_cased_string`. Assign it an empty string, then apply the `.join()` method with `snake_cased_char_list` passed as an argument.

# --hints--

You should join the characters in the `snake_cased_char_list` list into a string using the `.join()` method and assign the string to a variable named `snake_cased_string` after your `for` loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
