---
id: 657e928716b77b2277980276
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you are going to learn about list comprehensions in Python by building a program that can take a `camelCase` or `PascalCase` formatted string and convert that to a `snake_case` formatted string. 

In Python, a <dfn>List comprehension</dfn> is a construct that allows you to generate a new list by applying an expression to each item in an existing iterable and optionally filtering items with a condition. Apart from being briefer, list comprehensions often run faster.

The basic syntax of list comprehension is:

```py
[expression for item in iterable if condition]
```

You'll work on the implementation in the coming steps. 

Start defining a new function named `convert_to_snake_case()` that accepts a string named `pascal_or_camel_cased_string` as input. For now, add a `pass` statement inside the function.

# --hints--

You should define a new function named `convert_to_snake_case()`. Don't forget the colon at the end and use `pass` to fill the function body.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Your function should take a parameter named `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
