---
id: 657f04ed0035f47ed04d0f1f
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

In order to display the output of the `convert_to_snake_case()` function, you need to call the `main()` function. 

At the same level as the two existing functions, add a call to the `main()` function. You should see the given camel or pascal cased string converted to snake case upon execution.

# --hints--

You should add a call to the `main()` function.

```js
({
  test: () => {
    const pythonCode = `_Node(_code.replace('\\\\r', '')).has_call("main()")`;
    const result = runPython(pythonCode);
    assert(result);
  }
});
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
