---
id: 657f28a0482132aca51a9212
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Replace `pass` with a call to the `main()` function inside the body of the `if` statement. You should see the given camel or pascal cased string converted to snake case upon execution.

# --hints--

You should add a call to `main()` inside the `if` statement.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        assert.match(transformedCode, /\n +main\s*\(\s*\)/);
    }
})
```

You should not have `pass` in your `if` statement.

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+__name__\s+==\s+'__main__'\s*/);
        
        assert.notMatch(block_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
if __name__ == '__main__':
    pass
--fcc-editable-region--
```
