---
id: 66cc12fa504b0479dac479a0
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

The great thing about functions is that you can make multiple function calls with different arguments and the function should return the correct value each time. 

Remember that an argument is the value that you pass into the function when you call it.

Start by adding a `console.log` that calls the `calculateSum` function with the arguments `2` and `5`.

Then, add a second `console.log` that calls the `calculateSum` function with the arguments `10` and `10`.

Finally, add a third `console.log` that calls the `calculateSum` function with the arguments `5` and `5`.

# --hints--

You should have a `console.log` that calls the `calculateSum` function with the arguments `2` and `5`.

```js
assert.match(code, /console.log\(calculateSum\(2\s*,\s*5\)\)/);
```

You should have a `console.log` that calls the `calculateSum` function with the arguments `10` and `10`.

```js
assert.match(code, /console.log\(calculateSum\(10\s*,\s*10\)\)/);
```

You should have a `console.log` that calls the `calculateSum` function with the arguments `5` and `5`.

```js
assert.match(code, /console.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
