---
id: 68de7360e1832ea180577a8a
title: Build an Nth Fibonacci Number Calculator
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should define a list named `sequence` in the global scope, and it should be initialized with the values `[0, 1]`.  
2. You should create a function named `fibonacci`.  
3. The `fibonacci` function should accept one parameter, a positive integer `n`.  
4. Calling `fibonacci(n)` should use a dynamic programming approach to compute and return the `n`-th number from the Fibonacci sequence, where each number is the sum of the two preceding numbers.
5. Each computed number at the position `n` in the Fibonacci sequence should be stored in the `sequence` list at index `n - 1`.

# --hints--

You should have a list named `sequence` in the global scope initialized to `[0, 1]`.

```js
({ test: () => runPython(`assert isinstance(sequence, list) and sequence == [0, 1]`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]

# Test cases
print(f"fibonacci(0) = {fibonacci(0)}")   # Expected: 0
print(f"fibonacci(1) = {fibonacci(1)}")   # Expected: 1
print(f"fibonacci(2) = {fibonacci(2)}")   # Expected: 1
print(f"fibonacci(3) = {fibonacci(3)}")   # Expected: 2
print(f"fibonacci(5) = {fibonacci(5)}")   # Expected: 5
print(f"fibonacci(10) = {fibonacci(10)}") # Expected: 55
print(f"fibonacci(15) = {fibonacci(15)}") # Expected: 610
```
