---
id: 66bb3e20d3dc5b6d0a21f5dd
title: Build a Real Time Counter
challengeType: 14
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. **Do not copy this demo project**.

**User Stories:**

1. You should have an `h1` element with the text `Real-Time Character Counter`.
2. You should have a `textarea` element with the `id` of `text-input`.
3. There should be a `p` element with the `id` of `char-count`. It should initially contain the text `Character Count: 0/50`. This placeholder text should be replaced with the current count of characters in the `textarea` element.
4. You should attach an event listener to the `text-input` element that listens for the `input` event. 
5. Inside the event trigger function, you should get the value of the `textarea` element and store it in the variable `textInput`. You should also declare a variable `charCount` and set it equal to the length of the `textInput` variable.
6. Create a condition that checks if the current character count is greater than `50`. If it is, the user shouldn't be able to enter more characters. 
7. Get the character count display element by its ID and update the text to show the current character count using `innerHTML`.
8. If the character count is equal to `50`, the text should be displayed in `red`. 

**Note:** Be sure to link your stylesheet and the JavaScript file in your HTML.

# --hints--

There should be an `h1` element with the text of your topic.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

You should have a `textarea` element with the `id` of `text-input`. 

```js
assert($('textarea').attr('id') == 'text-input');
```

You should have a `p` element with the `id` of `char-count`. 

```js
assert($('p').attr('id') == 'char-count');
```

The `#char-count` `p` element should initially contain the text `Character Count: 0/50`.

```js
assert(document.querySelector('p')?.innerText, 'Character Count: 0/50');
```

You should attach an event listener to the `text-input` element.

```js
assert.match(code, /document\.getElementById\(\"textInput\"\)\.addEventListener/);
```

Your event listener on `#text-input` should listen for the `input` event.

```js
assert.match(code, /document\.getElementById\(\"textInput\"\)\.addEventListener\(\"input\"/);
```

Inside the event trigger function, you should get the value of the `textarea` element using its id and store it in the variable `textInput`. 

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);/);
```

Inside the event trigger function, you should declare a variable `charCount` and set it equal to the length of the `textInput` variable.

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;/);
```

Inside the event trigger function, you should create a condition that checks if the current character count is greater than `50`. 

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;\s+if\s*\(charCount\s*>\s*50\)\s*\{/);
```

Inside the event trigger function, if the character count is greater than `50`, you should use `substring` operation to limit the number of characters to `50` so the user can't enter more characters. Also, assign the substring to the `textInput` variable.

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;\s+if\s*\(charCount\s*>\s*50\)\s*\{\s+textInput\.value\s*=\s*textInput\.value\.substring\(0,\s*50\);/);
```

If the character count is greater than `50`, the value of `charCount` should be set to `50`.

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;\s+if\s*\(charCount\s*>\s*50\)\s*\{\s+textInput\.value\s*=\s*textInput\.value\.substring\(0,\s*50\);\s+charCount\s*=\s*50;/);
```

You should select the character count display element by its id `char-count` and store it in the variable `charCountDisplay`.

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;\s+if\s*\(charCount\s*>\s*50\)\s*\{\s+textInput\.value\s*=\s*textInput\.value\.substring\(0,\s*50\);\s+charCount\s*=\s*50;\s*}\s+let\s+charCountDisplay\s*=\s*document\.getElementById\("charCount"\);/);
```

You should update the `#char-count` text to show the current character count using `innerHTML` and using the string template `Character Count: ${charCount}/50`.

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;\s+if\s*\(charCount\s*>\s*50\)\s*\{\s+textInput\.value\s*=\s*textInput\.value\.substring\(0,\s*50\);\s+charCount\s*=\s*50;\s*}\s+let\s+charCountDisplay\s*=\s*document\.getElementById\("charCount"\);\s+charCountDisplay\.innerHTML\s*=\s*`Character Count:\s*\${charCount}\/50`;/);
```

If the character count is equal to `50`, the text should be displayed in `red`.

```js
assert.match(code, /document\.getElementById\("textInput"\)\.addEventListener\("input",\s*function\s*\(\)\s*{\s+let\s+textInput\s*=\s*document\.getElementById\("textInput"\);\s+let\s+charCount\s*=\s*textInput\.value\.length;\s+if\s*\(charCount\s*>\s*50\)\s*\{\s+textInput\.value\s*=\s*textInput\.value\.substring\(0,\s*50\);\s+charCount\s*=\s*50;\s*}\s+let\s+charCountDisplay\s*=\s*document\.getElementById\("charCount"\);\s+charCountDisplay\.innerHTML\s*=\s*`Character Count:\s*\${charCount}\/50`;\s+if\s*\(charCount\s*===\s*50\)\s*\{\s+charCountDisplay\.style\.color\s*=\s*"red";\s*}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("textInput").addEventListener("input", function () {
  let textInput = document.getElementById("textInput");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
