---
id: 683ecb010f4c6bafd428802c
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you'll validate a set of medical data to ensure that it complies with a set of rules.

Declare a variable named `medical_records` and assign it an empty list. In the following steps, you are going to use it to store your medical data.

# --hints--

You should declare a variable named `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

You should assign an empty list to your `medical_records` variable.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

