---
id: 6842aee5bcaf4ab8883f8191
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

As you learned in the previous lectures, a dictionary is a data structure that holds key-value pairs, where keys are immutable data types (such as strings) and unique within a dictionary:

```py
person = {
    'name': 'John',
    'age': 33
}
```

The `medical_records` list will store dictionaries, each of them representing a patient. Add a dictionary with a key `patient_id` and value of the string `P1001` to the `medical_records` list.

# --hints--

Your `medical_records` list should contain a single dictionary.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

You should have a `patient_id` key with the value of the string `P1001` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [
    
]
--fcc-editable-region--
```
