---
id: 6821ebdf237de8297eaee793
title: "JavaScript Challenge 17"
challengeType: 28
dashedName: javascript-challenge-17
---

# --description--

Description

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` should return `This is a secret message.`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {

  return message;
}
```
