---
id: 681cb1b0dab50c87ddb2e519
title: JavaScript Coding Challenge 8
challengeType: 28
dashedName: javascript-coding-challenge-8
---

# --description--

JavaScript Coding Challenge 8

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
