---
id: 681cb1b0dab50c87ddb2e51a
title: JavaScript Challenge 9
challengeType: 28
dashedName: javascript-challenge-9
---

# --description--

JavaScript Challenge 9

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
