---
id: 681cb1afdab50c87ddb2e515
title: JavaScript Coding Challenge 4
challengeType: 28
dashedName: javascript-coding-challenge-4
---

# --description--

JavaScript Coding Challenge 4

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
