---
id: 681cb1a2dab50c87ddb2e514
title: JavaScript Coding Challenge 3
challengeType: 28
dashedName: javascript-coding-challenge-3
---

# --description--

JavaScript Coding Challenge 3

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
