---
id: 681cb1b0dab50c87ddb2e518
title: JavaScript Coding Challenge 7
challengeType: 28
dashedName: javascript-coding-challenge-7
---

# --description--

JavaScript Coding Challenge 7

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
