---
id: 681cb1afdab50c87ddb2e516
title: JavaScript Coding Challenge 5
challengeType: 28
dashedName: javascript-coding-challenge-5
---

# --description--

JavaScript Coding Challenge 5

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
