---
id: 681cb1afdab50c87ddb2e517
title: JavaScript Challenge 6
challengeType: 28
dashedName: javascript-challenge-6
---

# --description--

JavaScript Challenge 6

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
