---
id: 6821ebf3237de8297eaee797
title: "JavaScript Challenge 21"
challengeType: 28
dashedName: javascript-challenge-21
---

# --description--

Description

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` should return `This is a secret message.`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {

  return message;
}
```
