---
id: 681cb1b0dab50c87ddb2e51b
title: JavaScript Challenge 10
challengeType: 28
dashedName: javascript-challenge-10
---

# --description--

JavaScript Challenge 10

# --hints--

Passing test

```js
assert.isTrue(true);
```

# --seed--

## --seed-contents--

```js
function seed(str) {

  return str;
}
```

# --solutions--

```js
function seed(str) {

  return str;
}
```
