---
id: 69a56b5069ca99f7317e6e19
title: Build a Cargo Manifest Validator
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

In this lab, you will leverage JavaScript to normalize and validate cargo manifests. A cargo manifest is a document that typically lists goods being transported (for example, by ship or train) and includes details about those goods.

Each cargo manifest will be represented as an object with the following structure:

```js
{
  containerId: 0,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb", 
  hazmat: false
}
```

Note that each cargo manifest object should contain:

- `containerId`: a unique integer value identifying the cargo container
- `destination`: a string denoting the cargo's target destination
- `weight`: a numerical float value representing the cargo's weight
- `unit`: a string describing the units for the cargo's weight property (either "kg" for kilograms or "lb" for pounds).
- `hazmat`: a boolean value indicating whether hazardous material handling is needed

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should implement the function `normalizeUnits(manifest)`, which converts a manifest object's weight to kilograms without mutating the source object (return a new object). Use the approximate conversion, 1 lb = 0.45 kg.
2. You should implement the function `validateManifest(manifest)`, which accepts a manifest object, validates it by checking its fields (except `containerId`) for missing or invalid values, then returns an object based on the manifest's validity.
    - A manifest's fields are invalid if the following conditions are met: 
        - `destination` is not a string, or is a string containing only whitespace
        - `weight` is not a positive number
        - `unit` is neither "kg" nor "lb"
        - `hazmat` is not a boolean
    - If all fields are valid, return a new empty object: `{}`.
    - If any fields are missing or invalid, return a new object that lists only those fields, with string values showing the type of problem: `"Missing"` if the field is not present, or `"Invalid"` if the field is present but invalid (its value is not correct).
        
        For example, if the manifest is missing the `destination` field and has an invalid `weight` field, return:

        ```js
        {
            destination: "Missing",
            weight: "Invalid"
        }
        ```

3. You should implement the function `processManifest(manifest)`, which accepts a manifest object, and:
    - If the manifest is valid, logs `Validation success: ${containerId}` and then logs the manifest's weight in kilograms as such, `Total weight: ${weight} kg`.
    - If the manifest is not valid, logs: `Validation error: ${containerId}` and then directly logs the object which is returned by calling `validateManifest()` on the original manifest.
    Note: each of these two cases should have two `console.log()` calls.

# --before-each--

```js
const _validManifests = [
    {
        containerId: 4,
        destination: "Monterey, California, USA",
        weight: 831,
        unit: "lb",
        hazmat: false
    },
    {
        containerId: 5,
        destination: "Montreal, Quebec, Canada",
        weight: 151,
        unit: "kg",
        hazmat: false
    }
];

const _invalidManifests = [
    {
        containerId: 6,
        destination: 123,
        weight: 0,
        unit: "pounds",
        hazmat: true
    },
    {
        containerId: 7,
        destination: "Buenos Aires, Argentina",
        weight: -21,
        unit: "KG"
    },
    {
        containerId: 8,
        destination: "Guadalajara, Jalisco, Mexico",
        weight: 9001,
        hazmat: "no"
    },
    {
        containerId: 9,
        destination: " ",
        weight: 350,
        unit: "kg",
        hazmat: 1
    }
];
```

# --hints--

You should have a function named `normalizeUnits()` which accepts one parameter, `manifest`.

```js
assert.isFunction(normalizeUnits);
const funcString = normalizeUnits.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

Your `normalizeUnits()` function should return a copy of the input manifest object with its weight value converted to kilograms. Use the approximate conversion, 1 lb = 0.45 kg.

```js
_validManifests.forEach((obj, i) => {
    const result = normalizeUnits(obj);
    assert.isObject(result, "normalizeUnits should return an object");
    assert.notStrictEqual(result, obj, "normalizeUnits must return a new object, not modify the original");
    const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
    const expected = {
        ...obj,
        weight: expectedWeight,
        unit: "kg"
    };

    assert.deepEqual(
        result,
        expected,
        `normalizeUnits did not return the expected normalized object for manifest at index ${i}`
    );
});
```

Your `normalizeUnits()` function should return a new copy of the input `manifest` object without mutating the original.

```js
const original = {
    containerId: 151,
    destination: "Saffron City",
    weight: 151,
    unit: "lb",
    hazmat: false
};

const expected = {
    containerId: 151,
    destination: "Saffron City",
    weight: 67.95,
    unit: "kg",
    hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
```

You should have a function named `validateManifest()` which accepts one parameter, `manifest`.

```js
assert.isFunction(validateManifest);
const funcString = validateManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

Your function `validateManifest()` should return an empty object `{}` if the input manifest object is valid.

```js
const solution = {}
for (const obj of _validManifests) {
    assert.deepEqual(
        validateManifest(obj),
        solution,
        "validateManifest() should return an empty object for valid manifests."
    );
}
```

Your function `validateManifest()` should return an object describing missing and/or invalid fields if the input manifest object is not valid.

```js
const solutions = [
    { destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
    { weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
    { unit: 'Missing', hazmat: 'Invalid' },
    { destination: 'Invalid', hazmat: 'Invalid' }
];

let i = 0;
for (const obj of _invalidManifests) {
    const result = validateManifest(obj);
    assert.deepEqual(
        result,
        solutions[i],
        "validateManifest() should return an object with errors for invalid manifests."
    );
    i += 1;
}
```

Your `validateManifest()` function should return a new object without mutating the original.

```js
const invalidOriginal = {
    containerId: 151,
    destination: "",
    weight: -151,
    unit: "kilograms",
    hazmat: "idk"
};

const validOriginal = {
    containerId: 151,
    destination: "Pallet Town",
    weight: 151,
    unit: "kg",
    hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
```

You should have a function named `processManifest()` which accepts one parameter, `manifest`.

```js
assert.isFunction(processManifest);
const funcString = processManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

Your function `processManifest()` should, if the input manifest object is valid, log a success message with the object's `containerId`, and then log the object's `weight` in kilograms. This should be two `console.log()` calls.

```js
const spy = __helpers.spyOn(console, "log");

try {
    for (const obj of _validManifests) {
        processManifest(obj);
    }
    const expectedCalls = [
        ["Validation success: 4"],
        ["Total weight: 373.95 kg"],
        ["Validation success: 5"],
        ["Total weight: 151 kg"],
    ]
    assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
    assert.fail(err);
} finally {
    spy.restore();
}
```

Calling `processManifest()` on a valid manifest object should first log the success message, `Validation success: ${containerId}`. 

```js
const testObj = {
    containerId: 97,
    destination: "Cerulean City",
    weight: 200,
    unit: "lb",
    hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
    processManifest(testObj);
    const expectedCalls = [
        ["Validation success: 97"],
        ["Total weight: 90 kg"],
    ]
    assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
    assert.fail(err);
} finally {
    spy.restore();
}
```

Calling `processManifest()` on a valid manifest object should also log its weight in kilograms in this format: `Total weight: ${weight} kg`. You should use the `normalizeUnits()` function for this conversion before logging.

```js
const testObj = {
    containerId: 97,
    destination: "Cerulean City",
    weight: 200,
    unit: "lb",
    hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
    processManifest(testObj);
    const expectedCalls = [
        ["Validation success: 97"],
        ["Total weight: 90 kg"],
    ]
    assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
    assert.fail(err);
} finally {
    spy.restore();
}
```

Your function `processManifest()` should, if the input manifest object is not valid, log an error message with the object's `containerId`, and then log the object returned by calling `validateManifest()` on the input object.

```js
const testObj = {
    containerId: 106,
    destination: "Lavender Town",
    weight: -66,
}

const spy = __helpers.spyOn(console, "log");

try {
    processManifest(testObj);
    const expectedCalls = [
        ["Validation error: 106"],
        [{
            weight: "Invalid",
            unit: "Missing",
            hazmat: "Missing"
        }]
    ];
    assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
    assert.fail(err);   
} finally {
    spy.restore();
}
```

Calling `processManifest()` on an invalid manifest object should first log the error message, `Validation error: ${containerId}`. 

```js
const testObj = {
    containerId: 106,
    destination: "Lavender Town",
    weight: -66,
}

const spy = __helpers.spyOn(console, "log");

try {
    processManifest(testObj);
    const expectedCalls = [
        ["Validation error: 106"],
        [{
            weight: "Invalid",
            unit: "Missing",
            hazmat: "Missing"
        }]
    ];
    assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
    assert.fail(err);   
} finally {
    spy.restore();
}
```

Calling `processManifest()` on an invalid manifest object should also log the object returned by calling `validateManifest()` on the original manifest object. Call `console.log()` directly on the returned object.

```js
const testObj = {
    containerId: 106,
    destination: "Lavender Town",
    weight: -66,
}

const spy = __helpers.spyOn(console, "log");

try {
    processManifest(testObj);
    const expectedCalls = [
        ["Validation error: 106"],
        [{
            weight: "Invalid",
            unit: "Missing",
            hazmat: "Missing"
        }]
    ];
    assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
    assert.fail(err);   
} finally {
    spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
    const normalized = {...manifest}
    if (normalized.unit === "lb") {
        normalized.weight = normalized.weight * 0.45;
        normalized.unit = "kg";
    }
    return normalized;
}

function validateManifest(manifest) {
    let result = {}

    if (manifest.destination === undefined) {
        result.destination = "Missing";
    }
    else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
        result.destination = "Invalid";
    }
    if (manifest.weight === undefined) {
        result.weight = "Missing";
    }
    else if (typeof manifest.weight !== "number" || manifest.weight <= 0) {
        result.weight = "Invalid";
    }
    if (manifest.unit === undefined) {
        result.unit = "Missing";
    }
    else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
        result.unit = "Invalid";
    }
    if (manifest.hazmat === undefined) {
        result.hazmat = "Missing";
    }
    else if (typeof manifest.hazmat !== "boolean") {
        result.hazmat = "Invalid";
    }

    return result;
}

function processManifest(manifest) {
    const validation = validateManifest(manifest);

    if (Object.keys(validation).length > 0) {
        console.log(`Validation error: ${manifest.containerId}`);
        console.log(validation);
    } else {
        const normalized = normalizeUnits(manifest);
        console.log(`Validation success: ${manifest.containerId}`);
        console.log(`Total weight: ${normalized.weight} kg`);
    }

}
```
