---
id: 6140c7e645d8e905819f1dd4
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Почніть зі стандартного шаблонного коду. Додайте свою декларацію `DOCTYPE`, свій елемент `html` з мовою встановлену на англійську, свої елементи `head` та `body`.

Додайте елемент `meta` для правильного `charset`, елемент `title` та елемент `link` для файлу `./styles.css`.

Встановіть `title` на `Ferris Wheel`.

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Ви повинні визначити тип документа як `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Ви повинні закрити декларацію `DOCTYPE` з допомогою `>` після типу.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Ваш елемент `html` повинен мати початковий теґ з атрибутом `lang` зі значенням `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Ваш елемент `html` повинен мати кінцевий теґ.

```js
assert.match(code, /<\/html\s*>/);
```

Ваша декларація `DOCTYPE` повинна бути на початку HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Ви повинні мати початковий теґ `head`.

```js
assert.match(code, /<head\s*>/i);
```

Ви повинні мати кінцевий теґ `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Ви повинні мати початковий теґ `body`.

```js
assert.match(code, /<body\s*>/i);
```

Ви повинні мати кінцевий теґ `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Ваш код повинен мати елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Ваш елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш проєкт повинен мати титул `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Пам’ятайте, що вживання великої літери та написання мають значення.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Ваш код повинен мати елемент `link`.

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
