---
id: 6140d0069049f5139d78da40
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Властивість `transform-origin` використовують для встановлення точки, навколо якої застосовується CSS-перетворення. Наприклад, під час виконання `rotate` (що ви зробите пізніше в цьому проєкті), `transform-origin` визначає, навколо якої точки обертається елемент.

Надайте селектору `.line` властивість `transform-origin` зі значенням `0% 0%`. Це зміщуватиме початкову точку на `0%` зліва та на `0%` зверху, встановлюючи її в верхньому лівому куті елемента.

# --hints--

Ваш селектор `.line` повинен мати властивість `transform-origin` зі значенням `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
