---
id: 6140d1a351e88f159ed54fca
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Використовуючи той самий шаблон, створіть окремий селектор для третього `.line`, четвертого `.line`, п'ятого `.line` та шостого `.line`.

Встановіть властивість `transform` для третього `.line` на `rotate(120deg)`, для четвертого на `rotate(180deg)`, для п’ятого на `rotate(240deg)` та для шостого на `rotate(300deg)`.

# --hints--

Ви повинні створити селектор `.line:nth-of-type(3)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

Ваш селектор `.line:nth-of-type(3)` повинен мати властивість `transform` зі значенням `rotate(120deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform === 'rotate(120deg)');
```

Ви повинні створити селектор `.line:nth-of-type(4)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

Ваш селектор `.line:nth-of-type(4)` повинен мати властивість `transform` зі значенням `rotate(180deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform === 'rotate(180deg)');
```

Ви повинні створити селектор `.line:nth-of-type(5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

Ваш селектор `.line:nth-of-type(5)` повинен мати властивість `transform` зі значенням `rotate(240deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform === 'rotate(240deg)');
```

Ви повинні створити селектор `.line:nth-of-type(6)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

Ваш селектор `.line:nth-of-type(6)` повинен мати властивість `transform` зі значенням `rotate(300deg)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform === 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
