---
id: 6140c7e645d8e905819f1dd4
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Почніть зі стандартного шаблонного коду. Додайте свою декларацію `DOCTYPE`, свій елемент `html` з мовою встановлену на англійську, свої елементи `head` та `body`.

Додайте елемент `meta` для правильного `charset`, елемент `title` та елемент `link` для файлу `./styles.css`.

Встановіть `title` на `Ferris Wheel`.

# --hints--

Ваш код повинен містити посилання `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Ви повинні додати пробіл після посилання `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Ви повинні визначити тип документа як `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Ви повинні закрити декларацію `DOCTYPE` з допомогою `>` після типу.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ваш елемент `html` повинен мати початковий теґ з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` повинен мати кінцевий теґ.

```js
assert(code.match(/<\/html\s*>/));
```

Ваша декларація `DOCTYPE` повинна бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви повинні мати початковий теґ `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий теґ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий теґ `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий теґ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` та `body` повинні бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Елемент `head` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Ваш код повинен мати елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Ваш елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш проєкт повинен мати титул `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

Пам’ятайте, що вживання великої літери та написання мають значення.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

Ваш код повинен мати елемент `link`.

```js
assert.match(code, /<link/)
```

Ви повинні мати один самозакривний елемент `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
