---
id: 6140c9d35015ae0ba0c250e8
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте `div` в елемент `body` та надайте йому `class` зі значенням `wheel`.

В новий `div` додайте шість елементів `span` з `class` зі значенням `line` та шість елементів `div` з `class` зі значенням `cabin`.

# --hints--

Ви повинні створити новий `div` в елементі `body`.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

Ваш `div` в елементі `body` повинен мати `class` зі значенням `wheel`.

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

Ваш елемент `.wheel` повинен мати шість елементів `span`.

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

Ваші шість елементів `span` в елементі `.wheel` повинні мати `class` зі значенням `line`.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

Ваш елемент `.wheel` повинен мати шість елементів `div`.

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

Ваші шість елементів `div` в елементі `.wheel` повинні мати `class` зі значенням `cabin`.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
