---
id: 6140cbeec34e970dfe75e710
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Створіть селектор для свого елемента `.wheel`. Почніть з встановлення `border` на `2px solid black`, `border-radius` на `50%` та `margin-left` на `50px`.

# --hints--

Ви повинні мати селектор `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

Ваш селектор `.wheel` повинен мати властивість `border` зі значенням `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

Ваш селектор `.wheel` повинен мати властивість `border-radius` зі значенням `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

Ваш селектор `.wheel` повинен мати властивість `margin-left` зі значенням `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
