---
id: 6140cdebd39d6a101e747529
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Створіть селектор для своїх елементів `.line`. Почніть з встановлення `background-color` на `black`, `width` на `50%` та `height` на `2px`.

# --hints--

Ви повинні мати селектор `.line`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

Ваш селектор `.line` повинен мати властивість `background-color` зі значенням `black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

Ваш селектор `.line` повинен мати властивість `width` зі значенням `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

Ваш селектор `.line` повинен мати властивість `height` зі значенням `2px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
