---
id: 614385513d91ae5c251c2052
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

添加一个带有文本 `Magazine` 的 `title` 元素，一个 `link` 元素为 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 字体样式表，一个 `link` 元素为 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome 样式表和一个 `link` 为 `./styles.css` 样式表。

你的字体样式表将加载三种不同的字体：`Anton`、`Baskervville` 和 `Raleway`。

# --hints--

你的代码应该包含三个自闭合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 3);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert(document.querySelectorAll('head > link').length === 3);
```

你的三个 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

你的其中一个链接元素应将 `href` 设置为 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`。

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

你的链接元素之一应将 `href` 设置为 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

你的 `link` 元素之一应该有一个 `href` 属性，其值为 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/)
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的项目应该有一个标题为 `Magazine`。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

请记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
