---
id: 6143862a5e54455d262c212e
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `body` 中，创建一个 `main` 元素。 然后在该元素中，创建一个 `class` 为 `heading` 的 `section` 元素。

# --hints--

你应该有一个 `main` 元素。

```js
assert.exists(document.querySelector('main'));
```

你的 `main` 元素应该在你的 `body` 元素中。

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

你应该有一个 `section` 元素。

```js
assert.exists(document.querySelector('section'));
```

你的 `section`元素应该在你的 `main` 元素中。

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

你的 `section` 元素应该将 `class` 设置为 `heading`。

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
