---
id: 6143869bb45bd85e3b1926aa
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在你的 `.heading` 元素中，创建一个 `header` 元素，将该元素的 `class` 设置为 `hero`。

在那个元素中，创建一个 `img` 元素，将 `src` 设置为 `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`，将 `alt` 设置为 `freecodecamp logo`，并将 `class` 设置为 `hero-img`。

`img` 元素上的 `loading` 属性可以设置为 `lazy`，以告诉浏览器在需要（例如，当用户将图像滚动到视图中时）的时候才获取图像资源。 另外一个好处是，延迟加载的元素在加载非延迟元素之前不会加载——这意味着互联网连接速度较慢的用户可以查看你的页面的内容，而无需等待图像加载。

给你的新 `img` 元素一个 `loading` 属性，设置为 `lazy`。

# --hints--

你应该创建一个 `header` 元素。

```js
assert.exists(document.querySelector('header'));
```

你的 `header` 元素应该在你的 `.heading` 元素中。

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

你的 `header` 元素应该将 `class` 设置为 `hero`。

```js
assert(document.querySelector('header')?.className === 'hero');
```

你应该创建一个 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素应该在你的 `header` 元素中。

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

你的 `img` 元素应该有一个 `src`，值为 `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`。

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

你的 `img` 元素应该有一个 `alt`，值为 `freecodecamp logo`。

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

你的 `img` 元素应该有一个 `loading` 属性，值为 `lazy`。

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

你的 `img` 元素应该有一个 `class` 为 `hero-img`。

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
