---
id: 614389f601bb4f611db98563
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在你的 `.author` 元素下方，创建一个带有 `social-icons` 类的新的 `div` 元素。

在新的 `div` 中添加五个 `a` 元素，并给它们添加以下 `href` 属性。

- 第一个 `a` 元素的 `href` 应为 `https://www.facebook.com/freecodecamp`。
- 第二个 `a` 元素的 `href` 应为 `https://twitter.com/freecodecamp`。
- 第三个 `a` 元素的 `href` 应为 `https://instagram.com/freecodecamp`。
- 第四个 `a` 元素的 `href` 应为 `https://www.linkedin.com/school/free-code-camp`。
- 第五个 `a` 元素的 `href` 应为 `https://www.youtube.com/freecodecamp`。

# --hints--

你应该创建一个新的 `div` 元素。

```js
assert(document.querySelectorAll('div')?.length === 2);
```

你的新 `div` 元素应该在你的 `.author` 元素之后。

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

你的新 `div` 元素应具有 `social-icons` 类。

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

你的 `.social-icons` 元素应该有五个 `a` 元素。

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

你的第一个 `a` 元素的 `href` 应为 `https://www.facebook.com/freecodecamp`。

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

你的第二个 `a` 元素的 `href` 应为 `https://twitter.com/freecodecamp`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

你的第三个 `a` 元素的 `href` 应为 `https://instagram.com/freecodecamp`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

你的第四个 `a` 元素的 `href` 应为 `https://www.linkedin.com/school/free-code-camp`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

你的第五个 `a` 元素的 `href` 应为 `https://www.youtube.com/freecodecamp`。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
