---
id: 61438b5b66d76a6264430f2a
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在每个新的 `a` 元素中，增加一个 `i` 元素，并为它们提供以下类：

- 你的第一个 `i` 元素应该有 `fab fa-facebook-f` 类。
- 你的第二个 `i` 元素应该有 `fab fa-twitter` 类。
- 你的第三个 `i` 元素应该有 `fab fa-instagram` 类。
- 你的第四个 `i` 元素应该有 `fab fa-linkedin-in` 类。
- 你的第五个 `i` 元素应该有 `fab fa-youtube` 类。

# --hints--

你应该有五个 `i` 元素。

```js
assert(document.querySelectorAll('i')?.length === 5);
```

每个 `a` 元素应该只有一个 `i` 元素。

```js
const aaaaa = [...document.querySelectorAll('.social-icons a')];
assert(aaaaa?.every(a => a?.children?.length === 1 && a?.children?.[0]?.localName === 'i'));
```

每个 `i` 元素都应该有一个 `class` 属性，其中包括 `fab`。

```js
const iiiii = [...document.querySelectorAll('i')];
assert(iiiii?.every(i => i?.classList?.contains('fab')));
```

第一个 `i` 元素应该有一个 `class` 属性，其中包括 `fa-facebook-f`。

```js
assert(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

第二个 `i` 元素应该有一个 `class` 属性，其中包括 `fa-twitter`。

```js
assert(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

第三个 `i` 元素应该有一个 `class` 属性，其中包括 `fa-instagram`。

```js
assert(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

第四个`i` 元素应该有一个`class` 属性，其中包括 `fa-linkedin-in`。

```js
assert(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

第五个 `i` 元素应该有一个 `class` 属性，其中包括 `fa-youtube`。

```js
assert(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
