---
id: 684cca09aac0fc035db4099f
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside your `p` element, add the text `Error! Something went wrong. Please try again.` where `Error!` should be wrapped inside of a `strong` element.

# --hints--

You should have a `strong` element inside your `p` element.

```js
assert.exists(document.querySelector("p strong"));
```

Your `strong` element should have the text `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

Your `p` element should have the correct text that ends with `Something went wrong. Please try again.`.

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      --fcc-editable-region--
      <p>

      </p>
      --fcc-editable-region--
    </div>
  </body>
</html>
```
