---
id: 684cd59ec94b5e063bf776f0
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

It would be nice to have a border around that light red background. 

To set the `border-width` in Tailwind CSS, you can use the `border-number` utility class where `number` will represent the lightness or thickness of your border. 

Here is an example of applying a `border-width` of `4px` on all sides:

```html
<p class="border-4">This is how to set the border-width.</p>
```

Inside your opening `div` tag, add the correct utility class for a `border-width` set to `2px` on all sides.

# --hints--

Your opening `div` tag should have a `border-2` class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
