---
id: 684d0eeb9c81e5109e52fe8b
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

The `div` element is starting to look better. But now the dismiss `button` is a little too close to the paragraph element. 

To create space between the flex items, you can use the following pattern: `gap-number`. `number` represents the fixed space between both rows and columns in a flex layout.

Here is an example of applying a `gap` to an element:

```html
<p class="flex gap-8">This is gap.</p>
```

Inside your opening `div` tag, use the correct utility class for setting the `gap` to `4`.

# --hints--

Your opening `div` tag should have the `gap-4` utility class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
