---
id: 684d160a2370b01310ae0fda
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

The new red color looks nice, but it would be better if the text were larger. To set the `font-size` for an element using Tailwind CSS, you can use the following pattern: `text-value`. 

Here is an example of setting the `font-size` for an element to `0.875rem` or (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

Inside your opening `p` element, use the correct utility class to set the `font-size` to `xl` which represents `1.25rem` or (`20px`).

# --hints--

Your opening `p` tag should have the `text-xl` utility class.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
