---
id: 61329d52e5010e08d9b9d66b
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Otro elemento `meta` importante para la accesibilidad y SEO es la definición de `description`. El valor del atributo `content` es usado por los motores de búsqueda para proporcionar una descripción de tu página.

Agrega un elemento `meta` con el atributo `name` establecido a `description`, y dale un útil atributo `content`.

# --hints--

Debes agregar un nuevo elemento `meta` al `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Debes dar al `meta` un atributo `name` de `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Debes darle al `meta` un atributo `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

El valor del atributo `content` no debe tener más de 165 caracteres. _Esta es la longitud máxima de la descripción de Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
