---
id: 61329b210dac0b08047fd6ab
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Continuando con los elementos `meta`, una definición de `viewport` le dice al navegador cómo renderizar la página. Incluyendo una mejor accesibilidad visual en móvil, y mejora _SEO_ (optimización del motor de búsqueda).

Agrega una definición de `viewport` con un atributo `content` que detalla el `width` y `initial-scale` de la página.

# --hints--

Debes crear un elemento `meta` dentro del elemento `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Debes dar al `meta` un atributo `name` de `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Debes dar al `meta` un atributo `content` de `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
