---
id: 613e275749ebd008e74bb62e
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Una propiedad útil de un _SVG_ (gráficos vectoriales escalables) es que contiene un atributo `path` que permite escalar la imagen sin afectar la resolución de la imagen resultante.

Actualmente, el `img` está asumiendo el tamaño predeterminado, el cual es demasiado grande. De manera adecuada, escala la imagen usando el `id` como selector y establece el `width` a `max(100px, 18vw)`.

# --hints--

Debes usar el selector `#logo` para apuntar al elemento `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Debes dar al `img` un `width` de `max(100px, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
