---
id: 613e2546d0594208229ada50
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Dentro del `header`, proporciona contexto sobre la página anidando un elemento `img`, `h1` y `nav`.

El `img` debe apuntar a `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`, y tener un `id` de `logo`.

El `h1` debe contener el texto `HTML/CSS Quiz`.

# --hints--

Debes agregar un elemento `img` al elemento `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Debes agregar un elemento `h1` al elemento `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Debes agregar un elemento `nav` al elemento `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Debes colocar los elementos `img`, `h1` y `nav` en ese orden.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Debes dar al elemento `img` un atributo `src` de `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Debes dar al elemento `img` un atributo `id` de `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Debes dar al elemento `h1` el texto `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
