---
id: 613e275749ebd008e74bb62e
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Una propiedad útil de un _SVG_ (gráficos vectoriales escalables) es que contiene un atributo `path` que permite escalar la imagen sin afectar la resolución de la imagen resultante.

Actualmente, el `img` está asumiendo el tamaño predeterminado, el cual es demasiado grande. CSS tiene una función `max`, que devuelve el conjunto más grande de una lista de valores separados por coma (CVS, valores separados por comas). Por ejemplo:

```css
img {
  width: max(250px, 25vw);
}
```

En el ejemplo de abajo, el ancho de la imagen será de 250px si el ancho del viewport es menor que 1000px. Si el ancho del viewport es mayor que 1000px, el ancho de la imagen será de 25vw. Esto es porque 25vw es igual al 25% del ancho del viewport.

Scale the image using its `id` as a selector, and setting the `width` to be the maximum of `10rem` or `18vw`.

# --hints--

Debes usar el selector `#logo` para seleccionar al elemento `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

You should give the `img` a `width` of `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
