---
id: 61408f155e798909b6908712
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Para habilitar la navegación en la página, agrega una lista desordenada con los tres elementos de la lista siguientes:

- `INFO`
- `HTML`
- `CSS`

El texto de los elementos de la lista debe estar envuelto en etiquetas de anclaje.

# --hints--

Debes anidar un elemento `ul` dentro del `nav`.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

Debes anidar tres elementos `li` dentro del elemento `ul`.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

Debes anidar un elemento `a` dentro de cada elemento `li`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

Debes dar al primer elemento `a` el texto `INFO`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

Debes dar al segundo elemento `a` el texto `HTML`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

Debes dar al tercer elemento `a` el texto `CSS`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
