---
id: 65821c1c40ccac44adf38e6b
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Create an empty list named `expenses`. You will use it to store each of your expenses.

# --hints--

You should declare a new variable named `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

You should assign an empty list to the `expenses` variable.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
