---
id: 666ab9013cd5563f2b39947a
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Now that you have explored lists, remove all of the code you wrote for your review so you can begin the project.

# --hints--

You should delete everything.

```js
({
    test: () => runPython(`
    assert not _Node(_code)
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)

my_list.pop()
print(my_list)

--fcc-editable-region--
```
