---
id: 61537c5f81f0cf325b4a854c
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `body` 元素中，创建一个 `class` 设置为 `header` 的 `div` 。

在 `.header` 元素内嵌套一个带有文本 `CSS FLEXBOX PHOTO GALLERY` 的 `h1` 元素。

# --hints--

你的 `body` 元素中应该有一个 `div` 元素。

```js
assert.exists(document.querySelector('body')?.querySelector('div'))
```

你的 `div` 元素应该将 `class` 设置为 `header`。

```js
assert(document?.querySelector('body')?.querySelector('div')?.classList?.contains('header'))
```

你的 `.header` 元素应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

你的 `h1` 应该包含文本 `CSS FLEXBOX PHOTO GALLERY`。 请记住，大小写很重要。

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'CSS FLEXBOX PHOTO GALLERY')
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
