---
id: 61537bb9b1a29430ac15ad38
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `head` 元素中，添加文本设置为 `Photo Gallery` 的 `title` 元素并且将 `styles.css` 文件添加到页面的 `link` 元素。

# --hints--

你的 `link` 元素的 `href` 属性值应为 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的项目应该拥有一个 `Photo Gallery` 的标题。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

请记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
