---
id: 61539f32a206bd53ec116465
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

请注意你的一些图像是如何失真的。 这是因为图像具有不同的纵横比。 你可以使用 `object-fit` 属性来确定图像的行为方式，而不是单独设置每个纵横比。

为 `.gallery img` 选择器提供一个 `object-fit` 属性，并将其设置为 `cover`。 这将告诉图像填充 `img` 容器，同时保持横纵比，以便裁剪到适合。

# --hints--

应将 `.gallery img` 选择器的 `object-fit` 属性设置为 `cover`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.objectFit === 'cover');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

--fcc-editable-region--
.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
--fcc-editable-region--
```
