---
id: 61537485c4f2a624f18d7794
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

从标准的 HTML 模板开始。 添加 `DOCTYPE` 声明、`html` 元素、 `head` 元素，以及一个 `body` 元素。

# --hints--

你的代码应包含 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应该在 `DOCTYPE` 声明后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应该将文档类型定义为 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你应该在 `>` 之后关闭 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个开始标签。

```js
assert(code.match(/<html\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 声明应该位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你应该有一个开始的 `head` 标签。

```js
assert(code.match(/<head\s*>/i));
```

你应该有一个结束的 `head` 标签。

```js
assert(code.match(/<\/head\s*>/i));
```

你应该有一个开始的 `body` 标签。

```js
assert(code.match(/<body\s*>/i));
```

你应该有一个结束的 `body` 标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 标签应该是兄弟关系。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
