---
id: 615392916d83fa4f02f7e2cf
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

你可能已经注意到你的图像已全部移动到同一行。

`flex-wrap` 属性确定当 flex 容器太小时项目的行为方式。 将此属性设置为 `wrap` 将允许项目换行到下一行/列（取决于主轴），`nowrap` 将阻止项目换行。 如果将其设置为 `nowrap`，则项目可能会收缩以适应或溢出。

将 `#gallery` 选择器的 `flex-wrap` 属性设置为 `wrap`。 你应该看到你的图片是四列布局。 这是因为你在前面的步骤中将它们的 `width` 设置成了 `25%`。

# --hints--

`#gallery` 选择器的 `flex-wrap` 属性应设置为 `wrap`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery')?.flexWrap === 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

.header {
  text-align: center;
  padding: 32px;
  background: #E0DDDD;
}

--fcc-editable-region--
#gallery {
  display: flex;
  flex-direction: row;
}
--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
