---
id: 61537a8054753e2f1f2a1574
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一个 `meta` 标签，将 `name` 设置为 `viewport`，将 `content` 设置为 `width=device-width, initial-scale=1`。

同时添加一个 `meta` 标签，将 `charset` 设置为 `UTF-8` 。

# --hints--

你应该拥有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

其中一个 `meta` 元素应该将 `name` 设置为 `viewport`，将 `content` 设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

你的另一个 `meta` 元素应该将 `charset` 的属性设置为 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
