---
id: 61537d86bdc3dd343688fceb
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

接下来，根据图片在文档中的顺序为 `img` 指定 `src` 属性。 第一个 `img` 的 `src` 的值应该为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。 其余的依次类推，把 `1` 替换成 `img` 在文档中的顺序。

# --hints--

九个 `img` 元素都应该具有 `src` 属性。

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

第一个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

第二个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg`。

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

第三个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg`。

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

第四个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg`。

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

第五个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg`。

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

第六个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg`。

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

第七个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg`。

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

第八个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg`。

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

第九个 `img` 元素的 `src` 属性的值应该设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg`。

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
