---
id: 61537c9eecea6a335db6da79
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `.header` 元素下面，创建一个 `class` 设置为 `gallery` 的 `div` 元素。 这个 `div` 将作为相册的容器。

在该 `.gallery` 元素中，创建九个 `img` 元素。

# --hints--

你应该在 `body` 元素中创建一个 `div` 元素。

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

新创建的 `div` 元素的 `class` 应该为 `gallery`。

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

新创建的 `div` 元素应该在 `.header` 元素后面。

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

`.gallery` 元素内应该有九个 `img` 元素。

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
