---
id: 615380dff67172357fcf0425
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

现在你的图片太大了。

创建一个 `.gallery img` 选择器来选取你的图片。 将它们的 `width` 都设置为 `100%`， `max-width` 都设置为 `350px`。

并将 `height` 属性设置为 `300px`，以保持图片大小一致。

# --hints--

你应该有一个 `.gallery img` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

你的 `.gallery img` 选择器应该有一个 `width` 属性，其值设置为 `100%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

你的 `.gallery img` 选择器应该有一个 `max-width` 属性，其值设置为 `350px`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

你的 `.gallery img` 选择器应该有一个 `height` 属性，其值设置为 `300px`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
