---
id: 61537c9eecea6a335db6da79
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `.header` 元素下面，创建一个 `id` 设置为 `gallery` 的 `div` 元素。

在该 `#gallery` 元素中，创建十个 `img` 元素。

# --hints--

你应该在 `body` 元素中创建第二个 `div` 元素。

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length === 2);
```

新的 `div` 元素应该在 `.header` 元素之后。

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[0]?.classList?.contains('header'));
```

新的 `div` 元素应该将 `id` 设置为 `gallery`。

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[1]?.id === 'gallery');
```

你的 `#gallery` 元素应该有十个 `img` 元素。

```js
assert(document.querySelector('#gallery')?.querySelectorAll('img')?.length === 10);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
  </body>
</html>
--fcc-editable-region--
```

```css

```
