---
id: 61537d86bdc3dd343688fceb
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

现在需要为每个 `img` 元素指定一个 `src` 属性。 你将在具有 `img` 的元素中用格式为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/#.jpg` 来替换 `#` 。

因此，第一个 `img` 元素将使用 `1.jpg`，第二个将使用`2.jpg` ，依此类推。

# --hints--

十个 `img` 元素都应该具有 `src` 属性。

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

第一个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

第二个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg`。

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

第三个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg`。

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

第四个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg`。

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

第五个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg`。

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

第六个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg`。

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

第七个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg`。

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

第八个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg`。

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

第九个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg`。

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

第十个 `img` 元素应该将 `src` 属性设置为 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg`。

```js
assert(document.querySelectorAll('img')?.[9]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
--fcc-editable-region--
    <div id="gallery">
      <img />
      <img />
      <img />
      <img />
      <img />
      <img />
      <img />
      <img />
      <img />
      <img />
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
