---
id: 615f171d05def3218035dc85
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

你的图像目前太大，你将无法看到你的 CSS 变化。 创建一个 `.gallery img` 选择器来定位你的图像。 将它们的 `width` 属性设置为 `100%`，`max-width` 设置为 `350px`，以便它们适当调整，不至于太大。

并将 `height` 属性设置为 `300px`，以保持图像大小一致。

# --hints--

你应该有一个 `.gallery img` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

你的 `.gallery img` 选择器应该将 `width` 属性设置为 `100%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

应将 `.gallery img` 选择器的 `max-width` 属性设置为 `350px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

你的 `.gallery img` 选择器应该将 `height` 属性设置为 `300px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
