---
id: 6494da0daf5df5197963671d
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

注意到蓝色的图片边框会超出红色的画廊边框。 这是因为浏览器计算容器元素大小的方式所致。

`box-sizing` 属性用于设置这种行为。 默认情况下，使用的是 `content-box` 模型。 在这个默认模型中，当某个元素有一个特定的宽度时，该宽度仅基于元素的内容来计算。 内边距和边框值会被加到总宽度中，因此元素宽度会增长以适应这些值。

尝试使用全局 `*` 选择器，将 `box-sizing` 设置为 `content-box`。 此时，你不会看到任何变化，因为你正在使用默认值。

# --hints--

你应该有一个 `*` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你的 `*` 选择器应该有一个 `box-sizing` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

你的 `*` 选择器的 `box-sizing` 属性应该设置为 `content-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
