---
id: 64dd076a1b33c86d84b4232f
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

为了更好地可视化你元素的尺寸，添加边框會很有帮助。

将你的 `.gallery` 元素的 `width` 设置为 `50%`，将 `border` 设置为 `5px solid red`。

然后，将 `img` 元素的 `width` 设置为 `100%`，将 `padding` 设置为 `5px`，将 `border` 设置为 `5px solid blue`。

# --hints--

你应该有一个 `.gallery` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

你的 `.gallery` 选择器应该有一个 `width` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

你的 `.gallery` 选择器应该有一个值为 `50%` 的 `width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

你的 `.gallery` 选择器应该有一个 `border` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

你的 `.gallery` 选择器应该有一个值为 `5px solid red` 的 `border` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

你应该有一个 `img` 选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

你的 `img` 选择器应该有一个 `width` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

你的 `img` 选择器应该有一个值为 `100%` 的 `width` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

你的 `img` 选择器应该有一个 `padding` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

你的 `img` 选择器应该有一个值为 `5px` 的 `padding` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

你的 `img` 选择器应该有一个 `border` 属性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

你的 `img` 选择器应该有一个值为 `5px solid blue` 的 `border` 属性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
