---
id: 65f01969115f933073b6be03
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

A vector is an object that has a length (or magnitude) and a direction and it cannot be expressed by a single number. In physics, vectors are commonly used to represent forces, velocities, accelerations, and other quantities.

In this project you are going to build a vector space, a set in which a series of operations is defined between the elements in that set. You'll learn all the details very soon.

For now, start the project by declaring an empty class named `Vector`.

# --hints--

You should declare an empty class named `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
