---
id: 65f0694fb2296f3caadf8347
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

The length, or norm, of a vector can be calculated as the square root of the sum of its components.
Delete `pass` and return `(self.x**2 + self.y**2)**0.5` from the `norm` method.

# --hints--

You should return `(self.x**2 + self.y**2)**0.5` from your `norm` method. Remember to delete `pass`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").find_body().is_equivalent("return (self.x**2 + self.y**2)**0.5")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
