---
id: 65f055b9190fc41ca35549b8
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Delete `pass` and assign `x` to the `x` attribute of the `Vector` object.

# --hints--

You should delete `pass` and assign `x` to `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").find_body().is_equivalent("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
