---
id: 65f055b9190fc41ca35549b8
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Python provides a variety of methods having both a heading and a trailing double underscore in their name. You have already met some, such as `__init__` and `__str__`. These methods having the name in the form of `__<name>__` are known as special methods.

They are called under the hood in specific situation (e.g. `__init__` during instantiation, `__str__` when the object is printed), affecting the behaviour of the class which defines them. Going on with this project, you are going to learn some of the most commonly used special methods.

For now, delete `pass` and assign `x` to the `x` attribute of the `Vector` object.

# --hints--

You should delete `pass` and assign `x` to `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").find_body().is_equivalent("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
