---
id: 65f3fda52c7be712bb63220b
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Verify that everything works fine by printing your `v1` object.

# --hints--

You should print `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
