---
id: 65f3fa60a93b84110b3f1708
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

You want to return a meaningful string representation. For example, in case of `v1`, you want to return a tuple with the values of the vector components: `(2, 3)`.

Replace `pass` with an f-string containing the two instance attributes separated by a comma within the curly braces and return it from the `__str__` method.

# --hints--

You should return `f'{self.x, self.y}'` from the `__str__` method.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").find_body().is_equivalent("return f'{self.x, self.y}'")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
