---
id: 65f41113845ed61e5065515d
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Modify the instantiation of `v1` and `v2` by using keyword arguments.

# --hints--

You should modify the assignment of `v1` to be `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

You should modify the assignment of `v2` to be `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)

v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
