---
id: 65f93e54a9121571dcdd3e79
title: Step 32
challengeType: 20
dashedName: step-32
---

# --description--

The output is correct, `repr` is giving you the string required to instantiate the objects. Now delete the second argument from both your `print` calls.

# --hints--

You should remove `f'\nrepr = {repr(v1)}'` from your first `print()` call.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 = {v1}')")`)) })
```

You should remove `f'\nrepr = {repr(v2)}'` from your second `print()` call.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v2 = {v2}')")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
```
