---
id: 65f9b486989cb90ff3e77ac8
title: Step 34
challengeType: 20
dashedName: step-34
---

# --description--

The `isinstance()` built-in function takes two arguments and returns a Boolean indicating if the object passed as the first argument is an instance of the class passed as the second argument.

```py
isinstance(7, int) # True
```

The `other` parameter of your `__add__` method represents the operand placed at the right side of the `+` operator in an addition operation.

You want to verify that the right-hand operand is an object of the same class of the left-hand operand (i.e. `self`). Create an `if` statement that checks that by combining the `not` operator with the `isinstance()` function. Pass `other` as the first argument, and use the `__class__` attribute of the current instance as the second parameter.

# --hints--

You should create an `if` statement that checks if `self` and `other` are not instances of the same class.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(other, self.__class__)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
