---
id: 65fa8d5cf0ef3b141010f5d8
title: Step 44
challengeType: 20
dashedName: step-44
---

# --description--

It's time to verify that the addition and subtraction operations work as expected. Declare another variable `v3` and assign it the sum of `v1` plus `v2`.

# --hints--

You should declare a variable `v3` and assign it the sum of `v1` and `v2`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_stmt("v3 = v1 + v2") or _Node(_code).has_stmt("v3 = v2 + v1")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Addition is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) + other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Subtraction is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) - other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
--fcc-editable-region--
```
