---
id: 65fc45b4853d71e3b05349a8
title: Step 49
challengeType: 20
dashedName: step-49
---

# --description--

Vectors can be multiplied by a *scalar*, i.e. a number that multiplies each single component. The result of scalar multiplication is a vector with the same orientation as the original vector but a different magnitude.

Start to implement the scalar multiplication by declaring a variable `kwargs`. Assign it a dictionary comprehension that iterates over the object attributes and for each key (`i`) creates a key-value pair, where the key is equal to the current key and the value is the product of the `i` attribute times `other`.

# --hints--

You should declare a variable `kwargs` and assign it a dictionary comprehension that iterates over the object attributes inside the `if` body.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("self.__dict__") or node.is_equivalent("vars(self)")
`)) })
```

The dictionary comprehension assigned to `kwargs` should use `i` to iterate over the object attributes.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

The dictionary comprehension assigned to `kwargs` should have `i` as the key.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

The dictionary comprehension assigned to `kwargs` should assign the product of the `i` attribute of `self` and `other` to each `i` key.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) * other") or node.is_equivalent("self.__getattribute__(i) * other")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if self.__class__ != other.__class__:
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if self.__class__ != other.__class__:
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __mul__(self, other):
        if isinstance(other, (int, float)):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
