---
id: 65fc8707f3af570332499991
title: Step 50
challengeType: 20
dashedName: step-50
---

# --description--

Still, inside the `if` body, return a new instance of the current class by using `__class__` and unpacking `kwargs` with the `**` operator. This will be the vector resulting from the scalar multiplication.

# --hints--

You should return a new instance of the current class by using `__class__` and unpacking `kwargs` with the `**` operator inside the `if` body.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].has_return("self.__class__(**kwargs)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if self.__class__ != other.__class__:
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if self.__class__ != other.__class__:
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __mul__(self, other):
        if isinstance(other, (int, float)):
            kwargs = {i: getattr(self, i) * other for i in vars(self)}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
