---
id: 65fc45b4853d71e3b05349a8
title: Step 49
challengeType: 20
dashedName: step-49
---

# --description--

Now it's time to implement the scalara multiplication. Replace `pass` with a variable `kwargs`. Assign it a dictionary comprehension that iterates over `__dict__` and for each key (`i`) creates a key-value pair, where the key is equal to the current key and the value is the product of the `i` attribute times `other`. Remember to use `__getattribute__` to access the value of `i`.

# --hints--

You should declare a variable `kwargs` and assign it a dictionary comprehension that iterates over `__dict__` inside the `if` body.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_iters()[0].is_equivalent("self.__dict__")
    `))
})
```

The dictionary comprehension assigned to `kwargs` should use `i` to iterate over `self.__dict__`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

The dictionary comprehension assigned to `kwargs` should have `i` as the key.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

The dictionary comprehension assigned to `kwargs` should assign the product of the `i` attribute of `self` and `other` to each `i` key in `self.__dict__`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").find_ifs()[0].find_bodies()[0].find_variable("kwargs").find_comp_expr().is_equivalent("self.__getattribute__(i) * other")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Addition is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) + other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Subtraction is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) - other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __mul__(self, other):
        if isinstance(other, (int, float)):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
