---
id: 65fd95c23beef982af29004c
title: Step 57
challengeType: 20
dashedName: step-57
---

# --description--

You want to compare two vectors, only when they belong to the same class. For that create an `if` statement that checks if `other` is not an instance of the same class of the current instance and raise a `TypeError` using the following f-string as the custom message: `f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'"`.

# --hints--

You should create an `if` statement that checks if `self` and `other` are not instances of the same class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__eq__").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(other, self.__class__)")`)) })
```

You should raise a `TypeError` and provide a custom error message using the provided string inside the `if` body.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__eq__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"Comparison operations are not defined between \\'{self.__class__.__name__}\\' and \\'{other.__class__.__name__}\\'")')`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Addition is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) + other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Subtraction is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) - other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __mul__(self, other):
        if isinstance(other, (int, float)):
            kwargs = {i: self.__getattribute__(i) * other for i in self.__dict__}
            return self.__class__(**kwargs)
        
        elif isinstance(other, self.__class__):
            args = [self.__getattribute__(i) * other.__getattribute__(i) for i in self.__dict__]
            return sum(args)
            
        raise TypeError(f"Multiplication is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
--fcc-editable-region--
    def __eq__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')
v5 = v1 * v2
print(f'v1 * v2 = {v5}')
```
