---
id: 66019319edb1cb4b57d3a793
title: Step 67
challengeType: 20
dashedName: step-67
---

# --description--

The last method you need is `__ge__`, which is called when the `>=` is used to compare two objects.

Define a `__ge__` method with `self` and `other` as the parameters and make it return the opposite of `self < other`.

# --hints--

You should define a `__ge__` method within the `R2Vector` class.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__ge__")`)) })
```

Your `__ge__` method should have two parameters, `self` and `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__ge__").has_args("self, other")`)) })
```

Your `__ge__` method should return the opposite of `self < other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__ge__").has_return("not self < other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Addition is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) + other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Subtraction is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) - other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __mul__(self, other):
        if isinstance(other, (int, float)):
            kwargs = {i: self.__getattribute__(i) * other for i in self.__dict__}
            return self.__class__(**kwargs)
        
        elif isinstance(other, self.__class__):
            args = [self.__getattribute__(i) * other.__getattribute__(i) for i in self.__dict__]
            return sum(args)
            
        raise TypeError(f"Multiplication is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")

    def __eq__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return all(self.__getattribute__(i) == other.__getattribute__(i) for i in self.__dict__)
        
    def __ne__(self, other):
        return not self == other

    def __lt__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return self.norm() < other.norm()

    def __gt__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return self.norm() > other.norm()

    def __le__(self, other):
        return not self > other
--fcc-editable-region--

--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')
v5 = v1 * v2
print(f'v1 * v2 = {v5}')
print(v1 != R2Vector(x=2, y=3))

```
