---
id: 6601ad0fe415985a5c83f3cc
title: Step 75
challengeType: 20
dashedName: step-75
---

# --description--

As a final step, call the `print` function and pass it the f-string `f'v1 x v2 = {v6}'` to see the output of the dot product.

With that, you have completed the vector space project. Well done!

# --hints--

You should print `f'v1 x v2 = {v6}'`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 x v2 = {v6}')")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Addition is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) + other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Subtraction is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) - other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __mul__(self, other):
        if isinstance(other, (int, float)):
            kwargs = {i: self.__getattribute__(i) * other for i in self.__dict__}
            return self.__class__(**kwargs)
        
        elif isinstance(other, self.__class__):
            args = [self.__getattribute__(i) * other.__getattribute__(i) for i in self.__dict__]
            return sum(args)
            
        raise TypeError(f"Multiplication is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")

    def __eq__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return all(self.__getattribute__(i) == other.__getattribute__(i) for i in self.__dict__)
        
    def __ne__(self, other):
        return not self == other

    def __lt__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return self.norm() < other.norm()

    def __gt__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return self.norm() > other.norm()

    def __le__(self, other):
        return not self > other

    def __ge__(self, other):
        return not self < other

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
        
    def cross(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Argument must be a '{self.__class__.__name__}' instance")
        kwargs = {}
        kwargs['x'] = self.y * other.z - self.z * other.y
        kwargs['y'] = self.z * other.x - self.x * other.z
        kwargs['z'] = self.x * other.y - self.y * other.x
        
        return self.__class__(**kwargs)
--fcc-editable-region--
v1 = R3Vector(x=2, y=3, z=1)
v2 = R3Vector(x=0.5, y=1.25, z=2)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')
v5 = v1 * v2
print(f'v1 * v2 = {v5}')
v6 = v1.cross(v2)
--fcc-editable-region--
```

# --solutions--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in self.__dict__.items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Addition is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) + other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Subtraction is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        kwargs = {i: self.__getattribute__(i) - other.__getattribute__(i) for i in self.__dict__}
        return self.__class__(**kwargs)

    def __mul__(self, other):
        if isinstance(other, (int, float)):
            kwargs = {i: self.__getattribute__(i) * other for i in self.__dict__}
            return self.__class__(**kwargs)
        
        elif isinstance(other, self.__class__):
            args = [self.__getattribute__(i) * other.__getattribute__(i) for i in self.__dict__]
            return sum(args)
            
        raise TypeError(f"Multiplication is not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")

    def __eq__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return all(self.__getattribute__(i) == other.__getattribute__(i) for i in self.__dict__)
        
    def __ne__(self, other):
        return not self == other

    def __lt__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return self.norm() < other.norm()

    def __gt__(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Comparison operations are not defined between '{self.__class__.__name__}' and '{other.__class__.__name__}'")
        return self.norm() > other.norm()

    def __le__(self, other):
        return not self > other

    def __ge__(self, other):
        return not self < other

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
        
    def cross(self, other):
        if not isinstance(other, self.__class__):
            raise TypeError(f"Argument must be a '{self.__class__.__name__}' instance")
        kwargs = {}
        kwargs['x'] = self.y * other.z - self.z * other.y
        kwargs['y'] = self.z * other.x - self.x * other.z
        kwargs['z'] = self.x * other.y - self.y * other.x
        
        return self.__class__(**kwargs)

v1 = R3Vector(x=2, y=3, z=1)
v2 = R3Vector(x=0.5, y=1.25, z=2)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')
v5 = v1 * v2
print(f'v1 * v2 = {v5}')
v6 = v1.cross(v2)
print(f'v1 x v2 = {v6}')
```
