---
id: 66598a6e182e398244476a95
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

A class can implement the `__getattribute__` method to override the default behavior when accessing an attribute. But in your case, that is not necessary, and you can further improve the readability of your code by using the `getattr()` built-in function instead.

`getattr()` takes an object as its first argument, and a string containing the attribute name as its second attribute.

Modify your `__str__` method to use `getattr()` instead of `__getattribute__()`.

# --hints--

You should modify your `__str__` method to use `getattr(self, i)` instead of `self.__getattribute__(i)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("str(tuple(getattr(self, i) for i in vars(self)))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in vars(self)))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())
print(f'v1 = {v1}')
print(f'v2 = {v2}')

```
