---
id: 6659875502b6d7765498f324
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

Use f-strings to modify your `print` calls so that they result in the output `v1 = (2, 3)` and `v2 = (2, 2, 3)`, respectively.

# --hints--

You should use an f-string to interpolate `v1` and generate the output `v1 = (2, 3)`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 = {v1}')")`)) })
```

You should use an f-string to interpolate `v2` and generate the output `v2 = (2, 2, 3)`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v2 = {v2}')")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())
--fcc-editable-region--
```
