---
id: 6659875502b6d7765498f324
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

Both the `norm` and `__str__` methods work fine but there's still something you can improve: readability. 

The `vars()` built-in function takes an object as its argument and returns the `__dict__` attribute of that object.

Instead of directly accessing the `__dict__` attribute of `self`, modify both `norm` and `__str__` to use the `vars()` function.

# --hints--

You should modify your `norm` method to use `vars(self)` instead of `self.__dict__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in vars(self).values())**0.5")`)) })
```

You should modify your `__str__` method to use `vars(self)` instead of `self.__dict__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("str(tuple(self.__getattribute__(i) for i in vars(self)))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return str(tuple(self.__getattribute__(i) for i in self.__dict__))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())
print(f'v1 = {v1}')
print(f'v2 = {v2}')

```
