---
id: 6552487e689f6e895f658717
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

Now you need to create a `new_char` variable at the end of your loop body. Set its value to `alphabet[new_index]`.

# --hints--

You should create a `new_char` variable.

```js
({ test: () => assert(__userGlobals.has("new_char")) })
```

You should set your `new_char` variable to `alphabet[new_index]`.

```js
({ test: () => assert.match(code, /new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]/s) })
```

Your new variable assignment should be at the end of your loop body.

```js
({ test: () => assert.match(code, /:(.*?)index.*\1new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]/s) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
