---
id: 65521fc818947e800bffe48a
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

Before printing the current character, declare a variable called `index` and assign the value returned by `alphabet.find(char)` to this variable.

# --hints--

You should declare a new variable named `index`.

```js
({ test: () => assert(__userGlobals.has("index")) })
```

You should assign `alphabet.find(char)` to your new variable.

```js
({ test: () => assert.match(code, /index\s*=\s*alphabet\.find\s*\(\s*char\s*\)/) })
```

You should declare your new `index` variable at the beginning of your `for` loop.

```js
({ test: () => assert.match(code, /text:\s+index\s*=\s*alphabet\.find\s*\(\s*char\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    print(char)
--fcc-editable-region--
```
