---
id: 6552114a32be887b19dd8fa8
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

Print your new variable.

# --hints--

You should print your `shifted` variable.

```js
({ test: () => assert.match(code, /print\s*\(\s*shifted\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
--fcc-editable-region--
```
