---
id: 6555e012a1496917c386b579
title: ステップ 85
challengeType: 20
dashedName: step-85
---

# --description--

Modify your last `print()` call to print `Decrypted text:` + `decryption`.

# --hints--

You should change your `print(decryption)` call so that it prints `Decrypted text:` followed by the value of the `decryption` variable using the `+` operator.

```js
({
  test: () => {
    assert.match(code, /print\s*\(\s*("|')Decrypted\stext:\s\1\s*\+\s*decryption\s*\)/)
    assert.isFalse(/print\s*\(\s*decryption\s*\)/.test(code))
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
encryption = encrypt(text, custom_key)
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
decryption = decrypt(encryption, custom_key)
print(decryption)
--fcc-editable-region--
```
