---
id: 646c5e727e56e743c9aed4a1
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Mit der CSS-Positionierung kannst du festlegen, wie ein Element im Browser positioniert werden soll. Es verfügt über eine `position`-Eigenschaft, die du auf `static`, `absolute`, `relative`, `sticky` oder `fixed` setzen kannst.

Sobald du die `position`-Eigenschaft des Elements festgelegt hast, kannst du das Element verschieben, indem du einen Pixel- oder Prozentwert für eine oder mehrere der `top`-, `right`-, `left`- oder `bottom`-Eigenschaften festlegst.

`static` ist die Standardpositionierung für alle Elemente. Wenn du sie einem Element zuweist, kannst du es nicht mehr mit `top`, `right`, `left` oder `bottom` verschieben.

Weise `.cat-head` eine `position`-Eigenschaft mit dem Wert `static` zu und setze dann die `top`- und `left`-Eigenschaften jeweils auf `100px`.


# --hints--

Dein `.cat-head`-Selektor sollte eine `position`-Eigenschaft auf `static` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Dein `.cat-head`-Selektor sollte eine `top`-Eigenschaft auf `100px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Dein `.cat-head`-Selektor sollte eine `left`-Eigenschaft auf `100px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
