---
id: 646ce9d790d2a44de5f99e04
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Erstelle zwei `div`-Elemente, das erste innerhalb des `cat-left-ear`-Elements und gib ihm die Klasse `cat-left-inner-ear`. Füge das zweite `div`-Element innerhalb des `cat-right-ear`-Elements und gib ihm die Klasse `cat-right-inner-ear`.

# --hints--

Du solltest dein vorhandenes `div`-Element mit der Klasse `cat-left-ear` nicht verändern.

```js
assert(document.querySelectorAll('div.cat-left-ear').length === 1);
```

Du solltest dein vorhandenes `div`-Element mit der Klasse `cat-right-ear` nicht verändern.

```js
assert(document.querySelectorAll('div.cat-right-ear').length === 1);
```

Du solltest ein `div`-Element innerhalb deines `.cat-left-ear`-Elementes haben.

```js
assert(document.querySelectorAll('.cat-left-ear div').length === 1);
```

Du solltest ein `div`-<Element innerhalb deines `.cat-right-ear` -Elementes haben.

```js
assert(document.querySelectorAll('.cat-right-ear div').length === 1);
```

Das neue `div`-Element innerhalb von `.cat-left-ear` sollte die Klasse `cat-left-inner-ear` haben.

```js
assert(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

Das neue `div` -Element innerhalb von `.cat-right-ear` sollte die Klasse `cat-right-inner-ear` haben.

```js
assert(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">

          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
