---
id: 646c48df8674cf2b91020ecc
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Füge ein `link`-Element innerhalb deines `head`-Elements hinzu. Setze für das `link`-Element das `rel`-Attribut auf `stylesheet` und das `href` auf `./styles.css`.

# --hints--

Du solltest ein `link`-Element haben.

```js
assert.match(code, /<link/)
```

Dein `link`-Element sollte ein `rel`- und `href`-Attribut enthalten, das auf `stylesheet` und `./styles.css` gesetzt ist.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
