---
id: 646c47867800472a4ed5d2ea
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Beginne mit der grundlegenden HTML-Struktur. Add a `DOCTYPE` reference of `html` and an `html` element with its `lang` attribute set to `en`. Also, add a `head` and a `body` element within the `html` element.

# --hints--

Du solltest eine `DOCTYPE`-Deklaration von `html` haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTMLs stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest ein öffnendes `html`-Tag mit dem `lang`-Attribut `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Du solltest ein abschließendes `html`-Tag haben.

```js
assert(code.match(/<\/html>/i));
```

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head>/i));
```

Du solltest ein abschließendes `head`-Tag haben.

```js
assert(code.match(/<\/head>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body>/i));
```

Du solltest ein abschließendes `body`-Tag haben.

```js
assert(code.match(/<\/body>/i));
```

Dein `body`-Element sollte nach dem `head`-Element stehen.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Your `head` and `body` elements should be inside the `html` element.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
