---
id: 646c59e942f35541923104bf
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Erstelle in deinem HTML-Code ein `main`-Element. Füge innerhalb dieses `main`-Elements ein `div`-Element mit der Klasse `cat-head` hinzu.

# --hints--

Du solltest ein `main`-Element haben.

```js
assert(document.querySelectorAll('main').length === 1);
```

Du solltest ein `div`-Element haben.

```js
assert(document.querySelectorAll('div').length === 1);
```

Dein `div`-Element sollte die Klasse `cat-head` enthalten.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

Dein `div`-Element sollte sich innerhalb deines `main`-Tags befinden.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
