---
id: 646c586be7180e40ddf74ff6
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Weise deinem `body`-Element eine `background-color` mit `#c9d2fc` als Wert zu.

# --hints--

Du solltest einen `body`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Dein `body`-Selektor sollte eine `background-color`-Eigenschaft auf `#c9d2fc` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(201, 210, 252)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```

