---
id: 646ceb843412c74edee27a79
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

Jetzt lernst du einen CSS-Trick, um Dreiecke zu zeichnen.

Dies wird verwendet, um die Ohren der Katze zu zeichnen.

Verwende einen Klassenselektor, um die `Höhe` und `Breite` Eigenschaften des `.cat-right-ear` Elementes auf `100px` zu setzen. Setze die `background-color` auf `white`. Setze die linken und rechten Ränder auf `35px solid blue`. Setze den oberen und unteren Rand auf `35px solid red`.

# --hints--

Du solltest einen `.cat-right-ear`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Dein `.cat-right-ear`-Selektor sollte eine `height`-Eigenschaft von `100px` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

Dein `.cat-right-ear`-Selektor sollte eine `width`-Eigenschaft von `100px` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

Bei deinem `.cat-right-ear`-Selektor sollte die ` background-color`-Eigenschaft auf `white` gesetzt sein. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-left`-Eigenschaft von `35px solid blue` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-right`-Eigenschaft von `35px solid blue` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-top`-Eigenschaft von `35px solid red` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-bottom`-Eigenschaft von `35px solid red` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
