---
id: 646cf2249f02ca5233d9af7c
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

Jetzt kannst du das rechte Ohr der Katze bearbeiten. Du hast bereits das HTML dafür.

Weise dem `.cat-right-ear`-Element mithilfe eines Klassenselektors einen linken und rechten Rand von jeweils `35px solid transparent` zu. Lege außerdem den unteren Rand auf `70px solid #5e5e5e` fest.

# --hints--

Du solltest einen `.cat-right-ear`-Selektor haben.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear'))
```

Dein `.cat-right-ear`-Selektor sollte eine `border-left`-Eigenschaft auf`35px solid transparent` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-right`-Eigenschaft auf `35px solid transparent` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-bottom`-Eigenschaft auf `70px solid #5e5e5e` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
