---
id: 646c48df8674cf2b91020ecc
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Füge ein `link`-Element innerhalb deines `head`-Elements hinzu. Setze für das `link`-Element das `rel`-Attribut auf `stylesheet` und das `href` auf `./styles.css`.

# --hints--

Du solltest ein `link`-Element haben.

```js
assert.match(code, /<link/)
```

Dein `link`-Element sollte ein `rel`-Attribut auf `stylesheet` gesetzt haben.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1/)
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `./styles.css` enthalten.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1\s+href=\1\.\/styles\.css\1\s*\/?>/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
