---
id: 646c48df8674cf2b91020ecb
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge in dein `head`-Element ein `meta`-Tag mit dem `charset`-Attribut von `utf-8` hinzu. Füge auch ein `title`-Element mit dem Text `fCC Cat Painting` hinzu.

# --hints--

Du solltest genau ein `meta`-Element hinzufügen.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Dein `meta`-Element sollte ein `charset`-Attribut auf `utf-8` gesetzt haben.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Du solltest genau ein `title`-Element hinzufügen.

```js
assert(document.querySelectorAll('title').length === 1);
```

Dein `title`-Element sollte den Text `fCC Cat Painting` enthalten. Beachte, dass Recht- und Großschreibung von Bedeutung sind.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
