---
id: 646c5d7057c45f432fcdd46c
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Weise dem `.cat-head`-Element mithilfe eines Klassenselektors eine Breite von `205px` und eine Höhe von `180px` zu. Gib ihm außerdem einen Rand von `1px solid #000` und einem `border-radius` von `46%`.

# --hints--

Du solltest einen `.cat-head`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Dein `.cat-head`-Selektor sollte eine `width` auf `205px` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

Dein `.cat-head`-Selektor sollte eine `height` auf `180px` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

Dein `.cat-head`-Selektor sollte eine `border` auf `1px solid #000` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

Dein `.cat-head`-Selektor sollte einen `border-radius` auf `46%` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
