---
id: 6476f5c17f99146071ee884c
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Du konntest sehen, dass nichts passiert. Das `.cat-head`-Element hat sich nicht bewegt, obwohl ein `top` und `left` jeweils auf `100px` gesetzt sind. Aber dies ist bei der `relative`-Positionierung nicht der Fall.

Wenn du den `relative`-Wert verwendest, wird das Element weiterhin entsprechend dem normalen Fluss des Dokuments positioniert, aber die `top`-, `left`-, `bottom`- und `right`-Werte werden aktiv.

Weise deinem `.cat-head` anstelle von `static` eine Position von `relative` zu und lasse die beiden `top`- und `left`-Eigenschaften unver채ndert.

# --hints--

Dein `.cat-head`-Selektor sollte eine `position`-Eigenschaft auf `relative` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzuf체gst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

Dein `.cat-head`-Selektor sollte eine `top`-Eigenschaft auf `100px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzuf체gst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Dein `.cat-head`-Selektor sollte eine `left`-Eigenschaft auf `100px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzuf체gst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
