---
id: 6476fd4213318f6ee211028a
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Ersetze nun den `position`-Eigenschaftswert deines `.cat-head` durch `fixed`. Belasse `top` und `left` wie sie sind.

Scrolle danach nach oben und unten, um zu sehen, wie der `fixed`-Wert funktioniert.

# --hints--

Dein `.cat-head`-Selektor sollte eine `position`-Eigenschaft auf `fixed` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

Dein `.cat-head`-Selektor sollte eine `top`-Eigenschaft auf `100px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Dein `.cat-head`-Selektor sollte eine `left`-Eigenschaft auf `100px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
