---
id: 6476f7a4827bcc61682f2347
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Die nächste Positionseigenschaft ist `absolute`. Wenn du den `absolute`-Wert für deine `position`-Eigenschaft verwendest, wird das Element aus dem normalen Fluss des Dokuments herausgenommen und seine Position wird dann durch die `top`-, `right`-, `bottom`- und `left`-Eigenschaften bestimmt.

Setze die Positionseigenschaft deines `.cat-head`-Elements auf `absolute` und dann die `top`- und `left`-Eigenschaften auf einen beliebigen Pixelwert.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

Dein `.cat-head`-Selektor sollte eine `position`-Eigenschaft auf `absolute` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Dein `.cat-head`-Selektor sollte eine `top`-Eigenschaft auf eine beliebige Zahl in `px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Dein `.cat-head`-Selektor sollte eine `left`-Eigenschaft auf eine beliebige Zahl in `px` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
