---
id: 6476fb7cbaafa36d65e9cf35
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` ist ein `position`-Eigenschaftswert, mit dem du ein Element auf der Seite fixieren kannst, egal wohin der Nutzer auf der Seite scrollt.

Du musst noch ein paar Markups machen, um zu sehen, wie die `fixed`-Positionierung funktioniert. Erstelle in deinem HTML ein `div`-Element mit der Klasse `box`.

# --hints--

Du solltest ein `div`-Element haben.

```js
assert(document.querySelectorAll('div').length === 2);
```

Dein `div`-Element sollte die Klasse `box` enthalten.

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
