---
id: 64770351e8586671ec0911f0
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Der letzte Wert der Positionseigenschaft ist `sticky`. `sticky`-Positionierung ist ein Hybrid aus der `relative`- und `fixed`-Positionierung. Sie ermöglicht es einem Element, an einer bestimmten Position innerhalb seines Elements oder Viewports zu **haften**, basierend auf der Scrollposition.

Ändere den Wert der `position`-Eigenschaft von `.cat-head` zu `sticky`, setze `top` auf `0` und entferne dann `left` mitsamt seinem Wert.

**Hinweis**: Um zu sehen, wie `sticky` funktioniert, musst du ein paar Texte vor und nach deinem `.cat-head` `div`-Element platzieren. If you scroll down after that, you'll see that the `.cat-head` gets stuck to the top and remains there.

# --hints--

Dein `.cat-head`-Selektor sollte eine `position`-Eigenschaft auf `sticky` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

Dein `.cat-head`-Selektor sollte eine `top`-Eigenschaft auf `0` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Du solltest die `left`-Eigenschaft und ihren Wert nicht in deinem Code haben.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
