---
id: 6477062778c85972eb648030
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Du solltest nun den Katzenkopf zentrieren.

Give the `.cat-head` element a `position` property set to `absolute`. Lege für die `right`-, `left`-, `top`- und `bottom`-Eigenschaften einen Wert von `0` fest und setze dann ihre `margin`-Eigenschaft an allen Seiten auf `auto`. Dies ist eine Methode, ein Element vertikal und horizontal mittels CSS-Positionierung zu zentrieren.

# --hints--

Dein `.cat-head`-Selektor sollte eine `position`-Eigenschaft auf `absolute` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Dein `.cat-head`-Selektor sollte eine `top`-Eigenschaft auf `0` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Dein `.cat-head`-Selektor sollte eine `left`-Eigenschaft auf `0` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

Dein `.cat-head`-Selektor sollte eine `right`-Eigenschaft auf `0` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

Dein `.cat-head`-Selektor sollte eine `bottom`-Eigenschaft auf `0` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

Dein `.cat-head`-Selektor sollte eine `margin`-Eigenschaft auf `auto` gesetzt haben. Stelle sicher, dass du ein Semikolon hinzufügst.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
