---
id: 647d855ac12cd436059acd39
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Verwende den Universalselektor, um deinem CSS `box-sizing: border-box;` hinzuzfügen. Dadurch wird sichergestellt, dass Elemente Padding und einen Rand in ihrer angegebenen Breite und Höhe enthalten.

# --hints--

Du solltest den Universalselektor (`*`) haben.

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

Dein Universalselektor sollte eine `box-sizing`-Eigenschaft auf `border-box` gesetzt haben.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
