---
id: 6676a8a8d9861319ee901a90
title: Schritt 24
challengeType: 0
dashedName: step-24
---

# --description--

Jetzt kannst du anfangen, an den Ohren deiner Katze zu arbeiten.

Bereinige deinen Kontrollcode, indem du den `.cat-right-ear` Selektor und alle zugehörigen Eigenschaften entfernst.

Weise dem `.cat-left-ear`-Element mithilfe eines Klassenselektors einen linken und rechten Rand von jeweils `35px solid transparent` zu. Lege außerdem den unteren Rand auf `70px solid #5e5e5e` fest.

# --hints--


Dein`.cat-right-ear`Selektor sollte keine`height`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

Dein`.cat-right-ear`Selektor sollte keine`width`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


Dein`.cat-right-ear`Selektor sollte keine`border-left`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

Dein`.cat-right-ear`Selektor sollte keine`border-right`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

Dein`.cat-right-ear`Selektor sollte keine`border-top`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

Dein`.cat-right-ear`Selektor sollte keine`border-bottom`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

Du solltest keinen `.cat-right-ear`-Selektor haben.

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Du solltest einen `.cat-left-ear`-Selektor haben.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

Dein`.cat-left-ear` sollte eine `border-left` Eigenschaft haben, die auf `35px solid transparent` gesetzt ist

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

Dein `.cat-left-ear`-Selektor sollte eine `border-right`-Eigenschaft haben, die auf `35px solid transparent` gesetzt ist

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

Dein `.cat-left-ear`-Selektor sollte eine `border-bottom`-Eigenschaft haben, die auf `70px solid #5e5e5e` gesetzt ist

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
