---
id: 6676a8b01e56ec1a1e07c202
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Beachte, dass du jetzt ein weißes Quadrat mit dicken Rändern hast. Die Ränder haben diagonale Kanten, mit denen Dreiecke erzeugt werden können.

In derselben Klasse passe die `-height ` und ` width ` auf `0` an. Setze die linke, rechte und obere Grenze auf `transparent`.

Entferne die `background-color` Eigenschaft, und du solltest ein Dreieck sehen.

# --hints--


Du solltest einen `.cat-right-ear`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Dein `.cat-right-ear`-Selektor sollte eine `height`-Eigenschaft von `0` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

Dein `.cat-right-ear`-Selektor sollte eine `width`-Eigenschaft von `0` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-left`-Eigenschaft haben, die auf`35px solid transparent` gesetzt ist. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-right`-Eigenschaft haben, die auf `35px solid transparent` gesetzt ist. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

Dein `.cat-right-ear`-Selektor sollte eine `border-top`-Eigenschaft haben, die auf`35px solid transparent` gesetzt ist. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


Dein `.cat-right-ear`-Selektor sollte eine `border-bottom`-Eigenschaft von `35px solid red` haben. Vergiss nicht, ein Semikolon hinzuzufügen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

Dein`.cat-right-ear`Selektor sollte keine`background-color`Eigenschaft haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
