---
id: 68adce01c0e1144d0a902959
title: "Python Challenge 26: IPv4 Validator"
challengeType: 29
dashedName: python-challenge-26
---

# --description--

Given a string, determine if it is a valid IPv4 Address. A valid IPv4 address consists of four integer numbers separated by dots (`.`). Each number must satisfy the following conditions:

- It is between 0 and 255 inclusive.
- It does not have leading zeros (except for 0 itself).
- Only numeric characters are allowed.

# --hints--

`is_valid_ipv4("192.168.1.1")` should return `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(is_valid_ipv4("192.168.1.1"))`)
}})
```

`is_valid_ipv4("0.0.0.0")` should return `true`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(is_valid_ipv4("0.0.0.0"))`)
}})
```

`is_valid_ipv4("255.01.50.111")` should return `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertFalse(is_valid_ipv4("255.01.50.111"))`)
}})
```

`is_valid_ipv4("255.00.50.111")` should return `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertFalse(is_valid_ipv4("255.00.50.111"))`)
}})
```

`is_valid_ipv4("256.101.50.115")` should return `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertFalse(is_valid_ipv4("256.101.50.115"))`)
}})
```

`is_valid_ipv4("192.168.101.")` should return `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertFalse(is_valid_ipv4("192.168.101."))`)
}})
```

`is_valid_ipv4("192168145213")` should return `false`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertFalse(is_valid_ipv4("192168145213"))`)
}})
```

# --seed--

## --seed-contents--

```js
def is_valid_ipv4(ipv4):

    return ipv4
```

# --solutions--

```js
def is_valid_ipv4(ipv4):
    parts = ipv4.split(".")    

    if len(parts) != 4:
        return False

    for part in parts:
        if not part.isdigit():
            return False

        num = int(part)
        if num < 0 or num > 255:
            return False

        if len(part) > 1 and part.startswith("0"):
            return False

    return True
```
