---
id: 68c1a929005bf54d342aa8d6
title: "Challenge 59: Space Week Day 5: Placeholder"
challengeType: 29
dashedName: challenge-59
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
