---
id: 68c1a929005bf54d342aa8d5
title: "Challenge 58: Space Week Day 4: Placeholder"
challengeType: 29
dashedName: challenge-58
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
