---
id: 68c497f3aaefc9fd9f1b0e26
title: "Challenge 62: Placeholder"
challengeType: 29
dashedName: challenge-62
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
