---
id: 68d2ba1468508398389487d0
title: "Challenge 75: Missile Strike"
challengeType: 29
dashedName: challenge-75
---

# --description--

Given a 2D array representing the location of your ship in the ocean, and an array with the coordinates (`[row, column]`) of a missile strike, return `"Hit"`, `"Miss"`, or `"Sunk"` using the following rules:

- The given 2D array will contain exactly one unsunk ship.
- Each cell in the 2D array will contain one of the following values:
  - `-`: Water (no ship).
  - `O`: A part of your ship that has not been struck.
  - `X`: A part of your ship that has already been struck.

- If the strike hits the water, return `"Miss"`.
- If the strike hits your ship, but at least one other part of the ship remains unstruck, return `"Hit"`.
- If the strike hits the last unstruck part of your ship, return `"Sunk"`.

For example, given:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

And `[2, 1]` for the coordinates of the missile strike, return `"Sunk"` because it hit the last unstruck part of your ship.

# --hints--

`strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` should return `"Sunk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Sunk")`)
}})
```

`strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` should return `"Miss"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strike([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Miss")`)
}})
```

`strike([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` should return `"Hit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strike([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Hit")`)
}})
```

`strike([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` should return `"Hit"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strike([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Hit")`)
}})
```

`strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` should return `"Sunk"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Sunk")`)
}})
```

`strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` should return `"Miss"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(strike([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Miss")`)
}})
```

# --seed--

## --seed-contents--

```py
def strike(map, coordinates):

    return map
```

# --solutions--

```py
def strike(map, coordinates):
    row, col = coordinates
    target = map[row][col]

    if target == "-":
        return "Miss"

    if target == "O":
        for r in range(len(map)):
            for c in range(len(map[r])):
                if (r != row or c != col) and map[r][c] == "O":
                    return "Hit"
        return "Sunk"

    return "Hit"
```
