---
id: 699c8e045ee7cb94ed2322d4
title: "Challenge 213: Word Length Converter"
challengeType: 29
dashedName: challenge-213
---

# --description--

Given a string of words, return a new string where each word is replaced by its length.

- Words in the given string will be separated by a single space
- Keep the spaces in the returned string.

For example, given `"hello world"`, return `"5 5"`.

# --hints--

`convert_words("hello world")` should return `"5 5"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("hello world"), "5 5")`)
}})
```

`convert_words("Thanks and happy coding")` should return `"6 3 5 6"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Thanks and happy coding"), "6 3 5 6")`)
}})
```

`convert_words("The quick brown fox jumps over the lazy dog")` should return `"3 5 5 3 5 4 3 4 3"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("The quick brown fox jumps over the lazy dog"), "3 5 5 3 5 4 3 4 3")`)
}})
```

`convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl")` should return `"5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(convert_words("Lorem ipsum dolor sit amet consectetur adipiscing elit donec ut ligula vehicula iaculis orci vel semper nisl"), "5 5 5 3 4 11 10 4 5 2 6 8 7 4 3 6 4")`)
}})
```

# --seed--

## --seed-contents--

```py
def convert_words(s):

    return s
```

# --solutions--

```py
def convert_words(s):

    return " ".join(str(len(word)) for word in s.split(" "))
```
