---
id: 69c6ff713a52713463aa7929
title: "Challenge 263: Binary Crossword"
challengeType: 29
dashedName: challenge-263
---

# --description--

Given a character, determine if its 8-bit binary representation can be found in the following grid, horizontally or vertically in either direction:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

For example, `"A"` has the binary representation `01000001`, which appears in the first row from left to right.

# --hints--

`is_in_crossword("I")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("I"), True)`)
}})
```

`is_in_crossword("D")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("D"), True)`)
}})
```

`is_in_crossword("0")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("0"), True)`)
}})
```

`is_in_crossword("u")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("u"), True)`)
}})
```

`is_in_crossword("Y")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Y"), False)`)
}})
```

`is_in_crossword("p")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("p"), False)`)
}})
```

`is_in_crossword("1")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("1"), False)`)
}})
```

`is_in_crossword("Q")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_in_crossword("Q"), False)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_in_crossword(char):

    return char
```

# --solutions--

```py
def is_in_crossword(char):
    grid = [
        "01000001",
        "01101111",
        "01000100",
        "01100101",
        "01010010",
        "01010100",
        "01101000",
        "10101110",
    ]
    target = format(ord(char), '08b')
    cols = [''.join(grid[r][c] for r in range(8)) for c in range(8)]
    return any(target in seq or target in seq[::-1] for seq in grid + cols)
```
