---
id: 69e2383af7832c8032603b8f
title: "Challenge 273: ISBN-13 Validator"
challengeType: 29
dashedName: challenge-273
---

# --description--

Given a string, determine if it is a valid ISBN-13 number.

A valid ISBN-13:

- Contains only digits and hyphens
- Has exactly 13 digits after removing hyphens
- Passes the following check:
  1. Multiply each digit by 1 or 3, alternating (multiply the first digit by 1, the second by 3, the third by 1, and so on).
  2. The sum of the results must be divisible by 10.

# --hints--

`is_valid_isbn_13("9780306406157")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9780306406157"), True)`)
}})
```

`is_valid_isbn_13("97803064061570")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("97803064061570"), False)`)
}})
```

`is_valid_isbn_13("978-0-13-595705-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-0-13-595705-9"), True)`)
}})
```

`is_valid_isbn_13("978-030-64061A-4")` should return `False`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("978-030-64061A-4"), False)`)
}})
```

`is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9")` should return `True`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertIs(is_valid_isbn_13("9-7-8-0-1-3-4-7-5-7-5-9-9"), True)`)
}})
```

# --seed--

## --seed-contents--

```py
def is_valid_isbn_13(s):

    return s
```

# --solutions--

```py
def is_valid_isbn_13(s):
    digits = s.replace("-", "")
    if not digits.isdigit() or len(digits) != 13:
        return False
    total = sum(int(d) * (1 if i % 2 == 0 else 3) for i, d in enumerate(digits))
    return total % 10 == 0
```
