---
id: 69e2383af7832c8032603b93
title: "Challenge 277: Placeholder"
challengeType: 29
dashedName: challenge-277
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
