---
id: 69f35a5bb823ed620fcb7cb8
title: "Challenge 279: Placeholder"
challengeType: 29
dashedName: challenge-279
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
