---
id: 69f35a5bb823ed620fcb7cbd
title: "Challenge 284: Placeholder"
challengeType: 29
dashedName: challenge-284
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
