---
id: 6a0dcd03ee4e68698080ef66
title: "Challenge 302: Placeholder"
challengeType: 29
dashedName: challenge-302
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
