---
id: 6a0dcc730cb92a616f86f0c4
title: "Challenge 300: Placeholder"
challengeType: 29
dashedName: challenge-300
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
