---
id: 6a0dcc730cb92a616f86f0c5
title: "Challenge 301: Placeholder"
challengeType: 29
dashedName: challenge-301
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
