---
id: 6a15cadf5f240d05a2649559
title: "Challenge 315: Summer Solstice"
challengeType: 29
dashedName: challenge-315
---

# --description--

Today is the summer solstice, the longest day of the year in the Northern Hemisphere and the shortest in the Southern. Given a latitude, return a string representing daytime and nighttime hours.

- The latitude will be between 90 (north pole) and -90 (south pole), inclusive
- The number of daytime hours = 12 + (latitude / 90) * 12
- Round the daytime hours to the nearest even number

Return a 24-character string using `"☀️"` for daytime hours and `"🌑"` for nighttime hours, where:

- Each character represents one hour, starting at midnight (hour 0)
- Sunrise and sunset fall symmetrically around noon

For example, a latitude of `0` (the equator) has 12 hours of daylight, so sunrise is at 6:00 AM and sunset is at 6:00 PM. Return: `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

# --hints--

`get_daytime_hours(0)` should return `"🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(0), "🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(90)` should return `"☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(90), "☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️")`)
}})
```

`get_daytime_hours(-90)` should return `"🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-90), "🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(-33)` should return `"🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-33), "🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

`get_daytime_hours(66.5)` should return `"🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(66.5), "🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑")`)
}})
```

`get_daytime_hours(40)` should return `"🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(40), "🌑🌑🌑☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️☀️🌑🌑🌑")`)
}})
```

`get_daytime_hours(-50)` should return `"🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(get_daytime_hours(-50), "🌑🌑🌑🌑🌑🌑🌑🌑🌑☀️☀️☀️☀️☀️☀️🌑🌑🌑🌑🌑🌑🌑🌑🌑")`)
}})
```

# --seed--

## --seed-contents--

```py
def get_daytime_hours(latitude):

    return latitude
```

# --solutions--

```py
def get_daytime_hours(latitude):
    daylight_hours = round((12 + (latitude / 90) * 12) / 2) * 2
    sunrise = 12 - daylight_hours // 2
    sunset = 12 + daylight_hours // 2
    result = ""
    for hour in range(24):
        result += "☀️" if sunrise <= hour < sunset else "🌑"
    return result
```
