---
id: 6a15cadf5f240d05a264955a
title: "Challenge 316: Leet Speak"
challengeType: 29
dashedName: challenge-316
---

# --description--

Given a lowercase string, return it translated into leet speak by replacing the letters below with their leet substitutions:

| Letter | Leet |
| - | - |
| `a` | `4` |
| `e` | `3` |
| `i` | `1` |
| `o` | `0` |
| `s` | `5` |
| `t` | `7` |

- Characters with no substitution are left unchanged.

# --hints--

`make_leet("cool")` should return `"c00l"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("cool"), "c00l")`)
}})
```

`make_leet("leet")` should return `"l337"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("leet"), "l337")`)
}})
```

`make_leet("hacker")` should return `"h4ck3r"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("hacker"), "h4ck3r")`)
}})
```

`make_leet("satellite")` should return `"5473ll173"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("satellite"), "5473ll173")`)
}})
```

`make_leet("abcdefghijklmnopqrstuvwxyz")` should return `"4bcd3fgh1jklmn0pqr57uvwxyz"`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(make_leet("abcdefghijklmnopqrstuvwxyz"), "4bcd3fgh1jklmn0pqr57uvwxyz")`)
}})
```

# --seed--

## --seed-contents--

```py
def make_leet(s):

    return s
```

# --solutions--

```py
def make_leet(s):
    leet_map = {"a": "4", "e": "3", "i": "1", "o": "0", "t": "7", "s": "5"}
    return "".join(leet_map.get(c, c) for c in s)
```
