---
id: 6a19b1062d1b153d8ac76d73
title: "Challenge 325: Placeholder"
challengeType: 29
dashedName: challenge-325
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
