---
id: 6a19b1062d1b153d8ac76d72
title: "Challenge 324: Duplicate Character Count"
challengeType: 29
dashedName: challenge-324
---

# --description--

Given two strings, return a count of characters from the second string that can be found in the first.

- Duplicate characters in the second string are counted separately.

# --hints--

`duplicate_character_count("aloha", "hei")` should return `1`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("aloha", "hei"), 1)`)
}})
```

`duplicate_character_count("jambo", "bonjour")` should return `4`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("jambo", "bonjour"), 4)`)
}})
```

`duplicate_character_count("hello", "hola")` should return `3`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello", "hola"), 3)`)
}})
```

`duplicate_character_count("ola", "hej")` should return `0`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ola", "hej"), 0)`)
}})
```

`duplicate_character_count("ciao", "konnichiwa")` should return `5`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("ciao", "konnichiwa"), 5)`)
}})
```

`duplicate_character_count("merhaba", "xin chao")` should return `2`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("merhaba", "xin chao"), 2)`)
}})
```

`duplicate_character_count("hello world", "hello to everyone around the world")` should return `26`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(duplicate_character_count("hello world", "hello to everyone around the world"), 26)`)
}})
```

# --seed--

## --seed-contents--

```py
def duplicate_character_count(str1, str2):

    return str1
```

# --solutions--

```py
def duplicate_character_count(str1, str2):
    return sum(1 for char in str2 if char in str1)
```
