---
id: 6a19dd69388e5a6e59e4f2ef
title: "Challenge 326: Placeholder"
challengeType: 29
dashedName: challenge-326
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
