---
id: 6a1d9f98e819ed70a0e994d9
title: "Challenge 327: Database Migration"
challengeType: 29
dashedName: challenge-327
---

# --description--

Given two database objects, return the second object with any missing properties from the first filled in.

- Fields that already exist in the record should not be overwritten.

# --hints--

`migrate_record({ username: "", posts: 0 }, { verified: True })` should return `{ username: "", posts: 0, verified: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({"username": "", "posts": 0}, {"verified": True}), {"username": "", "posts": 0, "verified": True})`)
}})
```

`migrate_record({ username: "", posts: 0 }, { verified: True })` should return `{ username: "", posts: 0, verified: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({"username": "", "posts": 0}, {"verified": True}), {"username": "", "posts": 0, "verified": True})`)
}})
```

`migrate_record({ username: "", posts: 0 }, { verified: True })` should return `{ username: "", posts: 0, verified: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({"username": "", "posts": 0}, {"verified": True}), {"username": "", "posts": 0, "verified": True})`)
}})
```

`migrate_record({ username: "", posts: 0 }, { verified: True })` should return `{ username: "", posts: 0, verified: True }`.

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(migrate_record({"username": "", "posts": 0}, {"verified": True}), {"username": "", "posts": 0, "verified": True})`)
}})
```

# --seed--

## --seed-contents--

```py
def migrate_record(schema, record):

    return schema
```

# --solutions--

```py
def migrate_record(schema, record):
    return {**schema, **record}
```
