---
id: 6a1d9f98e819ed70a0e994e1
title: "Challenge 335: Placeholder"
challengeType: 29
dashedName: challenge-335
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
