---
id: 6a22d77ddf034bc4e35b1d54
title: "Challenge 347: Placeholder"
challengeType: 29
dashedName: challenge-347
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
