---
id: 6a26df3e2988bcdded204896
title: "Challenge 360: Placeholder"
challengeType: 29
dashedName: challenge-360
---

# --description--

Placeholder description

# --hints--

Placeholder test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def placeholder(arg):

    return arg
```

# --solutions--

```py
def placeholder(arg):

    return arg
```
