---
id: 68f890b73bbfc9efa7379e7b
title: Debug Type Coercion Errors in a Buggy App
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

You've just joined a local web development shop, and your first assignment is to clean up some buggy code left behind by the previous developers.

They were attempting simple arithmetic operations, but something went wrong, the results don't make sense. Your task is to review, debug, and correct the code so it performs as expected and is easier to read.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a variable named `firstResult` that correctly adds the numbers `5` and `10` to produce the value `15`.
2. You should have a variable named `secondResult` that correctly subtracts `5` from `8` to produce the value `3`.
3. You should have a variable named `thirdResult` that produces the value `6` by adding two numbers. Replace the boolean currently used in the expression with a number.
4. You should have a variable named `fourthResult` that produces the value `8` by adding two numbers. Replace the boolean currently used in the expression with a number.
5. You should have a variable named `fifthResult` that correctly multiplies two numbers to get the product `20`. Replace the string currently used in the expression with a number.
6. You should have a variable named `sixthResult` that correctly adds two numbers to produce the value `22`. Replace the `null` currently used in the expression with a number.

# --hints--

You should have a variable called `firstResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.firstResult);
```

`firstResult` should have the numeric value `15`.

```js
assert.strictEqual(firstResult, 15);
```

You should have a variable called `secondResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.secondResult);
```

`secondResult` should have the numeric value `3`.

```js
assert.strictEqual(secondResult, 3);
```

You should have a variable called `thirdResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.thirdResult);
```

`thirdResult` should not include a boolean value (`true` or `false`).

```js
const explorer = await __helpers.Explorer(code);
const { thirdResult } = explorer.variables;
assert.notInclude(thirdResult?.value?.toString(), "true")
assert.notInclude(thirdResult?.value?.toString(), "false")
```

`thirdResult` should have the numeric value `6`.

```js
assert.strictEqual(thirdResult, 6);
```

You should have a variable called `fourthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.fourthResult);
```

`fourthResult` should not include a boolean value (`true` or `false`).

```js
const explorer = await __helpers.Explorer(code);
const { fourthResult } = explorer.variables;
assert.notInclude(fourthResult?.value?.toString(), "true")
assert.notInclude(fourthResult?.value?.toString(), "false")
```

`fourthResult` should have the numeric value `8`.

```js
assert.strictEqual(fourthResult, 8);
```

You should have a variable called `fifthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.fifthResult);
```

`fifthResult` should not include a string multiplication (e.g. "10" * 2).

```js
const explorer = await __helpers.Explorer(code);
const { fifthResult } = explorer.variables;
const expression = fifthResult?.value?.toString() || '';
assert.notMatch(expression, /(["'`]).*\1\s*\*/);
```

`fifthResult` should have the numeric value `20`.

```js
assert.strictEqual(fifthResult, 20);
```

You should have a variable called `sixthResult`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.sixthResult);
```

`sixthResult` should not include `null` in its expression.

```js
const explorer = await __helpers.Explorer(code);
const { sixthResult } = explorer.variables;
assert.notInclude(sixthResult?.value?.toString(), "null");
```

`sixthResult` should have the numeric value `22`.

```js
assert.strictEqual(sixthResult, 22);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
