---
id: 65942344ea9413e71f27135c
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Join the string representations of the elements using the `join()` method and the pipe character `'|'` as a separator.

# --hints--

You should use the `join` method to join the elements.

```js
assert.match(code, /join/);
```

The separator should be the pipe character `'|'`.

```js
assert.match(code, /'\|'\.join/);
```

You should join all elements in the iterable `str(item) for item in part`.

```js
assert.match(
  code,
  /("|')\|\1\.join\s*\(\s*str\s*\(\s*(\w+)\s*\)\s+for\s+\2\s+in\s+part\s*\)/m
);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
            row_list = []
--fcc-editable-region--
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                str(item) for item in part
--fcc-editable-region--
```
