---
id: 680900675ae3d54ee19590c7
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Create another `label` with the text `Real Name` and an `input` of type `password`. Give the `input` a `value` of `realName`. 

Then use the `onChange` event to update `realName`. The value for the `onChange` event should be an arrow function with `e` for the parameter. The arrow function should implicitly return `setRealName(e.target.value)`.

# --hints--

You should create a `label` element.

```js
const labels = document.querySelectorAll("label");
assert.lengthOf(labels, 2);
```

Your `label` element should have `Real Name` as its text.

```js
const labels = document.querySelectorAll("label");
assert.equal(labels[1]?.textContent.trim(), "Real Name");
```

You should create an `input` of type `password` inside your `label` element.

```js
const inputEl = document.querySelectorAll("label > input");
assert.equal(inputEl[1]?.getAttribute("type"), "password");
```

Your `input` element should have its value set to the `realName` state variable.

```js
assert.match(code, /value=\{realName\}/)
```

You should use the `onChange` event to set the value of the `realName` state variable so you can type in the `input` field.

```js
assert.match(code, /onChange=\{\(?e\)?\s*=>\s*setRealName\(e\.target\.value\)\}/)
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
        <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
        </label>
        --fcc-editable-region--

        --fcc-editable-region--
        </div>
      </form>
    </div>
  )
};
```
