---
id: 680900675ae3d54ee19590c8
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Next, create a `label` with the `className` of `section` and `column`, and a text `How did you get your powers?`. Inside the `label`, create a `select` element.


# --hints--

You should create a `label` element with a `className` of `section` and `column`.

```js
const labelEl = document?.querySelector("label.section.column");
assert.exists(labelEl);
```

Your `label` element should have the text `How did you get your powers?`.

```js
const labelEl = document?.querySelector("label.section.column");
assert.equal(labelEl?.textContent, "How did you get your powers?")
```

You should create a `select` element inside your `label` element.

```js
const selectEl = document?.querySelector("label > select");
assert.exists(selectEl);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = React.useState('');
  const [realName, setRealName] = React.useState('');
  const [powerSource, setPowerSource] = React.useState('');
  const [powers, setPowers] = React.useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
        <label>
          Hero Name
          <input
            type='text'
            value={heroName}
            onChange={e => setHeroName(e.target.value)}
          />
        </label>
        <label>
          Real Name
          <input
            type='password'
            value={realName}
            onChange={e => setRealName(e.target.value)}
          />
        </label>
        </div>
        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </div>
  )
};
```
