---
id: 61696ef7ac756c829f9e4048
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

In questo progetto lavorerai sullo stile della pagina all'interno di un file CSS esterno. Abbiamo già creato un file `styles.css` per te. Ma prima di poterlo usare, devi collegarlo alla pagina.

Annida un elemento `link` all'interno dell'elemento `head`. Dagli un attributo `rel` con valore di `stylesheet` e un attributo `href` con valore di `styles.css`.

# --hints--

Dovresti avere un elemento `link` auto-concludente.

```js
assert(document.querySelectorAll('link').length === 1);
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
