---
id: 61696ef7ac756c829f9e4048
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

In questo progetto lavorerai sullo stile della pagina all'interno di un file CSS esterno. Abbiamo già creato un file `styles.css` per te. Ma prima di poterlo usare, devi collegarlo alla pagina.

Annida un elemento `link` all'interno dell'elemento `head`. Give it a `rel` attribute set to `"stylesheet"` and an `href` attribute set to `"styles.css"`.

# --hints--

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Your `link` element should be a void element, it should not have a closing tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Your `link` element should have a `rel` attribute with the value `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Your `link` element should be inside the `head` element.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
