---
id: 616965351e74d4689eb6de30
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Puoi avere più elementi `meta` a chiusura automatica su una pagina web. Ciò che distingue un elemento `meta` dall'altro è l'attributo. Dovresti aggiungere un nuovo elemento meta per ogni attributo che vuoi specificare.

Aggiungi un altro elemento `meta` a chiusura automatica all'interno dell'elemento `head`. Assegnagli un attributo `name` con il valore `viewport` e un attributo `content` con il valore `width=device-width, initial-scale=1.0`, in modo che la pagina abbia lo stesso aspetto su tutti i dispositivi.

# --hints--

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Il nuovo elemento `meta` dovrebbe essere un elemento a chiusura automatica.

```js
assert(code.match(/<\/meta>/i) === null);
```

Il nuovo elemento `meta` dovrebbe avere un attributo `name` con il valore `viewport` e un attributo `content` con il valore `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
