---
id: 616965351e74d4689eb6de30
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Usa un tag `viewport` `<meta>` per assicurarti che la tua pagina venga visualizzata allo stesso modo su tutti i dispositivi.

Annida un elemento `meta` autoconcludente nell'elemento `head`. Assegnagli un attributo `name` con il valore `viewport` e un attributo `content` con il valore `width=device-width, initial-scale=1.0`.

# --hints--

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Il nuovo elemento `meta` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<\/meta>/i) === null);
```

Il nuovo elemento `meta` dovrebbe avere un attributo `name` con il valore `viewport`, e un attributo `content` con il valore `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
