---
id: 61695197ac34f0407e339882
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Come hai visto nei progetti precedenti, le pagine web dovrebbero iniziare con una dichiarazione `DOCTYPE html` seguita da un elemento `html`.

Aggiungi una dichiarazione `DOCTYPE html` in cima al documento seguita da un elemento `html`. Dai all'elemento `html` un attributo `lang` con `en` come valore.

# --hints--

Il codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un `>`.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

L'elemento `html` dovrebbe avere una tag di apertura con un attributo `lang` di `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/gi));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
