---
id: 616971b3cd990186b66c99a1
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Ora che il file CSS esterno è configurato, puoi iniziare ad agire sullo stile della pagina.

Come promemoria, ecco come fare riferimento a un elemento p e allinearlo a destra:

```css
p {
  text-align: right;
}
```

Crea una nuova regola CSS destinata alla classe `h1` e assegna alla proprietà `text-align` il valore `center`.

# --hints--

Dovresti usare un selettore `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

La regola CSS `h1` dovrebbe avere una proprietà `text-align` con il valore `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
