---
id: 61696ef7ac756c829f9e4048
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

In questo progetto lavorerai sullo stile della pagina all'interno di un file CSS esterno. Abbiamo già creato un file `styles.css` per te. Ma prima di poterlo usare, devi collegarlo alla pagina.

Annida un elemento `link` all'interno dell'elemento `head`. Dagli un attributo `rel` con valore di `stylesheet` e un attributo `href` con valore di `styles.css`.

# --hints--

Il codice dovrebbe avere un elemento `link`.

```js
assert(code.match(/<link/i)?.length === 1);
```

L'elemento `link` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<\/link>/i) === null);
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
const head = code.match(/<head>(.|\r|\n)*<\/head>/i)?.[0];
assert(head.match(/<link/i)?.length === 1)
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert(code.match(/<link[\s\S]*?rel=('|"|`)stylesheet\1/gi)?.length === 1);
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`. Fai attenzione alle maiuscole/minuscole quando crei il link a un file esterno.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
