---
id: 61695ab9f6ffe951c16d03dd
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Annida un elemento `head` all'interno dell'elemento `html`. Subito dopo l'elemento `head`, aggiungi un elemento `body`.

# --hints--

Dovresti avere un tag `head` di apertura.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag `head` di chiusura.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag `body` di apertura.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag `body` di chiusura.

```js
assert(code.match(/<\/body\s*>/i));
```

Dovrebbe esserci solo un tag di apertura `head`.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Dovrebbe esserci solo un tag di chiusura `head`.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Dovrebbe esserci solo un tag di apertura `body`.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Dovrebbe esserci solo un tag di chiusura `body`.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

L'elemento `head` dovrebbe essere vuoto.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

L'elemento `body` dovrebbe essere vuoto.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

L'elemento `body` dovrebbe essere posizionato dopo l'elemento `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
