---
id: 61695d1fbc003856628bf561
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Per comunicare ai browser come codificare i caratteri sulla pagina, assegna all'attributo `charset` il valore `utf-8`. L'`utf-8` è un set di caratteri universale che include quasi tutti i caratteri di tutte le lingue.

All'interno dell'elemento `head`, annida un elemento `meta` con l'attributo `charset` con il valore `utf-8`. Ricorda che gli elementi `meta` sono autoconcludenti e non hanno bisogno di un tag di chiusura.

# --hints--

Dovresti avere un elemento `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

L'elemento `meta` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<\/meta>/i) === null);
```

L'elemento `meta` dovrebbe avere un set di caratteri `charset` con il valore `utf-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
