---
id: 61695c4aad56f95497c19583
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Ricorda che l'elemento `title` fornisce ai motori di ricerca informazioni aggiuntive sulla pagina. Mostra anche il contenuto di quell'elemento `title` in altri due modi:

* nella barra del titolo quando la pagina è aperta
* nella scheda del browser della pagina quando si passa su di essa. Anche se questa scheda non è attiva, una volta che ci si passa sopra, viene visualizzato il testo di `title`.

All'interno dell'elemento `head`, annida un elemento `title` con il testo `Colored Markers`.

# --hints--

Dovresti avere un tag `title` di apertura.

```js
assert(code.match(/<title\s*>/i));
```

Dovresti avere un tag `title` di chiusura.

```js
assert(code.match(/<\/title\s*>/i));
```

Il progetto dovrebbe avere il titolo `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Controlla le maiuscole/minuscole e l'ortografia del testo del titolo.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
