---
id: 61695ab9f6ffe951c16d03dd
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Annida un elemento `head` all'interno dell'elemento `html`. Subito dopo l'elemento `head`, aggiungi un elemento `body`.

# --hints--

Dovresti avere un tag `head` di apertura.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag `head` di chiusura.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag `body` di apertura.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag `body` di chiusura.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere "fratelli".

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--

--fcc-editable-region--
</html>
```
