---
id: 616d50b93ba424d6282c99cf
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Come puoi vedere, il pennarello non sembra avere alcun colore. Il colore di sfondo è stato effettivamente applicato, ma poiché l'elemento `div` marker è vuoto, non possiede alcuna altezza come impostazione predefinita.

Nella regola CSS `.marker`, assegna alla proprietà `width` il valore `200px` e alla proprietà `height` il valore `25px`.

# --hints--

La regola CSS `.marker` dovrebbe avere una proprietà `width` con il valore `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

La regola CSS `.marker` dovrebbe avere una proprietà `height` con il valore `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
