---
id: 616d525007b8c5d8b3308b1c
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Il pennarello sarebbe più bello se fosse centrato sulla pagina. Un modo semplice per centrarlo è con la <dfn>proprietà shorthand</dfn> `margin`.

Nell'ultimo progetto, hai impostato separatamente l'area di margine degli elementi con proprietà come `margin-top` e `margin-left`. La proprietà shorthand `margin` rende più facile impostare aree di margine multiple allo stesso tempo.

Per centrare il pennarello sulla pagina, assegna alla proprietà `margin` il valore `auto`. In questo modo, assegni contemporaneamente il valore `auto` a `margin-top`, `margin-right`, `margin-bottom` e `margin-left`.

# --hints--

La regola CSS `.marker` dovrebbe avere una proprietà `margin` con il valore `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
