---
id: 616d3a67ccf800ad94ec89ae
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora aggiungerai alcuni elementi che alla fine trasformerai nei pennarelli colorati.

Prima di tutto, all'interno dell'elemento `body`, aggiungi un elemento `div` e assegna all'attributo `class` il valore `container`. Dovresti aggiungere un elemento `div` sotto il tuo elemento `h1`.

# --hints--

Dovresti avere un tag di apertura `div`.

```js
assert(code.match(/<div.*>/i));
```

Dovresti avere un tag di chiusura `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

L'elemento `div` dovrebbe essere all'interno dell'elemento `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

L'elemento `div` dovrebbe essere al di sotto dell'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

L'elemento `div` dovrebbe avere l'attributo `class` con il valore `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
