---
id: 616d4a84b756d9c4b8255093
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

It's time to add some color to the marker. Ricorda che un modo per aggiungere colore a un elemento è utilizzare una <dfn>parola chiave</dfn> come `black`, `cyan` o `yellow`.

Come promemoria, ecco come fare riferimento alla classe `freecodecamp`:

```css
.freecodecamp {

}
```

Crea una nuova regola CSS destinata alla classe `marker` e assegna alla proprietà `background-color` il valore `red`.

**Note:** You will not see any changes after adding the CSS.

# --hints--

You should create a class selector to target the `marker` class.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Your `.marker` CSS rule should have a `background-color` property set to `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
