---
id: 616d595270d902f0e7086e18
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Nonostante hai tre elementi `div` distinti, vengono visualizzati come un unico grande rettangolo. Dovresti aggiungere un po' di spazio tra di essi per rendere più facile la visualizzazione di ogni elemento.

Quando la proprietà shorthand `margin` ha due valori, il primo valore si riferisce alle proprietà `margin-top` e `margin-bottom`, e il secondo valore alle proprietà `margin-left` e `margin-right`.

Nella regola CSS `.marker`, assegna alla proprietà `margin` i valori `10px auto`.

# --hints--

La regola CSS `.marker` dovrebbe avere una proprietà `margin` con i valori `10px auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
