---
id: 616d55bd160a95e22453a081
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Ora che hai un pennarello centrato e colorato, è il momento di aggiungere gli altri evidenziatori.

Nell'elemento `div` con la classe `container`, aggiungi altri due elementi `div` e assegna loro la classe `marker`.

# --hints--

Il primo nuovo elemento `div` dovrebbe avere un tag di apertura.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

Il primo nuovo elemento `div` dovrebbe avere un tag di chiusura.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Il secondo nuovo elemento `div` dovrebbe avere un tag di apertura.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

Il secondo nuovo elemento `div` dovrebbe avere un tag di chiusura.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

I nuovi elementi `div` dovrebbero essere all'interno dell'elemento `div` con la classe `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

I nuovi elementi `div` dovrebbero avere entrambi la classe `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
