---
id: 61764c602bee6974e7790f35
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

A scuola, potresti aver imparato che rosso, giallo e blu sono colori primari e come creare nuovi colori mescolandoli. Tuttavia, questo è un modello obsoleto.

Attualmente esistono due principali modelli dei colori: il modello <dfn>additivo</dfn> RGB (rosso, verde, blu) utilizzato nei dispositivi elettronici, e il modello <dfn>detrattivo</dfn> CMYK (ciano, magenta, giallo, nero) utilizzato nella stampa. In questo progetto lavorerai con il modello RGB.

Innanzitutto, aggiungi `one` all'attributo `class` del primo elemento `div`.

```html
<div class="exampleClass1 exampleClass2">
  ...content
</div>
```

# --hints--

Dovresti aggiungere la classe `one` al primo elemento `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Il primo elemento `div` dovrebbe avere le classi `marker` e `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
