---
id: 617b65579ce424bf5f02ca73
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Esistono due principali modelli dei colori: il modello <dfn>additivo</dfn> RGB (rosso, verde, blu) utilizzato nei dispositivi elettronici, e il modello <dfn>detrattivo</dfn> CMYK (ciano, magenta, giallo, nero) utilizzato nella stampa.

In questo progetto lavorerai con il modello RGB. Ciò significa che un colore parte come nero, e cambia in base ai diversi livelli di rosso, verde e blu introdotti. La funzione CSS `rgb` è un modo semplice di visualizzare quest'aspetto.

Crea una nuova regola CSS che seleziona la classe `container` e assegna alla sua proprietà `background-color` il valore `rgb(0, 0, 0)` che corrisponde al nero.

# --hints--

Dovresti usare un selettore di classe per fare riferimento alla classe `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

La regola CSS `.container` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
