---
id: 617bb9fdef27bc6aa0470ac2
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

Per creare l'ultimo colore secondario, magenta, aggiorna la funzione `rgb` nella regola CSS `.three` per combinare puro blu e puro rosso.

# --hints--

La regola CSS `.three` dovrebbe avere una proprietà `background-color` con il valore `rgb(255, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 255, 0);
}

.two {
  background-color: rgb(0, 255, 255);
}

--fcc-editable-region--
.three {
  background-color: rgb(0, 0, 255);
}
--fcc-editable-region--

```
