---
id: 617bc3386dc7d07d6469bf20
title: Step 32
challengeType: 0
dashedName: step-32
---

# --description--

Nota che per creare l'arancione, hai agito sui valori `rgb` aumentando l'intensità del rosso e diminuendo l'intensità del verde. Questo perché l'arancione è la combinazione di rosso e giallo, e si trova tra i due sulla ruota dei colori.

Per creare il colore terziario verde primavera, combina il ciano con il verde. Aggiorna la funzione `rgb` nella regola CSS `.two` in modo che il verde sia al valore massimo, e imposta il blu a `127`.

# --hints--

La regola CSS `.two` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 255, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
