---
id: 617b97abd9f3f61d1590b815
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

Nel modello additivo dei colori RGB, i <dfn>colori primari</dfn> sono colori che, se combinati, creano bianco puro. Ma perché questo accada, ogni colore ha bisogno di essere alla sua massima intensità.

Prima di combinare i colori, reimposta il pennarello verde sul verde puro. Per la funzione `rgb` nella regola CSS `.two`, imposta il verde al valore massimo di `255`.

# --hints--

La regola CSS `.two` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
