---
id: 617b954d9f4f6217a749380e
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Ora aggiungi un po' più di spazio verticale tra i pennarelli e il bordo dell'elemento `container` in cui si trovano.

Nella regola CSS `.container`, usa la proprietà shorthand `padding` per aggiungere `10px` di padding superiore e inferiore, e imposta il padding di sinistra e destra su `0`. Il funzionamento è analogo alla proprietà shorthand `margin` che hai usato prima.

# --hints--

Non dovresti rimuovere la proprietà `background-color` e il suo valore dalla regola CSS `.container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

La regola CSS `.container` dovrebbe avere una proprietà `padding` con il valore `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
