---
id: 617b8e0d93a8d10d9a90e720
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Nota che il `background-color` del pennarello è ancora rosso. Questo perché hai impostato la funzione `rgb` sul suo valore massimo di `255` per il rosso, o 100% rosso, mentre entrambi i valori per il verde e per il blu sono a `0`.

Ora usa la funzione `rgb` per impostare gli altri colori.

Nella regola CSS `.two`, utilizza la funzione `rgb` per impostare il `background-color` al valore massimo per il verde e `0` per gli altri valori. E nella regola del CSS `.three`, utilizza la funzione `rgb` per impostare il `background-color` al valore massimo per il blu e `0` per gli altri valori.

# --hints--

La regola CSS `.two` non dovrebbe usare la parola chiave `green` per impostare la proprietà `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

La regola CSS `.two` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

La regola CSS `.three` non dovrebbe usare la parola chiave `blue` per impostare la proprietà `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

La regola CSS `.three` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
