---
id: 618a0927005553b74bfa05ae
title: Step 36
challengeType: 0
dashedName: step-36
---

# --description--

Una ruota dei colori è un cerchio in cui colori simili, o <dfn>tonalità</dfn>, sono vicini l'uno all'altro, e quelli diversi sono più lontani. Ad esempio, il rosso puro è tra le tonalità rosa e arancione.

Due colori che sono opposti tra loro sulla ruota dei colori sono chiamati <dfn>colori complementari</dfn>. Se due colori complementari vengono combinati, producono il grigio. Ma quando sono posizionati fianco a fianco, questi colori producono un forte contrasto visivo e appaiono più luminosi.

Nella funzione `rgb` per la regola CSS `.one`, imposta il rosso al valore massimo di `255` per ottenere il rosso puro. Nella funzione `rgb` per regola CSS `.two`, imposta il verde e il blu al valore massimo di `255` per ottenere il ciano.

# --hints--

La regola CSS `.one` dovrebbe avere una proprietà `background-color` con il valore `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

La regola CSS `.two` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 255, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
