---
id: 618a132676346ac9f7fd59dd
title: Step 39
challengeType: 0
dashedName: step-39
---

# --description--

Hai notato come i tuoi occhi sono naturalmente attratti al colore rosso al centro? Durante la progettazione di un sito, puoi utilizzare questo effetto per attirare l'attenzione su intestazioni importanti, pulsanti o link.

Ci sono diverse altre combinazioni di colori importanti al di fuori dei colori complementari, ma le imparerai un po' più avanti.

Per ora, utilizza la funzione `rgb` nella regola CSS `.two` per impostare il `background-color` sul nero.

# --hints--

La regola CSS `.two` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
