---
id: 618a00ed1ca871a2b3aca0cb
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

Ora che hai usato tutti i colori primari, secondari e terziari sulla ruota dei colori, sarà più facile capire altri concetti della teoria dei colori e come influiscono sul design.

Prima di tutto, nelle regole CSS `.one`, `.two` e `.three`, regola i valori nella funzione `rgb` in modo che il `background-color` di ogni elemento sia impostato sul nero puro. Ricorda che la funzione `rgb` utilizza il modello additivo dei colori, in cui i colori iniziano partono dal nero e cambiano con l'aumento dei valori di rosso, verde e blu.

# --hints--

La regola CSS `.one` dovrebbe possedere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

La regola CSS `.two` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

La regola CSS `.three` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
