---
id: 619b74fa777a2b2473c94f82
title: Step 45
challengeType: 0
dashedName: step-45
---

# --description--

Aggiorna il selettore di classe CSS `.two` in modo che selezioni la nuova classe `green`. Aggiorna anche il selettore di classe `.three` in modo che faccia riferimento alla nuova classe `blue`.

# --hints--

Il codice non dovrebbe più avere un selettore di classe `.two`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

Dovresti usare un selettore di classe che fa riferimento alla classe `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

La regola CSS `.green` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Il codice non dovrebbe più avere un selettore di classe `.three`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

Dovresti usare un selettore di classe che fa riferimento alla classe `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

La regola CSS `.blue` dovrebbe avere una proprietà `background-color` con il valore `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
