---
id: 619b7fd56aa2253778bcf5f7
title: Step 48
challengeType: 0
dashedName: step-48
---

# --description--

Il modello di colore <dfn>HSL</dfn>, o tonalità, saturazione e luminosità, è un altro modo per rappresentare i colori.

La funzione CSS hsl accetta 3 valori: un numero da 0 a 360 per la tonalità, una percentuale da 0 a 100 per la saturazione e una percentuale da 0 a 100 per la luminosità.

Se immagini una ruota dei colori, la tonalità rossa si trova a 0 gradi, il verde a 120 gradi e il blu a 240 gradi.

La saturazione è l'intensità di un colore dallo 0%, o grigio, al 100% per il colore puro. Devi aggiungere il segno di percentuale `%` ai valori di saturazione e luminosità.

La luminosità indica quanto appare brillante un colore, dallo 0%, completamente nero, al 100%, completamente bianco, mentre il 50% è neutrale.

Nella regola CSS `.blue`, utilizza la funzione `hsl` per cambiare la proprietà `background-color` in blu puro. Imposta la tonalità a `240`, la saturazione al `100%` e la luminosità al `50%`.

# --hints--

La regola CSS `.blue` dovrebbe avere una proprietà `background-color` con il valore `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
