---
id: 61a489b8579e87364b2d2cdb
title: Step 49
challengeType: 0
dashedName: step-49
---

# --description--

Hai imparato alcuni modi per impostare i colori piatti in CSS, ma puoi agire su un elemento anche usando una transizione di colore, o <dfn>gradiente</dfn>.

Un gradiente indica la variazione di un colore in un altro. La funzione CSS `linear-gradient` consente di controllare la direzione della transizione lungo una linea e i colori da utilizzare.

Una cosa da ricordare è che la funzione `linear-gradient` crea effettivamente un elemento `image`, ed è solitamente accoppiato con la proprietà `background` che può accettare un'immagine come valore.

Nella regola CSS `.red`, modifica la proprietà `background-color` in `background`.

# --hints--

La regola CSS `.red` dovrebbe avere una proprietà `background` con il valore `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
