---
id: 61a493ead935c148d2b76312
title: Step 50
challengeType: 0
dashedName: step-50
---

# --description--

La funzione `linear-gradient` è molto flessibile -- ecco la sintassi di base che userai in questo tutorial:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` è la direzione della linea utilizzata per la transizione. `color1` e `color2` sono gli argomenti, ovvero i colori che verranno utilizzati nella transizione stessa. Possono essere qualsiasi tipo di colore, ed essere indicati da parole chiave, valori esadecimali, `rgb` o `hsl`.

Applica al primo pennarello un gradiente rosso-verde lungo una linea di 90 gradi.

Nella regola CSS `.red` imposta la proprietà `background` su `linear-gradient()`, e assegna a `gradientDirection` il valore `90deg`.

# --hints--

La regola CSS `.red` dovrebbe avere una proprietà `background` con il valore `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
