---
id: 61a4a7877da33a73a1c1723e
title: Step 57
challengeType: 0
dashedName: step-57
---

# --description--

Ora che i colori-stop sono impostati, applicherai diverse sfumature di rosso a ogni argomento di colore nella funzione `linear-gradient`. Le sfumature sui bordi superiore e inferiore del pennarello saranno più scure, mentre il centro sarà più leggero, come se ci fosse una luce sopra di esso.

Per il primo argomento di colore, che è attualmente rosso puro, aggiornare la funzione `rgb` in modo che il valore per il rosso sia `122` il valore per il verde `74`e il valore per il blu `14`.

# --hints--

La regola CSS `.red` dovrebbe avere una proprietà `background` con il valore `linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(0, 255, 0) 50%, rgb(0, 0, 255) 100%)`.

```js
assert.include(['linear-gradient(rgb(122,74,14)0%,rgb(0,255,0)50%,rgb(0,0,255)100%)', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14)0%,rgb(0,255,0)50%,rgb(0,0,255)100%)repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(255, 0, 0) 0%, rgb(0, 255, 0) 50%, rgb(0, 0, 255) 100%);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
