---
id: 61a4ae5f29eb5584187201c3
title: Step 61
challengeType: 0
dashedName: step-61
---

# --description--

Per questo pennarello, utilizzerai i codici esadecimali per il gradiente.

Utilizza la funzione `linear-gradient` e imposta `gradientDirection` su `180deg`. Per il primo argomento di colore, usa un codice esadecimale con i valori `55` per il rosso, `68` per il verde e `0D` per il blu.

# --hints--

La regola CSS `.green` dovrebbe avere una proprietà `background` con il valore `linear-gradient(180deg, #55680D)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.green\{.*?background:linear-gradient\(180deg,#55680D\)/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background: #007F00;
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
