---
id: 61addaf7e83988b59a7d18ff
title: Step 82
challengeType: 0
dashedName: step-82
---

# --description--

La proprietà shorthand `border-left` consente di impostare contemporaneamente la larghezza, lo stile e il colore del bordo sinistro.

Ecco la sintassi:

```css
border-left: width style color;
```

Nella regola CSS `.sleeve`, sostituisci le proprietà `border-left-width`, `border-left-style`, e `border-left-color` con la proprietà shorthand `border-left`. I valori per la larghezza, lo stile e il colore del bordo sinistro dovrebbero essere gli stessi.

# --hints--

La regola CSS `.sleeve` non dovrebbe avere una proprietà `border-left-width` e il relativo valore.

```js
assert(!/border-left-width/g.test(code));
```

La regola CSS `.sleeve` non dovrebbe avere una proprietà `border-left-style` e il relativo valore.

```js
assert(!/border-left-style/g.test(code));
```

La regola CSS `.sleeve` non dovrebbe avere una proprietà `border-left-color` e il relativo valore.

```js
assert(!/border-left-color/g.test(code));
```

La regola CSS `.sleeve` dovrebbe avere una proprietà shorthand `border-left` con il valore `10px solid black`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left-width: 10px;
  border-left-style: solid;
  border-left-color: black;
}
--fcc-editable-region--

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
