---
id: 61b0a44a6b865738ba49b9fb
title: Step 86
challengeType: 0
dashedName: step-86
---

# --description--

L'ultima cosa che devi fare è aggiungere una leggera ombra ai pennarelli per farli sembrare ancora più realistici.

La proprietà `box-shadow` consente di applicare una o più ombre intorno a un elemento. Ecco la sintassi di base:

```css
box-shadow: offsetX offsetY color;
```

Ecco come funzionano i valori `offsetX` e `offsetY`:

* sia `offsetX` che `offsetY` accettano i valori numerici in `px` e in altre unità CSS
* un valore positivo `offsetX` muove l'ombra a destra e un valore negativo la sposta a sinistra
* un valore positivo `offsetY` sposta l'ombra verso il basso e un valore negativo la sposta verso l'alto
* se vuoi un valore di zero (`0`) per `offsetX`, `offsetY` o entrambi, non devi aggiungere un unità. Ogni browser capisce che zero significa nessun cambiamento.

L'altezza e la larghezza dell'ombra sono determinate dall'altezza e dalla larghezza dell'elemento a cui si applica. Puoi anche usare un valore opzionale `spreadRadius` per estendere la portata dell'ombra. Più informazioni su questo in seguito.

Inizia aggiungendo un'ombra semplice al pennarello rosso.

Nella regola CSS `.red`, aggiungi la proprietà `box-shadow` con i valori `5px` per `offsetX`, `5px` per `offsetY` e `red` per `color`.

# --hints--

La regola CSS `.red` dovrebbe avere una proprietà shorthand `box-shadow` con il valore `5px 5px red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
