---
id: 61b093179e7fc020b43b1bb3
title: Step 73
challengeType: 0
dashedName: step-73
---

# --description--

L'<dfn>opacità</dfn> descrive quanto un oggetto è opaco, o non trasparente. Per esempio, un muro è opaco e la luce non può passarvi attraverso. Un bicchiere invece è molto più trasparente, e puoi vedere attraverso il vetro dall'altra parte.

Con la proprietà CSS `opacity`, puoi controllare l'opacità di un elemento. Con il valore `0`, o 0%, l'elemento sarà completamente trasparente mentre il valore `1.0`, o 100%, l'elemento sarà completamente opaco come lo è di default.

La regola CSS `.sleeve` dovrebbe avere una proprietà `opacity` con il valore `0.5`.

# --hints--

La regola CSS `.sleeve` dovrebbe avere una proprietà `opacity` con il valore `0.5`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
