---
id: 61b093219e7fc020b43b1bb4
title: Step 74
challengeType: 0
dashedName: step-74
---

# --description--

Puoi impostare l'opacità di un elemento anche utilizzando il componente <dfn>alpha channel</dfn>. Simile alla proprietà `opacity`, l'alfa channel controlla la trasparenza/opacità di un colore.

Hai già impostato l'opacità del corpo con il nome di un colore e la proprietà `opacity`, ma puoi aggiungere un alfa channel alle altre proprietà di colore CSS.

All'interno della regola `.sleeve`, rimuovi la proprietà `opacity` e il relativo valore.

# --hints--

La regola CSS `.sleeve` non dovrebbe avere una proprietà `opacity` con il relativo valore.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
  opacity: 0.5;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
