---
id: 61b30a286c228bd0c493c09a
title: Step 89
challengeType: 0
dashedName: step-89
---

# --description--

E se volessi rendere l'ombra più estesa? Puoi farlo con il valore opzionale `spreadRadius`:

```css
box-shadow: offsetX offsetY blurRadius spreadRadius color;
```

Come per `blurRadius`, anche nel caso di `spreadRadius`, se non incluso, il valore predefinito è `0`.

Fai pratica aggiungendo un'ombra di 5 pixel direttamente attorno al pennarello blu.

Nella regola CSS `.blue`, aggiungi la proprietà `box-shadow` con i valori `0` per `offsetX`, `0` per `offsetY`, `0` per `blurRadius`, `5px` per `spreadRadius` e `blue` per `color`.

# --hints--

La regola CSS `.blue` dovrebbe avere una proprietà shorthand `box-shadow` con il valore `0 0 0 5px blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.boxShadow === 'blue 0px 0px 0px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
  box-shadow: 5px 5px 5px green;
}

--fcc-editable-region--
.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}
--fcc-editable-region--

```
