---
id: 66d6a7a3e1aa411e94bf2346
title: Build a House Painting
challengeType: 14
dashedName: build-a-house-painting
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `#house` `div`.
1. Your `#house` should have a `position` set to `relative` so its children can be positioned with respect to it.
1. Your `#house` should have a width of `500px` and a height of `400px`.
1. You should have a `#chimney`, `#decor`, `#window-1`, `#window-2`, and `#door` `div`s within the `#house`.
1. All of the immediate children of the `#house` should have a `position` of `absolute`.
1. All of the immediate children of the `#house` should have a width and height set.
1. Your `#chimney` should have a top value that puts it at the top of your `#house`.
1. Your `#chimney` should have a `z-index` that puts it behind the house.

**Note:** Be sure to link your stylesheet in your HTML to apply your CSS.

# --hints--

You should have a `#house` `div`.

```js
assert.exists(document.querySelector("div#house"));
```

Your `#house` should have a `position` set to `relative`.

```js
assert.equal(getComputedStyle(document.querySelector('#house'))?.getPropertyValue("position"), "relative");
```

Your `#house` should have a width of `500px`.

```js
console.log(getComputedStyle(document.querySelector('#house'))?.getPropertyValue("width"))
assert.equal(getComputedStyle(document.querySelector("#house"))?.getPropertyValue("width"), "500px")
```

Your `#house` should have a height of `400px`.

```js
assert.equal(getComputedStyle(document.querySelector('#house'))?.getPropertyValue("height"), "400px")
```

You should have a `#chimney` `div` within the `#house`.

```js
assert.exists(document.querySelector("div#house  div#chimney"));
```

`#chimney` should have the `position` property set to `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("position"), "absolute");
```

`#chimney` should have a width and height set.

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.notEqual(chimney.getPropertyValue("width"), "0px");
assert.notEqual(chimney.getPropertyValue("height"), "0px");
```

You should have a `#decor` `div` within the `#house`.

```js
assert.exists(document.querySelector("div#house  div#decor"));
```

`#decor` should have the `position` property set to `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#decor"))?.getPropertyValue("position"), "absolute");
```

`#decor` should have a width and height set.

```js
const decor = getComputedStyle(document.querySelector("#decor"));
assert.notEqual(decor.getPropertyValue("width"), "0px");
assert.notEqual(decor.getPropertyValue("height"), "0px");
```

You should have a `#window-1` `div` within the `#house`.

```js
assert.exists(document.querySelector("div#house  div#window-1"));
```

`#window-1` should have the `position` property set to `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#window-1"))?.getPropertyValue("position"), "absolute");
```

`#window-1` should have a width and height set.

```js
const window1 = getComputedStyle(document.querySelector("#window-1"));
assert.notEqual(window1.getPropertyValue("width"), "0px");
assert.notEqual(window1.getPropertyValue("height"), "0px");
```

You should have a `#window-2` `div` within the `#house`.

```js
assert.exists(document.querySelector("div#house  div#window-2"));
```

`#window-2` should have the `position` property set to `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#window-2"))?.getPropertyValue("position"), "absolute");
```

`#window-2` should have a width and height set.

```js
const window2 = getComputedStyle(document.querySelector("#window-2"));
assert.notEqual(window2.getPropertyValue("width"), "0px");
assert.notEqual(window2.getPropertyValue("height"), "0px");
```

You should have a `#door` `div` within the `#house`.

```js
assert.exists(document.querySelector("div#house  div#door"));
```

`#door` should have the `position` property set to `absolute`.

```js
assert.equal(getComputedStyle(document.querySelector("#door"))?.getPropertyValue("position"), "absolute");
```

`#door` should have a width and height set.

```js
const door = getComputedStyle(document.querySelector("#door"));
assert.notEqual(door.getPropertyValue("width"), "0px");
assert.notEqual(door.getPropertyValue("height"), "0px");
```

Your `#chimney` should have a top value that puts it at the top of your `#house`.

```js
const chimney = getComputedStyle(document.querySelector("#chimney"));
assert.equal(Number(chimney?.getPropertyValue("top").replace("px", "")), - Number(chimney?.getPropertyValue("height").replace("px", "")));
```

Your `#chimney` should have a `z-index` that puts it behind the house.

```js
const houseZIndex = getComputedStyle(document.querySelector("#house"))?.getPropertyValue("z-index");
const chimneyZIndex = getComputedStyle(document.querySelector("#chimney"))?.getPropertyValue("z-index");
if (houseZIndex === "auto") {
    assert.isBelow(Number(chimneyZIndex), 0)
} else {
    assert.isBelow(Number(chimneyZIndex), Number(houseZIndex));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>House Painting</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div id="house">
        <div id="chimney">
          <div id="smoke"></div>
        </div>
        <div id="decor"></div>
        <div class="window" id="window-1"></div>
        <div class="window" id="window-2"></div>
        <div id="door">
          <div id="door-knob"></div>
        </div>
        <div id="welcome">WELCOME</div>
    </div>
</body>
</html>
```

```css
* {
    box-sizing: border-box;
  }

  body {
    margin: 0;
    padding: 0;
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    background: linear-gradient(
      #b3e6ff,
      #b3e6ff 60%,
      green 60%
    );
  }

  #house {
    flex: 0 0 auto;
    width: 500px;
    height: 400px;
    position: relative;
    background-color: #ff9980;
    border: 6px solid #b35900;
  }

  #chimney {
    position: absolute;
    width: 90px;
    height: 100px;
    top: -100px;
    left: 65%;
    background: repeating-linear-gradient(
      #e6e6e6,
      #e6e6e6 10%,
      black 10%,
      black 12%
    );
    border: 2px solid black;
    z-index: -1;
  }

  #smoke {
    width: 40px;
    height: 40px;
    background-color: rgba(0,0,0,0.3);
    position: absolute;
    top: -20px;
    left: 20px;
    border-radius: 50%;
    box-shadow: 0 -10px 10px 10px rgba(0,0,0,0.3);
    border: none;
  }

  #decor {
    width: 100%;
    height: 110px;
    background-color: green;
    position: absolute;
    top: 0;
    background: repeating-linear-gradient(
      45deg,
      #b35900,
      #b35900 2%,
      transparent 2%,
      transparent 5%
    ), repeating-linear-gradient(
      -45deg,
      #b35900,
      #b35900 2%,
      #ff9980 2%,
      #ff9980 5%
    )
  }

  .window {
    width: 100px;
    height: 100px;
    background: linear-gradient(
      to right,
      #ffffb3,
      #ffffb3 48%,
      #b35900 48%,
      #b35900 52%,
      #ffffb3 52%,
      #ffffb3 100%
    );
    position: absolute;
    top: 40%;
    border: 6px solid #b35900;
  }

  #window-1 {
    left: 7.5%;
  }

  #window-2 {
    right: 7.5%;
  }

  #door {
    width: 130px;
    height: 180px;
    position: absolute;
    bottom: 0%;
    left: 185px;
    background-color: #e6e6e6;
    border: 6px solid #b35900;
    border-bottom: none;
  }

  #door-knob {
    width: 20px;
    height: 20px;
    background-color: #b35900;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    right: 5%;
  }

  #welcome {
    font-family: sans-serif;
    font-weight: bold;
    width: 130px;
    height: 40px;
    position: absolute;
    bottom: -47px;
    left: 178px;
    background-color: #85e085;
    border: 2px solid black;
    transform: skewX(-20deg);
    display: flex;
    align-items: center;
    justify-content: center;
  }
```
