---
id: 695858b856e7d10a3f5cc4d3
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will build a kitchen inventory tracker.

Get started by defining a variable called `eggs` and setting it to at least `1`.

# --hints--

You should have a variable called `eggs`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("eggs")`)) })
```

Your `eggs` variable should have an integer (whole number) value that is at least `1`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("eggs").is_integer() and _Node(_code).get_variable("eggs") >= 1`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
flour = 2
sugar = 3
```
