---
id: 6a3e71a0a3a44260e675b157
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

To solve the mystery of your missing egg and actually cook breakfast, you are going to write a new function. This function will act as your master recipe and will call other functions to get the job done.

Below your `print(eggs)` statement, define an empty function named `make_fried_egg` that accepts one parameter: `eggs`. 

This function will also introduce a concept called **name shadowing**. If you name a parameter the exact same name as a global variable (like `eggs`), the function will use its own local parameter `eggs` instead of the global one. The local variable "shadows" the global one!

Use the `pass` keyword inside the body for now.

# --hints--

You should define a function named `make_fried_egg`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_function("make_fried_egg")`)
  )
})
```

Your `make_fried_egg` function should have `eggs` as its parameter.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_function("make_fried_egg").has_args("eggs")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(eggs, 1)
print(eggs)
        
--fcc-editable-region--

--fcc-editable-region--
```
