---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Time to create a function to check what you have in your kitchen. In Python, you define functions using the `def` keyword, followed by the function name, parentheses `()`, and a colon `:`.

Here is an example of how to define a simple function:

```py
def my_function():
    pass
```

Define an empty function named `check_kitchen_stock()`. Use the `pass` keyword inside the function body for now so that Python doesn't throw a syntax error.

# --hints--

You should define a function named `check_kitchen_stock`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_function("check_kitchen_stock")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
