---
id: 6a3e58891e1ff13d07f4ab48
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Python reads the `def` block and remembers the function, but it does not run the code inside it immediately.

To execute the code and see your output, you need to **call** (or invoke) the function. Do this by typing the function's name followed by parentheses.

For example:

```py
my_function()
```

At the bottom of your code, outside of the function, call your `check_kitchen_stock` function.

# --hints--

You should call the `check_kitchen_stock` function at the end of your code. Make sure it is not indented, so it runs globally.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("check_kitchen_stock()")`))
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')
    
--fcc-editable-region--

--fcc-editable-region--
```
