---
id: 6a3e6c0eb5d050ce07781709
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

If the `if` condition evaluates to `False` (meaning you *do* have enough eggs), the function will skip the `if` block and continue running the rest of the code.

Now you need to process your data and return a new computed value. 

Outside of your `if` block (unindented, at the same level as the `if` keyword), use a `print()` call and an f-string to output: `'{num_eggs} egg(s) used out of {eggs} available.'`. 

Then, on the next line, `return` the remaining number of eggs by subtracting `num_eggs` from `eggs`.

# --hints--

You should have a print statement that prints `f'{num_eggs} egg(s) used out of {eggs} available.'` in your `use_eggs` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").has_stmt("print(f'{num_eggs} egg(s) used out of {eggs} available.')")`)
  )
})
```

You should return the remaining number of eggs by subtracting `num_eggs` from `eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").has_return("eggs - num_eggs")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs
--fcc-editable-region--
    
--fcc-editable-region--
```
