---
id: 6a3e6fd08125dfa615cb431b
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Now that your `use_eggs` function is complete, it is time to use it. When you call a function that requires **parameters**, you have to provide values for them inside the parentheses. These values are called **arguments**.

At the bottom of your code, call your `use_eggs` function and pass it two arguments: the global `eggs` variable, and the number `1` (representing the number of eggs you want to use).

# --hints--

You should call the `use_eggs` function and pass `eggs` and `1` as arguments.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("use_eggs(eggs, 1)")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs
        
--fcc-editable-region--
        
--fcc-editable-region--
```
