---
id: 6a3e6c0eb5d050ce07781709
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

If the `if` condition evaluates to `False` (meaning you *do* have enough eggs), the function will skip the `if` block and continue running the rest of the code.

Now you need to process your data and return a new computed value. 

Outside of your `if` block (unindented, at the same level as the `if` keyword), use a `print()` call and an f-string to output: `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'`.

Then, on the next line, `return` the remaining number of eggs by subtracting `eggs_to_use` from `available_eggs`.

# --hints--

You should have a print statement that prints `f'{eggs_to_use} egg(s) used out of {available_eggs} available.'` in your `use_eggs` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").has_stmt("print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')")`)
  )
})
```

You should return the remaining number of eggs by subtracting `eggs_to_use` from `available_eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").has_return("available_eggs - eggs_to_use")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs
--fcc-editable-region--
    
--fcc-editable-region--
```
