---
id: 6a3e58891e1ff13d07f4ab48
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

You have written the recipe, but right now it is just sitting in the cookbook. Python reads the `def` block and remembers the function, but it does not run the code inside it immediately.

To actually execute the code and see your print statements, you need to "call" (or invoke) the function. You can do this by typing the function's name followed by parentheses.

At the bottom of your code, outside of the function, call your `take_inventory` function.

# --hints--

You should call the `take_inventory` function at the end of your code. Make sure it is not indented, so it runs globally.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("take_inventory()")`))
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')
    
--fcc-editable-region--

--fcc-editable-region--
```
