---
id: 6a3ea8a55317ba7a08e21d5b
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Now you can finally answer the question from earlier! Why did the global number of eggs stay at `1`? Because when a function returns a value, it doesn't automatically overwrite your variables. You have to explicitly save that new value!

To update your actual kitchen inventory, you need to overwrite your global `eggs` variable with the new number of eggs returned by your `make_fried_egg` function.

At the very bottom of your code, call the `make_fried_egg` function and pass `eggs` as the argument. Assign the result of this function call back to the global `eggs` variable.

# --hints--

You should reassign the `eggs` variable to the result of `make_fried_egg(eggs)`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_variables("eggs")[-1].is_equivalent("eggs = make_fried_egg(eggs)")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)
        
def make_fried_egg(eggs):   
    available_eggs = use_eggs(eggs, 1)
    has_enough_eggs = available_eggs >= eggs  

    if has_enough_eggs:
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
