---
id: 6a3e71a0a3a44260e675b157
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

To solve the mystery of our missing egg and actually cook breakfast, we are going to write a new function. This function will act as our master recipe and will call other functions to get the job done.

Below your `print(eggs)` statement, define an empty function named `make_fried_egg` that accepts one parameter: `eggs`. 

Use the `pass` keyword inside the body for now.

# --hints--

You should define a function named `make_fried_egg`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg")`)
  )
})
```

Your make_fried_egg function should have eggs as its parameter.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_function("make_fried_egg").has_parameters(["eggs"])
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)
        
--fcc-editable-region--
        
--fcc-editable-region--
```
