---
id: 6a3e70b94f67d9ee3ecd8486
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

You just called `use_eggs(available_eggs, 1)`, which should have used one egg. So, the number of eggs in our kitchen should be `0` now, right?

You will find out! At the bottom of your code, use the `print()` function to print the value of the `available_eggs` variable. 

When you run the code, you will notice something surprising: the output is still `1`! Why didn't the number of eggs change? You will explore this mystery in the next few steps.

# --hints--

You should print the `available_eggs` variable at the end of your code.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("print(available_eggs)")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
        
--fcc-editable-region--

--fcc-editable-region--
```
