---
id: 6a3e524fc3394ef52ed5b797
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Functions have their own "local scope". This means variables created inside a function are like secret ingredients: they can only be accessed inside that specific function.

Inside your `take_inventory` function, replace the `pass` keyword with a new variable named `total_items` and assign it the sum of `available_eggs`, `available_flour`, and `available_sugar`.

# --hints--

You should define a variable named `total_items` inside your `take_inventory` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_variable("total_items")`)
  )
})
```

Your `total_items` variable should be the sum of `available_eggs`, `available_flour`, and `available_sugar`.

```js
({
  test: () => assert(
    runPython(`
      variable = _Node(_code).find_function("take_inventory").find_variables("total_items")[0]
      any(
        variable.is_equivalent(code)
        for code in [
          "total_items = available_eggs + available_flour + available_sugar",
          "total_items = available_flour + available_eggs + available_sugar",
          "total_items = available_sugar + available_flour + available_eggs",
          "total_items = available_eggs + available_sugar + available_flour",
          "total_items = available_flour + available_sugar + available_eggs",
          "total_items = available_sugar + available_eggs + available_flour"
        ]
      )
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
--fcc-editable-region--    
    pass
--fcc-editable-region--
```
