---
id: 6a3e524fc3394ef52ed5b797
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Functions have their own "local scope". This means variables created inside a function are like ingredients kept in that kitchen's pantry: they can only be accessed inside that specific function.

However, functions are not completely isolated. They can read "global" variables defined outside their scope. They can use those variables to perform calculations, make decisions, or create new local variables.

Inside your `check_kitchen_stock` function, replace the `pass` keyword with a new variable named `total_items` and assign it the sum of `available_eggs`, `available_flour`, and `available_sugar`.

# --hints--

You should define a variable named `total_items` inside your `check_kitchen_stock` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("check_kitchen_stock").has_variable("total_items")`)
  )
})
```

Your `total_items` variable should be the sum of `available_eggs`, `available_flour`, and `available_sugar`.

```js
({
  test: () => assert(
    runPython(`
      variable = _Node(_code).find_function("check_kitchen_stock").find_variables("total_items")[0]
      any(
        variable.is_equivalent(code)
        for code in [
          "total_items = available_eggs + available_flour + available_sugar",
          "total_items = available_flour + available_eggs + available_sugar",
          "total_items = available_sugar + available_flour + available_eggs",
          "total_items = available_eggs + available_sugar + available_flour",
          "total_items = available_flour + available_sugar + available_eggs",
          "total_items = available_sugar + available_eggs + available_flour"
        ]
      )
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
--fcc-editable-region--    
    pass
--fcc-editable-region--
```
