---
id: 6a3e5475c35dbc138d77c6f4
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now that you have calculated your total, you should display the inventory. In Python, you can use the `print()` function along with **f-strings** (formatted string literals) to easily output text and variables together.

For example:

```py
name = "Alice"
print(f"Hello, {name}!") # Output: Hello, Alice!
```

Inside your `take_inventory` function, below your `total_items` variable, add these exact four print statements (use the single quotes in the print statements):

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {available_eggs} eggs')`
3. `print(f'- {available_flour} flour')`
4. `print(f'- {available_sugar} sugar')`

# --hints--

You should have the `print(f'The kitchen has {total_items} total items:')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'The kitchen has {total_items} total items:')")`)
  )
})
```

You should have the `print(f'- {available_eggs} eggs')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'- {available_eggs} eggs')")`)
  )
})
```

You should have the `print(f'- {available_flour} flour')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'- {available_flour} flour')")`)
  )
})
```

You should have the `print(f'- {available_sugar} sugar')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'- {available_sugar} sugar')")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar  
--fcc-editable-region--
    
--fcc-editable-region--
```
