---
id: 6a3ea771b46a239c67628a07
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Your `make_fried_egg` function has successfully used an egg and checked if the cooking was successful. The last thing it needs to do is give back the updated egg count so the rest of the program knows about it.

At the very end of your `make_fried_egg` function (make sure it is not indented inside the `if` or `else` blocks), use the `return` keyword to return your `available_eggs` variable.

# --hints--

You should return `available_eggs` at the end of your `make_fried_egg` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").has_return("available_eggs")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)
        
def make_fried_egg(available_eggs):   
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')
       
--fcc-editable-region--
    
--fcc-editable-region--
```
