---
id: 6a3e70b94f67d9ee3ecd8486
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

You just called `use_eggs(eggs, 1)`, which should have used one egg. So, the number of eggs in our kitchen should be `0` now, right?

Let's find out! At the bottom of your code, use the `print()` function to print the value of the `eggs` variable. 

When you run the code, you will notice something surprising: the output is still `1`! Why didn't the number of eggs change? We will explore this mystery in the next few steps.

# --hints--

You should print the `eggs` variable at the end of your code.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("print(eggs)")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
        
--fcc-editable-region--
        
--fcc-editable-region--
```
