---
id: 6a3e4ae26f2b7e4b5b026cf8
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Now it's time to create a recipe (a function) to check what we have in our kitchen. In Python, functions are defined using the `def` keyword. 

Define an empty function named `take_inventory()`. Use the `pass` keyword inside the function body for now so that Python doesn't throw a syntax error.

# --hints--

You should define a function named `take_inventory`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

--fcc-editable-region--

--fcc-editable-region--
```
