---
id: 6a3eaa2caac22f67811fe40f
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

It's time for the moment of truth! You have successfully updated your global `eggs` variable with the new value returned by your `make_fried_egg` function.

You can check your kitchen inventory one last time to verify that the number of eggs has actually decreased. 

At the very end of your code, call the `take_inventory()` function again. When you run the code, you should see that the kitchen now has `0` eggs! 

Congratulations on completing the Python Functions and Scope workshop!

# --hints--

You should call the `take_inventory()` function a second time at the end of your code.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_calls("take_inventory")[1].is_equivalent("take_inventory()")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)
        
def make_fried_egg(eggs):   
    available_eggs = use_eggs(eggs, 1)
    has_enough_eggs = available_eggs >= eggs  

    if has_enough_eggs:
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')

    return available_eggs

eggs = make_fried_egg(eggs)    

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs
    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)

def make_fried_egg(eggs):
    available_eggs = use_eggs(eggs, 1)
    has_enough_eggs = available_eggs >= eggs
    if has_enough_eggs:
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')
    return available_eggs

eggs = make_fried_egg(eggs)
take_inventory()
```
