---
id: 6a3e5475c35dbc138d77c6f4
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now that you have calculated your total, you should display the inventory. In Python, you can use the `print()` function along with **f-strings** (formatted string literals) to easily output text and variables together.

An f-string allows you to insert variables directly into a string. You just place an `f` right before the opening quote, and wrap the variable names in curly braces `{}`. Notice how your function can read both the local variable (`total_items`) and the global variables defined outside the function (`eggs`, `flour`, `sugar`)!

Inside your `take_inventory` function, below your `total_items` variable, add these exact four print statements (use the single quotes in the print statements):

1. `print(f'The kitchen has {total_items} total items:')`
2. `print(f'- {eggs} eggs')`
3. `print(f'- {flour} flour')`
4. `print(f'- {sugar} sugar')`

# --hints--

You should have the `print(f'The kitchen has {total_items} total items:')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'The kitchen has {total_items} total items:')")`)
  )
})
```

You should have the `print(f'- {eggs} eggs')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'- {eggs} eggs')")`)
  )
})
```

You should have the `print(f'- {flour} flour')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'- {flour} flour')")`)
  )
})
```

You should have the `print(f'- {sugar} sugar')` statement inside your function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("take_inventory").has_call("print(f'- {sugar} sugar')")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    
--fcc-editable-region--
    
--fcc-editable-region--
```
