---
id: 6a3e67fc5ef4045201c28219
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Inside your `use_eggs` function, you need to check if you have enough eggs to cook with. You can do this using an `if` statement.

An `if` statement evaluates a condition. If the condition is true, the code inside the `if` block runs. 

Replace the `pass` keyword with an `if` statement that checks if `num_eggs` is greater than (`>`) `eggs`. Inside the `if` block, print the exact string `'The kitchen does not have enough eggs.'`.

# --hints--

You should have an `if` statement that checks if `num_eggs > eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").find_ifs()[0].find_conditions()[0].is_equivalent("num_eggs > eggs")`)
  )
})
```

You should print the exact message when `num_eggs` is greater than `eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").find_ifs()[0].has_call("print('The kitchen does not have enough eggs.')")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
