---
id: 6a3e6ad6b9554e31dcc30cd6
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

If we don't have enough eggs, we shouldn't try to cook! We need a way to stop the function from running any further. 

You can use the `return` keyword to do this. When a function reaches a `return` statement, it immediately stops executing and gives control (and the returned value) back to the place where the function was called.

Inside your `if` block, below the print statement, use the `return` keyword to return the value of `eggs`.

# --hints--

You should return `eggs` inside your `if` statement.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").find_ifs()[0].has_return("eggs")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
