---
id: 6a3e6ad6b9554e31dcc30cd6
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

If you don't have enough eggs, you shouldn't try to cook! You need a way to stop the function from running any further.

You can use the `return` keyword for this. When a function hits a `return` statement, it stops executing immediately and passes the returned value back to where it was called.

For example:

```py
def cook_dish():
    message = "The dish is ready to serve!"
    return message
    print("This line will never run.")

print(cook_dish())  # Output: The dish is ready to serve!
```

Inside your `if` block, below the print statement, use the `return` keyword to return the value of `available_eggs`.

# --hints--

You should return `available_eggs` inside your `if` statement.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs").find_ifs()[0].has_return("available_eggs")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
--fcc-editable-region--
        
--fcc-editable-region--
```
