---
id: 6a3e7284e89ce3367e709331
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now let's solve the mystery from earlier! When our `use_eggs` function returned the remaining number of eggs, we didn't save that new value anywhere. It just disappeared into the void! 

To actually update our inventory, we need to capture that returned value by assigning it to a variable. 

Inside your `make_fried_egg` function, replace the `pass` keyword. First, call `use_eggs(eggs, 1)` and assign the result to a variable named `new_eggs`. Then, create another variable named `has_enough_eggs` and assign it the result of the comparison `new_eggs < eggs`.



# --hints--

You should assign `new_eggs` to the result of `use_eggs(eggs, 1)` inside your `make_fried_egg` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_assignments("new_eggs")[0].is_equivalent("use_eggs(eggs, 1)")`)
  )
})
```

You should assign `has_enough_eggs` to the result of the comparison `new_eggs < eggs` inside your `make_fried_egg` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_assignments("has_enough_eggs")[0].is_equivalent("new_eggs < eggs")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)
        
def make_fried_egg(eggs):        
--fcc-editable-region--
    pass  
--fcc-editable-region--
```
