---
id: 6a3e5ccb99fd6053af9032a1
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Now let's create another function. This time, our function will take "parameters". Parameters are like placeholders for data that you pass into the function when you call it.

This function will also introduce a concept called **name shadowing**. If you name a parameter the exact same name as a global variable (like `eggs`), the function will use its own local parameter `eggs` instead of the global one. The local variable "shadows" the global one!

Below your `take_inventory` function define an empty function named `use_eggs` that accepts two parameters: `eggs` and `num_eggs`, in that order. 

Use the `pass` keyword inside the body for now.

# --hints--

You should define a function named `use_eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs")`)
  )
})
```

Your use_eggs function should have eggs and num_eggs as its parameters.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_function("use_eggs").has_parameters(["eggs", "num_eggs"])
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

--fcc-editable-region--

--fcc-editable-region--
```
