---
id: 6a3ea8a55317ba7a08e21d5b
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Now you can finally answer the question from earlier! Why did the global number of eggs stay at `1`? Because when a function returns a value, it doesn't automatically overwrite your variables. You have to explicitly save that new value!

You have noticed that you are using `available_eggs` variable in the `make_fried_egg` function, but it is not the same `available_eggs` variable that you defined at the top of your code. The `available_eggs` variable inside the function is a **local variable** that only exists while the function is running. When the function finishes, that local variable disappears and your global `available_eggs` variable remains unchanged.

To update your actual kitchen inventory, you need to overwrite your global `available_eggs` variable with the new number of eggs returned by your `make_fried_egg` function.

At the very bottom of your code, call the `make_fried_egg` function and pass `available_eggs` as the argument. Assign the result of this function call back to the global `available_eggs` variable.

# --hints--

You should reassign the `available_eggs` variable to the result of `make_fried_egg(available_eggs)`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_variables("available_eggs")[-1].is_equivalent("available_eggs = make_fried_egg(available_eggs)")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)
        
def make_fried_egg(available_eggs):   
    has_enough_eggs = available_eggs >= 1

    if has_enough_eggs:
        available_eggs = use_eggs(available_eggs, 1)
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')
    
    return available_eggs

--fcc-editable-region--

--fcc-editable-region--
```
