---
id: 6a3e6fd08125dfa615cb431b
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Now that your `use_eggs` function is complete, you can test it out. When calling a function that requires **parameters**, you pass the specific values inside the parentheses. These values are known as **arguments**.

Here is an example of how to call a function with arguments:

```py
def calculate_sum(a, b):
    print(a + b)
calculate_sum(3, 1) # Output: 4
```    

At the bottom of your code, call your `use_eggs` function and pass it two arguments: the global `available_eggs` variable and the number `1` (representing the number of eggs you want to use).

# --hints--

You should call the `use_eggs` function and pass `available_eggs` and `1` as arguments.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("use_eggs(available_eggs, 1)")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use
        
--fcc-editable-region--

--fcc-editable-region--
```
