---
id: 6a3e73eabddd614f91cc235a
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Now that you have checked if you have enough eggs, you can use that information to decide what to do.

Inside your `make_fried_egg` function, use an `if/else` statement to check the `has_enough_eggs` variable. 

- If `has_enough_eggs` is `True`, print `'Made a fried egg. Yummy!'`.
- Otherwise (the `else` block), print `'Could not make a fried egg. Not enough eggs!'`.

# --hints--

You should use an `if` statement to check `has_enough_eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")`)
  )
})
```

You should have an `else` block that prints the correct message when `has_enough_eggs` is False.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_elses()[0].find_prints()[0].is_equivalent("'Could not make a fried egg. Not enough eggs!'")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)
        
def make_fried_egg(eggs):   
    new_eggs = use_eggs(eggs, 1)
    has_enough_eggs = new_eggs < eggs  
       
--fcc-editable-region--

--fcc-editable-region--
```
