---
id: 6a3ea771b46a239c67628a07
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Your `make_fried_egg` function has successfully used an egg and checked if the cooking was successful. The last thing it needs to do is give back the updated egg count so the rest of the program knows about it.

At the very end of your `make_fried_egg` function (make sure it is not indented inside the `if` or `else` blocks), use the `return` keyword to return your `available_eggs` variable.

# --hints--

You should return `available_eggs` at the end of your `make_fried_egg` function.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").has_return("available_eggs")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(eggs, num_eggs):
    if num_eggs > eggs:
        print('The kitchen does not have enough eggs.')
        return eggs

    print(f'{num_eggs} egg(s) used out of {eggs} available.')
    return eggs - num_eggs

use_eggs(eggs, 1)
print(eggs)
        
def make_fried_egg(eggs):   
    available_eggs = use_eggs(eggs, 1)
    has_enough_eggs = available_eggs >= eggs

    if has_enough_eggs:
        print('Made a fried egg. Yummy!')
    else:
        print('Could not make a fried egg. Not enough eggs!')
       
--fcc-editable-region--
    
--fcc-editable-region--
```
