---
id: 6a3e73eabddd614f91cc235a
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Now that you have checked if you have enough eggs, you can use that information to decide what to do.

Inside your `make_fried_egg` function, use an `if/else` statement to check the `has_enough_eggs` variable. 

- If `has_enough_eggs` is True, reassign your available_eggs variable to the result of `use_eggs(eggs, 1)` and `print('Made a fried egg. Yummy!')`.
- Otherwise (the `else` block), `print('Could not make a fried egg. Not enough eggs!')`.

# --hints--

You should use an `if` statement to check `has_enough_eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")`)
  )
})
```

Inside your `if` block, you should reassign `available_eggs` to the result of `use_eggs(eggs, 1)`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(eggs, 1)")`)
  )
})
```

You should print the exact message when `has_enough_eggs` is True.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")`)
  )
})
```

You should have an `else` block that prints the correct message when `has_enough_eggs` is False.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")`)
  )
})
```

# --seed--

## --seed-contents--

```py
eggs = 1
flour = 2
sugar = 3

def take_inventory():
    total_items = eggs + flour + sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {eggs} eggs')
    print(f'- {flour} flour')
    print(f'- {sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(eggs, 1)
print(eggs)
        
def make_fried_egg(eggs):
    has_enough_eggs = eggs >= 1
    available_eggs = eggs
      
--fcc-editable-region--
    
--fcc-editable-region--
```
