---
id: 6a3e2b9b7eeca084dd4fb857
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

Welcome to the Python Functions and Scope workshop! You are going to learn about how functions work using a kitchen theme.

To get started, you should define some basic data for your kitchen inventory. Create three variables: `available_eggs`, `available_flour`, and `available_sugar`, and assign them the values `1`, `2`, and `3` respectively.

# --hints--

You should define a variable named `available_eggs` and assign it the value of `1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_variables("available_eggs")[0].is_equivalent("available_eggs = 1")`)
  )
})
```

You should define a variable named `available_flour` and assign it the value of `2`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_variables("available_flour")[0].is_equivalent("available_flour = 2")`)
  )
})
```

You should define a variable named `available_sugar` and assign it the value of `3`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_variables("available_sugar")[0].is_equivalent("available_sugar = 3")`)
  )
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
