---
id: 6a3e2b9b7eeca084dd4fb857
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

Welcome to the Python Functions and Scope workshop! You are going to learn about how functions work using a kitchen theme.

To get started, You shoulddefine some basic data for your kitchen inventory. Create three variables: `eggs`, `flour`, and `sugar`, and assign them the values `1`, `2`, and `3` respectively.

# --hints--

You should define a variable named `eggs` and assign it the value of `1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_assignments("eggs")[0].has_value("1")`)
  )
})
```

You should define a variable named `flour` and assign it the value of `2`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_assignments("flour")[0].has_value("2")`)
  )
})
```

You should define a variable named `sugar` and assign it the value of `3`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_assignments("sugar")[0].has_value("3")`)
  )
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
