---
id: 6a3e7284e89ce3367e709331
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now you are going to solve the mystery from earlier! When your `use_eggs` function returned the remaining number of eggs, you didn't save that new value anywhere. It just disappeared into the void! 

To actually update your inventory, you need to capture that returned value by assigning it to a variable. 

Inside your `make_fried_egg` function, replace the `pass` keyword. First, create a variable named `has_enough_eggs` and assign it the result of checking if `available_eggs` is greater than or equal to 1.



# --hints--

You should assign `has_enough_eggs` to the result of checking if `available_eggs` is greater than or equal to `1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_variables("has_enough_eggs")[0].is_equivalent("has_enough_eggs = available_eggs >= 1")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

take_inventory()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)
        
def make_fried_egg(available_eggs):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
