---
id: 6a3e5ccb99fd6053af9032a1
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Now create another function. This time, your function will take "parameters". Parameters are like placeholders for data that you pass into the function when you call it.

Below your `take_inventory` function, define an empty function named `use_eggs` that accepts two parameters: `available_eggs` and `eggs_to_use`, in that order. 

This function will also introduce a concept called **name shadowing**. If you give a parameter the exact same name as a global variable (like `available_eggs`), the function will use its own local parameter `available_eggs` instead of the global one. The local variable "shadows" the global one!

Use the `pass` keyword inside the body for now.

# --hints--

You should define a function named `use_eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("use_eggs")`)
  )
})
```

Your `use_eggs` function should have `available_eggs` and `eggs_to_use` as its parameters.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_function("use_eggs").has_args("available_eggs, eggs_to_use")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def take_inventory():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

take_inventory()

--fcc-editable-region--

--fcc-editable-region--
```
