---
id: 6a3e73eabddd614f91cc235a
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Now you are going to solve the mystery from earlier! When your `use_eggs` function returned the remaining number of eggs, you didn't save that new value anywhere. It just disappeared into the void! 

To actually update your inventory, you need to capture that returned value by assigning it to a variable.

Inside your `make_fried_egg` function, use an `if/else` statement to check the `has_enough_eggs` variable. 

- If `has_enough_eggs` is `True`, reassign your `available_eggs` variable to the result of `use_eggs(available_eggs, 1)` and `print('Made a fried egg. Yummy!')`.

- Otherwise (the `else` block), `print('Could not make a fried egg. Not enough eggs!')`.

# --hints--

You should use an `if` statement to check `has_enough_eggs`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_conditions()[0].is_equivalent("has_enough_eggs")`)
  )
})
```

Inside your `if` block, you should reassign `available_eggs` to the result of `use_eggs(available_eggs, 1)`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].find_variables("available_eggs")[0].is_equivalent("available_eggs = use_eggs(available_eggs, 1)")`)
  )
})
```

You should print the exact message ('Made a fried egg. Yummy!') when `has_enough_eggs` is True.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[0].has_call("print('Made a fried egg. Yummy!')")`)
  )
})
```

You should have an `else` block that prints the correct message ('Could not make a fried egg. Not enough eggs!') when `has_enough_eggs` is False.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("make_fried_egg").find_ifs()[0].find_bodies()[1].has_call("print('Could not make a fried egg. Not enough eggs!')")`)
  )
})
```

# --seed--

## --seed-contents--

```py
available_eggs = 1
available_flour = 2
available_sugar = 3

def check_kitchen_stock():
    total_items = available_eggs + available_flour + available_sugar
    print(f'The kitchen has {total_items} total items:')
    print(f'- {available_eggs} eggs')
    print(f'- {available_flour} flour')
    print(f'- {available_sugar} sugar')

check_kitchen_stock()

def use_eggs(available_eggs, eggs_to_use):
    if eggs_to_use > available_eggs:
        print('The kitchen does not have enough eggs.')
        return available_eggs

    print(f'{eggs_to_use} egg(s) used out of {available_eggs} available.')
    return available_eggs - eggs_to_use

use_eggs(available_eggs, 1)
print(available_eggs)
        
def make_fried_egg(available_eggs):
    has_enough_eggs = available_eggs >= 1
      
--fcc-editable-region--
    
--fcc-editable-region--
```
