---
id: 5dc1798ff86c76b9248c6eb3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Gli elementi d'intestazione da `h1` a `h6` sono utilizzati per caratterizzare l'importanza del contenuto che si trova al di sotto di essi. Più il numero è basso, maggiore è l'importanza, quindi gli elementi `h2` hanno meno importanza rispetto agli elementi `h1`. Usa un solo elemento `h1` per pagina e posiziona le intestazioni d'importanza inferiore sotto quelle di maggiore importanza.

Sotto l'elemento `h1`, aggiungi un elemento `h2` con il testo:

`Cat Photos`

# --hints--

L'elemento `h1` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('h1'));
```

L'elemento `h1` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/h1\>/));
```

Dovresti avere un solo elemento `h1`. Rimuovi quelli di troppo.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Il testo dell'elemento `h1` dovrebbe essere 'CatPhotoApp'. Hai omesso il testo o hai un refuso.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

L'elemento `h2` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('h2'));
```

L'elemento `h2` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere`/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/h2\>/));
```

Il testo dell'elemento `h2` dovrebbe essere 'Cat Photos'. Inserisci solo il testo `Cat Photos` tra i tag di apertura e chiusura `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

L'elemento `h2` dovrebbe essere al di sotto dell'elemento `h1`. L'elemento `h1` ha maggiore importanza e deve essere al di sopra dell'elemento `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

