---
id: 5dc17d3bf86c76b9248c6eb4
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Gli elementi `p` vengono usati per creare dei paragrafi di testo sui siti web. Crea un elemento `p` sotto l'elemento `h2` e dagli il seguente testo:

`See more cat photos in our gallery.`

# --hints--

Il tuo elemento `p` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<elementName>`.

```js
assert(document.querySelector('p'));
```

L'elemento `p` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/p\>/));
```

Il testo dell'elemento `p` dovrebbe essere `See more cat photos in our gallery.` Hai omesso il testo o hai un refuso.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

L'elemento `p` dovrebbe essere al di sotto dell'elemento `h2`. Sono nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

