---
id: 5dc24073f86c76b9248c6ebb
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Gli <dfn>attributi</dfn> HTML sono parole speciali utilizzate all'interno del tag di apertura di un elemento per controllarne il comportamento. L'attributo `src` in un elemento `img` specifica l'URL dell'immagine (dove l'immagine è localizzata). Un esempio di un elemento `img` con un attributo `src` è: `<img src="https://www.example.com/the-image.jpg">`.

Aggiungi un attributo `src` all'elemento `img` esistente con il seguente URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Il codice dovrebbe avere un elemento `img`. Potresti aver rimosso l'elemento `img` o non hai circondato il valore dell'attributo `src` con le virgolette.

```js
assert(document.querySelector('img'));
```

L'elemento `img` dovrebbe avere un attributo `src`. Hai omesso l'attributo o hai un refuso. Assicurati che ci sia uno spazio tra il nome dell'elemento e il nome dell'attributo.

```js
assert(document.querySelector('img').src);
```

L'attributo `src` dell'elemento `img` dovrebbe essere '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Hai omesso l'URL o hai un refuso. Le maiuscole/minuscole nell'URL sono importanti.

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Anche se hai impostato l'URL corretto per l'`src` dell'elemento `img`, è raccomandato circondare sempre il valore di un attributo con le virgolette.

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

