---
id: 5dc24614f86c76b9248c6ebd
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Puoi aggiungere un link a un'altra pagina con l'elemento di ancoraggio (`a`). Ad esempio, `<a href='https://freecodecamp.org'></a>` è il link a `freecodecamp.org`.

Aggiungi un elemento di ancoraggio con il link `https://freecatphotoapp.com` dopo il paragrafo. A questo punto, il link non verrà visualizzato nell’anteprima.

# --hints--

L'elemento di ancoraggio (`a`) dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('a'));
```

L'elemento di ancoraggio (`a`) dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/a\>/));
```

L'elemento di ancoraggio (`a`) dovrebbe essere al di sotto dell'elemento `p`. Sono nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

L'elemento di ancoraggio (`a`) non ha un attributo `href`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

L'elemento di ancoraggio (`a`) dovrebbe contenere il link `https://freecatphotoapp.com`. Hai omesso l'URL o hai un refuso.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Sebbene tu abbia impostato l'attributo `href` dell'elemento di ancoraggio ('a') con il link corretto, è raccomandato circondare sempre il valore di un attributo con virgolette.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

