---
id: 5ddbd81294d8ddc1510a8e56
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Il testo di un link deve essere posizionato tra i tag di apertura e chiusura di un elemento di ancoraggio (`a`). Ad esempio, `<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` è un link con il testo `click here to go to freeCodeCamp.org`.

Aggiungi il testo di ancoraggio `link to cat pictures` all'elemento di ancoraggio. Sarà il testo del link.

# --hints--

L'elemento di ancoraggio (`a`) dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('a'));
```

L'elemento di ancoraggio (`a`) dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/a\>/));
```

Il testo dell'elemento di ancoraggio (`a`) dovrebbe essere `link to cat pictures`. Assicurati di inserire il testo del link tra i tag di apertura e di chiusura dell'elemento di ancoraggio (`a`).

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

