---
id: 5dfa2407b521be39a3de7be1
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Aggiungi un attributo `target` con il valore `_blank` al tag di apertura dell'elemento di ancoraggio (`a`), in modo che il link si apra in una nuova scheda.

# --hints--

L'elemento `p` dovrebbe avere un elemento di ancoraggio annidato (`a`) con il testo `cat photos`. Potresti averlo cancellato o avere un refuso.

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

L'elemento di ancoraggio (`a`) non ha un attributo `target`. Verifica che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert(document.querySelector('a').hasAttribute('target'));
```

Il valore dell'attributo `target` dovrebbe essere `_blank`. Hai omesso il valore o hai un refuso. Ricorda che i valori degli attributi dovrebbero essere circondati da virgolette.

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

