---
id: 5dfa37b9eacea3f48c6300b0
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Dopo l'elemento `h3` con il testo `Things cats love:`, aggiungi un elemento lista non ordinata (`ul`). Nota che a questo punto non verrà visualizzato nulla.

# --hints--

L'elemento `ul` deve avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('ul'));
```

L'elemento `ul` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/ul>/));
```

Ci dovrebbe essere un elemento `ul` proprio sopra il tag di chiusura del secondo elemento `section`.

```js
const secondSectionLastElemNode = $('main > section')[1].lastElementChild;
assert(secondSectionLastElemNode.nodeName === 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

