---
id: 5dfa22d1b521be39a3de7be0
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Trasforma le parole `cat photos` all'interno dell'elemento `p` in un link usando lo stesso valore dell'attributo `href` al di sotto dell'elemento `p`. L'elemento `p` dovrebbe mostrare lo stesso testo nel browser, ma adesso le parole `cat photos` dovrebbero essere un link. Assicurati di rimuovere l'elemento `a` con il testo `cat photos` sulla riga sotto l'elemento `p`.

# --hints--

Il codice dovrebbe contenere solo un elemento di ancoraggio (`a`). Rimuovi gli elementi di ancoraggio di troppo.

```js
assert(document.querySelectorAll('a').length === 1);
```

L'elemento di ancoraggio (`a`) dovrebbe essere annidato all'interno dell'elemento `p`.

```js
assert($('p > a').length);
```

Il testo del link dovrebbe essere `cat photos`. Hai omesso il testo o hai un refuso.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com' &&
    nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Dopo aver annidato l'elemento di ancoraggio (`a`), l'unico contenuto dell'elemento `p` visibile nel browser dovrebbe essere `Click here to view more cat photos.` Controlla il testo, la spaziatura o la punteggiatura sia dell'elemento `p` che dell'elemento di ancoraggio annidato.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

