---
id: 5dfb6250eacea3f48c6300b2
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Dopo la lista non ordinata, aggiungi una nuova immagine con un attributo `src` con un valore di:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

E il valore dell'attributo `alt` di:

`A slice of lasagna on a plate.`

# --hints--

Ci dovrebbe essere un elemento `img` subito dopo il tag di chiusura `</ul>`.

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

La nuova immagine non ha un attributo `alt`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

Il valore dell'attributo `alt` della nuova immagine dovrebbe essere `A slice of lasagna on a plate.` Assicurati che il valore dell'attributo `alt` sia circondato da virgolette.

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

La nuova immagine non ha un attributo `src`. Verifica che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

Il valore dell'attributo `src` della nuova immagine dovrebbe essere `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Assicurati che il valore dell'attributo `src` sia circondato da virgolette.

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Anche se hai inserito l'URL corretto nell'attributo `src` della nuova immagine, è raccomandato inserire sempre il valore di un attributo tra virgolette.

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

