---
id: 5dfb6a35eacea3f48c6300b4
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

Un elemento di didascalia della figura (`figcaption`) viene usato per aggiungere una didascalia per descrivere l'immagine contenuta nell'elemento `figure`. Ad esempio, `<figcaption>A cute cat</figcaption>` aggiunge la didascalia `A cute cat`.

Dopo l'immagine annidata nell'elemento `figure`, aggiungi un elemento `figcaption` impostato su:

`Cats love lasagna.`

# --hints--

L'elemento `figcaption` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('figcaption'));
```

L'elemento `figcaption` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

L'elemento `figcaption` dovrebbe essere annidato nell'elemento `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

L’elemento `img` con la lasagna dovrebbe essere annidato nell’elemento `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

L'elemento `figcaption` annidato nell'elemento `figure` dovrebbe essere sotto l'elemento `img`. Sono nell'ordine sbagliato.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

Il testo dell'elemento `figcaption`dovrebbe essere `Cats love lasagna.` Hai omesso il valore o hai un refuso.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

