---
id: 5ef9b03c81a63668521804dc
title: Step 44
challengeType: 0
dashedName: step-44
---

# --description--

È possibile utilizzare i pulsanti di opzione per le domande in cui desideri una sola risposta tra opzioni multiple.

Ecco un esempio di un pulsante di opzione con l'opzione `cat`: `<input type="radio"> cat`. Ricorda che gli elementi `input` sono <dfn>self-closing</dfn>.

Prima dell'input di testo, aggiungi un pulsante di opzione con l'opzione impostata su:

`Indoor`

# --hints--

Dovresti creare un elemento `input` per il tuo pulsante di opzione. Controlla la sintassi.

```js
assert($('form > input').length >= 2);
```

L'elemento `input` dovrebbe avere un tag di apertura, ma non un tag di chiusura.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

Dovresti aver aggiunto un solo elemento `input` per il tuo pulsante di opzione. Rimuovi gli altri.

```js
assert($('form > input').length === 2);
```

Il nuovo elemento `input` dovrebbe essere sopra l'`input` esistente, che ha `text` come valore dell'attributo `type`. Sono nell'ordine sbagliato.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Il nuovo elemento `input` non ha un attributo `type`. Controlla che ci sia uno spazio dopo il nome del tag di apertura.

```js
assert($('input')[0].hasAttribute('type'));
```

Il nuovo elemento `input` dovrebbe avere solo un attributo `type`. Rimuovi quelli di troppo.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Il nuovo elemento `input` dovrebbe avere un attributo `type` con il valore `radio`. Hai omesso il valore o hai un refuso. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Sebbene tu abbia impostato `radio` come valore dell'attributo `type` del nuovo elemento `input`, è raccomandato inserire sempre il valore di un attributo tra virgolette.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Il testo `Indoor` del pulsante di opzione (`radio`) dovrebbe essere posizionato dopo di esso invece che prima.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

Il testo `Indoor` dovrebbe essere posizionato direttamente a destra del pulsante di opzione (`radio`). Hai omesso il testo o hai un refuso.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

