---
id: 5ef9b03c81a63668521804e5
title: Step 61
challengeType: 0
dashedName: step-61
---

# --description--

Per far sì che una casella di spunta o un pulsante di opzione risultino selezionati di default, è necessario aggiungere l'attributo `checked`. Non è necessario impostare un valore per l'attributo `checked`. Aggiungi la parola `checked` all'elemento `input`, assicurandoti che ci sia spazio tra esso e altri attributi.

Fai in modo che il primo pulsante di opzione e la prima casella di spunta siano selezionati di default.

# --hints--

Assicurati che ci siano ancora due pulsanti di opzione e tre caselle di spunta annidate nei rispettivi elementi `fieldset`.

```js
assert(
  $('input[type="radio"]').length === 2 &&
    $('fieldset > input[type="checkbox"]').length === 3
);
```

Il primo pulsante di opzione con l'`id` impostato su `indoor` dovrebbe avere l'attributo `checked`.

```js
assert($('input[type="radio"]')[0].hasAttribute('checked'));
```

Il secondo pulsante di opzione non deve avere l'attributo `checked`.

```js
assert(!$('input[type="radio"]')[1].hasAttribute('checked'));
```

Nella prima casella di spunta manca l'attributo `checked`.

```js
assert($('input[type="checkbox"]')[0].hasAttribute('checked'));
```

La seconda casella di spunta non deve avere l'attributo `checked`.

```js
assert(!$('input[type="checkbox"]')[1].hasAttribute('checked'));
```

La terza casella di spunta non deve avere l'attributo `checked`.

```js
assert(!$('input[type="checkbox"]')[2].hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

