---
id: 5ef9b03c81a63668521804db
title: Step 41
challengeType: 0
dashedName: step-41
---

# --description--

Per evitare che un utente invii il modulo senza tutte le informazioni richieste, è necessario aggiungere l'attributo `required` all'elemento `input`. Non è necessario impostare un valore per l'attributo `required`. Aggiungi la parola `required` all'elemento `input`, assicurandoti che ci sia spazio tra di esso e altri attributi.

# --hints--

Hai eliminato l'elemento `input` oppure ha una sintassi non valida. Tutti i valori degli attributi devono essere circondati da virgolette.

```js
assert($('input').length);
```

L'elemento `form` dovrebbe contenere solo l'elemento `input`. Rimuovi qualsiasi elemento HTML o testo aggiuntivo all'interno dell'elemento `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

L'elemento `input` dovrebbe avere un attributo `required`. Ricorda, basta aggiungere la parola `required` all'interno del tag dell'elemento `input`.

```js
assert($('input')[0].hasAttribute('required'));
```

Non bisogna assegnare un valore all'attributo `required`.

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

