---
id: 5ef9b03c81a63668521804e2
title: Step 54
challengeType: 0
dashedName: step-54
---

# --description--

I moduli usano comunemente delle caselle di spunta (checkbox) per le domande che possono avere più di una risposta. Ad esempio, ecco una casella di spunta con l'opzione `tacos`: `<input type="checkbox"> tacos`.

Sotto l'elemento `legend` che hai appena aggiunto, aggiungi un elemento `input` con `checkbox` come valore dell'attributo `type` e l'opzione:

`Loving`

# --hints--

L'elemento `input` della casella di spunta dovrebbe avere un tag di apertura, ma non un tag di chiusura.

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

Dovresti aver aggiunto solo un elemento di input per la casella di spunta. Rimuovi gli altri.

```js
assert($('fieldset > input').length === 1);
```

Il nuovo elemento `input` dovrebbe essere al di sotto dell'elemento `legend` con il testo `What's your cat's personality?`. Sono nell'ordine sbagliato.

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

Il nuovo elemento `input` non ha un attributo `type`. Controlla che ci sia uno spazio dopo il nome del tag di apertura.

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

L'elemento `input` dovrebbe avere un attributo `type` con il valore `checkbox`. Hai omesso il testo o hai un refuso. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

Sebbene tu abbia impostato `checkbox` come valore dell'attributo `type` del nuovo elemento `input`, è raccomandato inserire sempre il valore di un attributo tra virgolette.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Il testo `Loving` dovrebbe essere posizionato direttamente a destra della tua casella di spunta. Assicurati che ci sia uno spazio tra l'elemento e il testo. Hai omesso il valore o hai un refuso.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

