---
id: 5ef9b03c81a63668521804d1
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

Dopo l'elemento `figure` aggiungi un altro elemento `h3` con il testo:

`Top 3 things cats hate:`

# --hints--

Dovrebbe esserci un elemento `h3` proprio sopra il tag di chiusura del secondo elemento `section`. Assicurati che ci siano un tag di apertura e uno di chiusura.

```js
assert(
  document.querySelectorAll('main > section')[1].lastElementChild.nodeName ===
    'H3' && code.match(/<\/h3\>/g).length === 2
);
```

Il nuovo elemento `h3` dovrebbe contenere il testo `Top 3 things cats hate:`. Assicurati di includere i due punti alla fine del testo.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'top 3 things cats hate:'
);
```

Dovrebbe esserci un elemento `figure` sopra il nuovo elemento `h3`. Potresti aver cancellato accidentalmente l'elemento `figure`.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  .lastElementChild;
assert(
  secondSectionLastElemNode.nodeName === 'H3' &&
    secondSectionLastElemNode.previousElementSibling.nodeName === 'FIGURE'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

