---
id: 5efb23e70dc218d6c85f89b1
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Ci sono molti tipi di input che puoi creare usando l'attributo `type`. Puoi creare facilmente un campo di password, un pulsante di reset o un controllo per consentire agli utenti di selezionare un file dal proprio computer.

Crea un casella di testo per ottenere input di testo da un utente aggiungendo l'attributo `type` con il valore `text` all'elemento `input`.

# --hints--

Hai eliminato l'elemento `input` oppure ha una sintassi non valida. Se hai aggiunto degli attributi, assicurati che i loro valori siano circondati da virgolette.

```js
assert($('input').length);
```

L'elemento `form` dovrebbe contenere solo l'elemento `input`. Rimuovi qualsiasi elemento HTML o testo aggiuntivo tra i tag dell'elemento `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

L'elemento `input` non ha un attributo `type` con il valore `text`. Verifica che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert($('input')[0].hasAttribute('type'));
```

L'elemento `input` dovrebbe avere un attributo `type` con il valore `text`. Hai omesso il valore o hai un refuso. Ricorda che i valori degli attributi devono essere circondati da virgolette.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

Sebbene tu abbia impostato l'attributo `type` dell'elemento `input` come `text`, è sempre raccomandato inserire il valore di un attributo tra virgolette.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

