---
id: 5f07be6ef7412fbad0c5626b
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Prima di aggiungere qualsiasi nuovo contenuto, dovresti utilizzare un elemento `section` per separare i contenuti che hai già inserito da quelli futuri.

Prendi l'elemento `h2`, il commento, e gli elementi `p` e di ancoraggio (`a`) e annidali in un elemento `section`.

# --hints--

L'elemento `section` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('section'));
```

L'elemento `section` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

L'intero elemento `section` dovrebbe essere compreso tra i tag di apertura e chiusura dell'elemento `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

L'elemento `h2` esistente, il commento, l'elemento `p` e l'elemento di ancoraggio (`a`) dovrebbero essere tra i tag di apertura e chiusura dell'elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

L'elemento `h1` dovrebbe essere annidato nell'elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

