---
id: 5f07c98cdb9413cbd4b16750
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

È ora di aggiungere una nuova sezione. Aggiungi un secondo elemento `section` al di sotto dell'elemento `section` esistente.

# --hints--

L'elemento `section` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Dovresti aggiungere un solo tag di apertura `section`. Rimuovi quelli di troppo.

```js
assert(document.querySelectorAll('section').length === 2);
```

L'elemento `section` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Dovresti aggiungere un solo tag di chiusura `section`. Rimuovi quelli di troppo.

```js
assert(code.match(/<\/section>/g).length === 2);
```

Il secondo elemento `section` non dovrebbe essere annidato nel primo elemento `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Entrambi gli elementi `section` dovrebbero essere tra i tag di apertura e chiusura dell'elemento `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

