---
id: 5f07be6ef7412fbad0c5626b
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Prima di aggiungere qualsiasi nuovo contenuto, dovresti utilizzare un elemento `section` per separare i contenuti che hai già inserito da quelli futuri.

Prendi tutti gli elementi che si trovano attualmente all'interno dell'elemento `main` e annidali in un elemento `section`.

# --hints--

L'elemento `section` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<nomeElemento>`.

```js
assert(document.querySelector('section'));
```

L'elemento `section` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/section\>/));
```

L'intero elemento `section` dovrebbe essere compreso tra i tag di apertura e chiusura dell'elemento `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

L'elemento `h2` esistente, il commento, l'elemento `p` e l'elemento di ancoraggio (`a`) dovrebbero essere tra i tag di apertura e chiusura dell'elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElems = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElems.length === 3);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
    </main>
--fcc-editable-region--
  </body>
</html>
```

