---
id: 5f0d4ab1b435f13ab6550052
title: Step 52
challengeType: 0
dashedName: step-52
---

# --description--

Adesso, stai per aggiungere alcuni nuovi elementi di modulo `input`, quindi aggiungi un altro elemento `fieldset` direttamente sotto l'attuale elemento `fieldset`.

# --hints--

Il nuovo elemento `fieldset` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<nomeElemento>`.

```js
assert(document.querySelectorAll('fieldset').length >= 2);
```

Dovresti aggiungere solo un tag di apertura `fieldset`. Rimuovi quelli di troppo.

```js
assert(document.querySelectorAll('fieldset').length === 2);
```

Il nuovo elemento `fieldset` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/fieldset>/g).length >= 2);
```

Dovresti aggiungere solo un tag di chiusura `fieldset`. Rimuovi quelli di troppo.

```js
assert(code.match(/<\/fieldset>/g).length === 2);
```

Il secondo elemento `fieldset` non dovrebbe essere annidato nel primo elemento `fieldset`.

```js
const childrenOf1stFieldset = [
  ...document.querySelector('form > fieldset').children
];
const foundElems = childrenOf1stFieldset.filter((child) => {
  return child.nodeName === 'FIELDSET';
});
assert(foundElems.length === 0);
```

Entrambi gli elementi `fieldset` dovrebbero essere sopra il campo di testo. Sono nell'ordine sbagliato.

```js
const formChildren = $('form')[0].children;
assert(
  formChildren[0].nodeName === 'FIELDSET' &&
    formChildren[1].nodeName === 'FIELDSET' &&
    formChildren[2] &&
    formChildren[2].nodeName === 'INPUT' &&
    formChildren[2].getAttribute('type') === 'text'
);
```

Il nuovo elemento `fieldset` dovrebbe essere al di sotto dell'elemento `fieldset` esistente. Sono nell'ordine sbagliato.

```js
const fieldsetChildren = [...document.querySelectorAll('fieldset')].map(
  (elem) => elem.children
);
assert(fieldsetChildren[0].length > fieldsetChildren[1].length);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

