---
id: 6810869355fa2468f9ebaf9c
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Now print `encrypted_text` to see the output in the terminal.

# --hints--

You should print `encrypted_text`.

```js
assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`))
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)

--fcc-editable-region--
```
