---
id: 6818b80519e6c6c47a8bb6e8
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

As you learned in a previous lecture, functions can have parameters, which are variables that can be referenced within the function. Here's a `sum` function with two parameters, `num1` and `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

The message to encrypt and the shift are still hardcoded in your function. Give your function two parameters named `text` and `shift` and delete the declarations of both the `text` and `shift` variables from the `caesar` function body.

# --hints--

Your `caesar` function should have two parameters, `text` and `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

You should not have a `text` variable declared within the `caesar` function body.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

You should not have a `shift` variable declared within the `caesar` function body.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'Hello World'
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--
```
