---
id: 63b60aaaa65f8922bfce6b7e
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

В елементі `span` створіть елемент `label` для `entry-dropdown` та надайте йому текст `Add food or exercise:`. Потім створіть елемент `select` з `id` зі значенням `entry-dropdown` та `name` зі значенням `options`. Знизу додайте елемент `button` з `id` зі значенням `add-entry` та текст `Add Entry`.

Надайте елементу `button` атрибут `type` зі значенням `button`, щоб уникнути автоматичного надсилання форми.

# --hints--

Ви повинні додати елемент `label` до елемента `span`.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Новий елемент `label` повинен мати атрибут `for` зі значенням `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Новий елемент `label` повинен мати текст `Add food or exercise:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Ви повинні додати елемент `select` до елемента `span`.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Елемент `select` повинен йти після елемента `label`.

```js
assert(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName === 'LABEL');
```

Новий елемент `select` повинен мати атрибут `id` зі значенням `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Новий елемент `select` повинен мати атрибут `name` зі значенням `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Ви повинні додати елемент `button` до елемента `span`.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Елемент `button` повинен йти після елемента `select`.

```js
assert(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName === 'SELECT');
```

Новий елемент `button` повинен мати атрибут `id` зі значенням `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Новий елемент `button` повинен мати текст `Add Entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
