---
id: 63b607af6fcdb119aae9b16a
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Створіть елемент `fieldset` з `id` зі значенням `breakfast`.

В межах цього елементу створіть `legend` з текстом `Breakfast` та порожній `div` з `class` зі значенням `input-container`.

# --hints--

Ви повинні створити елемент `fieldset` у `form`.

```js
assert.exists(document.querySelector('form fieldset'));
```

Елемент `fieldset` повинен йти після елемента `input`.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

Елемент `fieldset` повинен мати атрибут `id` зі значенням `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Ви повинні створити елемент `legend` в межах `fieldset`.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

Елемент `legend` повинен мати текст `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Ви повинні створити елемент `div` в межах `fieldset`.

```js
assert.exists(document.querySelector('form fieldset div'));
```

Елемент `div` повинен йти після елемента `legend`.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

Елемент `div` повинен мати `class` зі значенням `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
