---
id: 63b60af1a0b9f7238a9dd294
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Меню вибору потрібні варіанти вибору їжі та вправ (елементів `fieldset`), які ви створили у попередніх кроках. Використайте елемент `option`, щоб створити новий варіант вибору для кожного `fieldset`. Атрибутом `value` кожного варіанту має бути `fieldset` `id`, а текстом кожного варіанту має бути текст з `legend`.

Встановіть варіант `Breakfast` як `selected`.

# --hints--

Ви повинні створити п’ять елементів `option` в межах елемента `select`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.length, 5);
```

Перший `option` повинен мати текст `Breakfast`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.textContent, 'Breakfast');
```

Перший `option` повинен мати атрибут `value` зі значенням `breakfast`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.value, 'breakfast');
```

Другий `option` повинен мати текст `Lunch`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.textContent, 'Lunch');
```

Другий `option` повинен мати атрибут `value` зі значенням `lunch`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.value, 'lunch');
```

Третій `option` повинен мати текст `Dinner`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.textContent, 'Dinner');
```

Третій `option` повинен мати атрибут `value` зі значенням `dinner`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.value, 'dinner');
```

Четвертий `option` повинен мати текст `Snacks`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.textContent, 'Snacks');
```

Четвертий `option` повинен мати атрибут `value` зі значенням `snacks`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.value, 'snacks');
```

П’ятий `option` повинен мати текст `Exercise`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.textContent, 'Exercise');
```

П’ятий `option` повинен мати атрибут `value` зі значенням `exercise`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.value, 'exercise');
```

Перший `option` повинен бути обраним за замовчуванням.

```js
assert.isTrue(document.querySelectorAll('.controls select option')?.[0]?.getAttributeNames()?.includes('selected'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">

              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
