---
id: 63b608ebf40c871d960fc004
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Вам потрібно ще два блоки коду `fieldset`: для `snacks` та `exercise`.

# --hints--

Ви повинні створити четвертий елемент `fieldset` у `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]);
```

Четвертий елемент `fieldset` повинен йти після третього елемента `fieldset`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.previousElementSibling?.tagName, "FIELDSET");
```

Четвертий елемент `fieldset` повинен мати атрибут `id` зі значенням `snacks`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.id, "snacks");
```

Четвертий елемент `fieldset` повинен містити елемент `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('legend'));
```

Четвертий елемент `legend` повинен мати текст `Snacks`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('legend')?.innerText, "Snacks");
```

Четвертий елемент `fieldset` повинен містити елемент `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('div'));
```

П’ятий елемент `div` повинен мати `class` зі значенням `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('div')?.className, "input-container");
```

Ви повинні створити п’ятий елемент `fieldset` у `form`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]);
```

П’ятий елемент `fieldset` повинен йти після четвертого елемента `fieldset`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.previousElementSibling?.tagName, "FIELDSET");
```

П’ятий елемент `fieldset` повинен мати атрибут `id` зі значенням `exercise`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.id, "exercise");
```

П’ятий елемент `fieldset` повинен містити елемент `legend`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('legend'));
```

П’ятий елемент `legend` повинен мати текст `Exercise`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('legend')?.innerText, "Exercise");
```

П’ятий елемент `fieldset` повинен містити елемент `div`.

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('div'));
```

П’ятий елемент `div` повинен мати `class` зі значенням `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
