---
id: 65825cda2668995d5168e37c
title: Step 37
challengeType: 20
dashedName: step-37
---

# --description--

Create another `elif` statement that checks if `choice == '4'`. Inside the new `elif`, create a variable `category` and assign it `input('Enter category to filter: ')` to filter the expense category.

# --hints--

You should have `elif choice == '4':` in your code.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')4\1\s*:/m));
  }
})
```

You should have `category = input('Enter category to filter: ')` in your new `elif` statement.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')4\1\s*:\s*category\s*=\s*input\s*\(\s*("|')Enter\scategory\sto\sfilter:\s\2\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)
--fcc-editable-region--           
        elif choice == '3':
            print('\nTotal Expenses: ', total_expenses(expenses))

--fcc-editable-region--
```
