---
id: 681cb1b2dab50c87ddb2e523
title: Python Coding Challenge 9
challengeType: 29
dashedName: python-coding-challenge-9
---

# --description--

Description

# --hints--
      
Passing test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def seed(s):
  
    return s

```

# --solutions--

```py
def seed(s):
  
    return s

```
