---
id: 681cb1b1dab50c87ddb2e51f
title: Python Coding Challenge 5
challengeType: 29
dashedName: python-coding-challenge-5
---

# --description--

Description

# --hints--
      
Passing test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def seed(s):
  
    return s

```

# --solutions--

```py
def seed(s):
  
    return s

```
