---
id: 681cb1b2dab50c87ddb2e521
title: Python Challenge 7
challengeType: 29
dashedName: python-challenge-7
---

# --description--

Description

# --hints--

Passing test

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertTrue(True)`)
}})
```

# --seed--

## --seed-contents--

```py
def seed(s):
  
    return s

```

# --solutions--

```py
def seed(s):
  
    return s

```
