---
id: 6821ec9d237de8297eaee79d
title: "Python Challenge 15"
challengeType: 29
dashedName: python-challenge-15
---

# --description--

Description

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` should return `This is a secret message.`

```js
({test: () => { runPython(`
from unittest import TestCase
TestCase().assertEqual(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.")`)
}})
```

# --seed--

## --seed-contents--

```py
def decode(message, shift):

    return message
```

# --solutions--

```py
def decode(message, shift):

    return message
```
