---
id: 646c48df8674cf2b91020ecb
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一個 `meta` 標籤，設置它的 `charset` 屬性爲 `utf-8`。 同時添加一個包含文本 `fCC Cat Painting` 的 `title` 元素。

# --hints--

你應該只添加一個 `meta` 元素。

```js
assert(document.querySelectorAll('meta').length === 1);
```

你的 `meta` 元素應該有一個 `charset` 屬性值爲 `utf-8`。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

你應該只添加一個 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

你的 `title` 元素應包含文本 `fCC Cat Painting`。 請注意拼寫和大小寫的問題。

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
