---
id: 646c5e727e56e743c9aed4a1
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

CSS 定位可讓你設定元素在瀏覽器中的位置。 它有一個 `position` 屬性，你可以設置其值爲 `static`、`absolute`、`relative`、`sticky` 或 `fixed`。

當你設置了元素的 `position` 屬性，你可以通過爲 `top`、`right`、`left`、 `bottom` 屬性中的一個或者多個設置像素值或百分比值來移動該元素。

`static` 是所有元素的默認定位。 如果你將元素設置爲這個定位，你將無法用 `top`、`right`、`left` 或 `bottom` 移動元素。

設置 `.cat-head` 元素的 `position` 屬性爲 `static`，分別設置 `top` 和 `left` 屬性爲 `100px`。


# --hints--

你的 `.cat-head` 選擇器應將 `position` 屬性設置爲 `static`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

你的 `.cat-head` 選擇器應將 `top` 屬性設置爲 `100px`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

你的 `.cat-head` 選擇器應將 `left` 屬性設置爲 `100px`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
