---
id: 646ceb843412c74edee27a79
title: 步驟 22
challengeType: 0
dashedName: step-22
---

# --description--

現在你將學習一個 CSS 技巧來繪製三角形。

這將用於繪製貓耳朵。

使用類選擇器，將 `.cat-right-ear` 元素的 `height` 和 `width` 屬性設置爲 `100px`。 將 `background-color` 設置爲 `white`。 將左側和右側的邊框設置爲 `35px solid blue`。 將頂部和底部的邊框設置爲 `35px solid red`。

# --hints--

你應該有一個 `.cat-right-ear` 選擇器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

你的 `.cat-right-ear` 選擇器應將 `height` 屬性設置爲 `100px`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

你的 `.cat-right-ear` 選擇器應將 `width` 屬性設置爲 `100px`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

你的 `.cat-right-ear` 選擇器應將 `background-color` 屬性設置爲 `white`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

你的 `.cat-right-ear` 選擇器應設置一個 `border-left` 屬性爲 `35px solid blue`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

你的 `.cat-right-ear` 選擇器應設置一個 `border-right` 屬性爲 `35px solid blue`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

你的 `.cat-right-ear` 選擇器應設置一個 `border-top` 屬性爲 `35px solid red`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

你的 `.cat-right-ear` 選擇器應設置一個 `border-bottom` 屬性爲 `35px solid red`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
