---
id: 646c5d7057c45f432fcdd46c
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

使用類選擇器，設置 `.cat-head` 元素的寬度爲 `205px`、高度爲 `180px`。 同時，設置它的邊框爲 `1px solid #000`、`border-radius` 爲 `46%`。

# --hints--

你應該有一個 `.cat-head` 選擇器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

你的 `.cat-head` 選擇器應將 `width` 設置爲 `205px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

你的 `.cat-head` 選擇器應將 `height` 設置爲 `180px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

你的 `.cat-head` 選擇器應將 `border` 設置爲 `1px solid #000`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

你的 `.cat-head` 選擇器應將 `border-radius` 設置爲 `46%`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
