---
id: 646c59e942f35541923104bf
title: 步驟6
challengeType: 0
dashedName: step-6
---

# --description--

回到你的 HTML 中，創建一個 `main` 元素。 在那個 `main` 元素內部，添加一個帶有類 `cat-head` 的 `div` 元素。

# --hints--

你應該有一個 `main` 元素。

```js
assert(document.querySelectorAll('main').length === 1);
```

你應該有一個 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 1);
```

你的 `div` 元素應該有類 `cat-head`。

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

你的 `div` 元素應該在 `main` 標籤內。

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
