---
id: 646ceb843412c74edee27a79
title: 步驟 22
challengeType: 0
dashedName: step-22
---

# --description--

你將把每隻耳朵畫成一個三角形。

使用類選擇器，給 `.cat-left-ear` 元素左右各添加 `35px solid transparent` 的邊框。 另外，把底邊設置爲 `70px solid #5e5e5e`。

# --hints--

你應該有一個 `.cat-left-ear` 選擇器。

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

你的 `.cat-left-ear` 選擇器應該有一個設置爲 `35px solid transparent` 的 `border-left` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

你的 `.cat-left-ear` 選擇器應該有一個設置爲 `35px solid transparent` 的 `border-right` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

你的 `.cat-left-ear` 選擇器應該有一個設置爲 `70px solid #5e5e5e` 的 `border-bottom` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
