---
id: 646c48df8674cf2b91020ecc
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `head` 元素中添加一個 `link` 元素。 對於那個 `link` 元素，將 `rel` 屬性設置爲 `stylesheet`，並將 `href` 設置爲 `./styles.css`。

# --hints--

你應該有一個 `link` 元素。

```js
assert.match(code, /<link/)
```

你的 `link` 元素應該有 `rel` 和 `href` 屬性，分別設置爲 `stylesheet` 和 `./styles.css`。

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
