---
id: 646c47867800472a4ed5d2ea
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

從基本的 HTML 結構開始。 添加一個 `html` 的 `DOCTYPE` 聲明，再添加一個 `html` 元素並設定它的屬性 `lang` 爲 `en`。 然後，在 `html` 元素中添加一個 `head` 和 `body` 元素。

# --hints--

你應該有 `DOCTYPE` 聲明爲 `html`。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你的 `DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個屬性 `lang` 設定爲 `en` 的 `html` 開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你應該有一個 `html` 結束標籤。

```js
assert(code.match(/<\/html>/i));
```

你應該有一個 `head` 開始標籤。

```js
assert(code.match(/<head>/i));
```

你應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head>/i));
```

你應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body>/i));
```

你應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body>/i));
```

你的 `body` 元素應該在 `head` 元素之後。

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

你的 `head` 和 `body` 元素應該在 `html` 元素裏面。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
