---
id: 6476f5c17f99146071ee884c
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以看到什麼也沒有發生。 儘管將 `top` 和 `left` 都設置爲 `100px`，但 `.cat-head` 元素並未移動。 但 `relative` 定位的情況並非如此。

當你使用 `relative` 值時，元素仍然按照文檔的正常流程定位，但是 `top`、`left`、`bottom` 和 `right` 值變爲活動狀態。

不要使用 `static`，而是將 `.cat-head` 設置爲 `relative` 定位，並將 `top` 和 `left` 屬性保持原樣。

# --hints--

你的 `.cat-head` 選擇器應將 `position` 屬性設置爲 `relative`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

你的 `.cat-head` 選擇器應將 `top` 屬性設置爲 `100px`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

你的 `.cat-head` 選擇器應將 `left` 屬性設置爲 `100px`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
