---
id: 6476f7a4827bcc61682f2347
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

下一個位置屬性是 `absolute`。 當你爲 `position` 屬性使用 `absolute` 值時，該元素將從文檔的正常流程中取出，然後其位置由 `top`、`right`、`bottom` 和 `left` 屬性決定。

將 `.cat-head` 元素的 position 屬性設置爲 `absolute`，然後將 `top` 和 `left` 屬性設置爲任意像素值。

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

你的 `.cat-head` 選擇器應將 `position` 屬性設置爲 `absolute`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

你的 `.cat-head` 選擇器應將 `top` 屬性設置爲以 `px` 結尾的任意數字。 確保你添加了分號。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

你的 `.cat-head` 選擇器應將 `left` 屬性設置爲以 `px` 結尾的任意數字。 確保你添加了分號。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
