---
id: 6476fc5cf14b276e6d04e82a
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

使用類選擇器爲你的 `.box` 元素指定寬度爲 `200px`、高度爲 `600px` 以及背景顏色爲 `#000`。 另外，爲其設置 `position` 爲 `absolute`、`top` 爲 `800px` 和 `left` > 爲 `650px`。

# --hints--

你應該有一個 `.box` 選擇器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

你的 `.box` 選擇器應將 `width` 屬性設置爲 `200px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

你的 `.box` 選擇器應將 `height` 屬性設置爲 `600px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

你的 `.box` 選擇器應將 `background-color` 屬性設置爲 `#000`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

你的 `.box` 選擇器應將 `position` 屬性設置爲 `absolute`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

你的 `.box` 選擇器應將 `left` 屬性設置爲 `650px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

你的 `.box` 選擇器應將 `top` 屬性設置爲 `800px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

