---
id: 6477062778c85972eb648030
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

現在你應該將貓頭居中。

將 `.cat-head` 元素的 `position` 屬性設置爲 `absolute`。 將 `right`、`left`、`top`、`bottom` 屬性設置值 `0` ，然後將其所有邊的 `margin` 屬性設置爲 `auto`。 這是使用 CSS 定位將元素垂直和水平居中的一種方法。

# --hints--

你的 `.cat-head` 選擇器應將 `position` 屬性設置爲 `absolute`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

你的 `.cat-head` 選擇器應將 `top` 屬性設置爲 `0`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

你的 `.cat-head` 選擇器應將 `left` 屬性設置爲 `0`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

你的 `.cat-head` 選擇器應將 `right` 屬性設置爲 `0`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

你的 `.cat-head` 選擇器應將 `bottom` 屬性設置爲 `0`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

你的 `.cat-head` 選擇器應將 `margin` 屬性設置爲 `auto`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
