---
id: 64770351e8586671ec0911f0
title: 步驟15
challengeType: 0
dashedName: step-15
---

# --description--

最後一個位置屬性值爲 `sticky`。 `sticky` 定位是 `relative` 定位和 `fixed` 定位的混合體。 它允許元素根據滾動位置**粘**到其包含元素或視口中的特定位置。

將 `.cat-head` 的 `position` 屬性值更改爲 `sticky`，將 `top` 設置爲 `0`，然後刪除 `left` 及其值。

**注意**：要了解 `sticky` 的工作原理，你必須在`.cat-head` `div` 之前和之後放置一些文本。 如果你在此之後向下滾動，你會看到 `.cat-head` 卡在頂部並保留在那裏。

# --hints--

你的 `.cat-head` 選擇器應將 `position` 屬性設置爲 `sticky`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

你的 `.cat-head` 選擇器應將 `top` 屬性設置爲 `0`。 確保你添加了分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

你的代碼中不應包含 `left` 屬性及其值。

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
