---
id: 6676a8b01e56ec1a1e07c202
title: 步驟-23
challengeType: 0
dashedName: step-23
---

# --description--

注意你現在有一個厚邊框的白色方塊。 邊框有對角線，可用於創建三角形。

在同一個類選擇器中，將 `height` 和 `width` 調整爲 `0`。 將左側、右側和頂部邊框設置爲 `transparent`。

刪除 `background-color` 屬性，你應該看到一個三角形。

# --hints--


你應該有 `.cat-right-ear` 選擇器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

你的 `.cat-right-ear` 選擇器應將 `height` 屬性設置爲 `0`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

你的 `.cat-right-ear` 選擇器應將 `width` 屬性設置爲 `0`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

你的 `.cat-right-ear` 選擇器應將 `border-left` 屬性設置爲 `35px solid transparent`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

你的 `.cat-right-ear` 選擇器應將 `border-right` 屬性爲 `35px solid transparent`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

你的 `.cat-right-ear` 選擇器應將 `border-top` 屬性設置爲 `35px solid transparent`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


你的 `.cat-right-ear` 選擇器應將 `border-bottom` 屬性爲 `35px solid red`。 不要忘記加上分號。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

你的 `.cat-right-ear` 選擇器不應該有 `background-color` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
