---
id: 684aaf9ec670c68d20efd0d0
title: Build a User Configuration Manager
challengeType: 27
dashedName: build-a-user-configuration-manager
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a dictionary named `settings` that stores the following user configuration preferences:

```md
 "theme": "light",
 "language": "English",
 "notifications": "enabled"
```

1. You should define a function named `add_setting` that takes a dictionary of settings and a tuple containing a key-value pair as arguments.

1. `add_setting(settings, (key, value))` should:

   - Convert `key` to lowercase.

   - If the `key` setting exists, return `Setting 'key' already exists! Cannot add a new setting with this name.`
   
   - If the `key` setting doesn't exist, add the key-value pair to the `settings` dictionary and return `Setting 'key' added with value 'value' successfully!`

1. You should define a function named `update_setting` that takes the a dictionary of settings and a tuple containing a key-value pair as arguments.

1. `update_setting(settings, (key, value))` should:
    
   - Convert `key` to lowercase.
   
   - If the `key` setting exists, update its value in the `settings` dictionary and return: `Setting 'key' updated to 'value' successfully!`
   
   - If the `key` setting doesn't exist, return `Setting 'key' does not exist! Cannot update a non-existing setting.`

1. You should define a function named `delete_setting` that takes the a dictionary of settings and a key as arguments.

1. `delete_setting(settings, key)` should:
    
   - Convert `key` to lowercase.
   
   - If the `key` setting exists, remove the key-value pair from `settings` and return `Setting 'key' deleted successfully!`
   
   - If the `key` setting does not exist, return `Setting not found!`

1. You should define a function named `view_settings` that takes a dictionary of settings as argument.

1. `view_settings(settings)` should:
   
   - Return `No settings available.` if the `settings` dictionary is empty.
   
   - If the dictionary contains any settings, return them formatted as follows:

```md
Current User Settings:
Theme: dark
Notifications: enabled
Volume: high

```

# --hints--

You should create a dictionary named `settings` with the correct values.

```js
assert.match(code, /settings\s*=\s*{[^}]*}/)
```

You should define a function named `add_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('add_setting')
    `)
})
```

The `add_setting` function should take two parameters.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` should convert the key to lowercase.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light"}
        result = add_setting(test_settings, ("VOLUME", "high"))
        assert "volume" in test_settings
        assert "VOLUME" not in test_settings
    `)
})
```

`add_setting` should return an error message when the key already exists.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light", "language": "English"}
        result = add_setting(test_settings, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting` should add a new key-value pair and return success message.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light"}
        result = add_setting(test_settings, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert test_settings["volume"] == "high"
    `)
})
```

You should define a function named `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

The `update_setting` function should take two parameters.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`update_setting` should convert the key to lowercase.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light", "volume": "low"}
        result = update_setting(test_settings, ("THEME", "dark"))
        assert test_settings["theme"] == "dark"
    `)
})
```

`update_setting` should update an existing key and return success message.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light", "language": "English"}
        result = update_setting(test_settings, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert test_settings["theme"] == "dark"
    `)
})
```

`update_setting` should return an error message when the key doesn't exist.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light"}
        result = update_setting(test_settings, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

You should define a function named `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

The `delete_setting` function should take two parameters.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` should convert the key to lowercase.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light", "language": "English"}
        result = delete_setting(test_settings, "THEME")
        assert "theme" not in test_settings
        assert "THEME" not in test_settings
    `)
})
```

`delete_setting` should remove an existing key and return success message.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light", "language": "English"}
        result = delete_setting(test_settings, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in test_settings
    `)
})
```

`delete_setting` should return an error message when the key doesn't exist.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "light"}
        result = delete_setting(test_settings, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

You should define a function named `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

The `view_settings` function should take one parameter.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` should return appropriate message for empty dictionary.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` should return formatted settings for non-empty dictionary.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(test_settings)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` should capitalize the first letter of each setting name.

```js
({
    test: () => runPython(`
        test_settings = {"theme": "dark", "volume": "high"}
        result = view_settings(test_settings)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

# remove at the end

print(add_setting(settings, ("volume", "high")))
print(update_setting(settings, ("theme", "dark")))
print(delete_setting(settings, "language"))
print(view_settings(settings))
```
