---
id: 684aaf9ec670c68d20efd0d0
title: Build a User Configuration Manager
challengeType: 27
dashedName: build-a-user-configuration-manager
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a dictionary named `settings` that stores the following user configuration preferences:

```md
 "theme": "light",
 "language": "English",
 "notifications": "enabled"
```

1. You should define a function named `add_setting` that takes a dictionary of settings and a tuple containing a key-value pair as arguments.

1. `add_setting(settings, (key, value))` should:

   - Convert `key` to lowercase.

   - If the `key` setting exists, return `Setting 'key' already exists! Cannot add a new setting with this name.`
   
   - If the `key` setting doesn't exist, add the key-value pair to the `settings` dictionary and return `Setting 'key' added with value 'value' successfully!`

1. You should define a function named `update_setting` that takes the a dictionary of settings and a tuple containing a key-value pair as arguments.

1. `update_setting(settings, (key, value))` should:
    
   - Convert `key` to lowercase.
   
   - If the `key` setting exists, update its value in the `settings` dictionary and return: `Setting 'key' updated to 'value' successfully!`
   
   - If the `key` setting doesn't exist, return `Setting 'key' does not exist! Cannot update a non-existing setting.`

1. You should define a function named `delete_setting` that takes the a dictionary of settings and a key as arguments.

1. `delete_setting(settings, key)` should:
    
   - Convert `key` to lowercase.
   
   - If the `key` setting exists, remove the key-value pair from `settings` and return `Setting 'key' deleted successfully!`
   
   - If the `key` setting does not exist, return `Setting not found!`

1. You should define a function named `view_settings` that takes a dictionary of settings as argument.

1. `view_settings(settings)` should:
   
   - Return `No settings available.` if the `settings` dictionary is empty.
   
   - If the dictionary contains any settings, return them formatted as follows:

```md
Current User Settings:
Theme: dark
Notifications: enabled
Volume: high

```

# --hints--


Sample hint text

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function(add_setting')
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

# remove at the end

print(add_setting(settings, ("volume", "high")))
print(update_setting(settings, ("theme", "dark")))
print(delete_setting(settings, "language"))
print(view_settings(settings))
```
