---
id: 615392916d83fa4f02f7e2cf
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Potresti aver notato che le immagini si sono tutte spostate nella stessa riga.

La proprietà `flex-wrap` determina come dovrebbero comportarsi gli oggetti quando il contenitore flex è troppo piccolo. Impostare questa proprietà su `wrap` permetterà agli oggetti di proseguire nella riga/colonna successiva (a seconda dell'asse principale), mentre `nowrap` impedirà agli oggetti essere posizionati in nuove righe o colonne. Con il valore `nowrap`, gli oggetti possono essere rimpiccioliti per adattarsi o debordare.

Assegna al selettore `#gallery` una proprietà `flex-wrap` con il valore `wrap`. Dovresti vedere le tue immagini prendere un layout a quattro colonne. Questo perché hai impostato la loro proprietà `width` al `25%` in uno step precedente.

# --hints--

Il selettore `#gallery` dovrebbe avere una proprietà `flex-wrap` con il valore `wrap`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery')?.flexWrap === 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

.header {
  text-align: center;
  padding: 32px;
  background: #E0DDDD;
}

--fcc-editable-region--
#gallery {
  display: flex;
  flex-direction: row;
}
--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
