---
id: 61537485c4f2a624f18d7794
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Inizia con il testo HTML standard. Aggiungi la dichiarazione `DOCTYPE` e gli elementi `html`, `head` e `body`.

# --hints--

Il codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un carattere `>`.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di apertura.

```js
assert(code.match(/<html\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag `head` di apertura.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag `head` di chiusura.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag `body` di apertura.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag `body` di chiusura.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L' elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
