---
id: 61537c9eecea6a335db6da79
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Al di sotto dell'elemento `.header`, crea un nuovo elemento `div` con un attributo `id` del valore `gallery`.

In quest'elemento `#gallery`, crea dieci elementi `img`.

# --hints--

Dovresti creare un secondo elemento `div` nell'elemento `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length === 2);
```

Il nuovo elemento `div` dovrebbe trovarsi dopo l'elemento `.header`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[0]?.classList?.contains('header'));
```

Il nuovo elemento `div` dovrebbe avere un `id` con il valore `gallery`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[1]?.id === 'gallery');
```

L'elemento `#gallery` dovrebbe contenere dieci elementi `img`.

```js
assert(document.querySelector('#gallery')?.querySelectorAll('img')?.length === 10);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
  </body>
</html>
--fcc-editable-region--
```

```css

```
