---
id: 6153897c27f6334716ee5abe
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Allinea al centro il testo di `.header`. Rendi il testo maiuscolo con la proprietà `text-transform` impostata su `uppercase`.

Dagli un padding di `32px` su tutti i lati. Imposta lo sfondo sul colore `#0a0a23` e il testo sul colore `#fff`.

Aggiungi un `border-bottom` con il valore `4px solid #fdb347`.


# --hints--

Dovresti avere un selettore `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Il selettore `.header` dovrebbe avere una proprietà `text-align` con il valore `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign === 'center');
```

Il selettore `.header` dovrebbe avere una proprietà `text-transform` con il valore `uppercase`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform === 'uppercase');
```

Il selettore `.header` dovrebbe avere una proprietà `padding` con il valore `32px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.padding === '32px');
```

Il selettore `.header` dovrebbe avere una proprietà `background-color` con il valore `#0a0a23`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor === 'rgb(10, 10, 35)');
```

Il selettore `.header` dovrebbe avere una proprietà `color` con il valore `#fff`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.color === 'rgb(255, 255, 255)');
```

Il selettore `.header` dovrebbe avere una proprietà `border-bottom` con il valore `4px solid #fdb347`.

```js
console.log(new __helpers.CSSHelp(document).getStyle('.header'))
assert(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom === '4px solid rgb(253, 179, 71)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
