---
id: 615f171d05def3218035dc85
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Al momento, le immagini sono troppo grandi e non sarai in grado di vedere i cambiamenti CSS.

Crea un selettore `#gallery img` per selezionare le immagini. Assegnagli una proprietà `width` con il valore `25%`.

Imposta anche la proprietà `height` su `300px` per mantenere uniforme la dimensione delle immagini.

# --hints--

Dovresti avere un selettore `#gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#gallery img'))
```

La regola `#gallery img` dovrebbe avere una proprietà `width` con il valore `25%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery img')?.width === '25%')
```

Il selettore `#gallery img` dovrebbe avere una proprietà `height` con il valore `300px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
