---
id: 6494da0daf5df5197963671d
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Nota come i bordi delle immagini blu ora vanno oltre il bordo della galleria rossa. Ciò è dovuto al modo in cui i browser calcolano la dimensione degli elementi contenitore.

La proprietà `box-sizing` viene utilizzata per definire questo comportamento. Per impostazione predefinita, viene utilizzato il modello `content-box`. Con questo modello, quando un elemento ha una larghezza specifica, tale larghezza viene calcolata solo in base al contenuto dell'elemento. I valori di padding e bordo vengono aggiunti alla larghezza totale, con l'elemento che cresce per adattarsi a questi valori.

Prova a impostare esplicitamente `box-sizing` su `content-box` con il selettore globale `*`. In questo momento, non vedrai alcun cambiamento, perché stai usando il valore predefinito.

# --hints--

You should have a `*` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Il selettore `*` dovrebbe avere una proprietà `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

Il selettore `*` dovrebbe avere la proprietà `box-sizing` impostata su `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
