---
id: 6823e169fda14755fbf00445
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now it is time to add another `h3` and `p` element to the page.

Add an `h3` element with the text `Introduction to CSS`.

Then below that `h3` element, add a `p` element with the following text:

```md
CSS represents the styles for a web page
```

# --hints--

You should add a second `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Your `h3` element's text should be `Introduction to CSS`. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /Introduction\s*to\s*CSS/i);
```

You should add a third `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Your `p` element should have the text `CSS represents the styles for a web page`. You can copy the text from the instructions. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /CSS\s*represents\s*the\s*styles\s*for\s*a\s*web\s*page/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Full Stack Curriculum</h2>
    <p>Learn all of the skills to become a full stack developer</p>

    <h3>Introduction to HTML</h3>
    <p>HTML represents the content and structure of a webpage</p>

    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>  
```
