---
id: 6823e169fda14755fbf00445
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now it is time to add another `h3` and `p` element to the page.

Add an `h3` element with the text `Introduction to CSS`.

Then below that `h3` element, add a `p` element with the text `CSS stands for Cascading Style Sheets and it represents the styles for a web page.`.

# --hints--

You should add a second `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Your `h3` element's text should be `Introduction to CSS`. You have either omitted the text, have a typo, or forgot the period.

```js
assert.equal(document.querySelector("h3")?.innerText, "Introduction to CSS");
```

You should add a third `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Your `p` element should have the text `CSS stands for Cascading Style Sheets and it represents the styles for a web page.`

```js
const lastParaEl = document.querySelector("p:last-of-type");
assert.strictEqual(lastParaEl?.innerText, "CSS stands for Cascading Style Sheets and it represents the styles for a web page.");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Certified Full Stack Developer Certification</h2>
    <p>In this certification, you will learn all of the skills to become a full stack developer.</p>

    <h3>Introduction to HTML</h3>
    <p>HTML stands for HyperText Markup Language and it represents the content and structure of a webpage.</p>

    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>  
```
