---
id: 6823ac607bfdbc46331b2559
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML stands for HyperText Markup Language and it represents the content and structure of a web page. In this workshop, you will write the HTML code for a partial curriculum outline.

This is your code editor where you can write HTML code. Write the following text in the editor:

```md
Welcome to freeCodeCamp
```

When are you done, click the "check your code" button to see if it's correct. 

# --hints--

You should have the text `Welcome to freeCodeCamp`. Double check for spelling.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and 
// keep camper retention this early on. 
assert.match(code, /Welcome\s*to\s*freeCodeCamp/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html> 
```

