---
id: 6823e637c1c0ed56f781b4fc
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Finally, JavaScript makes your webpage interactive — it lets you tell the page what to do when someone clicks a button, submits a form, or many other things.

For the last step of the workshop, add another `h3` and `p` element to the page describing JavaScript.

First, add an `h3` element with the text:

```md
Introduction to JavaScript
```

Then, below that `h3` element, add a `p` element with the following text:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

You should add a third `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Your `h3` element's text should be `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

You should add a fourth `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Your `p` element should have the text `JavaScript adds interactivity to a webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn all of the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
