---
id: 6823e637c1c0ed56f781b4fc
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

For the last step of the workshop, you will add another `h3` and `p` element to the page.

Add an `h3` element with the text `Introduction to JavaScript`.

Then below that `h3` element, add a `p` element with the text `JavaScript is used for the interactivity on the webpage.`.

# --hints--

You should add a third `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Your `h3` element's text should be `Introduction to JavaScript`. You have either omitted the text, have a typo, or forgot the period.

```js
assert.equal(document.querySelector("h3:last-of-type")?.innerText, "Introduction to JavaScript");
```

You should add a fourth `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Your `p` element should have the text `JavaScript is used for the interactivity on the webpage.`

```js
const lastParaEl = document.querySelector("p:last-of-type");
assert.strictEqual(lastParaEl?.innerText, "JavaScript is used for the interactivity on the webpage.");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Certified Full Stack Developer Certification</h2>
    <p>In this certification, you will learn all of the skills to become a full stack developer.</p>

    <h3>Introduction to HTML</h3>
    <p>HTML stands for HyperText Markup Language and it represents the content and structure of a webpage.</p>

    <h3>Introduction to CSS</h3>
    <p>CSS stands for Cascading Style Sheets and it represents the styles for a web page.</p>

    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>  
```

# --solutions--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Certified Full Stack Developer Certification</h2>
    <p>In this certification, you will learn all of the skills to become a full stack developer.</p>

    <h3>Introduction to HTML</h3>
    <p>HTML stands for HyperText Markup Language and it represents the content and structure of a webpage.</p>

    <h3>Introduction to CSS</h3>
    <p>CSS stands for Cascading Style Sheets and it represents the styles for a web page.</p>

    <h3>Introduction to JavaScript</h3>
    <p>JavaScript is used for the interactivity on the webpage.</p>
  </body>
</html> 
```
