---
id: 6823e637c1c0ed56f781b4fc
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

For the last step of the workshop, you will add another `h3` and `p` element to the page.

Add an `h3` element with the text `Introduction to JavaScript`.

Then below that `h3` element, add a `p` element with the following text:

```md
JavaScript is used for the interactivity on the webpage
```

# --hints--

You should add a third `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

Your `h3` element's text should be `Introduction to JavaScript`. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /Introduction\s*to\s*JavaScript/i);
```

You should add a fourth `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

Your `p` element should have the text `JavaScript is used for the interactivity on the webpage`

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /JavaScript\s*is\s*used\s*for\s*the\s*interactivity\s*on\s*the\s*webpage/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Full Stack Curriculum</h2>
    <p>Learn all of the skills to become a full stack developer</p>

    <h3>Introduction to HTML</h3>
    <p>HTML represents the content and structure of a webpage</p>

    <h3>Introduction to CSS</h3>
    <p>CSS represents the styles for a web page</p>

    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>  
```

# --solutions--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Full Stack Curriculum</h2>
    <p>Learn all of the skills to become a full stack developer</p>

    <h3>Introduction to HTML</h3>
    <p>HTML represents the content and structure of a webpage</p>

    <h3>Introduction to CSS</h3>
    <p>CSS represents the styles for a web page</p>

    <h3>Introduction to JavaScript</h3>
    <p>JavaScript is used for the interactivity on the webpage</p>
  </body>
</html> 
```
