---
id: 6823d6244511f252c8300eed
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

When you need to add a paragraph to a web page, you can use the `p` element like this:

```html
<p>I am a paragraph element.</p>
```

Below your `h2` element, add a `p` element with the following text:

```md
Learn all of the skills to become a full stack developer
```

# --hints--

Your `p` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Your `p` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/p\>/);
```

Your `p` element's text should be `Learn all of the skills to become a full stack developer`. You can copy the text from the instructions. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /Learn\s*all\s*of\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer/i);
```

Your `p` element should be below your `h2` element. 

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Full Stack Curriculum</h2>

--fcc-editable-region--
  </body>
</html>  
```
