---
id: 6823d9ac8bdc3853df65a1ff
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

There are a total of six heading elements in HTML. 

Here is the list of heading elements in order of most important to least:

```html
<h1>Heading level 1</h1>
<h2>Heading level 2</h2>
<h3>Heading level 3</h3>
<h4>Heading level 4</h4>
<h5>Heading level 5</h5>
<h6>Heading level 6</h6>
```

Most of the time you will be using `h1`, `h2` and `h3` elements with the occasional use of the `h4` and maybe `h5` elements. Very rarely will you be using an `h6` element. 

Below your `p` element, add an `h3` element with the text `Introduction to HTML`.

# --hints--

Your `h3` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector("h3"));
```

Your `h3` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h3\>/);
```

Your `h3` element's text should be `Introduction to HTML`. You have either omitted the text, have a typo, or forgot the period.

```js
assert.equal(document.querySelector("h3")?.innerText, "Introduction to HTML");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Certified Full Stack Developer Certification</h2>
    <p>In this certification, you will learn all of the skills to become a full stack developer.</p>

    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>  
```
