---
id: 6823e036ea4b71553558c01b
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Next, add another `p` element with the text `HTML stands for HyperText Markup Language and it represents the content and structure of a webpage.`.

# --hints--

You should have a second `p` element.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Your `p` element should have the text `HTML stands for HyperText Markup Language and it represents the content and structure of a webpage.`

```js
const lastParaEl = document.querySelector("p:last-of-type");
assert.strictEqual(lastParaEl?.innerText, "HTML stands for HyperText Markup Language and it represents the content and structure of a webpage.");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>Welcome to freeCodeCamp</h1>
    <h2>Certified Full Stack Developer Certification</h2>
    <p>In this certification, you will learn all of the skills to become a full stack developer.</p>

    <h3>Introduction to HTML</h3>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>  
```
