---
id: 6823ac607bfdbc46331b2559
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will begin practicing HTML by building out the first few sections of a curriculum outline. HTML stands for HyperText Markup Language and it represents the content and structure of a web page. 

The first element you will learn about is the `h1` element.

The `h1` element is a heading element and represents the main title for a web page. It is best practice to only use one `h1` element per page.

Here is an example:

```html
<h1>I am the main title</h1>
```

`h1` elements have an opening tag(`<h1>`) and a closing tag `</h1>`.

Add an `h1` element with the text of `Welcome to freeCodeCamp`.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--


Your `h1` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector("h1"));
```

Your `h1` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Your `h1` element's text should be `Welcome to freeCodeCamp`. You have either omitted the text, have a typo, or it is not between the `h1` element's opening and closing tags.

```js
assert.equal(document.querySelector("h1")?.innerText, "Welcome to freeCodeCamp");
```

You appear to be using a browser extension that is modifying the page. Be sure to turn off all browser extensions.

```js
if(__checkForBrowserExtensions){
  assert.isAtMost(document.querySelectorAll("script").length, 2);
  assert.equal(document.querySelectorAll("style").length, 1);
  assert.equal(document.querySelectorAll("link").length, 0);
}
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html> 
```

