---
id: 6823d9ac8bdc3853df65a1ff
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

There are six heading elements in HTML. Here are all of them in order from most important to least:

```html
<h1>Heading level 1</h1>
<h2>Heading level 2</h2>
<h3>Heading level 3</h3>
<h4>Heading level 4</h4>
<h5>Heading level 5</h5>
<h6>Heading level 6</h6>
```

Below your `p` element, add an `h3` element that displays the following text:

```md
Introduction to HTML
```

# --hints--

Your `h3` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector("h3"));
```

Your `h3` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Your `h3` element's text should be `Introduction to HTML`. You can copy the text from the instructions. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
