---
id: 6823e169fda14755fbf00445
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

You're getting the hang of it. Now it is time to add another `h3` and `p` element to the page.

First, add an `h3` element with the text:

```md
Introduction to CSS
```

Then, below that `h3` element, add a `p` element with the following text:

```md
CSS is used to style a webpage
```

# --hints--

You should add a second `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Your `h3` element's text should be `Introduction to CSS`. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

You should add a third `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Your `p` element should have the text `CSS is used to style a webpage`. You can copy the text from the instructions.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```
