---
id: 682ba2318000b62f179bdf04
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML is made of elements. The first one you will learn about is the `h1` element. It's the main
heading element of a web page and it looks like this:

```html
<h1>Welcome to freeCodeCamp</h1> 
```

`h1` elements have an opening tag (`<h1>`) and a closing tag (`</h1>`) with the text in between the tags.

Turn your `Welcome to freeCodeCamp` text into an `h1` element by adding opening and closing `h1` tags.

When done correctly, you should see an enlarged text display on the screen.

# --hints--

Your `h1` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector("h1"));
```

Your `h1` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Your `h1` element's text should be `Welcome to freeCodeCamp`. You have either omitted the text, have a typo, or it is not between the `h1` element's opening and closing tags.

```js
assert.equal(document.querySelector("h1")?.innerText, "Welcome to freeCodeCamp");
```

You appear to be using a browser extension that is modifying the page. Be sure to turn off all browser extensions.

```js
if(__checkForBrowserExtensions){
  assert.isAtMost(document.querySelectorAll("script").length, 2);
  assert.equal(document.querySelectorAll("style").length, 1);
  assert.equal(document.querySelectorAll("link").length, 0);
}
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    Welcome to freeCodeCamp
--fcc-editable-region--
  </body>
</html> 
```
