---
id: 682cd206883fc7b25eb539c5
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Notice that the HTML code you write in the editor shows up in the preview.

Below the text you added, type the following on line 2:

```md
Full Stack Curriculum
```

When you are done, click the "check your code" button to see if it's correct.

# --hints--

You should have the text `Full Stack Curriculum` in your editor. Double-check for spelling.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full\s*stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
Welcome to freeCodeCamp
--fcc-editable-region--

--fcc-editable-region--
```

