---
id: 683921c4769fd23dbadec2fe
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

While HTML defines the structure and content of a webpage, CSS is used to add style — things like colors, fonts, spacing, and layout.

Below the `h3` you just added, add another paragraph element with the text:

```md
CSS is used to style a webpage
```

# --hints--

You should add a second `h3` element to the page.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

Your `h3` element's text should be `Introduction to CSS`. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

You should add a third `p` element to the page.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Your `p` element should have the text `CSS is used to style a webpage`. You can copy the text from the instructions.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
