---
id: 6153893900438b4643590367
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

將 `.header` 文本居中對齊。 使用以 `uppercase` 爲值的 `text-transform` 屬性，使文本大寫。

將其所有方向的內邊距都設置爲 `32px`。 將背景顏色設置爲 `#0a0a23`，文本顏色設置爲 `#fff` 。

添加 `border-bottom`，值爲 `4px solid #fdb347`。

# --hints--

你應該有一個 `.header` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

`.header` 選擇器應該有一個值爲 `center` 的 `text-align` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

`.header` 選擇器應該有一個值爲 `uppercase` 的 `text-transform` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

`.header` 選擇器應該有一個值爲 `32px` 的 `padding` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

`.header` 選擇器應該有一個值爲 `#0a0a23` 的 `background-color` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

`.header` 選擇器應該有一個值爲 `#fff` 的 `color` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

`.header` 選擇器應該有一個值爲 `4px solid #fdb347` 的 `border-bottom` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
