---
id: 61537d86bdc3dd343688fceb
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

接下來，根據圖片在文檔中的順序爲 `img` 指定 `src` 屬性。 第一個 `img` 的 `src` 的值應該爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。 其餘的依次類推，把 `1` 替換成 `img` 在文檔中的順序。

# --hints--

九個 `img` 元素都應該具有 `src` 屬性。

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

第一個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

第二個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg`。

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

第三個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg`。

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

第四個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg`。

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

第五個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg`。

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

第六個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg`。

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

第七個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg`。

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

第八個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg`。

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

第九個 `img` 元素的 `src` 屬性的值應該設置爲 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg`。

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
