---
id: 61537c5f81f0cf325b4a854c
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `body` 元素中創建一個 `header` 元素，設置其 class 爲 `header` 。

在 `header` 元素內嵌套一個帶有文本 `css flexbox photo gallery` 的 `h1` 元素。

# --hints--

在 `body` 元素內部應該有一個 `header` 元素。

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

`header` 元素的 `class` 應該爲 `header`。

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

`header` 元素內部應該有 `h1` 元素。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

`h1` 元素內文本應該爲 `css flexbox photo gallery`。

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
