---
id: 615380dff67172357fcf0425
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

現在你的圖片太大了。

創建一個 `.gallery img` 選擇器來選取你的圖片。 將它們的 `width` 都設置爲 `100%`， `max-width` 都設置爲 `350px`。

並將 `height` 屬性設置爲 `300px`，以保持圖片大小一致。

# --hints--

你應該有一個 `.gallery img` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

你的 `.gallery img` 選擇器應該有一個 `width` 屬性，其值設置爲 `100%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

你的 `.gallery img` 選擇器應該有一個 `max-width` 屬性，其值設置爲 `350px`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

你的 `.gallery img` 選擇器應該有一個 `height` 屬性，其值設置爲 `300px`

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
