---
id: 6153897c27f6334716ee5abe
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

將 `.header` 文本居中對齊，使其所有邊都填充 `32px`，並將背景設置爲 `#E0DDDD`。

# --hints--

你應該擁有一個 `.header` 選擇器

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

應該將 `.header` 選擇器的 `text-align` 屬性設置爲 `center`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign === 'center');
```

應該將 `.header` 選擇器的 `padding` 屬性設置爲 `32px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.padding === '32px');
```

應該將 `.header` 選擇器的 `background-color` 屬性設置爲 `#E0DDDD`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor === 'rgb(224, 221, 221)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

--fcc-editable-region--

--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
