---
id: 6153908a366afb4d57185c8d
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

Flexbox 可以被認爲有兩個軸，主軸和交叉軸。 主軸由 `flex-direction` 屬性決定，有四個可選值：

- `row`（默認值）：主軸爲水平方向，起點在左端。
- `row-reverse`：主軸爲水平方向，起點在右端。
- `column`：主軸爲垂直方向，起點在上沿。
- `column-reverse`：主軸爲垂直方向，起點在下沿。

**注意**：軸和方向將因文字方向而異。 上面的值是左至右的文本方向。

嘗試不同的值來觀察佈局的變化。

練習完成時，將 `.gallery` 元素的 `flex-direction` 設置爲 `row`。

# --hints--

應該將 `.gallery` 選擇器的 `flex-direction` 屬性設置爲 `row`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection === 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
