---
id: 61537485c4f2a624f18d7794
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

從標準的 HTML 模板開始。 添加 `DOCTYPE` 聲明、`html` 元素、 `head` 元素，以及一個 `body` 元素。

# --hints--

你的代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 聲明後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該將文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `>` 之後關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開始標籤。

```js
assert(code.match(/<html\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 聲明應該位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個開始的 `head` 標籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `head` 標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開始的 `body` 標籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `body` 標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 標籤應該是兄弟關係。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
