---
id: 61537bb9b1a29430ac15ad38
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `head` 元素中，添加文本設置爲 `Photo Gallery` 的 `title` 元素並且將 `styles.css` 文件添加到頁面的 `link` 元素。

# --hints--

你的 `link` 元素的 `href` 屬性值應爲 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

你的代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的項目應該擁有一個 `Photo Gallery` 的標題。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
