---
id: 61537485c4f2a624f18d7794
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

從標準的 HTML 模板開始。 添加 `DOCTYPE` 聲明、`html` 元素、 `head` 元素，以及一個 `body` 元素。

將值爲 `en` 的 `lang` 屬性添加到 `<html>` 開始標籤裏。

# --hints--

代碼應包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 引用之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

應該定義文檔類型爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

應該在 `DOCTYPE` 聲明的類型後面加上一個 `>`。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 元素應該有一個開始標籤。

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

`<html>` 開始標籤應該有一個值爲 `en` 的 `lang` 屬性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

你的 `DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

應該有一個 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/i));
```

應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同級元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
