---
id: 61537a8054753e2f1f2a1574
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一個 `meta` 標籤，將 `name` 設置爲 `viewport`，將 `content` 設置爲 `width=device-width, initial-scale=1`。

同時添加一個 `meta` 標籤，將 `charset` 設置爲 `UTF-8` 。

# --hints--

你應該擁有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

其中一個 `meta` 元素應該將 `name` 設置爲 `viewport`，將 `content` 設置爲 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

你的另一個 `meta` 元素應該將 `charset` 的屬性設置爲 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
