---
id: 615392916d83fa4f02f7e2cf
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

你可能已經注意到你的圖像已全部移動到同一行。

`flex-wrap` 屬性確定當 flex 容器太小時項目的行爲方式。 將此屬性設置爲 `wrap` 將允許項目換行到下一行/列（取決於主軸），`nowrap` 將阻止項目換行。 如果將其設置爲 `nowrap`，則項目可能會收縮以適應或溢出。

將 `#gallery` 選擇器的 `flex-wrap` 屬性設置爲 `wrap`。 你應該看到你的圖片是四列布局。 這是因爲你在前面的步驟中將它們的 `width` 設置成了 `25%`。

# --hints--

`#gallery` 選擇器的 `flex-wrap` 屬性應設置爲 `wrap`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery')?.flexWrap === 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

.header {
  text-align: center;
  padding: 32px;
  background: #E0DDDD;
}

--fcc-editable-region--
#gallery {
  display: flex;
  flex-direction: row;
}
--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
