---
id: 6153a3952facd25a83fe8083
title: 步驟 20
challengeType: 0
dashedName: step-20
---

# --description--

`::after` 用來創建一個僞元素，作爲已選中元素的最後一個子元素。 你可以用它在最後一個圖像後添加一個空元素。 如果你設置其 `width` 爲和圖像相同，如果是兩列布局，它會將最後一張圖像擠到左邊。 現在它居中顯示，因爲你在 flex 容器上設置了 `justify-content: center`。

例如：

```CSS
.container::after {
  content: "";
  width: 860px;
}
```

使用 `::after` 僞元素在 `.gallery` 元素上創建一個新的選擇器。 添加 `content` 屬性設置爲空字符串 `""` 並設置 `width` 爲 `350px`。

# --hints--

`.gallery::after` 選擇器的 `content` 屬性應該設置爲空字符串 `""`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.content === "\"\"");
```

`.gallery::after` 選擇器的 `width` 屬性應該設置爲 `350px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.width === "350px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
