---
id: 61537c5f81f0cf325b4a854c
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `body` 元素中，創建一個 `class` 設置爲 `header` 的 `div` 。

在 `.header` 元素內嵌套一個帶有文本 `CSS FLEXBOX PHOTO GALLERY` 的 `h1` 元素。

# --hints--

你的 `body` 元素中應該有一個 `div` 元素。

```js
assert.exists(document.querySelector('body')?.querySelector('div'))
```

你的 `div` 元素應該將 `class` 設置爲 `header`。

```js
assert(document?.querySelector('body')?.querySelector('div')?.classList?.contains('header'))
```

你的 `.header` 元素應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

你的 `h1` 應該包含文本 `CSS FLEXBOX PHOTO GALLERY`。 請記住，大小寫很重要。

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'CSS FLEXBOX PHOTO GALLERY')
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
