---
id: 64dd076a1b33c86d84b4232f
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

爲了更好地可視化你元素的尺寸，添加邊框會很有幫助。

將你的 `.gallery` 元素的 `width` 設置爲 `50%`，將 `border` 設置爲 `5px solid red`。

然後，將 `img` 元素的 `width` 設置爲 `100%`，將 `padding` 設置爲 `5px`，將 `border` 設置爲 `5px solid blue`。

# --hints--

你應該有一個 `.gallery` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

你的 `.gallery` 選擇器應該有一個 `width` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

你的 `.gallery` 選擇器應該有一個值爲 `50%` 的 `width` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

你的 `.gallery` 選擇器應該有一個 `border` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

你的 `.gallery` 選擇器應該有一個值爲 `5px solid red` 的 `border` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

你應該有一個 `img` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

你的 `img` 選擇器應該有一個 `width` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

你的 `img` 選擇器應該有一個值爲 `100%` 的 `width` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

你的 `img` 選擇器應該有一個 `padding` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

你的 `img` 選擇器應該有一個值爲 `5px` 的 `padding` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

你的 `img` 選擇器應該有一個 `border` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

你的 `img` 選擇器應該有一個值爲 `5px solid blue` 的 `border` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
