---
id: 686386e76197458b7dfd52f7
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Inside the unordered list, create four list items. All of them should have the entity `&#10003;` as their text, with the following text inside a `span` in them, top down:

| List Item | Text   |
| ----------- | ------- |
| First List Item|`All Premium features`|
| Second List Item|`Up to 6 accounts`|
| Third List Item|`Individual playlists & libraries`|
| Fourth List Item|`Family Mix playlists`|

# --hints--

You should create four list items inside the `ul`.

```js
const targetUl = document.querySelectorAll("ul")[2]
const listItems = targetUl.querySelectorAll("li")

assert.lengthOf(listItems, 4)
```

Your first `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[2]
const firstListItem = targetUl.querySelectorAll("li")[0]
const firstSpan = firstListItem?.querySelector("span")

assert.isTrue(firstSpan?.textContent.includes("✓"))
```

Your first `li` element should the text `All Premium features` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[2]
const firstListItem = targetUl.querySelectorAll("li")[0]

assert.isTrue(firstListItem?.textContent?.includes("All Premium features"))
```

Your second `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[2]
const secondListItem = targetUl.querySelectorAll("li")[1]
const secondSpan = secondListItem?.querySelector("span")

assert.isTrue(secondSpan?.textContent.includes("✓"))
```

Your second `li` element should the text `Up to 6 accounts` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[2]
const secondListItem = targetUl.querySelectorAll("li")[1]

assert.isTrue(secondListItem?.textContent?.includes("Up to 6 accounts"))
```

Your third `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[2]
const thirdListItem = targetUl.querySelectorAll("li")[2]
const thirdSpan = thirdListItem?.querySelector("span")

assert.isTrue(thirdSpan?.textContent.includes("✓"))
```

Your third `li` element should the text `Individual playlists & libraries` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[2]
const thirdListItem = targetUl.querySelectorAll("li")[2]

assert.isTrue(thirdListItem?.textContent?.includes("Individual playlists & libraries"))
```

Your fourth `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[2]
const fourthListItem = targetUl.querySelectorAll("li")[3]
const fourthSpan = fourthListItem?.querySelector("span")

assert.isTrue(fourthSpan?.textContent.includes("✓"))
```

Your fourth `li` element should the text `Family Mix playlists` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[2]
const fourthListItem = targetUl.querySelectorAll("li")[3]

assert.isTrue(fourthListItem?.textContent?.includes("Family Mix playlists"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>
            <div>Most Popular</div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>Enjoy all of the features with a plan for up to 6 family members.</p>
            <ul>
            --fcc-editable-region--
            
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
