---
id: 686386e76197458b7dfd52f2
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Still inside the `div`, create a `p` element with the text `Enjoy the full music experience with unlimited access and downloads.`, and a `ul` element under it. Leave the `ul` element empty for now.

# --hints--

You should create a `p` element with the text `Enjoy the full music experience with unlimited access and downloads.` inside the `div` element.

```js
const targetDiv = document.querySelectorAll("div")[5]

assert.equal(targetDiv.querySelectorAll("p")[1]?.textContent, "Enjoy the full music experience with unlimited access and downloads.")
```

You should create an empty `ul` element under the `p` element.

```js
const targetDiv = document.querySelectorAll("div")[5]
const p = targetDiv.querySelectorAll("p")[1]
const ul = p?.nextElementSibling

assert.equal(ul?.tagName, "UL")
assert.isEmpty(ul?.innerHTML.trim())
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>Ad-supported streaming
              </li>
              <li>
                <span>&#10003;</span>Curated playlists
              </li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>

        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
