---
id: 686386e76197458b7dfd52f4
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now, create the CTA button for the premium pricing card by creating an anchor element with the text `Go Premium`. Assign it an `href` attribute of `#`.

# --hints--

You should create an anchor element.

```js
assert.lengthOf(document.querySelectorAll("a"), 2)
```

Your anchor element should have an `href` set to `#`.

```js
const links = document.querySelectorAll("a")
assert.equal(links[1]?.getAttribute("href"), "#")
```

Your anchor element should have the text `Go Premium`.

```js
const links = document.querySelectorAll("a")
assert.equal(links[1]?.textContent, "Go Premium")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          --fcc-editable-region--
            
          --fcc-editable-region--
        </div>
      </div>
    </main>
  </body>
</html>
```
