---
id: 686386e76197458b7dfd52f1
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Create another `div`. Inside it, create an `h2` element with the text `Premium`. Also, create a `p` element with the text `$9.99` and a `span` inside it with the text `/month`.

# --hints--

You should create another `div` under the `div` with the text `Most popular`.

```js
assert.lengthOf(document.querySelectorAll("div"), 6)
```

You should create an `h2` with the text `Premium` inside the `div`.

```js
const targetDiv = document.querySelectorAll("div")[5]

assert.equal(targetDiv.querySelector("h2")?.textContent, "Premium")
```

You should create a `p` element with the text `$9.99` inside the `div`.

```js
const targetDiv = document.querySelectorAll("div")[5]

assert.equal(targetDiv.querySelector("p")?.textContent?.includes("$9.99"), true)
```

You should create a `span` with the text `/month` inside the `p` element.

```js
const targetDiv = document.querySelectorAll("div")[5]
const targetP = targetDiv.querySelector("p")

assert.equal(targetP.querySelector("span")?.textContent?.includes("/month"), true)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>Ad-supported streaming
              </li>
              <li>
                <span>&#10003;</span>Curated playlists
              </li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>

        <div>
          <div>Most Popular</div>
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
      </div>
    </main>
  </body>
</html>
```
