---
id: 686386e76197458b7dfd52f0
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now, you should handle the HTML for the next pricing card, which is the **Premium** plan.

Create a `div` element to wrap the content for the plan. Inside it, create another `div` element with the text `Most Popular`.

# --hints--

You should create a `div` element.

```js
const wrapperDiv = document.querySelector("main > div");
const planDivs = wrapperDiv.querySelectorAll(":scope > div");

assert.lengthOf(planDivs, 2)
```

You should create an inner `div` with the text `Most popular`.

```js
const wrapperDiv = document.querySelector("main > div");
const targetPlanDiv = wrapperDiv.querySelectorAll(":scope > div")[1];
const targetDiv = targetPlanDiv.querySelector("div")

assert.equal(targetDiv?.textContent.trim(), "Most popular")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>Ad-supported streaming
              </li>
              <li>
                <span>&#10003;</span>Curated playlists
              </li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```
