---
id: 686386e76197458b7dfd52ee
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Inside the `ul`, create two list items. Inside the first list item, create a `span` element and add the HTML entity `&#10003;` as its text. Right after that entity, add the text `Ad-supported streaming`.

Inside the second list item, create a `span` and add the entity `&#10003;`, and after that, add the text `Curated playlists`.

# --hints--

You should have two list items inside your `ul` element.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Your first `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Your first `li` element should have the text `Ad-supported streaming` just after the `span` element.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(listItems[0]?.textContent.includes("Ad-supported streaming"), true)
```

Your second `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[1].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

Your second `li` element should have the text `Curated playlists` just after the `span` element.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[1].querySelector("span")

assert.equal(listItems[1]?.textContent.includes("Curated playlists"), true)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--

            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
