---
id: 686386e76197458b7dfd52ec
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Under the `h2`, create a `p` element with the text `$0`. Then, inside that `p` element and after the text, create a `span` element with the text `/month`.

# --hints--

You should create a `p` element with the text `$0`.

```js
assert.equal(document.querySelector("p")?.textContent.includes("$0"), true);
```

You should create a `span` element with the text `/month` inside the `p` element.

```js
assert.equal(document.querySelector("p > span")?.textContent, "/month")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--

            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
