---
id: 686386e76197458b7dfd52f3
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Create three list elements inside the `ul`. Inside the first one, create a `span` with the character entity `&#10003;` as the text, and add the text `Ad-free listening` outside the `span` but inside the `li`.

For the second `li`, create a span with the character entity `&#10003;` as its text, and the text `Offline playback` outside the `span` but inside the `li`.

The last `li` element should also have a `span` with the character entity `&#10003;`, and the text `Unlimited skips` outside the `span` but inside the `li`.

# --hints--

You should have three list elements inside you `ul`.

```js
const targetUl = document.querySelectorAll("ul")[1]

assert.lengthOf(targetUl.querySelectorAll("li"), 3)
```

Your first `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[1]
const firstListItem = targetUl.querySelectorAll("li")[0]
const firstSpan = firstListItem?.querySelector("span")

assert.isTrue(firstSpan?.textContent.includes("✓"))
```

Your first `li` element should the text `Ad-free listening` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[1]
const firstListItem = targetUl.querySelectorAll("li")[0]

assert.isTrue(firstListItem?.textContent?.includes("Ad-free listening"))
```

Your second `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[1]
const secondListItem = targetUl.querySelectorAll("li")[1]
const secondSpan = secondListItem?.querySelector("span")

assert.isTrue(secondSpan?.textContent.includes("✓"))
```

Your second `li` element should the text `Offline playback` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[1]
const secondListItem = targetUl.querySelectorAll("li")[1]

assert.isTrue(secondListItem?.textContent?.includes("Offline playback"))
```

Your third `li` element should have a `span` element with the HTML character entity `&#10003;` in it.

```js
const targetUl = document.querySelectorAll("ul")[1]
const thirdListItem = targetUl.querySelectorAll("li")[2]
const thirdSpan = thirdListItem?.querySelector("span")

assert.isTrue(thirdSpan?.textContent.includes("✓"))
```

Your third `li` element should the text `Unlimited skips` just after the `span`.

```js
const targetUl = document.querySelectorAll("ul")[1]
const thirdListItem = targetUl.querySelectorAll("li")[2]

assert.isTrue(thirdListItem?.textContent?.includes("Unlimited skips"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>Ad-supported streaming
              </li>
              <li>
                <span>&#10003;</span>Curated playlists
              </li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>

        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              --fcc-editable-region--

              --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
