---
id: 6864dcea9ad9f86942ff0158
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

It's time to move to the `div` element that wraps the three pricing cards. To arrange the cards, you can use the Tailwind CSS grid utility classes.

Here's how to assign a display of `grid` to an element using Tailwind:

```html
<div class="grid">
  <div>Grid Item</div>
  <div>Grid Item</div>
  <div>Grid Item</div>
</div>
```

Assign a class of `grid` to the `div` element that wraps the pricing cards.

# --hints--

Your `div` element should have the class `grid`.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("grid"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 and mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      --fcc-editable-region--
      <div>
      --fcc-editable-region--
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a>Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a>Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a>Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
