---
id: 6864e31ef17310b9e6aa376e
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

The grid items are too attached to each end of the browser, so you should work on some spacing for them.

Assign the class `max-w-6xl` to the `div` to to set a fixed width of `72rem` and prevent the items from stretching too wide.

As that will push all of them to the left, add `mx-auto` to the `div` to finally center the items.

# --hints--

Your `div` element should have the class `max-w-6xl`.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("max-w-6xl"))
```

Your `div` element should have the class `mx-auto`.

```js
const firstDiv = document.querySelector("div")

assert.isTrue(firstDiv.classList.contains("mx-auto"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 and mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
--fcc-editable-region--
      <div class="grid grid-cols-1 md:grid-cols-3 ">
--fcc-editable-region--
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a>Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a>Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a>Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
