---
id: 6864d9999f8a4e3f0d569a3f
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Next, make the `h1` text distinguishable by assigning it a `class` set to `text-3xl` to make it bigger, `text-center` to center it, `font-semibold` to make it bolder, `text-gray-900` to make it darker, and `md:text-5xl` to make it bigger on medium and larger screen sizes.

# --hints--

Your `h1` element should have the class `text-3xl`.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-3xl"))
```

Your `h1` element should have the class `text-center`.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-center"))
```

Your `h1` element should have the class `font-semibold`.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("font-semibold"))
```

Your `h1` element should have the class `text-gray-900`.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-gray-900"))
```

Your `h1` element should have the class `md:text-5xl`.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("md:text-5xl"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      --fcc-editable-region--
      <h1>Choose your listening plan</h1>
      --fcc-editable-region--
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a>Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a>Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a>Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
