---
id: 6864f8fef1584f68696e8e24
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

Now, you should move to the `Start listening` CTA to make it a proper button.

Assign it the utility classes `bg-indigo-100`, `px-4`, `py-2`, `text-center`, `font-semibold`, `text-indigo-700`, and `hover:bg-indigo-200`.

# --hints--

Your anchor element should have the class `bg-indigo-100`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("bg-indigo-100"))
```

Your anchor element should have the class `px-4`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("px-4"))
```

Your anchor element should have the class `py-2`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("py-2"))
```

Your anchor element should have the class `text-center`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("text-center"))
```

Your anchor element should have the class `font-semibold`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("font-semibold"))
```

Your anchor element should have the class `text-indigo-700`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("text-indigo-700"))
```

Your anchor element should have the class `hover:bg-indigo-200`.

```js
const firstATag = document.querySelector("a")

assert.isTrue(firstATag.classList.contains("hover:bg-indigo-200"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      <h1 class="mt-8 and mb-12 text-center text-3xl md:text-5xl font-semibold text-gray-900">Choose your listening plan</h1>
      <div class="grid grid-cols-1 md:grid-cols-3 max-w-6xl mx-auto gap-8 mt-16">
        <div class="bg-gray-100 ring-1 ring-gray-300 grid grid-rows-[1fr_auto] rounded-xl p-8 gap-6">
          <div class="grid grid-rows-[auto_auto_auto_1fr] gap-y-2">
            <h2 class="text-lg font-semibold text-indigo-600">Listener</h2>
            <p class="text-4xl font-bold text-gray-900">$0<span class="text-base font-medium text-gray-500">/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul class="mt-6 space-y-2 text-sm text-gray-700">
              <li>
                <span aria-hidden class="text-green-700 mr-2">&#10003;</span>Ad-supported streaming
              </li>
              <li>
                <span aria-hidden class="text-green-700 mr-2">&#10003;</span>Curated playlists
              </li>
            </ul>
          </div>
          --fcc-editable-region--
          <a href="#">Start listening</a>
          --fcc-editable-region--
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
