---
id: 6995d0325fd7da40854fac40
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In a previous lesson, you learned about binary search and how it differs from linear search.

In this workshop, you will put that into practice by building a binary search algorithm.

To begin, create a `binarySearch` function with a `searchList` and `value` parameters.

`searchList` is the array of elements the function will search within, and `value` is the target the function will try to find within the `searchList`.

# --hints--

You should create a `binarySearch` function.

```js
const explorer = await __helpers.Explorer(code);
const binarySearchFunc = explorer.allFunctions.binarySearch;
```

Your `binarySearch` function should have a `searchList` and `value` parameters.

```js

assert.isDefined(binarySearchFunc,"you have to define a function named binarySearch");
const paramNames = binarySearchFunc.parameters.map(p => p.name);
assert.deepEqual(paramNames, ["searchList", "value"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
