---
id: 6996c95e5bede03bdf2937a2
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

The next thing to do is to create a condition that will check if the target `value` is in the middle. 

Inside the `while` loop, create an `if` statement that checks if the target `value` is equal to `valueAtMiddle`. If it is, return the `pathToTarget` variable.

# --hints--

You should create an `if` statement with the condition `if (value === valueAtMiddle)`.

```js
assert.match(
  String(binarySearch),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

You should return the `pathToTarget` variable inside the `if` statement.

```js
assert.match(
  String(binarySearch),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
