---
id: 6996cfb6c01e9a802add38be
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Finally, if the `value` is not found at the middle and is also not greater than `valueAtMiddle`, then the `value` must be less than `valueAtMiddle`, and must be on the left. The final `else` block will handle this.

So, finish up the loop by removing the `break` keyword and adding an `else` block. Inside it, update the `high` variable by subtracting `1` from `mid`.

# --hints--

You should add an `else` block that updates the `high` variable to `mid - 1`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
// NOTE: This is the best we can currently do without Explorer AST support for identifying
// content in loop or inner blocks, which has been proposed for future enhancement.
// It matches the following variations:
// "else { high = mid - 1 }"
// "else { ANYTHING; high = mid-1; ANYTHING }"
assert.match(
  codeWithoutComments,
  /else\s*\{[\s\S]*?high\s*=\s*mid\s*-\s*1\s*;?[\s\S]*?\}/
);
```

You should remove the `break` keyword from the `while` loop.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
// NOTE: This is the best we can currently do without Explorer AST support for identifying
// content in loop or inner blocks, which has been proposed for future enhancement.
// It EXCLUDES the following variations:
// "break"
// "break;"
assert.notMatch(codeWithoutComments, /\bbreak\b/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
