---
id: 6996ccadf9048c79b857bb15
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

You can see that the first call returns `[3]`. That's because `3` is the middle value in `[1, 2, 3, 4, 5]`.

On the other hand, the second call returns an empty array. This happens because, after checking the middle value, the loop currently lacks the logic to narrow the search range by updating the `low` or `high` variables. It only acts as a single midpoint check and then finishes without further searching.

To allow the binary search to continue narrowing its search, add an `else if` statement that checks if `value` is greater than `valueAtMiddle`.

# --hints--

You should add an `else if` statement with the condition `value > valueAtMiddle`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
// Match the following variations:
// "else if (value > valueAtMiddle) { ANYTHING }"
// "else if (valueAtMiddle < value) { ANYTHING }"
assert.match(
  codeWithoutComments,
  /else\s+if\s*\((?:\s*value\s*>\s*valueAtMiddle\s*|\s*valueAtMiddle\s*<\s*value\s*)\)\s*\{[\s\S]*?\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
--fcc-editable-region--
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
