---
id: 6996c7363a4a437839f1f15d
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now that you've defined the boundaries of your search, it's time to create the loop that will perform the binary search algorithm.

Binary search works by repeatedly narrowing down the search space. This process continues as long as there is a valid range of elements to check.

You can express this with a `while` loop that will continue as long as your `low` pointer is less than or equal to your `high` pointer.

If `low` ever becomes greater than `high`, it means the search space has become empty, and the `value` is not in the array.

Inside the `binarySearch` function, create a `while` loop with a condition that checks if `low` is less than or equal to `high`.

# --hints--

You should create a `while` loop that runs while `low` is less than or equal to `high`. Leave the body of the loop empty for now.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(binarySearch?.toString()));
const validExpressions = ["while(low<=high){}", "while(high>=low){}"];
assert.isTrue(validExpressions.some(expr => strippedCode.includes(expr)));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
