---
id: 6996d18b31aab5cbbffbc731
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

One final small enhancement you can add is to also show the index at which the target `value` is found.

So, instead of returning only `pathToTarget` in the `if` statement, return an array containing `pathToTarget` and a template literal with the message `Value found at index ${mid}`.

With that, your binary search algorithm workshop is complete!

# --hints--

You should update the `return` statement in the `if` so the function returns an array with `pathToTarget` as the first element and a template literal for `Value found at index ${mid}` as the second element. Be sure to use backticks for the template literal.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch: binarySearchFn } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearchFn?.toString());
assert.match(
  codeWithoutComments,
  /return\s*\[\s*pathToTarget\s*,\s*(?:`Value found at index\s*\$\{\s*mid\s*\}`|['"]Value found at index\s*['"]\s*\.\s*concat\(\s*mid\s*\))\s*\]/
);
```

Your function should return the expected results.

```js
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5], 3), [[3], "Value found at index 2"]
);
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5, 9], 4), [[3, 5, 4], "Value found at index 3"]
);
assert.deepEqual(
  binarySearch([1, 3, 5, 9, 14, 22], 10), [[], "Value not found"]
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
--fcc-editable-region--
      return pathToTarget;
--fcc-editable-region--
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```

# --solutions--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return [pathToTarget, `Value found at index ${mid}`];
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
