---
id: 6996c8f94fe6b250f11a2a1e
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now, include `valueAtMiddle` in the `pathToTarget` array to track the steps taken during the search, regardless of whether it's the target.

# --hints--

You should use the `push()` method to add `valueAtMiddle` to the `pathToTarget` array.

```js
assert.match(
  String(binarySearch),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
