---
id: 6996c95e5bede03bdf2937a2
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

The next thing to do is to create a condition that will check if the target `value` is in the middle. 

Inside the `while` loop, create an `if` statement that checks if the target `value` is equal to `valueAtMiddle`. If it is, return the `pathToTarget` variable.

# --hints--

You should create an `if` statement with the condition `value === valueAtMiddle`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(binarySearch?.toString()));
// Match the following variations:
// "while(ANYTHING){ANYTHING;if(value===valueAtMiddle){ANYTHING}ANYTHING}}"
// "while(ANYTHING){ANYTHING;if(valueAtMiddle==value){ANYTHING}ANYTHING}}"
assert.match(
  strippedCode,
  /while\([^)]*\)\{[^{}]*if\((?:value={2,3}valueAtMiddle|valueAtMiddle={2,3}value)\)\{[^{}]*\}/
);
```

You should return the `pathToTarget` variable inside the `if` statement.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
// Match the following variations:
// "while (ANYTHING) { ANYTHING; if (ANYTHING) { ANYTHING; return pathToTarget; ANYTHING } ANYTHING }"
// "while(ANYTHING){ANYTHING;if(ANYTHING){ANYTHING;return(pathToTarget)ANYTHING}ANYTHING}}"
assert.match(
  codeWithoutComments,
  /while\s*\([^)]*\)\s*\{[^{}]*if\s*\([^)]*\)\s*\{[^{}]*\breturn(?:\s+pathToTarget|\s*\(\s*pathToTarget\s*\))\s*;?[^{}]*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
