---
id: 6996ca21a74104721f2c7e5d
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

You need to test out things so you will understand the flow of the algorithm at this initial stage.

To do that, you need to first break out of the loop. That's because the current implementation will only allow one iteration, so if the condition is not met, there will be an infinite loop.

Just after the `if`statement, use the `break` keyword to break out of the `while` loop. Then, after the `while` loop, return an empty array to signify that the value was not found.

# --hints--

You should use the `break` keyword to break out of the `while` loop after the `if` statement.

```js
assert.match(
  String(binarySearch),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

You should return an empty array after your `while` loop.

```js
assert.match(String(binarySearch), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
