---
id: 6996c60281feca6718cc15c1
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Inside your `binarySearch` function, create a `pathToTarget` variable and initialize it to an empty array, for a start.

# --hints--

You should declare a `pathToTarget` variable set to an empty array within your `binarySearch` function.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
