---
id: 615380dff67172357fcf0425
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--
In order to control the CSS Box Model, we have to create an `*` selector and use the `box-sizing` property. This property influences the sizes of the elements in your HTML file. From there, use the value,`content-box`. This is the default behavior and adds the padding and border to the width and height of the elements.  
Next, change the value from `content-box` to `border-box`. Unlike the former, the `border-box` value only changes an element's padding and border and keeps its width and height the same.  

# --hints--

You should have a `*` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Your `*` selector should have a `box-sizing` property set to `border-box` as the value.
When this property is set to `border-box`, an element's total width and height is formed by combining it with its padding and border.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
