---
id: 6493bc0d99879635209565aa
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Ooh you have a problem. See how the images overflow the container? That's because the box-sizing property is at the content-box value, aka the default. This doesn't include the padding and border when calculating an element's size. Now, to solve this issue, use border-box, the box-sizing property's second value. Unlike content-box, border-box excludes the margins and only adds the padding and border in calculating the size of the element. Try changing the box-sizing value from one to the other, and see the difference.
Lastly, give the img elements a uniform height of 300px.


# --hints--

This one: You should have a `.gallery img` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

Your `.gallery img` selector should have a `width` property set to `100%` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

Your `.gallery img` selector should have a `max-width` property set to `350px` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

Your `.gallery img` selector should have a `height` property set to `300px` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

Your `.gallery img` selector should have a `box-sizing` property set to `border-box` as the value. 

```js
assert( new __helpers.CSSHelp(document).getStyle('.gallery img')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
.gallery {
  border: 2px solid red;
  width: 50%;
  height: 50%;
}

img {
  width: 100%;
  height: 100%;
  padding: .5em;
  border: 1px solid green;
  box-sizing: content-box;
}
--fcc-editable-region--

--fcc-editable-region--

```
