---
id: 6493bc0d99879635209565aa
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Your images are too big. Now, before we try to change their sizing and positioning. We now have to ask ourselves "How can I change their size without making them too small?" If we change only the images' size and then add padding and border, the size is incorrect. They will overflow the container, or flow in other unexpected ways. The solution to this is the `box-sizing` property's second option, `border-box`. Unlike the `content-box` value, the border-box value will include padding and border in the size of the element, but *excludes* the margins, width, and height outside of the border.  From there, set the `height property` to 300px to keep your images a uniform size.


# --hints--

This one: You should have a `.gallery img` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

Your `.gallery img` selector should have a `width` property set to `100%` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

Your `.gallery img` selector should have a `max-width` property set to `350px` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

Your `.gallery img` selector should have a `height` property set to `300px` as the value.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

Your `.gallery img` selector should have a `box-sizing` property set to `border-box` as the value. 

```js
assert( new __helpers.CSSHelp(document).getStyle('.gallery img')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
.gallery {
  width: 50%;
  height: 50%;
  border: 2px solid red;
}
--fcc-editable-region--

--fcc-editable-region--

```
