---
id: 6494da0daf5df5197963671d
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Before you start sizing and positioning things, you need to address this question, "how is the size of my container being calculated?" There are two different ways to do this, and they can significantly impact how our content is displayed. 

Now remember, a block-level container consists of four parts: the content of the container itself, the padding around that content, the border around that padding and the margin around the border.  With that  in mind, this is where you meet the CSS property, `box-sizing`. It lets you specify how those parts are used in calculating a container's size.
If you do `box-sizing: content-box;`, you pretty much said, "the size of the box is solely the size of its content." Any padding (the space between the content and its border) and the border are not used in calculating the size of a box. This is the default for most browsers, but let's set it explicitly.

# --hints--

You should have a `img` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Your `img` selector should have a `box-sizing` property.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.boxSizing);
```

Your `img` selector should have its `box-sizing` property set to `content-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.boxSizing==='content-box');
```

Your `img` selector should have a `width` and `height` property, each set to `100%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.width==='100%' && new __helpers.CSSHelp(document).getStyle('img')?.height==='100%');
```

Your `img` selector should have a `padding` property.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Your `img` selector should have a `border` property.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.border);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
     <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">

  </body>
    
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--
```
