---
id: 6494da0daf5df5197963671d
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Before we get into sizing and positioning things, we need to address the question, "how is the size of my container being calculated?" There are two different ways to do this, and which we use can significantly impact how our content is displayed.

A block-level container consists of four parts: the content of the container itself, the padding around that content, the border around that padding and the margin around the border. The CSS property box-sizing lets us specify how those parts are used in calculating a container's size.

If we set box-sizing: content-box;, we are saying "the size of the box is solely the size of its content." Any padding (the space between the content and its border) and the border are not used in calculating the size of a box. This is the default for most browsers, but lets set it explicitly.
# --hints--

You should have a `img` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Your `img` selector should have a `box-sizing` property.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.boxSizing);
```

Your `img` selector should have its `box-sizing` property set to `content-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.boxSizing==='content-box');
```

Your `img` selector should have a `width` and `height` property, each set to `100%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.width==='100%' && new __helpers.CSSHelp(document).getStyle('img')?.height==='100%');
```

Your `img` selector should have a `padding` property.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Your `img` selector should have a `border` property.

```js
assert(new __helpers.CSSHelp(document).getStyle('img')?.border);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
