---
id: 60b69a66b6ddb80858c5159b
title: 步骤 37
challengeType: 0
dashedName: step-37
---

# --description--

是时候使用一些 FontAwesome 图标了。

`i` 元素用于习语文本，或与“正常”文本内容分开的文本。 这可能适用于_斜体_文本，例如科学术语，也可能适用于 FontAwesome 提供的图标。

在你的 `#white-paper` 元素中，添加四个 `i` 元素。 给它们一个 `fas fa-music` 的 `class`。

这个特殊的类是 FontAwesome 如何确定要加载哪个图标。 `fas` 表示图标的类别（此处为 FontAwesome Solid 字体），而 `fa-music` 选择特定图标。

# --hints--

你的 `#white-paper` 元素中应该有四个新的 `i` 元素。

```js
assert(document.querySelectorAll('#white-paper > i')?.length === 4);
```

你的所有 `i` 元素都应将 `class` 设置为 `fas fa-music`。

```js
const icons = document.querySelectorAll('#white-paper > i');
for (const icon of icons) {
  assert(icon.classList?.contains('fas'));
  assert(icon.classList?.contains('fa-music'));
};
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
      <div id="black-character">
        <div id="black-hat"></div>
        <div id="gray-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="white-paper">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

#white-hat {
  width: 0;
  height: 0;
  border-style: solid;
  border-width: 0 120px 140px 180px;
  border-top-color: transparent;
  border-right-color: transparent;
  border-bottom-color: GhostWhite;
  border-left-color: transparent;
  position: absolute;
  top: -140px;
  left: 0;
}

#black-mask {
  width: 100%;
  height: 50px;
  background-color: rgb(45, 31, 19);
  position: absolute;
  top: 0;
  left: 0;
  z-index: 1;
}

#gray-instrument {
  width: 15%;
  height: 40%;
  background-color: rgb(167, 162, 117);
  position: absolute;
  top: 50px;
  left: 125px;
  z-index: 1;
}

.black-dot {
  width: 10px;
  height: 10px;
  background-color: rgb(45, 31, 19);
  border-radius: 50%;
  display: block;
  margin: auto;
  margin-top: 65%;
}

#tan-table {
  width: 450px;
  height: 140px;
  background-color: #D2691E;
  position: absolute;
  top: 275px;
  left: 15px;
  z-index: 1;
}
```
