---
id: 60b69a66b6ddb80858c51580
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

`z-index` 属性用于为你的 HTML 元素创建“层”。 如果你熟悉图像编辑工具，你可能以前使用过图层。 这是一个类似的概念。

`z-index` 值较高的元素会显示在 `z-index` 值较低的元素之上。 这可以与上一课中的定位相结合，创造出独特的效果。

由于 `back-wall` 元素需要出现在你将要创建的其他元素的“后面”，因此为 `back-wall` 元素提供一个 `z-index` 为 `-1` 的属性。

# --hints--

`#back-wall` 选择器应该将 `z-index` 属性设置为 `-1`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.zIndex === '-1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
