---
id: 60b69a66b6ddb80858c51588
title: 步骤 18
challengeType: 0
dashedName: step-18
---

# --description--

使用 id 选择器，为 id 为 `white-hat` 的元素设置样式。 给它一个 `width` 和 `height` 均为 `0`，以及一个 `solid` 的 `border-style`。

# --hints--

应该使用 `#white-hat` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat'));
```

`#white-hat` 选择器应该将 `width` 属性设置为 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.width === '0px');
```

`#white-hat` 选择器应该将 `height` 属性设置为 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.height === '0px');
```

`#white-hat` 选择器应该有一个 `border-style` 属性设置为 `solid`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.borderStyle === 'solid');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

--fcc-editable-region--

--fcc-editable-region--

```
