---
id: 60b69a66b6ddb80858c51579
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一个 `meta` 标签，并将 `charset` 属性设置为 `utf-8`。 同时添加一个包含文本 `Picasso Painting` 的 `title` 元素。

# --hints--

应该只添加一个 `meta` 元素。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 元素应该有一个 `charset` 属性。

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

`charset` 属性应设置为 `utf-8`。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

应该只添加一个 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 元素应该包含文本 `Picasso Painting`。 请注意拼写和大小写的问题。

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
