---
id: 60b69a66b6ddb80858c51581
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

在你的 `#back-wall` 元素的下方，创建一个 `div`，其包含 `class` 为 `characters`。 这是你将创建绘画角色的地方。

# --hints--

应该只添加一个新的 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 2);
```

新的 `div` 元素应该在 `#back-wall` 元素之后。

```js
assert(document.querySelector('#back-wall')?.nextElementSibling?.localName === 'div');
```

新的 `div` 元素应该将 `class` 设置为 `characters`。

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('characters'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
