---
id: 60b69a66b6ddb80858c51586
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

使用 id 选择器为 id 为 `offwhite-character` 的元素创建规则。 给它一个 `width` 为 `300px`、一个 `height` 为 `550px` 和一个 `background-color` 为 `GhostWhite`。

# --hints--

应该使用 `#offwhite-character` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character'));
```

`#offwhite-character` 应该有一个 `width` 属性设置为 `300px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.width === '300px');
```

`#offwhite-character` 应该有一个 `height` 属性设置为 `550px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.height === '550px');
```

`#offwhite-character` 应该有一个 `background-color` 属性设置为 `GhostWhite`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.backgroundColor === 'ghostwhite');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

--fcc-editable-region--

--fcc-editable-region--

```
