---
id: 60b69a66b6ddb80858c51584
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

这个角色需要眼睛。 在 `#black-mask` 元素中创建两个 `div` 元素。 按顺序为它们提供 `eyes left` 和 `eyes right` 类。

# --hints--

应该在 `#black-mask` 元素中创建 2 个 `div` 元素。

```js
assert(document.querySelectorAll('#black-mask div').length === 2);
```

第一个新 `div` 应该有 `eyes` 和 `left` 类。

```js
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('left'));
```

第二个新 `div` 应该有 `eyes` 和 `right` 类。

```js
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
        <div id="gray-instrument"></div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
