---
id: 60b69a66b6ddb80858c51578
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

首先设置你的 HTML 结构。 添加一个 `<!DOCTYPE>` 声明和一个 `html` 元素，其中 `lang` 属性设置为 `en`。 在 `html` 元素中，添加一个 `head` 元素和一个 `body` 元素。

# --hints--

代码应包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

应该在 `DOCTYPE` 引用之后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

应该将文档类型定义为 `html`。

```js
assert(code.match(/html/gi));
```

应该在类型之后使用 `>` 关闭 `DOCTYPE` 声明。

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 元素应该有一个 `lang` 属性为 `en` 的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

应该有一个 `head` 的开始标签。

```js
assert(code.match(/<head\s*>/i));
```

应该有一个 `head` 的结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html

--fcc-editable-region--

--fcc-editable-region--

```

```css

```
