---
id: 60b69a66b6ddb80858c5157f
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

通常，HTML 以自上而下的方式呈现。 代码顶部的元素位于页面顶部。 但是，很多时候你可能希望将元素移动到不同的位置。 你可以使用 `position` 属性来执行此操作。

将 `#back-wall` 元素的 `position` 属性设置为 `absolute`。 `absolute` 位置将元素从自上而下的文档流中取出，并允许你相对于其容器调整它。

手动定位元素时，你可以使用 `top`、`left`、`right` 和 `bottom` 移动其布局。 设置 `#back-wall` 的 `top` 值为 `0`，`left` 值为 `0`。


# --hints--

`#back-wall` 选择器应该将 `position` 属性设置为 `absolute`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.position === 'absolute');
```

`#back-wall` 选择器应该将 `top` 属性设置为 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.top === '0px');
```

`#back-wall` 选择器应该将 `left` 属性设置为 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.left === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
