---
id: 60b69a66b6ddb80858c5157a
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

FontAwesome 是一个 SVG 图标库，其中许多都可以免费使用。 你将在此项目中使用其中一些图标，因此你需要将外部样式表链接到你的 HTML。

添加一个 `link` 元素，其中 `rel` 为 `stylesheet`，`href` 为 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

# --hints--

你应该有两个 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 2);
```

`link` 元素应该具有 `stylesheet` 的 `rel`。

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('rel') === 'stylesheet');
```

`link` 元素应该有一个 `href` 为 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
