---
id: 60b69a66b6ddb80858c51595
title: 步骤 31
challengeType: 0
dashedName: step-31
---

# --description--

使用 id 选择器，为 id 为 `tan-table` 的元素设置样式。 给它一个 `width` 值为 `450px`、一个 `height` 值为 `140px`，和一个 `background-color` 值为 `#D2691E`。

# --hints--

应该有一个 `#tan-table` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('#tan-table'));
```

`#tan-table` 选择器应该有一个 `width` 属性设置为 `450px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#tan-table')?.width === '450px');
```

`#tan-table` 选择器应该将 `height` 属性设置为 `140px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#tan-table')?.height === '140px');
```

`#tan-table` 选择器应该将 `background-color` 属性设置为 `#D2691E`。

```js
assert (new __helpers.CSSHelp(document).getStyle('#tan-table')?.backgroundColor === 'rgb(210, 105, 30)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

#white-hat {
  width: 0;
  height: 0;
  border-style: solid;
  border-width: 0 120px 140px 180px;
  border-top-color: transparent;
  border-right-color: transparent;
  border-bottom-color: GhostWhite;
  border-left-color: transparent;
  position: absolute;
  top: -140px;
  left: 0;
}

#black-mask {
  width: 100%;
  height: 50px;
  background-color: rgb(45, 31, 19);
  position: absolute;
  top: 0;
  left: 0;
  z-index: 1;
}

#gray-instrument {
  width: 15%;
  height: 40%;
  background-color: rgb(167, 162, 117);
  position: absolute;
  top: 50px;
  left: 125px;
  z-index: 1;
}

.black-dot {
  width: 10px;
  height: 10px;
  background-color: rgb(45, 31, 19);
  border-radius: 50%;
  display: block;
  margin: auto;
  margin-top: 65%;
}

  --fcc-editable-region--

  --fcc-editable-region--

```
