---
id: 60b80da8676fb3227967a731
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

链接你的 CSS 文件，即使你还没有编写任何 CSS。

添加一个 `link` 元素，其中 `rel` 值为 `stylesheet`，`href` 值为 `styles.css`。

# --hints--


你的代码应该有一个 `link` 元素。

```js
assert.match(code, /<link/)
```

你应该有一个自闭合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
