---
id: 6853e58d52ec64349a5760ca
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will learn how to implement the shortest path algorithm using Dijkstra's algorithm. You will write a Python function that computes the shortest path between node in a graph, and also returns the path taken.

To get started, define a variable named `INF` and assign it the value `float("inf")`.

# --hints--

You should define the variable `INF` at the top of your file.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`))
  }
})
```

You should assign `float("inf")` to the variable `INF`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).find_variable("INF").is_assigned("float('inf')")`))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
