---
id: 686b78bc13d518e1267448e4
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

In addition to tracking distances, we also need to keep track of the actual paths taken to reach each node

For this, we will use a list named `paths` that stores the path taken to reach each node.

Each entry in the `paths` list starts as a list containing just the node number itself. Use a list comprehension to create a list of lists like this:

```py
  routes = [[node_post]]
```

Create a variable named `paths` and initialize it using an empty list. Inside the list add a list `[node_no]`. 

# --hints--

You should have a variable called `paths`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("paths")`))
  }
})
```

You should initialize `paths` with an empty list.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).find_variable("paths").is_assigned("[]")`))
  }
})
```

You should have `[node_no]` inside the list.

```js
({
  test: () => {
    assert(runPython(`"[node_no]" in _Node(_code).find_variable("paths").source()`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  
--fcc-editable-region--
```
