---
id: 686b7c80bb309ae443946fac
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

We need to create one path list for each node in the graph. To do this, we’ll use a `for` loop inside the list comprehension that goes through all node numbers.

Inside the list comprehension, immediately after `[node_no]`, add a `for` loop that iterates over the range of `n`.

# --hints--

You should create a `for` loop.

```js
({
  test: () => {
    assert(runPython(`(_Node(_code).has_for_loop())`))
  }
})
```

You should have the `for` loop immediately after `[node_no]`.

```js
({
  test: () => {
    assert(runPython(`any("[node_no] for" in comp.source for comp in _Node(_code).find_list_comprehensions())`))
  }
})
```

The `for` loop should iterate over the range of `n`.

```js
({
  test: () => {
    assert(runPython(`any("range(n)" in loop.source for loop in _Node(_code).find_for_loops())`))
  }
})
```

You should have the `for` loop inside the list comprehension.

```js
({
  test: () => {
    assert(runPython(`any(loop.is_in_list_comp() for loop in _Node(_code).find_for_loops())`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0

--fcc-editable-region--
  paths = [[node_no]]

--fcc-editable-region--
```
