---
id: 687b48ccb46c5e6126d30bc6
title: Step 28
challengeType: 20
dashedName: step-28
---

# --description--

Finally, print the results for the current node.

Your output should show two things, the distance from the `start_node` to `node_no` and the path taken to get there.

Format the output clearly so the user can easily read both the distance and the full path.

# --hints--

You should print the distance from the `start_node` to `node_no`.

```js
({
  test: () => {
    assert(runPython(`
_Node(_code).find_calls("print").some(lambda call:
    "distance" in call.code and "start_node" in call.code and "node_no" in call.code
)
    `))
  }
})
```

You should use an f-string to format the output.

```js
({
  test: () => {
    assert(runPython(`
_Node(_code).find_calls("print").some(lambda call:
    call.code.strip().startswith("f")
)
    `))
  }
})
```

You should include both the distance and the path in the output.

```js
({
  test: () => {
    assert(runPython(`
_Node(_code).find_calls("print").some(lambda call:
    "path" in call.code
)
    `))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        new_distance = distances[current] + distance
        if new_distance < distances[node_no]:
          distances[node_no] = new_distance
          paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
  targets = [target_node] if target_node is not None else range(n)
  for node_no in targets:
    if node_no == start_node or distances[node_no] == INF:
      continue
    path = " -> ".join(str(n) for n in paths[node_no])
    
--fcc-editable-region--
```
