---
id: 6a89c69cbe1bb69845992d8b
title: Build a Range of Numbers Generator
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

In this lab, you will build a `range_of_numbers` function that uses recursion to generate a list of numbers within a specified range.

For example:

- `range_of_numbers(3, 9)` should return `[3, 4, 5, 6, 7, 8, 9]`.
- `range_of_numbers(5, 5)` should return `[5]`.

Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a function named `range_of_numbers` with two parameters named `start_num` and `end_num`.
1. The function should return a list of consecutive integers that begins with `start_num` and ends with `end_num`, including both values.
1. The value of `start_num` will always be less than or equal to `end_num`.
1. Your function should use recursion by calling itself.
1. Your function should not use `for` or `while` loops, comprehensions, or the `range()` function.
1. When `start_num` equals `end_num`, the function should return a list containing `start_num`. This is the base case.
1. For the recursive case, the function should call itself with an argument that moves toward the base case, then add the current number to the returned list.
1. The function should not use global variables to store the returned list.

# --hints--

You should have a function named `range_of_numbers` with parameters named `start_num` and `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` should return `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` should return `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` should return `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` should return `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

Your `range_of_numbers` function should call itself.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

Your `range_of_numbers` function should not use loops, comprehensions, or the `range()` function.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Calling `range_of_numbers(1, 3)` should not affect a later call. After that call, `range_of_numbers(6, 9)` should still return `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
