---
id: 5d822fd413a79914d39e98d0
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Crea un elemento `div` en el `body` con una clase de `background-buildings`. Este será un contenedor para un grupo de edificios.

# --hints--

Debe crear un elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

El elemento `div` debe estar dentro del elemento `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

El elemento `div` debe tener una clase de `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

