---
id: 5d822fd413a79914d39e98d1
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Asigne a su elemento `.background-buildings` un `width` y `height` de `100%` para convertirlo en el ancho y alto completos de su padre, el `body`.

# --hints--

Debe utilizar la clase `background-buildings` para seleccionar el elemento correcto.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

El elemento `.background-buildings` debe tener un `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

El elemento `.background-buildings` debe tener un `height` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

