---
id: 5d822fd413a79914d39e98cd
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

En CSS, puedes orientar todo con un asterisco. Agregue un borde a todo usando el selector `*` y dándole un `border` de `1px solid black`. Este es un truco que ayuda a visualizar dónde están los elementos y su tamaño. Eliminará esto más tarde.

# --hints--

Debe utilizar el selector `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Debe usar la propiedad `border` para aplicar estilo a todos los elementos.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Todos los elementos deben tener un borde `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
