---
id: 5d822fd413a79914d39e98ce
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

También agregue un `box-sizing` de `border-box` a todo. Esto hará que el borde que agregó no agregue ningún tamaño a sus elementos.

# --hints--

Debe usar la propiedad `box-sizing`.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

Debe utilizar el selector `*` existente.

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

Todos los elementos deben tener un `box-sizing` de `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

