---
id: 5d822fd413a79914d39e98cf
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Puedes ver el `body` (es el cuadro más interno de tu página); el cuadro que lo rodea es el elemento `html`. Haz que tu `body` llene toda la ventana gráfica dándole una `height` de `100vh`. Elimine el `margin` predeterminado del `body` estableciendo el `margin` en `0`. Finalmente, establezca la propiedad `overflow` en `hidden` para ocultar las barras de desplazamiento que aparecen cuando algo se extiende más allá de la ventana gráfica.

# --hints--

Debe utilizar el selector `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Tu `body` debe tener una `height` de `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Su `body` debe tener un `margin` de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Su `body` debe tener la propiedad `overflow` configurada en `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

