---
id: 5d822fd413a79914d39e98d2
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Anida un `div` con una clase de `bb1` en el contenedor de edificios en segundo plano. Abra su archivo `styles.css` y asigne a `.bb1` un `width` de `10%` y `height` de `70%`. "bb" significa "edificio de fondo", este será su primer edificio.

# --hints--

Debe crear un nuevo elemento `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Debe asignar al nuevo `div` una clase de `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Debe utilizar un selector de clase `.bb1` para aplicar estilo al elemento.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Debe asignar al elemento `.bb1` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Debe asignar al elemento `.bb1` un `height` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

