---
id: 5d822fd413a79914d39e98cc
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro de `head`, añada una etiqueta`meta` con un `charset` de `UTF-8`, un `title` con el título `City Skyline` y un elemento `link` que vincula su archivo `styles.css`.

# --hints--

Puedes crear una etiqueta `meta` dentro de la etiqueta `head` element.

```js
assert.exists(document.querySelector('head > meta'));
```

Debe asignar a la etiqueta `meta` un `charset` de `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Tu código debe tener una etiqueta `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

La etiqueta `title` debe estar dentro de la etiqueta `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Tu proyecto debe tener estar titulado como `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Recuerda, las mayúsculas, las minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Tu código debe tener una etiqueta `link`.

```js
assert.match(code, /<link/)
```

No debes cambiar las etiquetas existentes `head`. Asegúrate de no cerrar la etiqueta de cierre.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Debes tener una etiqueta `link` de cierre automático.

```js
assert(document.querySelectorAll('link').length === 1);
```

Su etiqueta `link` debe estar dentro de su etiqueta `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Tu etiqueta `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu etiqueta `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
