---
id: 5d822fd413a79914d39e98c9
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Bienvenido al proyecto Skyline de variables de CSS! Para iniciar, añada la declaración `!DOCTYPE html` en la parte superior del documento para que el navegador sepa qué tipo de documento está utilizando.

# --hints--

Su código debería contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debe incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento de la forma `html`.

```js
assert(code.match(/html/gi));
```

Debe cerrar la declaración `DOCTYPE` con un `>` después de escribirla.

```js
assert(code.match(/html\s*>/gi));
```

Tu declaración de `DOCTYPE` debe estar al principio de tu HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
