---
id: 5d822fd413a79914d39e98e7
title: Paso 31
challengeType: 0
dashedName: step-31
---

# --description--

Dar a los seis nuevos elementos estos valores: `width` y `height` `10%` y `60%` a `.fb1`, `10%` y `40%` a `.fb2`, `10%` y `35%` a `.fb3`, `8%` y `45%` a `.fb4`, `10%` y `33%` a `.fb5`, y `9%` y `38%` a `.fb6`.

# --hints--

Debe crear un selector `.fb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

Debe asignar al selector `.fb1` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

Debe asignar al selector `.fb1` un `height` de `60%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

Debe crear un selector `.fb2`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

Debes darle al selector `.fb2` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

Debe asignar al selector `.fb2` un `height` de `40%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

Debe crear un selector `.fb3`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

Debe asignar al selector `.fb3` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

Debe asignar al selector `.fb3` un `height` de `35%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

Debe crear un selector `.fb4`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

Debes darle al selector `.fb4` un `width` de `8%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

Debes darle al selector `.fb4` una `height` de `45%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

Debe crear un selector `.fb5`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

Debes darle al selector `.fb5` un `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

Debe asignar al selector `.fb5` un `height` de `33%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

Debe crear un selector `.fb6`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

Debe asignar al selector `.fb6` un `width` de `9%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

Debes darle al selector `.fb6` una `height` de `38%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--

```

