---
id: 5d822fd413a79914d39e98cb
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

A continuación, añadir etiquetas `head` y `body` dentro del elemento `html`.

# --hints--

Debes tener solo una etiqueta de apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta de cierre `head` tag.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Las etiquetas `head` y `body` deben ser hermanas.

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

La etiqueta `head` debe estar dentro de la etiqueta `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

La etiqueta `body` debe estar dentro de la etiqueta `html`.

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">

</html>
--fcc-editable-region--

```

