---
id: 5d822fd413a79914d39e98d3
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Anida cuatro elementos `div` en el contenedor `.bb1`. Dales las clases `bb1a`, `bb1b`, `bb1c` y `bb1d` en ese orden. Este edificio tendrá cuatro secciones.

# --hints--

Debe crear cuatro nuevos elementos `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Debe asignar a uno de los nuevos elementos `div` una clase de `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Debe asignar a uno de los nuevos elementos `div` una clase de `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Debe asignar a uno de los nuevos elementos `div` una clase de `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Debe asignar a uno de los nuevos elementos `div` una clase de `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Debe colocar los nuevos elementos `div` en el orden correcto.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Debe colocar los nuevos elementos `div` dentro del elemento `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

