---
id: 6153a3952facd25a83fe8083
title: Hatua ya 23
challengeType: 0
dashedName: step-23
---

# --description--

Kipengele cha pseudo cha `::after` huunda kipengele ambacho ni mtoto wa mwisho wa kipengele kilichochaguliwa. Unaweza kuitumia kuongeza kipengee tupu baada ya picha ya mwisho. Ukiipa `width` sawa na picha itasukuma picha ya mwisho kushoto wakati ghala iko katika mpangilio wa safu wima mbili. Kwa sasa, iko katikati kwa sababu umeweka `justify-content: center` kwenye flex container.

Mfano:

```CSS
.container::after {
  content: "";
  width: 860px;
}
```

Unda kichaguzi kipya kwa kutumia kipengele cha pseudo cha `::after` kwenye kipengele cha `.gallery`. Ongeza sifa ya `content` iliyowekwa kuwa na string tupu `""` na `350px` iliyowekwa kwa sifa ya `width`.

# --hints--

Kichaguzi chako cha `.gallery::after` kinapaswa kuwa na string tupu `""` ilioyowekwa kwa sifa ya `content`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.content === "\"\"");
```

Kichaguzi chako cha `.gallery::after` kinapaswa kuwa na `350px` iliyowekwa kwa sifai ya `width`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.width === "350px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
