---
id: 6153893900438b4643590367
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Pangilia maandishi yako ya `.header` katikati. Fanya maandishi kuwa herufi kubwa kwa kutumia kipengele cha `text-transform` chenye `uppercase` kama thamani.

Ipe padding ya `32px` pande zote. Weka mandhari ya nyuma kuwa `#0a0a23` na maandishi kuwa `#fff` kama thamani za rangi.

Ongeza `border-bottom` yenye thamani ya `4px solid #fdb347`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Kichaguzi chako cha `.header` kinafaa kuwa na sifa ya `text-align` yenye `center` kama thamani yake.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Kichaguzi chako cha `.header` kinafaa kuwa na sifa ya `text-transform` yenye `uppercase` kama thamani yake.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Kichaguzi chako cha `.header` kinafaa kuwa na sifa ya `padding` yenye thamani ya `32px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Kichaguzi chako cha `.header` kinafaa kuwa na sifa ya `background-color` yenye `#0a0a23` kama thamani yake.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Kichaguzi chako cha `.header` kinafaa kuwa na sifa ya `color` yenye `#fff` thamani yake.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Kichaguzi chako cha `.header` kinafaa kuwa na sifa ya `border-bottom` yenye `4px solid #fdb347` kama thamani yake.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
