---
id: 61537c5f81f0cf325b4a854c
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Ongeza kipengele cha `header` ndani ya kipengele cha `body` na ukipe class ya `header`.

Ndani ya `header`, unda `h1` yenye `css flexbox photo gallery` kama maandishi.

# --hints--

Unapaswa kuwa na kipengele cha `header` ndani ya kipengele chako cha `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Kipengele chako cha `header` kinafaa kuwa na `class` yenye thamani ya `header`.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Kipengele chako cha `header` kinafaa kuwa na kipengele cha `h1` ndani yake.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `css flexbox photo gallery`.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
