---
id: 61537485c4f2a624f18d7794
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Anza na boilerplate ya kawaida ya HTML. Ongeza tamko la `DOCTYPE`, kipengele cha `html`, kipengele cha `head` na kipengele cha `body`.

Ongeza sifa ya `lang` kwenye lebo ya ufunguzi `<html>` na `en` iliyowekwa kama thamani.

# --hints--

Msimbo wako unapaswa kuwa na rejeleo la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Unapaswa kujumuisha nafasi baada ya rejeleo la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Unapaswa kufafanua `html` kama aina ya hati.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Unapaswa kufunga tamko la `DOCTYPE` kwa `>` baada ya aina.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na tagi ya kufungua.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na tagi ya kufunga.

```js
assert(code.match(/<\/html\s*>/));
```

Tagi yako ya kufungua `<html>` inapaswa kuwa na sifa ya `lang` na `en` kama thamani.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Unapaswa kuwa na tagi ya `head` ya kufungua.

```js
assert(code.match(/<head\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Unapaswa kuwa na tagi ya `body` ya kufungua.

```js
assert(code.match(/<body\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Vipengele vya `head` na `body` vinapaswa kuwa ndugu.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Kipengele cha `head` kinafaa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Kipengele cha `body` kinafaa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
