---
id: 615f171d05def3218035dc85
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

Picha zako ni kubwa mno. Unda kichaguzi cha `.gallery img` ili kuzilenga. Zipe zote `width` ya `100%` na `max-width` ya `350px` ili zipungue inavyohitajika lakini zisiwe kubwa sana.

Pia, weka sifa ya `height` kuwa `300px` ili kuweka picha zako ziwe na ukubwa sawa.

# --hints--

Unapaswa kuwa na kichaguzi cha `.gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

Kichaguzi chako cha `.gallery img` kinafaa kuwa na sifa ya `width` iliyowekwa kuwa `100%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

Kichaguzi chako cha `.gallery img` kinafaa kuwa na sifa ya `max-width` iliyowekwa kuwa `350px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

Kichaguzi chako cha `.gallery img` kinafaa kuwa na sifa ya `height` iliyowekwa kuwa `300px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
