---
id: 62a257659d0d1e2456f24ba2
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Before you start writing your project code, you should move it to its own file to keep things organized. 

Remove your `console.log("Hello World");` line. Then give your now empty `script` element a `src` attribute set to `./script.js`.

# --hints--

You should not have a `console.log("Hello World");` line in your code.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Your `script` element should have a `src` attribute set to `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
