---
id: 61fd5a93fd62bb35968adeab
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

`DOCTYPE` と、このドキュメントが英語だという指定を含む `html` 要素、`head` 要素、そして `body` 要素で HTML を書き始めましょう。

`head` 要素内には、`charset` と `viewport` を適切に設定する `meta` 要素と、適切なタイトルを含む `title` 要素、そしてスタイルシートへの `link` 要素を追加してください。

# --hints--

コードには `<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

コードには `html` 要素が 1 つ必要です。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

コードの `html` 要素内に `head` 要素が必要です。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

コードの `html` 要素内に `body` 要素が必要です。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

`head` 要素は `body` 要素の前に置く必要があります。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

`meta` 要素のひとつは `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もうひとつの `meta` 要素は `charset` 属性を `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` にはテキストが必要です。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

コードには `link` 要素が 1 つ必要です。

```js
assert(/<link/.test(code))
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
