---
id: 61fd778081276b59d59abad6
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

最初の `tr` 内に `th` 要素を追加して、`Cash This is the cash we currently have on hand.` というテキストを設定してください。 そのテキストの `Cash` 以外の文を `span` 要素で囲み、`class` を `description` に設定してください。

th の後に、`td` 要素を 3 つ追加して、`$25`、`$30`、`$28` というテキストを (この順で) 設定してください。 3 つ目の `td` 要素の `class` 属性を `current` に設定してください。

# --hints--

1 つ目の `tr` 要素内に `th` 要素が 1 つ必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th'));
```

`th` 要素のテキストは `Cash This is the cash we currently have on hand.` にする必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th')?.innerText === 'Cash This is the cash we currently have on hand.');
```

`This is the cash we currently have on hand.` というテキストを `span` 要素で囲む必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.textContent === 'This is the cash we currently have on hand.');
```

`span` 要素の `class` 属性を `description` に設定する必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.classList?.contains('description'));
```

`td` 要素が 3 つ必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td').length === 3);
```

1 つ目の `td` 要素には `$25` というテキストが必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[0]?.textContent === '$25');
```

2 つ目の `td` 要素には `$30` というテキストが必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[1]?.textContent === '$30');
```

3 つ目の `td` 要素には `$28` というテキストが必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.textContent === '$28');
```

3 つ目の `td` 要素の `class` を `current` に設定する必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
