---
id: 61fd6cc9475a784b7776233e
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

既存の `div` 要素の下に新しい `div` 要素を追加して、`class` を `table-wrap` に設定してください。 これがテーブル (表) のコンテナーになります。

その中に `table` 要素を 3 つ追加してください。 後ほど CSS を使ってこの 3 つが 1 つのテーブルに見えるようにスタイルを設定しますが、個別のテーブルを使うことで、スクリーンリーダーがドキュメントフローを理解するのに役立ちます。

# --hints--

新しい `div` 要素を作成してください。

```js
assert(document.querySelectorAll('div')?.length === 2);
```

新しい `div` 要素の `class` を `table-wrap` に設定する必要があります。

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

`.table-wrap` 要素は既存の `div` の後に置く必要があります。

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

`.table-wrap` 要素内には `table` 要素が 3 つ必要です。

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
