---
id: 61fd6343fb42013d99bcd7f3
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`body` 要素内で、`main` 要素の中に `section` 要素をネストしてください。

# --hints--

`body` 要素内の最初の要素は `main` 要素である必要があります。

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

`main` 要素の中に `section` 要素が 1 つ必要です。

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
