---
id: 61fd71d596e8f253b9408b39
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`tr` 要素は、テーブルの行を表すために使われます。 `thead` 要素内に `tr` 要素を追加してください。 その `tr` 要素内に、`td` 要素を 1 つと `th` 要素を 3 つ追加してください。

`td` 要素はデータセルを、`th` 要素はヘッダーセルを表します。

# --hints--

`thead` 要素内には `tr` 要素が 1 つ必要です。

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

`tr` 要素の最初の子要素は `td` 要素である必要があります。

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

`tr` 要素内には、`td` 要素の後に 3 つの `th` 要素が必要です。

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
