---
id: 61fd7a160ed17960e971f28b
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

4 つ目の `tr` 要素内に `th` 要素を追加して、テキストを `Total Assets` としてください。 `Assets` というテキストを `span` 要素で囲み、その `class` 属性を `sr-only` に設定してください。

その後に、`td` 要素を 3 つ追加して、`$579`、`$736`、`$809` というテキストを (この順で) 設定してください。 3 つ目の `td` 要素の `class` 属性を `current` に設定してください。

# --hints--

4 つ目の `tr` 要素内に `th` 要素が 1 つ必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th'));
```

`th` 要素のテキストは `Total Assets` にする必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th')?.innerText === 'Total Assets');
```

`Assets` というテキストを `span` 要素で囲む必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th > span')?.textContent === 'Assets');
```

`span` 要素の `class` 属性を `sr-only` に設定する必要があります。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th > span')?.classList?.contains('sr-only'));
```

`td` 要素が 3 つ必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td').length === 3);
```

1 つ目の `td` 要素には `$579` というテキストが必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[0]?.textContent === '$579');
```

2 つ目の `td` 要素には `$736` というテキストが必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[1]?.textContent === '$736');
```

3 つ目の `td` 要素には `$809` というテキストが必要です。

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[2]?.textContent === '$809');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
