---
id: 61fd719788899952e67692b9
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`thead` と `tbody` 要素は、テーブルのどの部分がヘッダーで、どの部分が主要なデータ、つまりコンテンツであるかを表します。

1 つ目の `table` の `caption` 要素の下に、`thead` と `tbody` を追加してください。

# --hints--

1 つ目の `table` 要素内に `thead` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

1 つ目の `table` 要素内に `tbody` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

`thead` 要素は `caption` 要素のすぐ下に置く必要があります。

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

`tbody` 要素は `thead` 要素のすぐ下に置く必要があります。

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
