---
id: 61fd6b7c83dbf54a08cf0498
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

Within your `div` element, add three `span` elements. Give each of them a `class` attribute set to `year`, and add the following text (in order): `2019`, `2020`, and `2021`.

# --hints--

`div` 要素内に `span` 要素が 3 つ必要です。

```js
assert(document.querySelector('div')?.children?.length === 3);
```

各 `span` 要素の `class` 属性を `year` に設定する必要があります。

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

1 つ目の `span` には `2019` というテキストが必要です。

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

2 つ目の `span` には `2020` というテキストが必要です。

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

3 つ目の `span` には `2021` というテキストが必要です。

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
