---
id: 61fd75ea7f663457612dba02
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

先ほど追加した各 `th` 要素内に、`class` を `sr-only year` に設定した `span` 要素をネストしてください。 span 要素には次のテキストを (この順で) 追加してください: `2019`、`2020`、`2021`

また、3 つ目の `th` 要素の `class` 属性を `current` に設定してください。

`td` 要素は空のままにしてください。 この要素は、テーブルを 4 列レイアウトにすることでヘッダーを正しい列と関連付ける目的で追加しています。

# --hints--

各 `th` 要素内に `span` 要素が 1 つ必要です。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

新しく追加した各 `span` 要素の `class` 属性は `sr-only year` に設定する必要があります。

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

1 つ目の `span` 要素は `2019` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

2 つ目の `span` 要素は `2020` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

3 つ目の `span` 要素は `2021` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

3 つ目の `th` 要素の `class` を `current` に設定する必要があります。

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

`td` 要素は空にしてください。

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
