---
id: 61fd8fd08af43372f02952d0
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

各 `th` 要素に `span` 要素を追加して、クラスを `sr-only` に設定し、`2019`、`2020`、`2021` というテキストをこの順で追加してください。

# --hints--

各 `th` 要素内に `span` 要素が 1 つ必要です。

```js
const ths = [...document.querySelectorAll('table')?.[1]?.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

新しく追加した各 `span` 要素の `class` 属性は `sr-only` に設定する必要があります。

```js
const ths = [...document.querySelectorAll('table')?.[1]?.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
});
```

1 つ目の `span` 要素は `2019` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('table')?.[1]?.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

2 つ目の `span` 要素は `2020` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('table')?.[1]?.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

3 つ目の `span` 要素は `2021` というテキストを持つ必要があります。

```js
assert(document.querySelectorAll('table')?.[1]?.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

`td` 要素は空にしてください。

```js
assert(document.querySelectorAll('table')?.[1]?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelectorAll('table')?.[1]?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
            <caption>Liabilities</caption>
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
            <tbody>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
