---
id: 61fd66c687e610436494c6f1
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`section` 要素内に、`h1` 要素とその中にネストされた `span` 要素を追加してください。

# --hints--

`section` 要素には `h1` 要素が 1 つ必要です。

```js
assert(document.querySelector('section')?.children?.[0]?.localName === 'h1');
```

`h1` 要素内に `span` 要素が 1 つ必要です。

```js
assert(document.querySelector('h1')?.children?.[0]?.localName === 'span');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
