---
id: 66faa20000000000000000b5
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

The class is just a blueprint — Phaser needs a `Phaser.Game` instance to actually mount the scene to the page.

Below the class, declare a `config` object with:

- `type: Phaser.AUTO` (lets Phaser pick WebGL or Canvas)
- `width: 480` and `height: 360`
- `parent: "game-container"` (the `<div>` already in the HTML)
- `pixelArt: true` (keeps 16×16 art crisp instead of blurry)
- `physics` configured with `default: "arcade"` and `arcade: { gravity: { y: 0 } }` (top-down games have no gravity)
- `scene: [MainScene]`

Then declare a `game` constant assigned to `new Phaser.Game(config)`.

# --hints--

You should declare a `config` object.

```js
assert.match(code, /const\s+config\s*=\s*\{/);
```

`config` should set `type` to `Phaser.AUTO`.

```js
assert.match(code, /type\s*:\s*Phaser\.AUTO/);
```

`config` should set `width: 480` and `height: 360`.

```js
assert.match(code, /width\s*:\s*480\b/);
assert.match(code, /height\s*:\s*360\b/);
```

`config` should set `parent` to `"game-container"`.

```js
assert.match(code, /parent\s*:\s*(['"`])game-container\1/);
```

`config` should enable `pixelArt: true`.

```js
assert.match(code, /pixelArt\s*:\s*true\b/);
```

`config` should set `physics` to use arcade with zero gravity.

```js
assert.match(code, /physics\s*:\s*\{[\s\S]*?default\s*:\s*(['"`])arcade\1[\s\S]*?arcade\s*:\s*\{\s*gravity\s*:\s*\{\s*y\s*:\s*0\s*\}\s*\}/);
```

`config` should register `MainScene` in its `scene` array.

```js
assert.match(code, /scene\s*:\s*\[\s*MainScene\s*\]/);
```

You should declare a `game` constant assigned to `new Phaser.Game(config)`.

```js
assert.match(code, /const\s+game\s*=\s*new\s+Phaser\.Game\(\s*config\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {

  }

  create() {

  }

  update() {

  }
}

--fcc-editable-region--

--fcc-editable-region--

```
