---
id: 66faa20000000000000000b9
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Now switch to `create`. Phaser has loaded the assets; this is where you assemble the scene.

Inside `create`, draw a dark slate-blue background that fills the play area: call `this.add.rectangle(240, 180, 480, 360, 0x223344)`. The first two arguments are the rectangle's center; the next two are its width and height; the last is the fill color. `0x223344` is a hex literal — Phaser expects colors as numbers, not CSS strings.

# --hints--

You should call `this.add.rectangle` to draw a 480×360 background centered at `(240, 180)` with color `0x223344`.

```js
assert.match(code, /this\.add\.rectangle\(\s*240\s*,\s*180\s*,\s*480\s*,\s*360\s*,\s*0x223344\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "slime",
      "/curriculum-assets/phaser/spritesheets/slime.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {

  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
