---
id: 66faa20000000000000000c8
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

The coin sheet has 8 frames showing a full spin. Below `walk-right`, register the coin animation: `key: "coin-spin"`, frames from `start: 0, end: 7` of `"coin"`, `frameRate: 12`, `repeat: -1`.

# --hints--

You should call `this.anims.create` with the `key` `"coin-spin"`.

```js
assert.match(code, /this\.anims\.create\(\s*\{[\s\S]*?key\s*:\s*(['"`])coin-spin\1/);
```

The `coin-spin` animation should use `start: 0, end: 7` of `"coin"`.

```js
assert.match(code, /key\s*:\s*(['"`])coin-spin\1[\s\S]*?frames\s*:\s*this\.anims\.generateFrameNumbers\(\s*(['"`])coin\2\s*,\s*\{\s*start\s*:\s*0\s*,\s*end\s*:\s*7\s*\}\s*\)/);
```

The `coin-spin` animation should set `frameRate: 12` and `repeat: -1`.

```js
assert.match(code, /key\s*:\s*(['"`])coin-spin\1[\s\S]*?frameRate\s*:\s*12[\s\S]*?repeat\s*:\s*-1/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "slime",
      "/curriculum-assets/phaser/spritesheets/slime.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    this.player.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(false);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-right", true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
