---
id: 66faa20000000000000000b8
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

The slime is the third sprite. Its sheet at `/curriculum-assets/phaser/spritesheets/slime.png` is a 64×16 strip — 4 frames of 16×16 showing the slime squishing in place.

Below the coin load, call `this.load.spritesheet` with the key `"slime"`, the slime path, and the same frame config.

# --hints--

You should call `this.load.spritesheet` with the key `"slime"` and the slime sprite-sheet path.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])slime\1\s*,\s*(['"`])\/curriculum-assets\/phaser\/spritesheets\/slime\.png\2/);
```

The slime sprite sheet should be loaded with `frameWidth: 16` and `frameHeight: 16`.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])slime\1[\s\S]*?frameWidth\s*:\s*16[\s\S]*?frameHeight\s*:\s*16/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {

  }

  update() {

  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
