---
id: 66faa20000000000000000ba
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Time to spawn the hero. Below the background rectangle, create the player by assigning `this.physics.add.sprite(240, 180, "hero", 0)` to `this.player`.

The first two arguments place the sprite at the center of the canvas. `"hero"` is the texture key from `preload`. The fourth argument, `0`, is the starting frame — the still-facing-down pose at index 0 of the sheet.

`this.physics.add.sprite` (rather than `this.add.sprite`) gives the sprite an arcade-physics body, so it can move with `setVelocity` later.

# --hints--

You should assign `this.physics.add.sprite(240, 180, "hero", 0)` to `this.player`.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\(\s*240\s*,\s*180\s*,\s*(['"`])hero\1\s*,\s*0\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "slime",
      "/curriculum-assets/phaser/spritesheets/slime.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.add.rectangle(240, 180, 480, 360, 0x223344);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {

  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
