---
id: 66faa20000000000000000c3
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Detect the left arrow. Below `setVelocity(0)`, add an `if` block:

- when `this.cursors.left.isDown` is true, call `this.player.setVelocityX(-120)`, `this.player.setFlipX(false)`, and `this.player.anims.play("walk-left", true)`.

The `true` second argument to `anims.play` is `ignoreIfPlaying`. Without it, the animation restarts at frame 0 every tick the key is held — and the sprite looks frozen.

# --hints--

Your `update` should branch on `this.cursors.left.isDown`.

```js
assert.match(code, /if\s*\(\s*this\.cursors\.left\.isDown\s*\)/);
```

Pressing the left arrow should call `this.player.setVelocityX(-120)`.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*?this\.player\.setVelocityX\(\s*-120\s*\)/);
```

Pressing the left arrow should call `this.player.setFlipX(false)`.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*?this\.player\.setFlipX\(\s*false\s*\)/);
```

Pressing the left arrow should play the `walk-left` animation with `ignoreIfPlaying`.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*?this\.player\.anims\.play\(\s*(['"`])walk-left\1\s*,\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "slime",
      "/curriculum-assets/phaser/spritesheets/slime.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
  }

  update() {
    this.player.setVelocity(0);
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
