---
id: 66faa20000000000000000d8
title: Step 39
challengeType: 0
dashedName: step-39
---

# --description--

Inside `collectCoin`, after the score update, add the level-cleared check:

```js
if (this.score >= 5) {
  this.cleared = true;
  this.add.text(240, 180, "Level Cleared!", {
    fontSize: "32px",
    fill: "#ffff00"
  }).setOrigin(0.5);
  this.input.keyboard.enabled = false;
}
```

`setOrigin(0.5)` centers the text at `(240, 180)` instead of anchoring its top-left there. Disabling `this.input.keyboard.enabled` freezes the player so the celebration stands still.

# --hints--

Your `collectCoin` should branch on `this.score >= 5`.

```js
assert.match(code, /if\s*\(\s*this\.score\s*>=\s*5\s*\)/);
```

On clear, you should set `this.cleared = true`.

```js
assert.match(code, /this\.score\s*>=\s*5[\s\S]*?this\.cleared\s*=\s*true/);
```

On clear, you should add a centered `"Level Cleared!"` text at `(240, 180)`.

```js
assert.match(code, /this\.add\.text\(\s*240\s*,\s*180\s*,\s*(['"`])Level Cleared!\1[\s\S]*?fontSize\s*:\s*(['"`])32px\2[\s\S]*?fill\s*:\s*(['"`])#ffff00\3[\s\S]*?\)\s*\.setOrigin\(\s*0\.5\s*\)/);
```

On clear, you should disable the keyboard with `this.input.keyboard.enabled = false`.

```js
assert.match(code, /this\.input\.keyboard\.enabled\s*=\s*false/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "slime",
      "/curriculum-assets/phaser/spritesheets/slime.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12,
      repeat: -1
    });
    this.coins = this.physics.add.group();
    [
      [80, 80],
      [400, 80],
      [80, 280],
      [400, 280],
      [240, 80]
    ].forEach(([x, y]) => {
      const coin = this.coins.create(x, y, "coin");
      coin.setScale(2);
      coin.anims.play("coin-spin");
    });
    this.anims.create({
      key: "slime-idle",
      frames: this.anims.generateFrameNumbers("slime", { start: 0, end: 3 }),
      frameRate: 4,
      repeat: -1
    });
    this.slime = this.physics.add.sprite(240, 320, "slime");
    this.slime.setScale(2);
    this.slime.setCollideWorldBounds(true);
    this.slime.setVelocityX(60);
    this.slime.setBounce(1, 0);
    this.slime.anims.play("slime-idle");
    this.physics.add.overlap(this.player, this.coins, this.collectCoin, null, this);
    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Coins: 0/5", {
      fontSize: "16px",
      fill: "#ffffff"
    });
    this.physics.add.overlap(this.player, this.slime, this.hitSlime, null, this);
    this.cleared = false;
  }

  update() {
    this.player.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(false);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-right", true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    } else {
      this.player.anims.play("idle", true);
    }
    if (this.slime.body.velocity.x > 0) {
      this.slime.setFlipX(true);
    } else if (this.slime.body.velocity.x < 0) {
      this.slime.setFlipX(false);
    }
  }

  collectCoin(player, coin) {
    this.tweens.add({
      targets: coin,
      scale: 0,
      duration: 200,
      onComplete: () => coin.destroy()
    });
    this.score += 1;
    this.scoreText.setText("Coins: " + this.score + "/5");
--fcc-editable-region--

--fcc-editable-region--
  }

  hitSlime() {
    this.cameras.main.flash(150, 255, 0, 0);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
