---
id: 66faa20000000000000000b6
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Time to load the hero. The hero sprite sheet lives at `/curriculum-assets/phaser/spritesheets/hero-walk.png`. It's a 64×64 PNG split into 16 frames — four directions of four-frame walk cycles, all 16×16.

Inside `preload`, call `this.load.spritesheet("hero", "/curriculum-assets/phaser/spritesheets/hero-walk.png", { frameWidth: 16, frameHeight: 16 })`.

# --hints--

You should call `this.load.spritesheet` with the key `"hero"` and the hero sprite-sheet path.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])hero\1\s*,\s*(['"`])\/curriculum-assets\/phaser\/spritesheets\/hero-walk\.png\2/);
```

The hero sprite sheet should be loaded with `frameWidth: 16` and `frameHeight: 16`.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])hero\1[\s\S]*?frameWidth\s*:\s*16[\s\S]*?frameHeight\s*:\s*16/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Animated Character Explorer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {

  }

  update() {

  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);

```
