---
id: 6920d189e36b6b60a6330fee
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now, inside the loop, log the `char` variable to the console.

# --hints--

You should have a `console.log` statement with `char` as its argument inside the loop.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);

assert.match(normalizedFunctionCode, /console\.log\(char\)/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
