---
id: 6920c9cda1c7f056f75e4b1a
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Now it’s time to write your first `for...of` loop. In the preceding lessons, you learned that you can use a `for...of` loop to iterate over values from an iterable (e.g., strings or arrays).

Here is an example for a `for...of` loop:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Add a `for...of` loop with an empty code block inside your function. It should iterate over every character in the `str` argument.

# --hints--

Your `printCharacters` function should have a `for...of` loop inside its code block.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Your `for...of` loop should iterate over the `str` argument of your function.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
