---
id: 6921d90b5ce50185505d2004
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Since the `getMatchedWordCount` function needs to return a numerical count, you will start by setting up a counter variable.

Create a variable named `count` using the `let` keyword, initialize it with the value `0`, and add a `return` statement that returns the `count` variable.

# --hints--

You should declare the `count` variable inside the function `getMatchedWordCount`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables?.count);
```

To allow the `count` variable to be incremented later, declare it with the `let` keyword.

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables?.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

Your function should contain an assignment of the value `0` to the `count` variable.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables?.count?.value?.matches('0'));
```

Your function should contain a `return` statement that returns the `count` variable.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

Your `getMatchedWordCount` function should return the value `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
