---
id: 6921e10ae2681205b516fbe6
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Now it's time to add the logic to your loop that increments `count` appropriately.

Inside the loop, use a conditional statement to increment `count` by `1` if the variable `word` equals the variable `match`. Otherwise, leave the value of `count` unchanged.

# --hints--

`getMatchedWordCount([], "bar")` should return `--fcc-expected--`, but it returned `--fcc-actual--`.

```js
assert.equal(getMatchedWordCount([], "bar"), 0);
```

`getMatchedWordCount(["foo"], "bar")` should return `--fcc-expected--`, but it returned `--fcc-actual--`.

```js
assert.equal(getMatchedWordCount(["foo"], "bar"), 0);
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` should return `--fcc-expected--`, but it returned `--fcc-actual--`.

```js
assert.equal(getMatchedWordCount(["foo", "bar", "foo"], "bar"), 1);
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` should return `--fcc-expected--`, but it returned `--fcc-actual--`.

```js
assert.equal(getMatchedWordCount(["foo", "bar", "foo"], "foo"), 2);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
