---
id: 6921d688b8d58581d34459b9
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

In the next few steps, you will build a function that counts how often a string appears in an array of strings.

To start, define an empty function named `getMatchedWordCount` with the parameters `sentence` and `match` in that order.

# --hints--

You should create a `getMatchedWordCount` function.

```js
assert.isFunction(getMatchedWordCount);
```

Your `getMatchedWordCount` function should have two parameters: `sentence` and `match`, in that order.

```js
const explorer = await __helpers.Explorer(code);
assert.equal(explorer.allFunctions.getMatchedWordCount?.parameters[0]?.toString(), 'sentence');
assert.equal(explorer.allFunctions.getMatchedWordCount?.parameters[1]?.toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
