---
id: 6921d688b8d58581d34459b9
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Now you’ll work on the previously mentioned function that counts how often a string appears in an array of strings.

To start, define an empty function named `getMatchedWordCount` with the parameters `sentence` and `match` in that order.

# --hints--

You should create a `getMatchedWordCount` fuction.

```js
assert.isFunction(getMatchedWordCount);
```

Your `getMatchedWordCount` function should have two parameters: `sentence` and `match`, in that order.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
