---
id: 66d75dd0aa65a71600dc669b
title: Build an Inventory Management Program
challengeType: 14
dashedName: build-an-inventory-management-program
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should declare an empty array named `inventory`, that will store product objects having a key `name` with the value of a string, and a key `quantity` with the value of an integer number.
1. You should create a function named `findProductIndex` that takes the product name as its argument and returns the index of the corresponding product object inside the `inventory` array. If the product is not found, the function should return `-1`.
1. You should create a function named `addProduct` that takes a product object as its argument.
1. If the product is already present in the inventory, the `addProduct` function should update its `quantity` value and log to the console the product name followed by a space and `quantity updated`.
1. If the product is not present in the inventory, the `addProduct` function should push the product to the `inventory` array and log to the console the product name followed by a space and `added to inventory`. 
1. You should create a function named `removeProduct` that takes the product name and quantity as its arguments.
1. The `removeProduct` function should subtract the passed quantity from the corresponding product object inside the inventory and log the string `Remaining <product-name> pieces: <product-quantity>` to the console, where `<product-name>` should be replaced by the product name, and `<product-quantity>` should be replaced by the product quantity.
1. If the quantity after the subtraction is zero, `removeProduct` should remove the product object from the inventory. Instead, if the quantity in the inventory is not enough to perform the subtraction, the `removeProduct` function should log the string `Not enough <product-name> available, remaining pieces: <product-quantity>` to the console.
1. If the product to remove is not present in the inventory, the `removeProduct` function should log `<product-name> not found` to the console. 

# --hints--

You should declare an empty array named `inventory`.

```js
assert.match(code, /(let|const)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

You should create a function named `findProductIndex`.

```js
assert.isFunction(findProductIndex)
```

Your `findProductIndex` function should take the product name as its argument and should return the index of the corresponding product object inside the `inventory` array.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("sugar"), 2)
```

Your `findProductIndex` function should return `-1` when the product object is not found inside the `inventory` array.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

You should create a function named `addProduct`.

```js
assert.isFunction(addProduct)
```

When the product passed to `addProduct` is present in the `inventory` array, the function should update its `quantity` value.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flour', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}]);
assert.equal(JSON.stringify(inventory), expected)
```

When the product passed to `addProduct` is present in the `inventory` array, the function should log to the console the product name followed by a space and `quantity updated`.

```js
assert(false);
```

When the product passed to `addProduct` is not present in the inventory, the `addProduct` function should push the product to the `inventory` array.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'salt', quantity: 8});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}]);
assert.equal(JSON.stringify(inventory), expected)
```

When the product passed to `addProduct` is not present in the inventory, the `addProduct` function should log to the console the product name followed by a space and `added to inventory`.

```js
assert(false);
```

You should create a function named `removeProduct`.

```js
assert.isFunction(removeProduct);
```

when the product to remove is not present in the inventory, the `removeProduct` function should log `<product-name> not found` to the console.

```js
assert(false);
```



# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
