---
id: 61969d66cfcdba137d021558
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

لمنع الجبل من دفع عنصر `.ground`، أضبط `position` لمنعه من أخذ مساحة في تصميم الصفحة.

# --hints--

يجب أن تعطي `.left-mountain` خاصية `position` بقيمة `absolute`. توقّعت `--fcc-actual--` أن تكون بقيمة `--fcc-expected--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.position, 'absolute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
