---
id: 61967e74a8e3690ab6292daa
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أستهدف عنصر `body` وحدد قيمة خاصية `background` ألى التدريج الخطي (linear gradient) بزاوية 45 دراجة (degrees) مع اتجاه عقارب الساعة (clockwise), بَدْء من `rgb(118, 201, 255)` و إنهاء من `rgb(247, 255, 222)`.

# --hints--

يجب عليك استخدام منتقي العنصر `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب عليك إضافة `background` داخل منتقي العنصر `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

يجب عليك تعيين `background` إلى`linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
