---
id: 61969e7451455614217e901b
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

لجعل عناصر الجبال تبدو أشبه بالجبل حقاٌ، يمكنك استخدام خاصية التحويل `skew`، التي تأخذ حجيتين (arguments). الأولى هي زاوية لجزّ (shear) المحور الأفقي (x-axis)، والثانية هي زاوية لجزّ (shear) المحور الرأسي (y-axis).

استخدم خاصية `transform` لجعل الجبل منحرف بـ `0deg` في المحور الأفقي (x-axis)، و `44deg` في المحور الرأسي (y-axis).

# --hints--

يجب عليك إعطاء عنصر `.left-mountain` خاصية الـ `transform`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform);
```

قم باستخدام `skew` على `transform`.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.transform, 'skew');
```

يجب عليك إعطاء `.left-mountain` خاصية `transform` بقيمة `skew(0deg, 44deg)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.getPropVal('transform', true), 'skew(0deg,44deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
