---
id: 619c155df0063a0a3fec0e32
title: الخطوة 45
challengeType: 0
dashedName: step-45
---

# --description--

اعطي عنصر `.face` خاصية `width` بقيمة `60%`، وخاصية`height` بقيمة `70%`، وخاصية `background-color` بقيمة `white`.

# --hints--

يجب عليك استخدام منتقي `.face`.

```js
assert.match(code, /\.face\s*\{/);
```

يجب عليك إعطاء عنصر `.face` خاصية `width`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.face')?.width);
```

يجب أن تعطي عنصر `.face` خاصية `width` بقيمة `--fcc-expected--`, ولكن وجدت بقيمة `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.face')?.width, '60%');
```

يجب عليك إعطاء عنصر `.face` خاصية `height`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.face')?.height);
```

يجب أن تعطي عنصر `.face` خاصية `height` بقيمة `--fcc-expected--`, ولكن وجدت بقيمة `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.face')?.height, '70%');
```

يجب عليك إعطاء عنصر `.face` خاصية `background-color`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.face')?.backgroundColor);
```

يجب أن تعطي عنصر `.face` خاصية `background-color` بقيمة `--fcc-expected--`, ولكن وجدت بقيمة `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.face')?.backgroundColor, 'white');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin">
      <div class="penguin-head">
        <div class="face"></div>
        <div class="face"></div>
      </div>
      <div class="penguin-body"></div>
    </div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
  z-index: 4;
  position: relative;
}

.penguin * {
  position: absolute;
}

.penguin-head {
  width: 50%;
  height: 45%;
  background: linear-gradient(
    45deg,
    gray,
    rgb(239, 240, 228)
  );
  border-radius: 70% 70% 65% 65%;
  top: 10%;
  left: 25%;
  z-index: 1;
}

--fcc-editable-region--

--fcc-editable-region--

.penguin-body {
  width: 53%;
  height: 45%;
  background: linear-gradient(
    45deg,
    rgb(134, 133, 133) 0%,
    rgb(234, 231, 231) 25%,
    white 67%
  );
  border-radius: 80% 80% 100% 100%;
  top: 40%;
  left: 23.5%;
}

.penguin-body::before {
  content: "";
  position: absolute;
  width: 50%;
  height: 45%;
  background-color: gray;
  top: 10%;
  left: 25%;
  border-radius: 0% 0% 100% 100%;
  opacity: 70%;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
