---
id: 594810f028c0303b75339acd
title: 'Збиткові, досконалі та надлишкові числа'
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

Натуральні числа можна класифікувати на три категорії на основі їхніх властивих дільників.

Нехай $P(n)$ буде сумою властивих дільників числа `n`, де властивими дільниками є всі натуральні числа `n`, окрім самого числа `n`.

Якщо `P(n) < n`, то `n` класифікується як `deficient` (збиткове).

Якщо `P(n) === n`, то `n` класифікується як `perfect` (досконале).

Якщо `P(n) > n`, то `n` класифікується як `abundant` (надлишкове).

**Наприклад**: властивими дільниками числа `6` є `1`, `2` та `3`. `1 + 2 + 3 = 6`, тому `6` класифікується як досконале число.

# --instructions--

Реалізуйте функцію, яка розраховує кількість цілих чисел від `1` до `num` (включно) в кожній з трьох категорій. Виведіть результат у вигляді масиву в наступному форматі: `[збиткове, досконале, надлишкове]`.

# --hints--

`getDPA` має бути функцією.

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)` має повернути масив.

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` має повернути масив довжиною `3`.

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)` має повернути `[3758, 3, 1239]`.

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)` має повернути `[7508, 4, 2488]`.

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)` має повернути `[15043, 4, 4953]`.

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
