---
id: 599d1566a02b571412643b84
title: Ефіопське множення
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

Ефіопське множення — це метод множення цілих чисел з використанням додавання, подвоєння та ділення на два.

**Метод:**

<ol>
  <li>Візьміть два числа, які потрібно перемножити, і напишіть їх зверху двох стовпчиків</li>
  <li>У стовпчику зліва діліть останнє число навпіл декілька разів, відкидаючи будь-які остачі, і записуйте результат нижче в цьому ж стовпчику, поки не отримаєте <code>1</code></li>
  <li>У стовпчику справа множте останнє число на два та записуйте результат знизу, поки не отримаєте результат в тому ж рядку, що й <code>1</code> зліва</li>
  <li>Перегляньте отриману таблицю і відкиньте будь-який рядок справа, де значення зліва є парним</li>
  <li>Додайте значення справа, щоб отримати добуток двох початкових чисел</li>
</ol>

**Наприклад:** `17 × 34`

<pre>17   34
</pre>

Ділимо на два перший стовпчик:

<pre>17   34
8
4
2
1
</pre>

Множимо на два другий стовпчик:

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

Закреслюємо рядки, в яких перша комірка парна:

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

Підсумовуємо решту чисел в стовпчику справа:

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

Отже, якщо помножити `17` на `34` за ефіопським методом, то отримаємо `578`.

# --instructions--

Вам потрібно визначити три функції/методи/процедури/підпрограми:

<ol>
  <li>для ділення цілого числа на два</li>
  <li>для множення цілого числа на два</li>
  <li>для визначення, чи ціле число є парним</li>
</ol>

Використайте ці функції, щоб створити функцію, яка виконує ефіопське множення.

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` має бути функцією.

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` має повернути `578`.

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` має повернути `1058`.

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` має повернути `324`.

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` має повернути `5488`.

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` має повернути `4662`.

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
