---
id: 5a23c84252665b21eecc8041
title: Сума множини
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

Обчисліть **n**<sup>-ий</sup> член <em>множини</em>, тобто суму перших **n** членів відповідної <em>послідовності</em>. Неформально таке значення (або його обмеження, коли **n** прагне до нескінченності) називають *сумою множин*, звідси й назва завдання. Для цього завдання використайте: $S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$.

# --instructions--

Напишіть функцію, яка приймає $a$ та $b$ як параметри і повертає суму членів послідовності від $a$ до $b$.

# --hints--

`sum` має бути функцією.

```js
assert(typeof sum == 'function');
```

`sum(1, 100)` має повернути число.

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)` має повернути `1.6349839001848923`.

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)` має повернути `0.009262256361481223`.

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)` має повернути `0.044086990748706555`.

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)` має повернути `0.08619778593108679`.

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)` має повернути `1.5497677311665408`.

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
