---
id: bad87fee1348bd9aedf08829
title: 创建一个输入框
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c2EVnf6'
forumTopicId: 16823
---

# --description--

现在让我们来创建一个`form`表单。

`input`输入框可以让你轻松获得用户的输入。

你可以像这样创建一个文本输入框：

`<input type="text">`

注意：`input`输入框是没有结束标记的。

# --instructions--

在列表下面创建一个`type`属性为`text`的`input`输入框。

# --hints--

网页中有一个`type`属性为`text`的`input`输入框。

```js
assert($('input[type=text]').length > 0);
```

# --solutions--

