---
id: bad87fee1348bd9aedf08817
title: '用 # 号来创建链接占位符'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
---

# --description--

有时你想为网站添加一个 `a` 元素，但还不确定要将它链接到哪里。这时，你可以使用链接占位符。

链接占位符还有一个应用场景，就是实现在点击 `a` 元素时执行 `JavaScript` 代码。我们会在后面的课程中学到。

# --instructions--

目前，`href` 的属性值是 "`https://freecatphotoapp.com`"。请将 `href` 属性的值替换为`#`，以此来创建链接占位符。

例如: `href="#"`

# --hints--

`a` 的 `href` 属性值应为 "#"。

```js
assert($('a').attr('href') === '#');
```

# --solutions--

