---
id: 65661b72d6745ebec6a96923
title: ステップ 17
challengeType: 20
dashedName: step-17
---

# --description--

In the same way, modify the remaining two lists considering that the `C-D` distance is `7`.

# --hints--

`my_graph["C"]` should be a list of tuples.

```js
({ test: () => assert(runPython(`
    type(my_graph["C"]) is list and all(type(i) is tuple for i in my_graph["C"])
  `))
})
```

`my_graph["C"]` should be a list of tuples where the first item in the tuple is the connected node and the second item is the distance.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph["D"]` should be a list of tuples.

```js
({ test: () => assert(runPython(`
    type(my_graph["D"]) is list and all(type(i) is tuple for i in my_graph["D"])
  `))
})
```

`my_graph["D"]` should be a list of tuples where the first item in the tuple is the connected node and the second item is the distance.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
