---
id: 5f331e55dfab7a896e53c3a1
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` є одним з елементів, які надають додаткову інформацію, яку не видно на вебсторінці, але є важливою для пошукових систем або для того, як відображається сторінка.

Всередині елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `utf-8`, щоб повідомити браузеру як кодувати символи сторінки. Зауважте, що елементи `meta` є самозакривними.

# --hints--

Ви повинні мати теґ `meta`.

```js
assert(code.match(/<meta.*?\/?>/is));
```

Ваш теґ `meta` повинен мати атрибут `charset`.

```js
assert(code.match(/<meta charset=/i));
```

Ваш атрибут `charset` повинен мати значення `utf-8`.

```js
assert(code.match(/charset=('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
