---
id: 5f332b23c2045fb843337579
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Оскільки доданий в минулому кроці елемент `p` надає додаткову інформацію про кафе, вкладіть елементи `h1` та `p` в елемент `header`.

# --hints--

Ви повинні мати початковий теґ `<header>`.

```js
assert(code.match(/<header>/i));
```

Ви повинні мати кінцевий теґ `</header>`.

```js
assert(code.match(/<\/header>/i));
```

Ваш елемент `h1` повинен бути вкладеним в елементі `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

Ваш елемент `p` повинен бути вкладеним в елементі `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
</html>
```

