---
id: 5f332b23c2045fb843337579
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Щоб відвідувачі знали, що кафе було засновано в 2020, додайте елемент `p` під елементом `h1` з текстом `Est. 2020`.

# --hints--

Ви повинні мати початковий теґ `<p>`.

```js
assert(code.match(/<p>/i));
```

Ви повинні мати кінцевий теґ `</p>`.

```js
assert(code.match(/<\/p>/i));
```

Ви не повинні змінювати наявний елемент `h1`. Переконайтеся, що не видалили кінцевий теґ.

```js
assert($('h1').length === 1);
```

Ваш елемент `p` повинен бути під елементом `h1`.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

Ваш елемент `p` повинен мати текст `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

