---
id: 5f33071498eb2472b87ddee4
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Як ви дізналися в останніх кількох кроках застосунку «Cat Photo», існує базова структура, яка необхідна для початку створення вебсторінки.

Додайте теґ `<!DOCTYPE html>` та елемент `html` з атрибутом `lang` зі значенням `en`.

# --hints--

Ви повинні мати декларацію `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ви повинні мати початковий теґ `<html>` з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ви повинні мати кінцевий теґ `<html>`. Пам'ятайте, що кінцеві теґи мають `/`, що супроводжує початкову дужку `<`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
