---
id: 5f33294a6af5e9188dbdb8f3
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Назва кафе – `CAMPER CAFE`. Додайте елемент `h1` в межах елемента `body`. Напишіть назву кафе великими літерами, щоб вона виділялась.

# --hints--

Ви повинні мати кінцевий теґ `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Ви повинні мати кінцевий теґ `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Ви не повинні змінювати наявний елемент `body`.

```js
assert($('body').length === 1);
```

Ваш елемент `h1` повинен бути вкладеним в елементі `body`.

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

Ваш елемент `h1` повинен мати текст `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
