---
id: 5f3313e74582ad9d063e3a38
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте елемент `head` в межах елемента `html`, щоб можна було додати елемент `title`. Текст елемента `title` повинен бути `Cafe Menu`.

# --hints--

Ви повинні мати початковий теґ `<head>`.

```js
assert(code.match(/<head>/i));
```

Ви повинні мати кінцевий теґ `</head>`.

```js
assert(code.match(/<head>/i));
```

Ви повинні мати початковий теґ `<title>`.

```js
assert(code.match(/<title>/i));
```

Ви повинні мати кінцевий теґ `</title>`.

```js
assert(code.match(/<\/title>/i));
```

Ваш елемент `title` повинен бути вкладеним в елементі `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Ваш елемент `title` повинен мати текст `Cafe Menu`. Можливо, доведеться перевірити написання.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
