---
id: 5f33310c1851c6c4da013250
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Час додати трішки вмісту меню. Додайте елемент `main` під наявним елементом `header`. В кінцевому підсумку там буде ціна напоїв та десертів, які пропонує кафе.

# --hints--

Ваш код повинен мати початковий теґ `<main>`.

```js
assert(code.match(/<main>/i));
```

Ваш код повинен мати кінцевий теґ `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Ви не повинні змінювати наявний елемент `header`. Переконайтеся, що випадково не видалили кінцевий теґ.

```js
assert($('header').length === 1);
```

Ваш теґ `main` повинен йти після теґа `header`.

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
