---
id: 5f33294a6af5e9188dbdb8f3
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Час додати деякий вміст меню. Додайте елемент `main` в межі наявного елемента `body`. В кінцевому підсумку там буде ціна напоїв та десертів, які пропонує кафе.

# --hints--

Ваш код повинен мати початковий теґ `<main>`.

```js
assert(code.match(/<main>/i));
```

Ваш код повинен мати кінцевий теґ `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Ви не повинні змінювати наявний елемент `body`. Переконайтеся, що випадково не видалили кінцевий теґ.

```js
assert($('body').length === 1);
```

Ваш теґ `main` повинен бути в межах теґу `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
