---
id: 5f332a88dc25a0fd25c7687a
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Назва кафе – `CAMPER CAFE`. Додайте елемент `h1` в межах елемента `main`. Напишіть назву кафе великими літерами, щоб вона виділялась.

# --hints--

Ви повинні мати початковий теґ `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Ви повинні мати кінцевий теґ `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Ви не повинні змінювати наявний елемент `main`.

```js
assert($('main').length === 1);
```

Ваш елемент `h1` повинен бути вкладеним в елементі `main`.

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

Ваш елемент `h1` повинен мати текст `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

