---
id: 5f344fbc22624a2976425065
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Створіть елемент `h2` в елементі `section` та надайте йому текст `Coffee`.

# --hints--

Ви повинні мати початковий теґ `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

Ви повинні мати кінцевий теґ `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

Ви не повинні змінювати наявний елемент `section`. Переконайтеся, що не видалили кінцевий теґ.

```js
assert($('section').length === 1);
```

Ваш елемент `h2` повинен бути в межах елемента `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Ваш елемент `h2` повинен мати текст `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
