---
id: 5f3477cbcb6ba47918c1da92
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Щоб стилізація сторінки була схожою на мобільному пристрої та комп'ютері, потрібно додати елемент `meta` зі спеціальним атрибутом `content`.

Додайте наступне в межах елемента `head`:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Ваш код повинен мати два елементи `meta`.

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

Ваш елемент `meta` повинен мати атрибут `name` зі значенням `viewport`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

Ваш елемент `meta` повинен мати атрибут `content` зі значенням `width=device-width, initial-scale=1.0`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
