---
id: 5f3477ae9675db8bb7655b30
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

В попередніх кроках ви використовували <dfn>селектор типу</dfn> для стилізації елемента `h1`. Відцентруйте елементи `h2` та `p`, додавши новий селектор типу до кожного з наявних елементів `style`.

# --hints--

Ви не повинні змінювати наявний селектор `h1`.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Ви не повинні додавати новий теґ `style`. Додайте нові правила CSS до наявного тегу `style`.

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

Ви повинні додати новий селектор `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Ви повинні додати новий селектор `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Ваш елемент `h1` повинен мати `text-align` зі значенням `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Ваш елемент `h2` повинен мати `text-align` зі значенням `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Ваш елемент `p` повинен мати `text-align` зі значенням `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

