---
id: 5f3477ae8466a9a3d2cc953c
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Оскільки тепер ви маєте CSS в файлі `styles.css`, вилучіть елемент `style` та весь його вміст. Текст, що був зцентрований, зміститься ліворуч одразу після вилучення.

# --hints--

Ви не повинні мати жодних теґів `style` в своєму коді.

```js
assert(!code.match(/style/i));
```

Ви не повинні мати селекторів CSS в файлі HTML.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

