---
id: 5f344fad8bf01691e71a30eb
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

До цього моменту ви мали обмеження на презентацію та вигляд створеного вмісту. Щоб взяти все під контроль, додайте елемент `style` в межах елемента `head`.

# --hints--

Ваш код повинен мати початковий теґ `<style>`.

```js
assert(code.match(/<style\s*>/i));
```

Ваш код повинен мати кінцевий теґ `</style>`.

```js
assert(code.match(/<\/style\s*>/));
```

Ваш елемент `style` повинен бути вкладеним в елементі `head`.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
