---
id: 5f3477ae34c1239cafe128be
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Тепер у вас є три селектори типу з такою ж стилізацією. Ви можете додати таку ж групу стилів до багатьох елементів, створюючи список селекторів. Кожен селектор розділений комами, ось так:

```css
selector1, selector2 {
  property: value;
}
```

Видаліть три наявні селектори типу та замініть їх одним списком селектора, що центрує текст елементів `h1`, `h2` та `p`.

# --hints--

Ви повинні використати один селектор типу для всіх трьох елементів `h1, h2, p`. Переконайтеся, що саме такий порядок.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Ви повинні мати лише один селектор в елементі `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

