---
id: 5f3477cb2e27333b1ab2b955
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Now you need to link the `styles.css` file, so the styles will be applied again. Inside the `head` element, add a `link` element. Give it a `rel` attribute with the value of `"stylesheet"` and a `href` attribute with the value of `"styles.css"`.

# --hints--

Ваш код повинен мати елемент `link`.

```js
const link = document.querySelector('link');
assert(link);
```

Ви не повинні змінювати наявний елемент `head`. Переконайтеся, що не видалили кінцевий теґ.

```js
assert($('head').length === 1);
```

Ви повинні мати один самозакривний елемент `link`.

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
const link = document.querySelector('head > link');
assert(link);
```

Ваш елемент `link` повинен містити атрибут `rel` зі значенням `stylesheet`.

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

Ваш елемент `link` повинен містити атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
