---
id: 5f344fc1520b6719f2e35605
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

В меню буде дві секції: для кави та десертів. Додайте елемент `section` в межах елемента `main`, щоб мати місце для розміщення всіх доступних напоїв.

# --hints--

Ви повинні мати початковий теґ `<section>`.

```js
assert(code.match(/<section\s*>/i));
```

Ви повинні мати кінцевий теґ `</section>`.

```js
assert(code.match(/<\/section\s*>/i));
```

Ви не повинні змінювати наявний елемент `main`. Переконайтеся, що не видалили кінцевий теґ.

```js
assert($('main').length === 1);
```

Ваш елемент `section` повинен бути в межах елемента `main`.

```js
const main = document.querySelectorAll('main')?.[0];
assert(main.children[0].tagName === 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
