---
id: 5f356ed60a5decd94ab66986
title: Крок 22
challengeType: 0
removeComments: false
dashedName: step-22
---

# --description--

Коментарі в CSS виглядають так:

```css
/* comment here */
```

В аркуші стилів прокоментуйте лінію, яка містить властивість `background-color` зі значенням, щоб був помітний ефект від стилізації лише елемента `div`. Це знову зробить фон білим.

# --hints--

Ви повинні прокоментувати лінію `background-color: burlywood;` в CSS.

```js
assert(code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i));
```


Ваш `body` повинен мати білий фон.

```js
const bodyCSS = $('body').css('background-color');
assert(bodyCSS === "rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

div {
  width: 300px;
}
```

