---
id: 5f356ed63c7807a4f1e6d054
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Тепер метою є зробити так, щоб `div` не займав всю ширину сторінки. CSS властивість `width` чудово підходить для цього. Створіть новий селектор типу в таблиці стилів, що надає елементу `div` ширину зі значенням `300px`.

# --hints--

Ви повинні мати селектор типу `div`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('div');
assert(hasDiv);
```

Ви повинні встановити властивість `width` на `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

Ваш `div` повинен мати ширину зі значенням 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle('div')?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

