---
id: 5f35e5c4321f818cdc4bed30
title: Крок 29
challengeType: 0
dashedName: step-29
---

# --description--

Має хороший вигляд. Час додати деякі елементи меню. Додайте порожній елемент `article` під заголовком `Coffee`. Він міститиме смак та ціну кожного напою, який ви пропонуєте.

# --hints--

Ви повинні мати початковий теґ `<article>`.

```js
assert(code.match(/<article>/i));
```

Ви повинні мати кінцевий теґ `</article>`.

```js
assert(code.match(/<\/article>/i));
```

Ви не повинні змінювати наявний елемент `h2`.

```js
assert($('h2').length === 1);
```

Ваш елемент `article` повинен йти після елемента `h2`.

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

