---
id: 5f356ed6199b0cdef1d2be8f
title: Крок 26
challengeType: 0
dashedName: step-26
---

# --description--

До цього часу для стилізації елементів ви використовували лише селектори типу. <dfn>Селектор класу</dfn> визначається за назвою з крапкою одразу перед нею, наприклад:

```css
.class-name {
  styles
}
```

Замініть наявний селектор `div` на селектор класу, змінюючи `div` на клас `menu`.

# --hints--

Ви повинні мати селектор класу `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Ви не повинні мати селектор `div`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('div');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
div {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

