---
id: 5f356ed60785e1f3e9850b6e
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Тепер легко побачити, що текст відцентрований всередині елемента `#menu`. Наразі ширина елемента `#menu` вказана в пікселях (`px`).

Change the `width` property's value to be `80%`, to make it 80% the width of its parent element (`body`).

# --hints--

Ви повинні встановити властивість `width` на `80%`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert(hasWidth);
```

Ви не повинні мати властивість `width` зі значенням `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(!hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--
#menu {
  width: 300px;
  background-color: burlywood;
}
--fcc-editable-region--
```

