---
id: 5f356ed69db0a491745e2bb6
title: Крок 27
challengeType: 0
dashedName: step-27
---

# --description--

Щоб застосували стилізацію класу до елемента `div`, додайте атрибут `class` до початкового теґа елемента `div` та встановіть значення `menu`.

# --hints--

Ваш `div` досі повинен відтворювати. Переконайтеся, що ви не спотворили теґ `<div>`.

```js
assert($('div').length === 1);
```

Ваш елемент `div` повинен мати клас `menu`.

```js
assert($('div').attr('class').includes('menu'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div>
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

