---
id: 5f356ed656a336993abd9f7c
title: Крок 25
challengeType: 0
dashedName: step-25
---

# --description--

Потім потрібно зцентрувати `div` горизонтально. Це можна зробити, встановивши властивості `margin-left` та `margin-right` на `auto`. Думайте про поле як невидимий простір навколо елементу. Використовуючи ці дві властивості, зцентруйте елемент `div` в межах елемента `body`.

# --hints--

Ви повинні встановити властивість `margin-left` на `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert(hasMargin);
```

Ви повинні встановити властивість `margin-right` на `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert(hasMargin);
```

Ви повинні встановити властивості `margin-left` та `margin-right` вашого `div` на `auto`.

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('div')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('div')?.getPropertyValue('margin-left');
assert(divMarginRight === 'auto');
assert(divMarginLeft === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
div {
  width: 80%;
  background-color: burlywood;
}
--fcc-editable-region--
```
