---
id: 5f3c866de7a5b784048f94b1
title: Крок 36
challengeType: 0
dashedName: step-36
---

# --description--

Схоже до того, що ви хочете, але було б добре, якби смак та ціна були на одній лінії. Елементи `p` є елементами <dfn>на рівні блоків</dfn>, отож вони займають всю ширину батьківського елемента.

Щоб вони були на одній лінії, вам потрібно застосувати певну стилізацію для елементів `p`, щоб вони поводились як <dfn>вбудовані</dfn> елементи. Додайте атрибут `class` зі значенням `item` до першого елемента `article` під заголовком `Coffee`.

# --hints--

Ви повинні застосувати клас `item` до елемента `article`.

```js
assert(code.match(/<article\s*class=('|")item\1\s*>/i))
```

Ви повинні мати лише один елемент класу `item`.

```js
assert($('.item').length === 1);
```

Ваш перший елемент `article` повинен мати клас `item`.

```js
assert($('article')[0].className === 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
