---
id: 5f3c866d28d7ad0de6470505
title: Крок 32
challengeType: 0
dashedName: step-32
---

# --description--

Цієї миті смаки та ціни знаходяться один над одним та розміщені відповідно до елементів `p`. Було б чудово, якби смак був зліва, а ціна справа.

Додайте назву класу `flavor` до елемента `p` `French Vanilla`.

# --hints--

Ви повинні додати клас `flavor` до свого елемента `p`.

```js
assert(code.match(/<p\s*class=('|")flavor\1\s*>/i));
```

Ви повинні мати лише один елемент з класом `flavor`.

```js
assert($('.flavor').length === 1);
```

Ваш клас `flavor` повинен бути на елементі `p` з текстом `French Vanilla`.

```js
assert($('.flavor')[0].innerText.match(/French Vanilla/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            <p>French Vanilla</p>
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
