---
id: 5f716ad029ee4053c7027a7a
title: Крок 49
challengeType: 0
dashedName: step-49
---

# --description--

Вкладіть два елементи `p` всередині елемента `article`. Текст першого повинен бути `Donut`, текст другого – `1.50`. Розмістіть їх в одній лінії та переконайтесь, що між ними немає місця.

# --hints--

Ви не повинні змінювати наявний елемент `article`.

```js
assert($('article').length === 6);
```

Ваш новий елемент `article` повинен мати два елементи `p`.

```js
assert($('article').last().children('p').length === 2);
```

Ваш перший елемент `p` повинен мати текст `Donut`.

```js
assert($('article').last().children('p')[0].innerText.match(/Donut/i));
```

Ваш другий елемент `p` повинен мати текст `1.50`.

```js
assert($('article').last().children('p')[1].innerText.match(/1\.50/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

