---
id: 694bf6ea530e19a9c48c59f0
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

If you concatenate two strings like `'John' + 'Doe'`, the result is `'JohnDoe'` with no space. To fix this, you need to concatenate a string containing a space (`' '`) between them.


Update your `full_name` variable so it concatenates `first_name`, a space, and `last_name`.

# --hints--

You should add a space (`' '`) between `first_name` and `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
