---
id: 694bf6eb530e19a9c48c59f2
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

You cannot concatenate a string and a number directly in Python; if you try, you will get a `TypeError`.
To fix this, you must convert the number to a string first using the `str()` function. For example: `str(28)` becomes `"28"`.

# --instructions--

Create a variable `employee_info`. Assign it the value of `full_name` concatenated with `" is "`, `employee_age` (converted to a string), and `" years old"`. Print the result.

# --hints--

You should convert `employee_age` to a string using `str(employee_age)`.

```js
assert(locals.get('employee_info') === "John Doe is 28 years old");
```

# --seed--

## --seed-contents--

```python

full_name = "John Doe"
employee_age = 28

# --fcc-editable-region--
# Create employee_info below

# --fcc-editable-region--
```
