---
id: 694bf6eb530e19a9c48c59f6
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

You can slice from any part of the string, not just the beginning. And it is helpful in many cases.

Create a variable `year_code` and assign it the slice of `employee_code` from index `4` to `8`. This will extract `2026`.

Then create a variable `initials` and assign it the slice of `employee_code` from index `9` to `11`. This will extract `JD`.

Finally, print both variables to the terminal.

# --hints--

You should create a variable `year_code` that slices `employee_code` from index `4` to `8`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

You should create a variable `initials` that slices `employee_code` from index `9` to `11`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

You should print `year_code` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

You should print `initials` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_info = 'Experience: ' + str(5) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
