---
id: 694bf6eb530e19a9c48c59f4
title: Introduction to f-strings
challengeType: 20
dashedName: step-6
---

# --description--

Concatenating many strings using `+` and converting numbers using `str()` can get messy and hard to read.
Python 3.6 introduced **f-strings** to solve this. By adding the letter `f` before the opening quote, you can put variables directly inside curly braces `{}`.

Example:

```python
name = "John"
print(f"Hello {name}") # Output: Hello John
```


We have added `position` and `salary` variables for you.

Create a variable `employee_card`. Use an f-string to format it exactly like this:
`"Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}"`

Then print `employee_card`.

# --hints--

You should create `employee_card` with the correct formatted string.

```js
const test = await runPython(`
    employee_card == "Employee: John Doe | Age: 28 | Position: Data Analyst | Salary: $75000"
`);
assert(test);
```

You should print `employee_card`.

```js
const test = await runPython(`
    _Node(_code).has_call("print(employee_card)")
`);
assert(test);
```

# --seed--

## --seed-contents--

```python
first_name = "John"
last_name = "Doe"
full_name = first_name + " " + last_name
address = "123 Main Street"
address += ", Apartment 4B"
employee_age = 28
employee_info = full_name + " is " + str(employee_age) + " years old"
print(employee_info)
years_experience = 5
experience_info = "Experience: " + str(years_experience) + " years"
print(experience_info)

position = "Data Analyst"
salary = 75000

--fcc-editable-region--
# Create employee_card below

--fcc-editable-region--
```
