---
id: 694bf6eb530e19a9c48c59f5
title: String Slicing
challengeType: 20
dashedName: step-7
---

# --description--

You often need to extract a specific portion of a string. In Python, this is called **slicing**.


The syntax is `string[start:stop]`.

* `start`: The index where the slice begins (inclusive).
* `stop`: The index where the slice ends (**exclusive**).

For example, if `text = "Python"`, then `text[0:2]` gives `"Py"`.


We have defined an `employee_code` for you.
Create a variable `department` and assign it the slice of `employee_code` from index `0` to `3`. Then print `department`.

# --hints--

# --hints--

You should create a variable `department` containing the first 3 characters of `employee_code`.

```js
const test = await runPython(`
    department == "DEV"
`);
assert(test);
```

You should print `department`.

```js
const test = await runPython(`
    _Node(_code).has_call("print(department)")
`);
assert(test);
```

# --seed--

## --seed-contents--

```python
first_name = "John"
last_name = "Doe"
full_name = first_name + " " + last_name
address = "123 Main Street"
address += ", Apartment 4B"
employee_age = 28
employee_info = full_name + " is " + str(employee_age) + " years old"
print(employee_info)
years_experience = 5
experience_info = "Experience: " + str(years_experience) + " years"
print(experience_info)
position = "Data Analyst"
salary = 75000
employee_card = f"Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}"
print(employee_card)

employee_code = "DEV-2024-JD-001"

--fcc-editable-region--
# Create department below

--fcc-editable-region--
```
