---
id: 694bf6eb530e19a9c48c59f1
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

When you want to add content to the end of an existing string variable, you can use the **augmented assignment** operator `+=`.
This is shorter than writing `var = var + "new text"`.


An `address` variable has been created for you. Use `+=` to add the string `", Apartment 4B"` to the `address`.

# --hints--

You should use the `+=` operator to add `", Apartment 4B"` to the `address` variable.

```js
({ test: () => true })
```

You should have the correct string in `address`.

```js
({ test: () => true })
assert(test);
```

# --seed--

## --seed-contents--

```python
first_name = "John"
last_name = "Doe"
full_name = first_name + " " + last_name
address = "123 Main Street"

--fcc-editable-region--
# Update the address below

--fcc-editable-region--
```
