---
id: 694bf6eb530e19a9c48c59f3
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Let's practice `str()` conversion one more time. Just like with age, you must convert any numeric variable to a string before concatenating it with other text.

# --instructions--

Create a variable `experience_info`. Assign it the value of `"Experience: "` concatenated with `years_experience` (converted to a string) and `" years"`. Print the result.

# --hints--

You should convert `years_experience` to a string using `str(years_experience)`.

```js
assert(locals.get('experience_info') === "Experience: 5 years");
```

# --seed--

## --seed-contents--

```python

years_experience = 5

# --fcc-editable-region--
# Create experience_info below

# --fcc-editable-region--
```
