---
id: 694bf3af05d10fd89dab05cc
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In Python, you can combine strings using the `+` operator. This is called concatenation.
For example:

```python
greeting = "Hello" + "World"
print(greeting) # Output: HelloWorld
```

# --instructions--

Create a variable `full_name` by concatenating `first_name` and `last_name`. Then print `full_name`.

# --hints--

You should create a variable named `full_name`.

```js
assert(locals.get('full_name'));
```

You should assign the concatenation of `first_name` and `last_name` to `full_name`.

```js
assert.equal(locals.get('full_name'), "JohnDoe");
```

# --seed--

## --seed-contents--

```python

first_name = "John"
last_name = "Doe"

# --fcc-editable-region--
# Create the variable and then print it

# --fcc-editable-region--
```
