---
id: 694bf3af05d10fd89dab05cc
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

In Python, you can combine strings using the `+` operator. This is called concatenation. Here is an example:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Create a variable `full_name` by concatenating `first_name` and `last_name`. Then print `full_name`.

# --hints--

You should have a variable named `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

You should assign a string formed by concatenating `first_name` and `last_name` to your `full_name` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

You should print the `full_name` variable.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
