---
id: 694bf6eb530e19a9c48c59f7
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

You can also use negative numbers to slice from the end of a string.

* `-1` refers to the last character.
* `-2` refers to the second-to-last character, and so on.

To get the last 3 characters of a string, you can use `string[-3:]`.

# --instructions--

Create a variable `last_three`. Use negative indexing to extract the last 3 characters from `employee_code` (which represent the ID number).
Print `last_three`.

# --hints--

You should use negative indexing `[-3:]` to create `last_three`.

```js
assert(locals.get('last_three') === "001");
```

# --seed--

## --seed-contents--

```python

employee_code = "DEV-2024-JD-001"

# --fcc-editable-region--
# Create last_three below

# --fcc-editable-region--
```
