---
id: 694bf3af05d10fd89dab05cc
title: Concatenate Strings
challengeType: 20
dashedName: step-1
---

# --description--

In Python, you can combine strings using the `+` operator. This is called concatenation.


For example:

```python
greeting = "Hello" + "World"
print(greeting) # Output: HelloWorld
```

Create a variable `full_name` by concatenating `first_name` and `last_name`. Then print `full_name`.

# --hints--

You should create a variable named `full_name`.

```js
const test = await runPython(`
    _Node(_code).has_variable("full_name")
`);
assert(test);
```

You should print the `full_name` variable.

```js
const test = await runPython(`
    _Node(_code).has_call("print(full_name)")
`);
assert(test);
```

# --seed--

## --seed-contents--

```python
first_name = "John"
last_name = "Doe"

--fcc-editable-region--
# Create the variable and then print it

--fcc-editable-region--
```
