---
id: 694bf6eb530e19a9c48c59f3
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Let's practice `str()` conversion one more time. Just like with age, you must convert any numeric variable to a string before concatenating it with other text.


Create a variable `experience_info`. Assign it a string formed by concatenating `'Experience: '`, the integer `5` (converted to a string), and `' years'`. Print the result to the terminal.

# --hints--

You should create a variable `experience_info` with the value `'Experience: 5 years'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(5) + ' years'")`)) })
```

You should print `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
