---
id: 694bf6eb530e19a9c48c59f4
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Concatenating many strings using `+` and converting numbers using `str()` can get messy and hard to read.
Python 3.6 introduced **f-strings** to solve this. By adding the letter `f` before the opening quote, you can put variables directly inside curly braces `{}`.

Example:

```python
name = "John"
print(f"Hello {name}") # Output: Hello John
```

# --instructions--

We have added `position` and `salary` variables for you.

Create a variable `employee_card`. Use an f-string to format it exactly like this:
`"Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}"`

Then print `employee_card`.

# --hints--

You should create an f-string starting with `f"` or `f'`.
You should insert `full_name`, `employee_age`, `position`, and `salary` inside curly braces `{}`.

```js
assert(locals.get('employee_card') === "Employee: John Doe | Age: 28 | Position: Data Analyst | Salary: $75000");
```

# --seed--

## --seed-contents--

```python

full_name = "John Doe"
employee_age = 28
position = "Data Analyst"
salary = 75000

# --fcc-editable-region--
# Create employee_card below

# --fcc-editable-region--
```
