---
id: 694bf6ea530e19a9c48c59f0
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

If you simply add two names like `"John" + "Doe"`, the result is `"JohnDoe"` with no space.
To fix this, you need to concatenate a string containing a space `" "` between them.

# --instructions--

Update your `full_name` variable so it concatenates `first_name`, a space `" "`, and `last_name`.

# --hints--

You should add a space `" "` between `first_name` and `last_name`.

```js
assert(locals.get('full_name') === "John Doe");
```

# --seed--

## --seed-contents--

```python

first_name = "John"
last_name = "Doe"
full_name = first_name + last_name
print(full_name)

# --fcc-editable-region--
# Update the full_name you had created in the previous step

# --fcc-editable-region--
```
