---
id: 694bf6eb530e19a9c48c59f5
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

You often need to extract a specific portion of a string. In Python, this is called **slicing**.

The syntax is `string[start:stop]`.

* `start`: The index where the slice begins (inclusive).
* `stop`: The index where the slice ends (**exclusive**).

For example, if `text = "Python"`, then `text[0:2]` gives `"Py"`.

# --instructions--

We have defined an `employee_code` for you.
Create a variable `department` and assign it the slice of `employee_code` from index `0` to `3`. Then print `department`.

# --hints--

You should slice `employee_code` from index `0` to `3`.

```js
assert(locals.get('department') === "DEV");
```

# --seed--

## --seed-contents--

```python

employee_code = "DEV-2024-JD-001"

# --fcc-editable-region--
# Create department below

# --fcc-editable-region--
```
