---
id: 694bf6eb530e19a9c48c59f4
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Concatenating many strings using `+` and converting numbers using `str()` can get messy and hard to read.


Python 3.6 introduced **f-strings** to solve this. By adding the letter `f` before the opening quote, you can put variables directly inside curly braces `{}`.

Example:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```


Create a variable `employee_card` and assign it an f-string that displays `Employee:` followed by the value of the `full_name` variable.

# --hints--

You should create `employee_card` with the correct string.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_info = 'Experience: ' + str(5) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
