---
id: 694bf6eb530e19a9c48c59f6
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

You can slice from any part of the string, not just the beginning.

For the code `"DEV-2024-JD-001"`:

* `"2024"` starts at index 4 and ends before index 8.
* `"JD"` starts at index 9 and ends before index 11.

# --instructions--

1. Create a variable `year_code` and assign it the slice of `employee_code` from index `4` to `8`.
2. Create a variable `initials` and assign it the slice from index `9` to `11`.
3. Print both variables.

# --hints--

You should slice `employee_code` from index `4` to `8` for `year_code`.

```js
assert(locals.get('year_code') === "2024");
```

You should slice `employee_code` from index `9` to `11` for `initials`.

```js
assert(locals.get('initials') === "JD");
```

# --seed--

## --seed-contents--

```python

employee_code = "DEV-2024-JD-001"
department = employee_code[0:3]
print("Department: " + department)


# --fcc-editable-region--
# Create year_code and initials below

# --fcc-editable-region--
```
