---
id: 694bf6eb530e19a9c48c59f5
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

When working with strings, you'll often need to extract a specific portion of a string. This is called **slicing**.

The syntax is `string[start:stop]`, where:

* `start` is the index where the slice begins (**inclusive**).
* `stop` is the index where the slice ends (**exclusive**).

For example, if `text = 'Python'`, then `text[0:2]` gives `'Py'`.

Define `employee_code` as `'DEV-2026-JD-001'`. After that, create a variable `department` and assign it the slice of `employee_code` from index `0` to `3`. Then print `department` to the terminal.

# --hints--
You should have a variable named `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

You should assign the string `DEV-2026-JD-001` to your `employee_code` variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

You should have a variable named `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

You should slice the first three characters of `employee_code` and assign them to your `department` variable.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

You should print `department` to the terminal.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
