---
id: 694bf6ea530e19a9c48c59f0
title: Concatenate Strings with Spaces
challengeType: 20
dashedName: step-2
---

# --description--

If you simply add two names like `"John" + "Doe"`, the result is `"JohnDoe"` with no space.
To fix this, you need to concatenate a string containing a space `" "` between them.


Update your `full_name` variable so it concatenates `first_name`, a space `" "`, and `last_name`.

# --hints--

You should add a space `" "` between `first_name` and `last_name`.

```js
const test = await runPython(`
    full_name == "John Doe"
`);
assert(test);
```

# --seed--

## --seed-contents--

```python
first_name = "John"
last_name = "Doe"

--fcc-editable-region--
full_name = first_name + last_name
print(full_name)
--fcc-editable-region--
```
