---
id: 695fcfbba648b3a96a9495a4
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

The output in the terminal is getting crowded. Before you continue, it's time to clean it up.

Remove all the `print()` statements from your code.

# --hints--

You should not have any `print` calls in your code.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
