---
id: 69727ea5c9606bedfe36d3d2
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Now let's add the final details to the card.

Create a variable named `position` with the value of the string `Data Analyst` and a variable named `salary` with the value of the integer `75000`.

Then, update your `employee_card` f-string to include the age, position, and salary. It should follow this exact format:

`Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}`

Finally, print `employee_card` to see the result.

# --hints--

You should create a variable `position` with the value `'Data Analyst'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

You should create a variable salary with the value 75000.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

You should update employee_card to match the required format.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

You should print employee_card.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_info = 'Experience: ' + str(5) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
