---
id: 657effaa2a5e0277d71f9cbe
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

By this point, the variable `snake_cased_char_list` holds the list of converted characters. To combine these characters into a single string, you can utilize the `.join()` method.

The `join` method works by concatenating each element of a list into a string, separated by a designated string, known as the separator.

```py
result_string = ''.join(characters)
```

The example above joins together the elements of the `characters` list into a single string where each element is concatenated together using an empty string as the separator.

Now, right after the `for` loop, use the `.join()` method to join the elements in `snake_cased_char_list` using an empty string as the separator. Assign the result to a new variable named `snake_cased_string`.


# --hints--

You should join the characters in the `snake_cased_char_list` list into a string using the `.join()` method and assign the string to a variable named `snake_cased_string` after your `for` loop.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
