---
id: 684ccbdc1ef89e042b744e09
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

The last part of the markup to add would be the dismiss button. 

Below your `p` element, add a `button` element with an `aria-label` set to `"Dismiss"`. Inside that `button` element, add a `span` element with an `aria-hidden` attribute. The `span` element text should be `&times;`. 

*NOTE*: Since this workshop will only be focused on working with Tailwind CSS styles, the `button` will not have any type of interactivity added to it. 

# --hints--

You should have a `button` element.

```js
assert.exists(document.querySelector("button"));
```
 
Your `button` element should have an `aria-label` attribute.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

Your `button` element should have an `aria-label` attribute set to `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

Your `button` element should have a `span` element nested inside.

```js
assert.exists(document.querySelector("button span"));
```

Your `span` element should have an `aria-hidden` attribute.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

Your `span` element should have the text `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```
