---
id: 684cc59eed2706fa50e8933c
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to start working with the basics of Tailwind CSS by building an error message component. 

For the first few steps, you will focus on building out the initial HTML markup. 

Inside of the `body` element, add a `div` element. Then inside of that `div` element, add a `p` element.

# --hints--

You should have a `div` element.

```js
assert.exists(document.querySelector("div"));
```

You should have a `p` element.

```js
assert.exists(document.querySelector("p"));
```

Your `p` element should be inside of the `div` element.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
