---
id: 684cd027744c2a0520ba4d58
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now that all of the markup has been added, it is time to start applying Tailwind CSS to the elements. Tailwind uses utility classes which are applied directly to the markup.  These classes include styling for fonts, colors, Flexbox, Grid and more. 

For the next few steps, you will focus on adding some basic styling for your `div` element. The first style you will need to apply is a background color. To apply a background color in Tailwind CSS, you can use the `bg` prefix, followed by the color name and number. The number represents how light or dark the color will be. 

Here is an example of applying a dark blue background color to a paragraph element.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

In your opening `div` tag, add a class with the utility class `bg-red-100` which represents a light red background color.

# --hints--

Your opening `div` tag should have a class attribute.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

Your class attribute should have the utility class `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
