---
id: 684d09cad4ec900e41db2e74
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Right now the `div` element is placed in the upper left hand corner of the preview screen. But it would be nice to add some space to the top of the element. To apply `margin`, to an element using Tailwind CSS, you can use the following pattern: `m-number`. The number represents the fixed pixels to apply to all four sides.

If you want to apply margin to only one side you can use any of the following patterns: `mt-number`(margin top), `mb-number`(margin bottom), `mr-number`(margin right) and `ml-number`(margin left).

Here is an example of apply `margin-bottom` to an element:

```html
<p class="mb-8">This is margin.</p>
```

Inside your opening `div` tag, use the correct utility class to apply `margin-top` of `4px`.

# --hints--

Your opening `div` tag should have the `mt-4` utility class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
