---
id: 684d17a3b350bb13d3c19c0f
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Now it is time to apply some styles to the `button` element. 

Start by adding a class attribute to your `button` element. Next, apply the correct utility class for setting a red text color with the shade set to `700`.

Then, inside your opening `button` element, use the correct utility class to set the `font-size` to `2xl` which represents `1.5rem` or (`24px`).

# --hints--

Your `button` element should have a class attribute.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

Your opening `button` tag should have the `text-red-700` utility class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

Your opening `button` tag should have the `text-2xl` utility class.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
