---
id: 684d0d369e2b3b0f288d2c20
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

It would be nice to horizontally center the `div` element. To apply `margin` to two sides at once, you can use the following patterns: `mx-number` or `my-number`. 

`mx` sets the margin for left and right while `my` sets the margin for the top and bottom.

Here is an example of applying the top and bottom margins to an element:

```html
<p class="my-8">This is margin.</p>
```

Inside your opening `div` tag, use the correct utility class to set margin left and right to `auto`.

# --hints--

Your opening `div` tag should have the `mx-auto` utility class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
