---
id: 684d03df2a0ddf0c735865f0
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Now you see that the element width is half the size of the `body`. This works well for larger devices but not as well for smaller devices.

Tailwind CSS follows a mobile first approach in design. So if you want to create responsive designs for other devices, you can use Tailwind's responsive variants like `md` and `lg` like this:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

The following example will set the width to `66%` of the parent element for device sizes of `1024px` and higher. 

Inside your opening `div` tag, modify your `w-1/2` so it is only applied to device sizes of `768px` and higher. You can prefix this class with the `md` variant. 

# --hints--

Your opening `div` tag should have the `md:w-1/2` utility class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

You should no longer have the `w-1/2` class in your opening `div` tag.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
