---
id: 68501d32b69f9d7756334ffc
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

If you enlarge the preview window, you should see the `div` element centered horizontally on the page. But for smaller devices, the width for the `div` element spans the entire width of the page.  

Update the `mx-auto` class so it is only applied for `md` devices and larger.

# --hints--

Your opening `div` tag should have the `md:mx-auto` utility class.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

You should no longer have the `mx-auto` class in your opening `div` tag.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
