---
id: 642db8c409d9991d0b3b2f0d
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

O projeto começa com um contêiner HTML básico e com o CSS correspondente. Sua primeira tarefa será gerar células para sua planilha por meio de programação.

O objeto global `window` representa a janela (ou aba) do navegador. Ele tem uma propriedade `onload` que permite definir um comportamento quando a janela carregar a página inteira, incluindo folhas de estilo e scripts.

Comece definindo a propriedade `onload` de `window` como uma arrow function sem parâmetros. Na função, declare uma variável `container` e atribua a ela o valor do elemento usando o `id` de `container`.

# --hints--

Você deve acessar a propriedade `onload` do elemento `window`.

```js
assert.match(code, /window\.onload/);
```

Você deve definir a propriedade `onload` como sendo uma função.

```js
assert.isFunction(window.onload);
```

Você deve usar a sintaxe das arrow functions.

```js
assert.match(code, /window\.onload\s*=\s*\(/);
```

A função `onload` não deve receber nenhum parâmetro.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)/);
```

Você deve declarar uma variável `container` na função de `onload`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*(?:let|var|const)\s+container/);
```

A variável `container` deve ser declarada com `const`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container/);
```

Você deve usar `document.getElementById()`

```js
assert.match(code, /document\.getElementById\(/);
```

Você deve obter o elemento através do `id` `container`.

```js
assert.match(code, /document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

Você deve atribuir a regra `#container` ao elemento `container`.

```js
assert.match(code, /window\.onload\s*=\s*\(\s*\)\s*=>\s*\{\s*const\s+container\s*=\s*document\.getElementById\(\s*('|"|`)container\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
