---
id: 642ddfdea4200e313f80a4b6
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Lembre-se de que o objeto `document` tem um método `.createElement()`, que permite que você crie elementos do HTML dinamicamente.

Na função `createLabel`, declare uma variável `label` e atribua a ela um novo elemento `div`.

# --hints--

Você deve declarar uma variável `label` na função `createLabel`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

A variável `label` deve ser declarada com `const`.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

Você deve usar o método `.createElement()` no objeto `document`.

```js
assert.match(code, /document\.createElement\(/);
```

Você deve passar a string `div` para o método `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

Você deve atribuir o novo elemento `div` a `label`.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
