---
id: 642df9df4b5216350de7b0d2
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Você precisará de uma função para gerar um intervalo de números.

Declare uma função `range` vazia, que recebe os parâmetros `start` e `end`. Use o construtor de `Array()` e retorne implicitamente um array vazio.

# --hints--

Você deve declarar uma variável `range`.

```js
assert.match(code, /(?:let|var|const)\s+range/);
```

Você deve usar `const` para declarar a variável `range`.

```js
assert.match(code, /const\s+range/);
```

A variável `range` deve ser uma função.

```js
assert.isFunction(range);
```

A função `range` deve usar a sintaxe das arrow functions.

```js
assert.match(code, /const\s+range\s*=\s*\(.*\)\s*=>/);
```

A função `range` deve receber `start` como primeiro parâmetro.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start/)
```

A função `range` deve receber `end` como segundo parâmetro.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

A função `range` deve usar um retorno implícito. Lembre-se de que isso significa que você não usará chaves.

```js
assert.notMatch(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

A função `range` deve usar o construtor de `Array()`. Construtores primitivos não precisam da palavra-chave `new`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(/);
```

Você não deve passar nada para o construtor de `Array()`.

```js
assert.match(code, /const\s+range\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*Array\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
