---
id: 642def66e6a60432c9a0371e
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Defina o `className` do elemento `label` como `label` e defina `textContent` como o parâmetro `name`.

# --hints--

Você deve acessar a propriedade `className` do elemento `label`.

```js
assert.match(code, /label\.className/);
```

Você deve definir a propriedade `className` como a string `label`.

```js
assert.match(code, /label\.className\s*=\s*('|"|`)label\1/);
```

Você deve acessar a propriedade `textContent` do elemento `label`.

```js
assert.match(code, /label\.textContent/);
```

Você deve definir a propriedade `textContent` como o valor de `name`.

```js
assert.match(code, /label\.textContent\s*=\s*name/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");

  }
}
--fcc-editable-region--
```
