---
id: 642e02be7845f13b014cd2b0
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Agora que você tem uma função `range`, você também pode usá-la para criar um intervalo de letras.

Declare uma função `charRange` usando a palavra-chave `const` e a sintaxe das arrow functions. Ela deve receber os parâmetros `start` e `end`. A função deve retornar implicitamente o resultado da chamada de `range()`, com `start` e `end` como argumentos.

# --hints--

Você deve declarar uma variável `charRange`.

```js
assert.match(code, /(?:let|const|var)\s+charRange/);
```

A variável `charRange` deve ser declarada com `const`.

```js
assert.match(code, /const\s+charRange/);
```

A variável `charRange` deve ser uma função.

```js
assert.isFunction(charRange);
```

A função `charRange` deve usar a sintaxe das arrow functions.

```js
assert.match(code, /const\s+charRange\s*=\s*\(.*\)\s*=>/);
```

A função `charRange` deve receber `start` como primeiro parâmetro.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start/);
```

A função `charRange` deve receber `end` como segundo parâmetro.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)/);
```

A função `charRange` deve usar um retorno implícito.

```js
assert.notMatch(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*{/);
```

A função `charRange` deve chamar a função `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(/);
```

Você deve passar `start` e `end` como argumentos para a chamada de `range`.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\s*,\s*end\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--

--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
