---
id: 643498755d54c6279ba09078
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

A maioria dos programas de planilha inclui funções integradas para cálculo.

Declare uma função `sum` (soma) que receba um parâmetro `nums`, que será um array de números. Ela deve retornar o resultado da chamada de `reduce` no array para somar todos os números.

# --hints--

Você deve declarar uma variável `sum`.

```js
assert.match(code, /(?:let|const|var)\s+sum/);
```

Você deve usar `const` para declarar a variável `sum`.

```js
assert.match(code, /const\s+sum/);
```

A variável `sum` deve ser uma função.

```js
assert.isFunction(sum);
```

A função `sum` deve usar a sintaxe das arrow functions.

```js
assert.match(code, /const\s+sum\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

A função `sum` deve receber um parâmetro `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>/);
```

A função `sum` deve usar um retorno implícito.

```js
assert.notMatch(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

A função `sum` deve retornar o resultado da chamada de `.reduce()` em `nums`.

```js
assert.match(code, /const\s+sum\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*nums\.reduce\(/);
```

A função `sum` deve retornar a soma de todos os números em `nums`.

```js
const numbers = [1, 2, 3, 4, 5];
assert.equal(sum(numbers), 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
