---
id: 6434552bcc0a951a0a99df3b
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

A função `range` espera por números, mas os valores de `start` e `end` serão strings (especificamente, eles serão caracteres únicos, como `A`).

Converta os valores de `start` e `end` na chamada de `range()` para números usando neles o método `.charCodeAt()`, passando o número `0` como argumento para esse método.

# --hints--

Você deve usar o método `.charCodeAt()`.

```js
assert.match(code, /\.charCodeAt\(/);
```

Você deve chamar o método `.charCodeAt()` em `start`.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Você deve passar `0` para o método `.charCodeAt()` de `start`.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Você deve chamar o método `.charCodeAt()` em `end`.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Você deve passar `0` para o método `.charCodeAt()` de `end`.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Você deve usar o método `.charCodeAt()` diretamente na chamada de `range` chamada.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
