---
id: 64345c560591891f64976f7a
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora que as funções auxiliares foram concluídas, de volta ao manipulador de eventos `onload`, você deve declarar uma variável `letters`. Atribua a ela o resultado da chamada de `charRange()` com as letras `A` e `J` como argumentos.

# --hints--

Você deve declarar uma variável `letters`.

```js
assert.match(code, /(?:let|const|var)\s+letters/);
```

Você deve usar `const` para declarar a variável `letters`.

```js
assert.match(code, /const\s+letters/);
```

Você deve atribuir a chamada de `charRange()` à variável `letters`.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(/);
```

Você deve passar `A` como o primeiro argumento para a chamada de `charRange()`.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1/);
```

Você deve passar `J` como o segundo argumento para a chamada de `charRange()`.

```js
assert.match(code, /const\s+letters\s*=\s*charRange\(\s*('|"|`)A\1\s*,\s*('|"|`)J\2\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
