---
id: 6434750c53db16218f41e6e1
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Lembre-se de que `range()` retorna um array. Assim, você pode encadear métodos de array diretamente à chamada da função.

Chame `range()` com `1` e `99` como argumentos e encadeie o método `.forEach()`. Passe para o método `.forEach()` uma função de callback vazia que receba `number` como parâmetro.

# --hints--

Você deve chamar a função `range()`.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

Você deve passar `1` como o primeiro argumento para a chamada de `range()`.

```js
assert.match(code, /range\(\s*1/);
```

Você deve passar `99` como o segundo argumento para a chamada de `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

Você deve encadear o método `.forEach()` à chamada de `range()`.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

Você deve passar uma função de callback para `.forEach()` usando a sintaxe das arrow functions.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

A função de callback deve ter `number` como o único parâmetro.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
