---
id: 643475e13dc727231acd0f72
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

Agora, na chamada aninhada de `.forEach()`, declare uma variável `input`. Use o método `.createElement()` do objeto `document` para criar um elemento `input`. Defina o atributo `type` como `text` e o atributo `id` como `letter + number`.

# --hints--

Você deve declarar uma variável `input`.

```js
assert.match(code, /(?:var|let|const)\s+input/)
```

Você deve usar `const` para declarar a variável `input`.

```js
assert.match(code, /const\s+input/)
```

Você deve chamar o método `.createElement()` no objeto `document`.

```js
assert.lengthOf(code.match(/document\.createElement\(/g), 2)
```

Você deve passar a string `input` para o método `.createElement()`.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Você deve atribuir o novo elemento `input` à variável `input`.

```js
assert.match(code, /const\s+input\s*=\s*document\.createElement\(\s*('|"|`)input\1\s*\)/)
```

Você deve acessar a propriedade `type` do elemento `input`.

```js
assert.match(code, /input\.type/);
```

Você deve definir o atributo `type` do elemento `input` como `text`.

```js
assert.match(code, /input\.type\s*=\s*('|"|`)text\1/)
```

Você deve acessar a propriedade `id` do elemento `input`.

```js
assert.match(code, /input\.id/);
```

Você deve definir o atributo `id` do elemento `input` como `letter + number`.

```js
assert.match(code, /input\.id\s*=\s*letter\s\+\snumber/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {

    })
--fcc-editable-region--
  })
}
```
