---
id: 6437133052eaf04d7300e622
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Declare uma função `average`, que receba um array de números como o parâmetro `nums`. Ela deve retornar a média de todos os números do array.

A média pode ser calculada dividindo a soma de todos os números pelo comprimento do array. Lembre-se de que você tem uma função `sum` (soma) que você pode usar.

# --hints--

Você deve declarar uma variável `average`.

```js
assert.match(code, /(?:let|const|var)\s+average/);
```

Você deve usar `const` para declarar a variável `average`.

```js
assert.match(code, /const\s+average/);
```

A variável `average` deve ser uma função.

```js
assert.isFunction(average);
```

A função `average` deve usar a sintaxe das arrow functions.

```js
assert.match(code, /const\s+average\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

A função `average` deve ter um parâmetro `nums`.

```js
assert.match(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)/);
```

A função `average` deve usar um retorno implícito.

```js
assert.notMatch(code, /const\s+average\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*{/);
```

A função `average` deve retornar o valor da média do array `nums`.

```js
assert.equal(average([1,2,3]), 2);
assert.equal(average([1,2,3,4,5]), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
