---
id: 66f26c32ec6f90df01a44f60
title: Build a Contact Form
challengeType: 14
dashedName: lab-contact-form
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `div` with a class of `form-container` to hold your form.

2. You should have a `form` element within the `.form-container` `div`.

3. Within the `form` element, you should have the following elements and input fields:

   - A `h2` element with some text.
   - A text input field for the name with the `type` set to `text` and `id`, `name`, `required` attributes.
   - An email input field with the `type` set to `email` and `id`, `name`, `required` attributes.
   - A textarea for the message with `id`, `name`, and the `required` attribute.
   - A `button` element with a `type` attribute of `submit`.

4. Each `input` and `textarea` elements should have their corresponding `label` element with a `for` attribute set to the element's `id`.

5. The `form-container` `div` should have a background color, and have values for `border-radius`, `padding`, and `width`.

6. The `label` elements should have a margin and a font color.

7. The `input` and `textarea` elements should have values for:
   
   - `width`.
   - `padding`.
   - Margin on the bottom.

8. The `button` element should have a:
   
   - Background color
   - Font size apart from the default.

9.  The button should have a hover effect that changes the background color.

**Note:** Be sure to link your stylesheet in your HTML and apply your CSS.

# --hints--

You should have a `div` with a class of `form-container`.

```js
const el = document.querySelector('.form-container');
assert(!!el && el.tagName === 'DIV');
```

You should have a `form` element within the `.form-container` `div`.

```js
const el = document.querySelector('.form-container form');
assert(!!el && el.tagName === 'FORM');
```

Within the `form`, you should have an `h2` element with some text.

```js
const el = document.querySelector('.form-container form h2');
assert(!!el && el.innerText.length > 0);
```

Within the `form`, you should have a text input field with the type set to `text`.

```js
const el = document.querySelector('.form-container form input[type="text"]');
assert(!!el && el.tagName === 'INPUT');
```

Your `type=text` input field should have an `id`, a `name`, and a `required` attribute.

```js
const el = document.querySelector('.form-container form input[type="text"]');

assert(!!el && el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Within the `form`, you should have an email input field with the `type` set to `email`.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert(!!el && el.tagName === 'INPUT');
```

Your `type=email` input field should have an `id`, a `name`, and a `required` attribute.

```js
const el = document.querySelector('.form-container form input[type="email"]');
assert(!!el && el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Within the `form`, you should have a `textarea` with an `id`, `name`, and a `required` attribute.

```js
const el = document.querySelector('.form-container form textarea');
assert(!!el && el.tagName === 'TEXTAREA');
assert(!!el && el.required);
assert.isNotEmpty(el?.getAttribute('id'));
assert.isNotEmpty(el?.getAttribute('name'));
```

Within the `form`, you should have a `button` element with a `type` attribute of `submit`.

```js
const el = document.querySelector('.form-container form button');
assert(!!el && el.tagName === 'BUTTON');
assert(!!el && el.getAttribute('type') === 'submit');
```

You should have a `label` element for the `#name` `input` with a `for` attribute matching the `id` value on the corresponding `input`.

```js
// fetch id for the type=text input
const nameInput = document.querySelector('.form-container form input[type="text"]')?.getAttribute('id');

// fetch all labels
const labels = document.querySelectorAll('.form-container form label');

// check if labels exist
assert(labels.length > 0);

let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === nameInput) {
        found = true;
        break;
    }
}

// assert that the correct label was found
assert(found);
```

You should have a `label` element for the `#email` `input` with a `for` attribute matching the `id` value on the corresponding `input`.

```js
// fetch id for the type=email input
const emailInput = document.querySelector('.form-container form input[type="email"]')?.getAttribute('id');

// fetch all labels
const labels = document.querySelectorAll('.form-container form label');

// check if labels exist
assert(labels.length > 0);

let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === emailInput) {
        found = true;
        break;
    }
}

// assert that the correct label was found
assert(found);
```

You should have a `label` element for the `#message` `textarea` with a `for` attribute matching the `id` value on the `textarea`.

```js
// fetch id for the textarea
const messageInput = document.querySelector('.form-container form textarea')?.getAttribute('id');

// fetch all labels
const labels = document.querySelectorAll('.form-container form label');

// check if labels exist
assert(labels.length > 0);

let found = false;

for (let label of labels) {
    if (label.getAttribute('for') === messageInput) {
        found = true;
        break;
    }
}

// assert that the correct label was found
assert(found);

```

The labels shouldn't be empty.

```js
// fetch all labels inside the form
const labels = document.querySelectorAll('.form-container form label');

// check if labels exist
assert(labels.length > 0);

for (let label of labels) {
    // check if the label's text content is not empty or just whitespace
    const labelText = label.textContent.trim();
    assert(labelText.length > 0);
}

```

The `form-container` `div` should have a background color.

```js
assert(new __helpers.CSSHelp(document).getStyle('.form-container')?.backgroundColor);
```

The `form-container` `div` should have a visible border that is also rounded.

```js
assert(new __helpers.CSSHelp(document).getStyle('.form-container')?.border);
assert(new __helpers.CSSHelp(document).getStyle('.form-container')?.borderRadius);
```

The `form-container` `div` should have padding.

```js
assert(new __helpers.CSSHelp(document).getStyle('.form-container')?.padding);
```

The `form-container` `div` should have a width.

```js
assert(new __helpers.CSSHelp(document).getStyle('.form-container')?.width);
```

Your labels should have a margin and font colour.

```js
assert(new __helpers.CSSHelp(document).getStyle('label')?.margin);
assert(new __helpers.CSSHelp(document).getStyle('label')?.color);
```

Your `input` elements should have a width, padding, and a margin on the bottom.

```js
// Get the input styles once
const inputStyles = new __helpers.CSSHelp(document).getStyle('input');

// Check for width
const width = inputStyles?.width;
assert(width);

// Check for padding (or individual paddings if padding is not defined)
const padding = inputStyles?.padding;
const paddingTop = inputStyles?.paddingTop;
const paddingBottom = inputStyles?.paddingBottom;
const paddingLeft = inputStyles?.paddingLeft;
const paddingRight = inputStyles?.paddingRight;

assert(
  padding || (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

// Check for margin-bottom
const marginBottom = inputStyles?.marginBottom;
assert(marginBottom);

```

Your `textarea` element should have a width, padding, margin on the bottom.

```js
// Get the input styles once
const inputStyles = new __helpers.CSSHelp(document).getStyle('textarea');

// Check for width
const width = inputStyles?.width;
assert(width);

// Check for padding (or individual paddings if padding is not defined)
const padding = inputStyles?.padding;
const paddingTop = inputStyles?.paddingTop;
const paddingBottom = inputStyles?.paddingBottom;
const paddingLeft = inputStyles?.paddingLeft;
const paddingRight = inputStyles?.paddingRight;

assert(
  padding || (paddingTop && paddingBottom && paddingLeft && paddingRight)
);

// Check for margin-bottom
const marginBottom = inputStyles?.marginBottom;
assert(marginBottom);
```

Your `button` element should have a background color.

```js
assert(new __helpers.CSSHelp(document).getStyle('button')?.backgroundColor);
```

Your `button` element should have a font size.

```js
assert(new __helpers.CSSHelp(document).getStyle('button')?.fontSize);
```

Your `button` element should have a hover effect that changes the background color.

```js
const aHoverColor = new __helpers.CSSHelp(document).getStyle('button:hover')?.getPropertyValue('color');
const buttonOriginalColor = new __helpers.CSSHelp(document).getStyle('button')?.getPropertyValue('color');
assert(aHoverColor !== buttonOriginalColor);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Contact Form</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Styled Form Project</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="form-container">
        <form>
            <h2>Contact Us</h2>
            <label for="name">Name:</label>
            <input type="text" id="name" name="name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" required>

            <label for="message">Message:</label>
            <textarea id="message" name="message" required></textarea>

            <button type="submit">Submit</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background: linear-gradient(45deg, #83a4d4, #b6fbff);
    font-family: Arial, sans-serif;
}

.form-container {
    background-color: white;
    border-radius: 10px;
    padding: 20px;
    width: 300px;
    text-align: center;
}

form h2 {
    margin-bottom: 20px;
    color: #333;
}

label {
    display: block;
    margin: 10px 0 5px;
    color: #555;
}

input{
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

textarea {
    width: 100%;
    padding: 10px;
    margin-bottom: 10px;
    border: 1px solid #ddd;
    border-radius: 5px;
    box-sizing: border-box;
    resize: vertical;
}

button {
    background-color: #4CAF50;
    color: white;
    padding: 10px 15px;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    font-size: 1em;
}

button:hover {
    background-color: #45a049;
}
```
