---
id: 6a7061b30dc6edf4528c6a07
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Inside `normalizeDirectory`, declare four empty objects: `departmentsById`, `programsById`, `instructorsById`, and `instructorsByEmail`.

These objects are lookup tables: they use an ID or another unique value, such as an email, as a key. This lets you access a record directly instead of searching through an array.

# --before-each--

```js
globalThis.getLookupVariables = () =>
  __helpers.Explorer(code).then((explorer) => {
    const fn = explorer.functions.normalizeDirectory;
    assert.exists(fn);
    return fn.variables;
  });
```

# --after-each--

```js
delete globalThis.getLookupVariables;
```

# --hints--

You should declare `departmentsById` as an empty object.

```js
const variables = await globalThis.getLookupVariables();
assert.exists(variables.departmentsById);
assert.isTrue(variables.departmentsById.value.matches('{}'));
```

You should declare `programsById` as an empty object.

```js
const variables = await globalThis.getLookupVariables();
assert.exists(variables.programsById);
assert.isTrue(variables.programsById.value.matches('{}'));
```

You should declare `instructorsById` as an empty object.

```js
const variables = await globalThis.getLookupVariables();
assert.exists(variables.instructorsById);
assert.isTrue(variables.instructorsById.value.matches('{}'));
```

You should declare `instructorsByEmail` as an empty object.

```js
const variables = await globalThis.getLookupVariables();
assert.exists(variables.instructorsByEmail);
assert.isTrue(variables.instructorsByEmail.value.matches('{}'));
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
--fcc-editable-region--
  
--fcc-editable-region--
}

normalizeDirectory(rawData);

```
