---
id: 6a70b2e4b65f00de6f414041
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

Each department has a `programs` array nested inside it. Add a second `for` loop inside the existing loop to iterate over `departmentData.programs`, from `j = 0` up to `departmentData.programs.length`.

Inside this new loop, declare a variable called `program` set to `departmentData.programs[j]` and log it to the console.

You should also remove the `console.log(departmentsById[departmentId])` line from your code. 

# --hints--

You should delete `console.log(departmentsById[departmentId]);`.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*departmentsById\s*\[\s*departmentId\s*\]\s*\)/
);
```

You should use a `for` loop.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+j\s*=\s*0\s*;/
);
```

Your `for` loop should run while `j` is less than `departmentData.programs.length`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+j\s*=\s*0\s*;\s*j\s*<\s*departmentData\.programs\.length\s*;/
);
```

You should declare a variable called `program` set to `departmentData.programs[j]`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /(?:const|let)\s+program\s*=\s*departmentData\s*\.\s*programs\s*\[\s*j\s*\]/
);
```

You should call `console.log()` with `program` as the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*program\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };
      
    --fcc-editable-region--
    console.log(departmentsById[departmentId]);

    --fcc-editable-region--
  }
}

normalizeDirectory(rawData);
```
