---
id: 6a708940c9b142fce3c4b9f8
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

`departmentEntries` is an array. You need to process each `[key, value]` pair individually to extract the department name and its data, then build your lookup tables as you go.

Remove `console.log(departmentEntries)`. Then, write a `for` loop that runs from `i = 0` up to `departmentEntries.length`.

# --hints--

You should remove the `console.log(departmentEntries)` line from your code.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*departmentEntries\s*\)/
);
```

You should use a `for` loop.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(/
);
```

Your `for` loop should start with `i` equal to `0`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+i\s*=\s*0\s*;/
);
```

Your `for` loop should run while `i` is less than `departmentEntries.length`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*departmentEntries\.length\s*;/
);
```

Your `for` loop should increment `i` by `1` on each iteration.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*departmentEntries\.length\s*;\s*(?:i\s*\+\+|\+\+\s*i|i\s*\+=\s*1|i\s*=\s*i\s*\+\s*1|i\s*=\s*1\s*\+\s*i)\s*\)/,
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  --fcc-editable-region--
  console.log(departmentEntries);
  
  --fcc-editable-region--

}

normalizeDirectory(rawData);

```
