---
id: 6a70b10d9a1a0ff3cb4775b8
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

Now that you have a `departmentId`, use it to store the department's data inside `departmentsById`.

Assign `departmentsById[departmentId]` an object containing:

`id`: set to `departmentId`,  
`code`: set to `departmentData.Dept_code`,  
`name`: set to `departmentName`  
`programIds`: an empty array (you'll fill this in later steps)

Lastly, log the `departmentsById[departmentId]` to the console.

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, "log");
const getLogs = () => globalThis.spy.calls.map((call) => call?.[0]);
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should assign an object to `departmentsById[departmentId]`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /departmentsById\s*\[\s*departmentId\s*\]\s*=\s*{/ 
);
```

The object should have an `id` property set to `departmentId`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /id\s*:\s*departmentId/
);
```

The object should have a `code` property set to `departmentData.Dept_code`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /code\s*:\s*departmentData\s*\.\s*Dept_code/
);
```

The object should have a `name` property set to `departmentName`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /name\s*:\s*departmentName/
);
```

The object should have a `programIds` property set to an empty array.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programIds\s*:\s*\[\s*\]/
);
```

You should call `console.log()` with `departmentsById[departmentId]` as the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*departmentsById\s*\[\s*departmentId\s*\]\s*\)/
);
```

Each logged department object should correctly map the expected `id`, `code`, `name`, and empty `programIds` values.

```js
const expectedDepartments = Object.entries(rawData).map(
  ([departmentName, departmentData], index) => ({
    id: departmentData.Dept_code || `dept-${index + 1}`,
    code: departmentData.Dept_code,
    name: departmentName,
    programIds: [],
  }),
);

assert.deepEqual(getLogs(), expectedDepartments);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;
      
    --fcc-editable-region--
    
    --fcc-editable-region--
  }
}

normalizeDirectory(rawData);

```
