---
id: 6a705f3f4d36006ee0d90f87
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Declare a function called `normalizeDirectory` that takes `rawData` as a parameter. 

Then below that function declaration, call `normalizeDirectory` and pass in `rawData` for the argument. 

# --hints--

You should declare a function named `normalizeDirectory`.

```js
assert.isFunction(normalizeDirectory);
```

`normalizeDirectory` function should take `rawData` as a parameter.

```js
const explorer = await __helpers.Explorer(code);
const fn = explorer.allFunctions.normalizeDirectory;
assert.strictEqual(fn?.parameters.length, 1);
assert.isTrue(fn?.parameters[0]?.matches('rawData'));
```

You should call `normalizeDirectory` function with `rawData` as a argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /(?:^|[;}])\s*normalizeDirectory\s*\(\s*rawData\s*\)\s*(?:;|$)/m,
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
