---
id: 6a70a88fe27c85f91d5137e8
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Each entry in `departmentEntries` is itself a small array in the shape `[key, value]`. Instead of accessing these by index (`entry[0]`, `entry[1]`), you can use array destructuring to unpack them into named variables in one line.

Destructure `departmentEntries[i]` into `departmentName` and `departmentData`. Then use two separate
`console.log()` statements: one to log `departmentName` and another to log `departmentData`.

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, 'log');
const getLogs = () => globalThis.spy.calls.map((call) => call?.[0]);
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should destructure `departmentEntries[i]` into `departmentName` and `departmentData`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s*\[\s*departmentName\s*,\s*departmentData\s*\]\s*=\s*departmentEntries\s*\[\s*i\s*\]/,
);
```

You should call `console.log()` with `departmentName` as the argument.

```js
const [departmentName] = Object.entries(rawData)[0];
const logs = getLogs();

assert.isTrue(
  logs.some((log) => JSON.stringify(log) === JSON.stringify(departmentName)),
);
```

You should call `console.log()` with `departmentData` as the argument.

```js
const [, departmentData] = Object.entries(rawData)[0];
const logs = getLogs();

assert.isTrue(
  logs.some((log) => JSON.stringify(log) === JSON.stringify(departmentData)),
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    --fcc-editable-region--

    --fcc-editable-region--
  }
}

normalizeDirectory(rawData);
```

