---
id: 6a782e189b32c29bf131e69c
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now you will loop through the departments in `rawData`. Use a `for...in` loop to iterate over the keys of `rawData`, where each key is a department name. Use `dept` as the loop variable and log `dept` to the console inside the loop.

Here is a reminder on how to work with a `for...in` loop:

```js
const menu = {
  "Italian": { dishes: ["Pizza", "Pasta"] },
  "Mexican": { dishes: ["Tacos", "Burritos"] }
};

for (const cuisine in menu) {
  console.log(cuisine);
}
// "Italian"
// "Mexican"
```

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, "log");
const getLogs = () => globalThis.spy.calls.map(call => call?.[0]);
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should create a `for...in` loop that iterates over `rawData`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(?:const|let)\s+dept\s+in\s+rawData\s*\)\s*\{[\s\S]*\}/
);
```

You should log the `dept` to the console inside the loop.

```js
const expectedKeys = Object.keys(rawData);
const logs = getLogs();

assert.isTrue(expectedKeys.every(key => logs.some(log => JSON.stringify(log) === JSON.stringify(key))));

```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function getInstructorByEmail(email) {

  --fcc-editable-region--
  
  --fcc-editable-region--
  return null;
}

console.log(getInstructorByEmail("fatima@example.com"));
```
