---
id: 6a78549572797d79eedc16e5
title: Step 25
challengeType: 1
dashedName: step-25
---

# --description--

Your four lookup tables are now fully built. Have `normalizeDirectory` return an object containing 
all four objects: `departmentsById`, `programsById`, `instructorsById`, and `instructorsByEmail`.

Since the property names match the variable names, you can use object property shorthand: write `{ departmentsById }` instead of `{ departmentsById: departmentsById }`.

# --hints--

Your returned object should contain the completed `departmentsById` lookup table.

```js
const actual = normalizeDirectory(rawData);
assert.isObject(actual);
const expected = Object.fromEntries(
  Object.entries(rawData).map(([name, dept], i) => {
    const id = dept.Dept_code || 'dept-' + (i + 1);
    return [
      id,
      {
        id,
        code: dept.Dept_code,
        name,
        programIds: dept.programs.map((p) => p.id),
      },
    ];
  }),
);
assert.deepEqual(actual.departmentsById, expected);
```

Your returned object should contain the completed `programsById` lookup table.

```js
const actual = normalizeDirectory(rawData);
const expected = Object.fromEntries(
  Object.entries(rawData).flatMap(([name, dept], i) =>
    dept.programs.map((program) => [
      program.id,
      { ...program, departmentId: dept.Dept_code || 'dept-' + (i + 1) },
    ]),
  ),
);
assert.deepEqual(actual?.programsById, expected);
```

Your returned object should contain the completed `instructorsById` lookup table.

```js
const actual = normalizeDirectory(rawData);
const expected = Object.fromEntries(
  Object.entries(rawData).flatMap(([name, dept], i) =>
    dept.programs.flatMap((program) =>
      program.instructors.map((instructor) => [
        instructor.id,
        {
          ...instructor,
          programId: program.id,
          departmentId: dept.Dept_code || 'dept-' + (i + 1),
        },
      ]),
    ),
  ),
);
assert.deepEqual(actual?.instructorsById, expected);
```

Your returned object should contain the completed `instructorsByEmail` lookup table.

```js
const actual = normalizeDirectory(rawData);
const expected = Object.fromEntries(
  Object.values(rawData).flatMap((dept) =>
    dept.programs.flatMap((program) =>
      program.instructors
        .filter((instructor) => instructor.email)
        .map((instructor) => [instructor.email, instructor.id]),
    ),
  ),
);
assert.deepEqual(actual?.instructorsByEmail, expected);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

--fcc-editable-region--

--fcc-editable-region--
}

normalizeDirectory(rawData);

```
