---
id: 6a7848aad918b10afae7d1f0
title: Step 20
challengeType: 1
dashedName: step-20
---

# --description--

Now store each program in `programsById`, keyed by its own id.

The stored program should have its original properties plus a new `departmentId` property linking it to its department. Object spread syntax (`...`) copies an object's own enumerable properties into a new object, where you can add extra properties.

Start by removing the `console.log(program)` line from your code.

Then, assign `programsById[program.id]` an object that spreads `program` and adds `departmentId`.

Lastly, log `programsById[program.id]` to the console to inspect the stored object.

The console will show each stored program with its original properties and the added `departmentId`.

# --hints--

You should delete `console.log(program)`.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*program\s*\)/
);
```

You should assign an object to `programsById[program.id]`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programsById\s*\[\s*program\.id\s*\]\s*=\s*{/
);
```

The object should spread `program`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programsById\s*\[\s*program\.id\s*\]\s*=\s*{\s*\.\.\.\s*program/
);
```

The object should include `departmentId`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programsById\s*\[\s*program\.id\s*\]\s*=\s*{\s*\.\.\.\s*program\s*,\s*departmentId/
);
```

You should call `console.log()` with `programsById[program.id]` as the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*programsById\s*\[\s*program\.id\s*\]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

--fcc-editable-region--
      console.log(program);

--fcc-editable-region--
    }
  }
}

normalizeDirectory(rawData);
```
