---
id: 6a7840732007006f073005d8
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now, test the function with the emails `"alejandro@example.com"` and `"sofia@example.com"`. Log each result as you did for `"fatima@example.com"`.

The console will show the three instructor names returned by the function: `Fatima`, `Alejandro`, and `Sofia`.

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, "log");
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should have a second `console.log()`.

```js
assert.isAtLeast(spy.calls.length, 2);
```

You should call the `getInstructorByEmail` function with `'alejandro@example.com'` for the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /getInstructorByEmail\s*\(\s*["']alejandro@example\.com["']\s*\)/
);
```

Your `getInstructorByEmail("alejandro@example.com")` function call should be inside the second `console.log()`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*getInstructorByEmail\s*\(\s*(['"])alejandro@example\.com\1\s*\)\s*\)/
);
```

You should have a third `console.log()`.

```js
assert.isAtLeast(spy.calls.length, 3);
```

You should call the `getInstructorByEmail` function with `'sofia@example.com'` for the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /getInstructorByEmail\s*\(\s*["']sofia@example\.com["']\s*\)/
);
```

Your `getInstructorByEmail("sofia@example.com")` function call should be inside the third `console.log()`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*getInstructorByEmail\s*\(\s*(['"])sofia@example\.com\1\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      for (const instructor of program['instructors']) {
        if (instructor.email === email) {
          return instructor.name; 
        }
      }
    }
  }
  return null;
}

console.log(getInstructorByEmail("fatima@example.com"));
--fcc-editable-region--

--fcc-editable-region--
```
