---
id: 6a78513949298537bfa222b6
title: Step 23
challengeType: 1
dashedName: step-23
---

# --description--

In this step, you will store each `instructor` in `instructorsById`, keyed by its own id. 
Just like you did for `programs`, assign `instructorsById[instructor.id]` an object that 
spreads `instructor` and adds `programId` (set it to `program.id`) and `departmentId`.

Before assigning the object, delete `console.log(instructor)` from the previous step.

Log `instructorsById[instructor.id]` to the console to inspect the stored object.

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, 'log');
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should delete `console.log(instructor);`.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*instructor\s*\)/,
);
```

You should assign an object to `instructorsById[instructor.id]`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /instructorsById\s*\[\s*instructor\.id\s*\]\s*=\s*\{/,
);
```

The object should spread `instructor`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /instructorsById\s*\[\s*instructor\.id\s*\]\s*=\s*\{[^}]*\.\.\.\s*instructor\b/,
);
```

You should log each stored instructor with its original properties, `programId` set to `program.id`, and the correct `departmentId`.

```js
const logged = spy.calls.flat();
for (const [name, department] of Object.entries(rawData)) {
  for (const program of department.programs) {
    for (const instructor of program.instructors) {
      assert.deepInclude(logged, {
        ...instructor,
        programId: program.id,
        departmentId: department.Dept_code,
      });
    }
  }
}
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

--fcc-editable-region--
        console.log(instructor);

--fcc-editable-region--
      }

    }
  }
}

normalizeDirectory(rawData);
```
