---
id: 6a86e8c4888db27d84cbd157
title: Step 38
challengeType: 1
dashedName: step-38
---

# --description--

Now, to find the rooms that belong to the requested building, write a `for` loop over `programValues`, 
from `i = 0` up to `programValues.length`.

# --hints--

You should use a `for` loop.

```js
const explorer = await __helpers.Explorer(code);
const fn = explorer.functions.listRoomsByBuilding;

const fnSource = fn.toString();

assert.match(
  __helpers.removeJSComments(fnSource),
  /for\s*\(\s*let\s+i\s*=\s*0\s*;/
);
```

Your `for` loop should run while `i` is less than `programValues.length`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*programValues\.length\s*;/
);
```

Your `for` loop should increment `i` by `1` on each iteration.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*programValues\.length\s*;\s*(?:i\s*\+\+|\+\+\s*i|i\s*\+=\s*1|i\s*=\s*i\s*\+\s*1|i\s*=\s*1\s*\+\s*i)\s*\)/,
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};


function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

  return {
    departmentsById,
    programsById,
    instructorsById,
    instructorsByEmail,
  };
}

const normalizedData = normalizeDirectory(rawData);


function getInstructorByEmail(email, normalizedData) {
  const instructorId = normalizedData.instructorsByEmail[email];

  if (!instructorId) {
    return "Instructor not found";
  }

  const instructor = normalizedData.instructorsById[instructorId];
  const department = normalizedData.departmentsById[instructor.departmentId];

  return {
    name: instructor.name ?? "Unknown Instructor",
    officeHours: instructor.officeHours ?? "Not Available",
    department: department.name ?? "Unknown Department",
  };
}


function listRoomsByBuilding(buildingCode, normalizedData) {

  const programValues = Object.values(normalizedData.programsById);
  const rooms = [];

--fcc-editable-region--
  
--fcc-editable-region--
}
```
