---
id: 68e420f3d4bc7ae590dc7957
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Excellent! You can see how the `order` property rearranged the sections. Now let's use `flex-grow` to make the main content area expand to fill available space, and add responsive behavior for larger screens.

First, let's make the main content grow to fill the available vertical space by giving it `flex-grow: 1`. This will push the footer to the bottom of the viewport.

Then, add a media query for screens `768px` and wider where we:
1. Change the container to `flex-direction: row` for a horizontal layout
2. Reorder the sections using `order` properties:
   - Header: `order: 1` (spans full width at top)
   - Sidebar: `order: 2` (left side)
   - Main Content: `order: 3` (right side, grows to fill space)
   - Footer: `order: 4` (spans full width at bottom)

For the larger screen layout, also set:
- `.dashboard-container` to `flex-wrap: wrap`
- `.dashboard-header` and `.dashboard-footer` to `flex-basis: 100%` (full width)
- `.sidebar` to `flex-basis: 250px` (fixed width)
- `.main-content` to `flex-grow: 1` (fill remaining space)

# --hints--

Your `.main-content` should have `flex-grow: 1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.main-content')?.flexGrow, '1');
```

You should have a media query for `(min-width: 768px)`.

```js
const cssRules = [...document.styleSheets[0].cssRules];
const mediaQuery = cssRules.find(rule => 
  rule.type === CSSRule.MEDIA_RULE && 
  rule.conditionText === '(min-width: 768px)'
);
assert.exists(mediaQuery);
```

Inside the media query, `.dashboard-container` should have `flex-direction: row`.

```js
const cssRules = [...document.styleSheets[0].cssRules];
const mediaQuery = cssRules.find(rule => 
  rule.type === CSSRule.MEDIA_RULE && 
  rule.conditionText === '(min-width: 768px)'
);
const containerRule = [...mediaQuery.cssRules].find(rule => 
  rule.selectorText === '.dashboard-container'
);
assert.include(containerRule?.cssText || '', 'flex-direction: row');
```

Inside the media query, `.dashboard-container` should have `flex-wrap: wrap`.

```js
const cssRules = [...document.styleSheets[0].cssRules];
const mediaQuery = cssRules.find(rule => 
  rule.type === CSSRule.MEDIA_RULE && 
  rule.conditionText === '(min-width: 768px)'
);
const containerRule = [...mediaQuery.cssRules].find(rule => 
  rule.selectorText === '.dashboard-container'
);
assert.include(containerRule?.cssText || '', 'flex-wrap: wrap');
```

Inside the media query, `.sidebar` should have `order: 2`.

```js
const cssRules = [...document.styleSheets[0].cssRules];
const mediaQuery = cssRules.find(rule => 
  rule.type === CSSRule.MEDIA_RULE && 
  rule.conditionText === '(min-width: 768px)'
);
const sidebarRule = [...mediaQuery.cssRules].find(rule => 
  rule.selectorText === '.sidebar'
);
assert.include(sidebarRule?.cssText || '', 'order: 2');
```

Inside the media query, `.main-content` should have `order: 3`.

```js
const cssRules = [...document.styleSheets[0].cssRules];
const mediaQuery = cssRules.find(rule => 
  rule.type === CSSRule.MEDIA_RULE && 
  rule.conditionText === '(min-width: 768px)'
);
const mainRule = [...mediaQuery.cssRules].find(rule => 
  rule.selectorText === '.main-content'
);
assert.include(mainRule?.cssText || '', 'order: 3');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Advanced Flexbox Dashboard</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="dashboard-container">
      <header class="dashboard-header">
        <h1>Dashboard</h1>
      </header>
      <nav class="sidebar">
        <h2>Navigation</h2>
      </nav>
      <main class="main-content">
        <h2>Main Content</h2>
      </main>
      <footer class="dashboard-footer">
        <p>© 2025 Dashboard</p>
      </footer>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  line-height: 1.6;
  color: #333;
  background: #f4f6f9;
}

.dashboard-container {
  display: flex;
  flex-direction: column;
  min-height: 100vh;
}

.dashboard-header {
  order: 1;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #2c3e50;
  color: white;
}

.main-content {
  order: 2;
  padding: 1rem;
  border: 2px solid #ddd;
  background: white;
}

.sidebar {
  order: 3;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #34495e;
  color: white;
}

.dashboard-footer {
  order: 4;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #7f8c8d;
  color: white;
}

--fcc-editable-region--

--fcc-editable-region--
```
