---
id: 68e420f3d4bc7ae590dc7955
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Great! Now you have the basic structure. Let's make the dashboard container a flexbox and set up the initial layout.

The dashboard should be a flexible container that arranges its children (header, sidebar, main, footer) in a column by default, but we'll later use the `order` property to rearrange them for different screen sizes.

First, add some content to your sections to see the layout in action:

1. Inside the `header`, add an `h1` with the text "Dashboard"
2. Inside the `sidebar`, add an `h2` with the text "Navigation" 
3. Inside the `main` element, add an `h2` with the text "Main Content"
4. Inside the `footer`, add a `p` with the text "© 2025 Dashboard"

Then, make the `.dashboard-container` a flex container with `display: flex` and `flex-direction: column`. Also add a `min-height: 100vh` so it fills the viewport.

# --hints--

Your `header` should contain an `h1` element with the text "Dashboard".

```js
const header = document.querySelector('header.dashboard-header');
const h1 = header?.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'Dashboard');
```

Your `sidebar` should contain an `h2` element with the text "Navigation".

```js
const sidebar = document.querySelector('nav.sidebar');
const h2 = sidebar?.querySelector('h2');
assert.strictEqual(h2?.textContent?.trim(), 'Navigation');
```

Your `main` should contain an `h2` element with the text "Main Content".

```js
const main = document.querySelector('main.main-content');
const h2 = main?.querySelector('h2');
assert.strictEqual(h2?.textContent?.trim(), 'Main Content');
```

Your `footer` should contain a `p` element with the text "© 2025 Dashboard".

```js
const footer = document.querySelector('footer.dashboard-footer');
const p = footer?.querySelector('p');
assert.strictEqual(p?.textContent?.trim(), '© 2025 Dashboard');
```

Your `.dashboard-container` should have `display: flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-container')?.display, 'flex');
```

Your `.dashboard-container` should have `flex-direction: column`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-container')?.flexDirection, 'column');
```

Your `.dashboard-container` should have `min-height: 100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-container')?.minHeight, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Advanced Flexbox Dashboard</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="dashboard-container">
      --fcc-editable-region--
      <header class="dashboard-header">
        
      </header>
      <nav class="sidebar">
        
      </nav>
      <main class="main-content">
        
      </main>
      <footer class="dashboard-footer">
        
      </footer>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  line-height: 1.6;
  color: #333;
  background: #f4f6f9;
}

--fcc-editable-region--

--fcc-editable-region--
```
