---
id: 68e420f3d4bc7ae590dc7956
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Perfect! Now you can see the basic column layout. Let's use the `order` property to create a mobile-first responsive design where elements reorder themselves at different screen sizes.

On mobile devices, we want this order:
1. Header (order: 1)
2. Main Content (order: 2)  
3. Sidebar/Navigation (order: 3)
4. Footer (order: 4)

But on larger screens, we'll rearrange to put the sidebar before main content for a traditional dashboard layout.

First, let's set up the mobile-first ordering by adding `order` properties to each section:

- `.dashboard-header` should have `order: 1`
- `.main-content` should have `order: 2`  
- `.sidebar` should have `order: 3`
- `.dashboard-footer` should have `order: 4`

Also add some basic styling to make the sections visible:
- Give each section `padding: 1rem` and `border: 2px solid #ddd`
- Give `.dashboard-header` a `background: #2c3e50` and `color: white`
- Give `.sidebar` a `background: #34495e` and `color: white`
- Give `.main-content` a `background: white`
- Give `.dashboard-footer` a `background: #7f8c8d` and `color: white`

# --hints--

Your `.dashboard-header` should have `order: 1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-header')?.order, '1');
```

Your `.main-content` should have `order: 2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.main-content')?.order, '2');
```

Your `.sidebar` should have `order: 3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sidebar')?.order, '3');
```

Your `.dashboard-footer` should have `order: 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-footer')?.order, '4');
```

All sections should have `padding: 1rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-header')?.padding, '16px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.sidebar')?.padding, '16px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.main-content')?.padding, '16px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-footer')?.padding, '16px');
```

Your `.dashboard-header` should have a dark background and white text.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-header')?.backgroundColor, 'rgb(44, 62, 80)');
assert.equal(new __helpers.CSSHelp(document).getStyle('.dashboard-header')?.color, 'white');
```

Your `.sidebar` should have a dark background and white text.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sidebar')?.backgroundColor, 'rgb(52, 73, 94)');
assert.equal(new __helpers.CSSHelp(document).getStyle('.sidebar')?.color, 'white');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Advanced Flexbox Dashboard</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="dashboard-container">
      <header class="dashboard-header">
        <h1>Dashboard</h1>
      </header>
      <nav class="sidebar">
        <h2>Navigation</h2>
      </nav>
      <main class="main-content">
        <h2>Main Content</h2>
      </main>
      <footer class="dashboard-footer">
        <p>© 2025 Dashboard</p>
      </footer>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  line-height: 1.6;
  color: #333;
  background: #f4f6f9;
}

.dashboard-container {
  display: flex;
  flex-direction: column;
  min-height: 100vh;
}

--fcc-editable-region--

--fcc-editable-region--
```
