---
id: 68e420f3d4bc7ae590dc7958
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Great! Now you have a responsive dashboard that reorders itself. Let's add widget cards to the main content area and use `align-content` to control how multiple lines of cards are distributed.

This step will demonstrate the power of `align-content` for multi-line flex layouts - exactly like your pricing plans lab!

First, let's add some widget cards to the main content area. Replace the `<h2>Main Content</h2>` in your HTML with:

```html
<h2>Dashboard Widgets</h2>
<div class="widgets-grid">
  <div class="widget">
    <h3>Sales</h3>
    <p>$12,450</p>
  </div>
  <div class="widget">
    <h3>Users</h3>
    <p>1,240</p>
  </div>
  <div class="widget">
    <h3>Orders</h3>
    <p>856</p>
  </div>
  <div class="widget">
    <h3>Revenue</h3>
    <p>$8,930</p>
  </div>
  <div class="widget">
    <h3>Conversion</h3>
    <p>4.2%</p>
  </div>
  <div class="widget">
    <h3>Traffic</h3>
    <p>15,678</p>
  </div>
</div>
```

Then create CSS for the widgets grid:
- `.widgets-grid` should be a flex container with `display: flex`, `flex-wrap: wrap`, and `align-content: space-around`
- Set a minimum height with `min-height: 300px` so you can see the `align-content` effect
- Each `.widget` should have `flex: 1 1 200px` (grow, shrink, basis), `margin: 10px`, `padding: 1rem`, `background: #fff`, `border-radius: 8px`, and `box-shadow: 0 2px 4px rgba(0,0,0,0.1)`

# --hints--

Your main content should contain an element with class `widgets-grid`.

```js
assert.exists(document.querySelector('.main-content .widgets-grid'));
```

Your `.widgets-grid` should contain 6 `.widget` elements.

```js
const widgets = document.querySelectorAll('.widgets-grid .widget');
assert.equal(widgets.length, 6);
```

Your `.widgets-grid` should have `display: flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widgets-grid')?.display, 'flex');
```

Your `.widgets-grid` should have `flex-wrap: wrap`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widgets-grid')?.flexWrap, 'wrap');
```

Your `.widgets-grid` should have `align-content: space-around`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widgets-grid')?.alignContent, 'space-around');
```

Your `.widgets-grid` should have `min-height: 300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widgets-grid')?.minHeight, '300px');
```

Your `.widget` elements should have `flex: 1 1 200px`.

```js
const widgetStyle = new __helpers.CSSHelp(document).getStyle('.widget');
assert.equal(widgetStyle?.flex, '1 1 200px');
```

Your `.widget` elements should have a white background.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widget')?.backgroundColor, 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Advanced Flexbox Dashboard</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="dashboard-container">
      <header class="dashboard-header">
        <h1>Dashboard</h1>
      </header>
      <nav class="sidebar">
        <h2>Navigation</h2>
      </nav>
      <main class="main-content">
        --fcc-editable-region--
        <h2>Main Content</h2>
        --fcc-editable-region--
      </main>
      <footer class="dashboard-footer">
        <p>© 2025 Dashboard</p>
      </footer>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  line-height: 1.6;
  color: #333;
  background: #f4f6f9;
}

.dashboard-container {
  display: flex;
  flex-direction: column;
  min-height: 100vh;
}

.dashboard-header {
  order: 1;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #2c3e50;
  color: white;
}

.main-content {
  order: 2;
  padding: 1rem;
  border: 2px solid #ddd;
  background: white;
  flex-grow: 1;
}

.sidebar {
  order: 3;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #34495e;
  color: white;
}

.dashboard-footer {
  order: 4;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #7f8c8d;
  color: white;
}

@media (min-width: 768px) {
  .dashboard-container {
    flex-direction: row;
    flex-wrap: wrap;
  }
  
  .dashboard-header {
    order: 1;
    flex-basis: 100%;
  }
  
  .sidebar {
    order: 2;
    flex-basis: 250px;
  }
  
  .main-content {
    order: 3;
    flex-grow: 1;
  }
  
  .dashboard-footer {
    order: 4;
    flex-basis: 100%;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
