---
id: 68e420f3d4bc7ae590dc7959
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Perfect! You now have a responsive dashboard with widget cards that wrap to multiple lines. Let's explore different `align-content` values to see how they affect the distribution of your widget rows.

Currently your widgets use `align-content: space-around`. Let's experiment with different values to understand their effects:

1. First, increase the `min-height` of `.widgets-grid` to `400px` to give more space for the align-content effect
2. Change `align-content` from `space-around` to `space-between`
3. Add some navigation links to the sidebar to make it more realistic

Add this content inside your `.sidebar` (replace the existing h2):

```html
<h2>Navigation</h2>
<ul class="nav-menu">
  <li><a href="#">Dashboard</a></li>
  <li><a href="#">Analytics</a></li>
  <li><a href="#">Reports</a></li>
  <li><a href="#">Settings</a></li>
</ul>
```

Then style the navigation with:
- `.nav-menu` should have `list-style: none`
- `.nav-menu li` should have `margin: 0.5rem 0`
- `.nav-menu a` should have `color: #ecf0f1`, `text-decoration: none`, and `padding: 0.5rem 0`
- `.nav-menu a:hover` should have `color: #3498db`

This completes your advanced flexbox dashboard that demonstrates `order`, `flex-grow`, and `align-content` - all the properties covered in your pricing plans lab!

# --hints--

Your `.widgets-grid` should have `min-height: 400px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widgets-grid')?.minHeight, '400px');
```

Your `.widgets-grid` should have `align-content: space-between`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.widgets-grid')?.alignContent, 'space-between');
```

Your sidebar should contain a `ul` element with class `nav-menu`.

```js
assert.exists(document.querySelector('.sidebar .nav-menu'));
```

Your `.nav-menu` should contain 4 `li` elements.

```js
const navItems = document.querySelectorAll('.nav-menu li');
assert.equal(navItems.length, 4);
```

Your `.nav-menu` should have `list-style: none`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.nav-menu')?.listStyle, 'none');
```

Your `.nav-menu a` elements should have `color: rgb(236, 240, 241)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.nav-menu a')?.color, 'rgb(236, 240, 241)');
```

Your `.nav-menu a` elements should have `text-decoration: none`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.nav-menu a')?.textDecoration, 'none solid rgb(236, 240, 241)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Advanced Flexbox Dashboard</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="dashboard-container">
      <header class="dashboard-header">
        <h1>Dashboard</h1>
      </header>
      <nav class="sidebar">
        --fcc-editable-region--
        <h2>Navigation</h2>
        --fcc-editable-region--
      </nav>
      <main class="main-content">
        <h2>Dashboard Widgets</h2>
        <div class="widgets-grid">
          <div class="widget">
            <h3>Sales</h3>
            <p>$12,450</p>
          </div>
          <div class="widget">
            <h3>Users</h3>
            <p>1,240</p>
          </div>
          <div class="widget">
            <h3>Orders</h3>
            <p>856</p>
          </div>
          <div class="widget">
            <h3>Revenue</h3>
            <p>$8,930</p>
          </div>
          <div class="widget">
            <h3>Conversion</h3>
            <p>4.2%</p>
          </div>
          <div class="widget">
            <h3>Traffic</h3>
            <p>15,678</p>
          </div>
        </div>
      </main>
      <footer class="dashboard-footer">
        <p>© 2025 Dashboard</p>
      </footer>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  line-height: 1.6;
  color: #333;
  background: #f4f6f9;
}

.dashboard-container {
  display: flex;
  flex-direction: column;
  min-height: 100vh;
}

.dashboard-header {
  order: 1;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #2c3e50;
  color: white;
}

.main-content {
  order: 2;
  padding: 1rem;
  border: 2px solid #ddd;
  background: white;
  flex-grow: 1;
}

.sidebar {
  order: 3;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #34495e;
  color: white;
}

.dashboard-footer {
  order: 4;
  padding: 1rem;
  border: 2px solid #ddd;
  background: #7f8c8d;
  color: white;
}

.widgets-grid {
  display: flex;
  flex-wrap: wrap;
  align-content: space-around;
  min-height: 300px;
}

.widget {
  flex: 1 1 200px;
  margin: 10px;
  padding: 1rem;
  background: #fff;
  border-radius: 8px;
  box-shadow: 0 2px 4px rgba(0,0,0,0.1);
}

@media (min-width: 768px) {
  .dashboard-container {
    flex-direction: row;
    flex-wrap: wrap;
  }
  
  .dashboard-header {
    order: 1;
    flex-basis: 100%;
  }
  
  .sidebar {
    order: 2;
    flex-basis: 250px;
  }
  
  .main-content {
    order: 3;
    flex-grow: 1;
  }
  
  .dashboard-footer {
    order: 4;
    flex-basis: 100%;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
