---
id: 68e420f3d4bc7ae590dc7954
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Welcome to building an advanced responsive dashboard! In this workshop, you'll create a professional dashboard layout that demonstrates advanced flexbox properties including `order`, `flex-grow`, and `align-content`.

Your dashboard will feature:

- A responsive header with navigation
- A flexible sidebar that reorders on mobile
- A main content area with widget cards
- Multi-line card layouts that adapt to screen size

Let's start by creating the basic HTML structure. You'll build a dashboard with a header, sidebar, main content area, and footer.

Create the basic HTML structure with a `div` element that has a `class` of `dashboard-container`. Inside this container, add four sections:

1. A `header` element with `class="dashboard-header"`
2. A `nav` element with `class="sidebar"`  
3. A `main` element with `class="main-content"`
4. A `footer` element with `class="dashboard-footer"`

# --hints--

You should have a `div` element with a `class` of `dashboard-container`.

```js
assert.exists(document.querySelector('div.dashboard-container'));
```

Your `dashboard-container` should contain a `header` element with `class="dashboard-header"`.

```js
assert.exists(document.querySelector('.dashboard-container header.dashboard-header'));
```

Your `dashboard-container` should contain a `nav` element with `class="sidebar"`.

```js
assert.exists(document.querySelector('.dashboard-container nav.sidebar'));
```

Your `dashboard-container` should contain a `main` element with `class="main-content"`.

```js
assert.exists(document.querySelector('.dashboard-container main.main-content'));
```

Your `dashboard-container` should contain a `footer` element with `class="dashboard-footer"`.

```js
assert.exists(document.querySelector('.dashboard-container footer.dashboard-footer'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Advanced Flexbox Dashboard</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  line-height: 1.6;
  color: #333;
  background: #f4f6f9;
}
```
