---
id: 6626b4c58c027d86478ff5eb
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

The teacher is really happy with the program you have created so far. But now they want to have an easy way to check if a student has a passing grade. A passing grade is anything that is not an `"F"`.

Complete the function `hasPassingGrade` that takes a student score as a parameter and returns `true` if the student has a passing grade and `false` if they do not.

A few function calls have been provided for you to test your function as you build it out.

**Tips**

- Use the `getGrade` function to get the student's grade. Then check if the grade is passing or not.

# --hints--

Your `hasPassingGrade` function should return `true` if the grade is `"A"`.

```js
assert(hasPassingGrade(100) === true);
```

Your `hasPassingGrade` function should return `false` if the grade is `"F"`.

```js
assert(hasPassingGrade(53) === false);
```

Your `hasPassingGrade` function should return `true` for all passing grades.

```js
assert(hasPassingGrade(87) === true);
assert(hasPassingGrade(99) === true);
assert(hasPassingGrade(66) === true);
assert(hasPassingGrade(73) === true);
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (let i = 0; i < scores.length; i++) {
    sum += scores[i];
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90 && score < 100) {
    return "A";
  } else if (score >= 80 && score < 90) {
    return "B";
  } else if (score >= 70 && score < 80) {
    return "C";
  } else if (score >= 60 && score < 70) {
    return "D";
  } else {
    return "F";
  }
}
--fcc-editable-region--
function hasPassingGrade(score) {
  
}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
