---
id: 6626a060c4006f793e10cb33
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now the teacher needs your help converting the student score to a letter grade. 

Your task is to create a `getGrade` function that takes a `score` as a parameter and returns a string representing a letter grade corresponding to the score.

Here are the scores and their corresponding letter grades:

| Score Range | Grade   |
| ----------- | ------- |
| `100`       | `"A++"` |
| `90 - 99`   | `"A"`   |
| `80 - 89`   | `"B"`   |
| `70 - 79`   | `"C"`   |
| `60 - 69`   | `"D"`   |
| `0 - 59`    | `"F"`   |

For example, if the student's score is `88`, the function should return `"B"`.
If the student's score is `57`, the function should return `"F"`.

A few function calls have been provided for you to test your function as you build it out.

**Tips**

- In the previous project, you learned about conditional statements(`if`, `else if`, and `else`). You can use these statements to check the score and return the corresponding grade.
- You also learned about comparison operators (`>`, `<`, `>=`, `<=`, `==`, `===`, `!=`, `!==`). You can use these operators to compare the score with the score ranges and return the corresponding grade.

# --hints--

Your `getGrade` function should return `"A++"` if the score is `100`.

```js
assert.strictEqual(getGrade(100), "A++");
```

Your `getGrade` function should return `"A"` if the score is `94`.

```js
assert.strictEqual(getGrade(94), "A");
```

Your `getGrade` function should return `"B"` if the score is between `80` and `89`.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Your `getGrade` function should return `"C"` if the score is `78`.

```js
assert.strictEqual(getGrade(75), "C");
```

Your `getGrade` function should return `"D"` if the score is between `60` and `69`.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Your `getGrade` function should return `"F"` if the score is `57`.

```js
assert.strictEqual(getGrade(57), "F");
```

Your `getGrade` function should return `"F"` if the score is between `0` and `59`.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (let i = 0; i < scores.length; i++) {
    sum += scores[i];
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
