---
id: 662693f82c91a66be46c881b
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this mini project, you will review JavaScript fundamentals by building a gradebook app. 

A teacher has finished grading their students' tests and needs your help to calculate the average score for the class.

Complete the `getAverage` function which takes in an array of test scores and returns the average score.

The average is calculated by adding up all the scores and dividing by the total number of scores.

```js
average = sum of all scores / total number of scores
```

A couple of test calls have been provided for you so you can test out your function as you build it out.

**Tips**

- You learned about loops in the previous project. Use a loop to iterate over the `scores` array and add up all the scores.
- You also learned about the `length` property. You can use that for your total number of scores.
- Test your code often and break up the problem into smaller parts to make it easier to solve.

# --hints--

Your `getAverage` function should return a number.

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` should return `71.4`.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.4);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` should return `85.6`.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.6);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` should return `93.4`.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 93.4);
```

Your `getAverage` function should return the average score of the test scores.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 54);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
