---
id: 6626b8dcf5057f896f948440
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now that the teacher has all of the information they need, they want to be able to message the student with the results. 

Complete the `studentMsg` function that takes an array of total scores and a student score as parameters. The function should return a string representing a message to the student.

If the student passed the course, the string should follow this format:

```js
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

If the student failed the course, the string should follow this format:

```js
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

The `average-goes-here` should be the average of the total scores. The `grade-goes-here` should be the student's grade.

**Tips**

- Use the `getAverage` function to calculate the average of the total scores.
- Use the `getGrade` function to get the student's grade.
- Use string concatenation to build the message. Remember that you can use the `+` operator to concatenate strings. 

# --hints--

Your function call of `studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` should return the following message: `"Class average: 67.7. Your grade: F. You failed the course."`:

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 67.7. Your grade: F. You failed the course.");
```

Your function call of `studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 100)` should return the following message: `"Class average: 67.7. Your grade: A++. You passed the course."`:

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 100), "Class average: 67.7. Your grade: A++. You passed the course.");
```

Your `studentMsg` function should return a the correct message based on the student's score and the class average.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.125. Your grade: A++. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.125. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (let i = 0; i < scores.length; i++) {
    sum += scores[i];
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90 && score < 100) {
    return "A";
  } else if (score >= 80 && score < 90) {
    return "B";
  } else if (score >= 70 && score < 80) {
    return "C";
  } else if (score >= 60 && score < 70) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
--fcc-editable-region--

console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (let i = 0; i < scores.length; i++) {
    sum += scores[i];
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90 && score < 100) {
    return "A";
  } else if (score >= 80 && score < 90) {
    return "B";
  } else if (score >= 70 && score < 80) {
    return "C";
  } else if (score >= 60 && score < 70) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(grade) ? "passed" : "failed"
  } the course.`;
}

```
