---
id: 66faa80000000000000000ab
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
Time to give the player input. Tapping a gem selects it (scale 1.15 pop). Tapping a second gem either deselects (same gem), rejects (non-adjacent), or requests a swap (cardinal adjacent). Diagonal taps don't count - only horizontal/vertical neighbors.

After `/* create-input */`, add:

```js
this.busy = false;
this.selected = null;
for (let r = 0; r < this.rows; r++) {
  for (let c = 0; c < this.cols; c++) {
    this.gems[r][c].setInteractive();
  }
}
this.input.on("gameobjectdown", (pointer, gem) => {
  if (this.busy) return;
  if (!this.selected) {
    this.selected = gem;
    this.tweens.add({
      targets: gem,
      scale: 1.15,
      duration: 120,
      ease: "Cubic.easeOut"
    });
    return;
  }
  this.tweens.add({ targets: this.selected, scale: 1, duration: 120 });
  const prev = this.selected;
  this.selected = null;
  if (gem === prev) return;
  const dr = Math.abs(gem.getData("row") - prev.getData("row"));
  const dc = Math.abs(gem.getData("col") - prev.getData("col"));
  if (dr + dc !== 1) return;
  const a = { row: prev.getData("row"), col: prev.getData("col") };
  const b = { row: gem.getData("row"), col: gem.getData("col") };
  this.swap(a, b);
});
```

`this.swap` doesn't exist yet - we'll define it next step. JavaScript event handlers don't fire until after `create()` returns, so this works as long as `this.swap` is defined *somewhere* in `create()` before the first pointer event arrives.

# --hints--
Initialize `this.busy = false` and `this.selected = null`.

```js
assert.match(code, /this\.busy\s*=\s*false[\s\S]*this\.selected\s*=\s*null/);
```

Loop `gem.setInteractive()` over every grid cell.

```js
assert.match(code, /this\.gems\[r\]\[c\]\.setInteractive\(\s*\)/);
```

Listen for `"gameobjectdown"` on `this.input`.

```js
assert.match(code, /this\.input\.on\(\s*["\']gameobjectdown["\']\s*,\s*\(\s*pointer\s*,\s*gem\s*\)\s*=>/);
```

Pop the first selection with `tweens.add({ targets: gem, scale: 1.15 })`.

```js
assert.match(code, /targets:\s*gem\s*,\s*scale:\s*1\.15/);
```

Reject non-adjacent taps with `if (dr + dc !== 1) return;`.

```js
assert.match(code, /if\s*\(\s*dr\s*\+\s*dc\s*!==\s*1\s*\)\s*return/);
```

Call `this.swap(a, b)` on cardinal-adjacent taps.

```js
assert.match(code, /this\.swap\(\s*a\s*,\s*b\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillRoundedRect(0, 0, 48, 48, 10);
    g.lineStyle(2, 0xffffff, 0.5);
    g.strokeRoundedRect(0, 0, 48, 48, 10);
    g.generateTexture("gem", 48, 48);
    g.destroy();
    this.gems = [];
    this.gemTypes = [];
    for (let r = 0; r < this.rows; r++) {
      this.gems[r] = [];
      this.gemTypes[r] = [];
      for (let c = 0; c < this.cols; c++) {
        let type;
        do {
          type = Phaser.Math.Between(0, this.gemColors.length - 1);
        } while (
          (c >= 2 &&
            this.gemTypes[r][c - 1] === type &&
            this.gemTypes[r][c - 2] === type) ||
          (r >= 2 &&
            this.gemTypes[r - 1][c] === type &&
            this.gemTypes[r - 2][c] === type)
        );
        this.gemTypes[r][c] = type;
        const sprite = this.add
          .sprite(this.gemX(c), this.gemY(r), "gem")
          .setTint(this.gemColors[type])
          .setData({ row: r, col: c, type, special: null });
        this.gems[r][c] = sprite;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene]
};

new Phaser.Game(config);
```
