---
id: 66faa80000000000000000a7
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the match-3 puzzle workshop. Across 50 steps you'll build a complete match-3: an 8x8 grid of gem sprites, swap-and-cascade mechanics, three special-gem rules (stripe, color-bomb, wrapped), score, levels, time-attack mode, audio cues, post-FX bloom and vignette, a tutorial overlay, an idle hint system, mode select, pause menu, and `prefers-reduced-motion` damping.

No arcade physics this time - everything is tween-driven. You'll see why `tweens.chain` is the most important tool in your kit by the end of this workshop.

Start with a `MainScene` extending `Phaser.Scene` with a constructor that calls `super("MainScene")`, an empty `preload`, an empty `create`, and an `update(time, delta)`.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {}

  create() {}

  update(time, delta) {}
}
```

The 800x600 `config` is already wired in below.

# --hints--
Define `class MainScene extends Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

Call `super("MainScene")` from the constructor.

```js
assert.match(code, /super\(\s*["\']MainScene["\']\s*\)/);
```

Define empty `preload()` and `create()` methods.

```js
assert.match(code, /preload\s*\(\s*\)\s*\{[\s\S]*?\}\s*[\s\S]*create\s*\(\s*\)\s*\{[\s\S]*?\}/);
```

Define `update(time, delta)`.

```js
assert.match(code, /update\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene]
};

new Phaser.Game(config);
```
