---
id: 66faa80000000000000000ba
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--
When a color-bomb activates, it wipes EVERY gem of its own color from the board. Modify `activateSpecial` to add a `color-bomb` branch that scans the entire grid for matching type and adds them all to `targets`.

Insert into `activateSpecial`, after the `stripe-v` branch:

```js
} else if (special === "color-bomb") {
  const colorTarget = gem.getData("type");
  for (let r = 0; r < this.rows; r++) {
    for (let c = 0; c < this.cols; c++) {
      if (this.gemTypes[r][c] === colorTarget && this.gems[r][c]) {
        targets.push({ row: r, col: c, type: colorTarget });
      }
    }
  }
}
```

A color-bomb on a 5-color board with even distribution clears ~13 gems per activation, scoring `13 * 10 * combo`. Combine with a stripe (chain via `targets.forEach`) and you can clear 30+ gems in one chain.

# --hints--
Add an `else if (special === "color-bomb")` branch.

```js
assert.match(code, /else\s+if\s*\(\s*special\s*===\s*["\']color-bomb["\']\s*\)/);
```

Read the target color via `gem.getData("type")`.

```js
assert.match(code, /const\s+colorTarget\s*=\s*gem\.getData\(/);
```

Iterate every grid cell and push matching-type cells.

```js
assert.match(code, /if\s*\(\s*this\.gemTypes\[r\]\[c\]\s*===\s*colorTarget\s*&&\s*this\.gems\[r\]\[c\]\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillRoundedRect(0, 0, 48, 48, 10);
    g.lineStyle(2, 0xffffff, 0.5);
    g.strokeRoundedRect(0, 0, 48, 48, 10);
    g.generateTexture("gem", 48, 48);
    g.destroy();
    this.gems = [];
    this.gemTypes = [];
    for (let r = 0; r < this.rows; r++) {
      this.gems[r] = [];
      this.gemTypes[r] = [];
      for (let c = 0; c < this.cols; c++) {
        let type;
        do {
          type = Phaser.Math.Between(0, this.gemColors.length - 1);
        } while (
          (c >= 2 &&
            this.gemTypes[r][c - 1] === type &&
            this.gemTypes[r][c - 2] === type) ||
          (r >= 2 &&
            this.gemTypes[r - 1][c] === type &&
            this.gemTypes[r - 2][c] === type)
        );
        this.gemTypes[r][c] = type;
        const sprite = this.add
          .sprite(this.gemX(c), this.gemY(r), "gem")
          .setTint(this.gemColors[type])
          .setData({ row: r, col: c, type, special: null });
        this.gems[r][c] = sprite;
      }
    }
    this.busy = false;
    this.selected = null;
    for (let r = 0; r < this.rows; r++) {
      for (let c = 0; c < this.cols; c++) {
        this.gems[r][c].setInteractive();
      }
    }
    this.input.on("gameobjectdown", (pointer, gem) => {
      if (this.busy) return;
      if (!this.selected) {
        this.selected = gem;
        this.tweens.add({
          targets: gem,
          scale: 1.15,
          duration: 120,
          ease: "Cubic.easeOut"
        });
        return;
      }
      this.tweens.add({ targets: this.selected, scale: 1, duration: 120 });
      const prev = this.selected;
      this.selected = null;
      if (gem === prev) return;
      const dr = Math.abs(gem.getData("row") - prev.getData("row"));
      const dc = Math.abs(gem.getData("col") - prev.getData("col"));
      if (dr + dc !== 1) return;
      const a = { row: prev.getData("row"), col: prev.getData("col") };
      const b = { row: gem.getData("row"), col: gem.getData("col") };
      this.swap(a, b);
    });
    this.findMatches = () => {
      const cells = new Map();
      const runs = [];
      for (let r = 0; r < this.rows; r++) {
        let runStart = 0;
        let runType = this.gemTypes[r][0];
        for (let c = 1; c <= this.cols; c++) {
          const t = c < this.cols ? this.gemTypes[r][c] : -2;
          if (t === runType && runType >= 0) continue;
          if (c - runStart >= 3 && runType >= 0) {
            runs.push({ orient: "h", row: r, col: runStart, len: c - runStart, type: runType });
            for (let k = runStart; k < c; k++) {
              cells.set(`${r},${k}`, { row: r, col: k, type: runType });
            }
          }
          runStart = c;
          runType = t;
        }
      }
      for (let c = 0; c < this.cols; c++) {
        let runStart = 0;
        let runType = this.gemTypes[0][c];
        for (let r = 1; r <= this.rows; r++) {
          const t = r < this.rows ? this.gemTypes[r][c] : -2;
          if (t === runType && runType >= 0) continue;
          if (r - runStart >= 3 && runType >= 0) {
            runs.push({ orient: "v", row: runStart, col: c, len: r - runStart, type: runType });
            for (let k = runStart; k < r; k++) {
              cells.set(`${k},${c}`, { row: k, col: c, type: runType });
            }
          }
          runStart = r;
          runType = t;
        }
      }
      return { cells: [...cells.values()], runs };
    };
--fcc-editable-region--
    this.activateSpecial = (gem, allCells) => {
      const special = gem.getData("special");
      if (!special) return;
      gem.setData({
        row: gem.getData("row"),
        col: gem.getData("col"),
        type: gem.getData("type"),
        special: null
      });
      const sr = gem.getData("row");
      const sc = gem.getData("col");
      const targets = [];
      if (special === "stripe-h") {
        for (let c = 0; c < this.cols; c++) {
          if (this.gems[sr][c]) {
            targets.push({ row: sr, col: c, type: this.gemTypes[sr][c] });
          }
        }
      } else if (special === "stripe-v") {
        for (let r = 0; r < this.rows; r++) {
          if (this.gems[r][sc]) {
            targets.push({ row: r, col: sc, type: this.gemTypes[r][sc] });
          }
        }
      }
      targets.forEach((cell) => {
        const key = `${cell.row},${cell.col}`;
        if (allCells.has(key)) return;
        allCells.set(key, cell);
        const chainGem = this.gems[cell.row][cell.col];
        if (chainGem && chainGem.getData("special")) {
          this.activateSpecial(chainGem, allCells);
        }
      });
    };
--fcc-editable-region--
    this.cascade = (matches, depth) => {
      this.score ??= 0;
      this.combo ??= 0;
      this.combo = depth + 1;
      this.events.emit("comboUpdate", this.combo);
      const allCells = new Map(
        matches.cells.map((c) => [`${c.row},${c.col}`, c])
      );
      [...allCells.values()].forEach((cell) => {
        const g2 = this.gems[cell.row][cell.col];
        if (g2 && g2.getData("special")) this.activateSpecial(g2, allCells);
      });
      this.score += allCells.size * 10 * this.combo;
      this.events.emit("scoreUpdate", this.score);
      const spared = new Set();
      matches.runs.forEach((run) => {
        if (run.len < 4) return;
        const key = `${run.row},${run.col}`;
        if (spared.has(key)) return;
        const existingGem = this.gems[run.row][run.col];
        if (existingGem && existingGem.getData("special")) return;
        spared.add(key);
        const gem = existingGem;
        if (!gem) return;
        let kind;
        if (run.len >= 5) {
          kind = "color-bomb";
        } else {
          kind = run.orient === "h" ? "stripe-h" : "stripe-v";
        }
        gem.setData({
          row: run.row,
          col: run.col,
          type: run.type,
          special: kind
        });
        this.tweens.add({
          targets: gem,
          scale: 1.4,
          duration: 160,
          yoyo: true,
          ease: "Cubic.easeOut"
        });
      });
      [...allCells.values()].forEach(({ row, col }) => {
        if (spared.has(`${row},${col}`)) return;
        const gem = this.gems[row][col];
        this.gems[row][col] = null;
        this.gemTypes[row][col] = -1;
        if (!gem) return;
        this.tweens.add({
          targets: gem,
          scale: 0,
          duration: 200,
          ease: "Cubic.easeIn",
          onComplete: () => gem.destroy()
        });
      });
      this.time.delayedCall(220, () => {
        for (let c = 0; c < this.cols; c++) {
          let writeRow = this.rows - 1;
          for (let r = this.rows - 1; r >= 0; r--) {
            if (this.gemTypes[r][c] === -1) continue;
            if (writeRow !== r) {
              this.gemTypes[writeRow][c] = this.gemTypes[r][c];
              this.gems[writeRow][c] = this.gems[r][c];
              this.gemTypes[r][c] = -1;
              this.gems[r][c] = null;
              const moved = this.gems[writeRow][c];
              moved.setData({
                row: writeRow,
                col: c,
                type: this.gemTypes[writeRow][c],
                special: moved.getData("special")
              });
              this.tweens.add({
                targets: moved,
                y: this.gemY(writeRow),
                duration: 280,
                ease: "Cubic.easeIn"
              });
            }
            writeRow--;
          }
          for (let r = writeRow; r >= 0; r--) {
            const type = Phaser.Math.Between(0, this.gemColors.length - 1);
            this.gemTypes[r][c] = type;
            const sprite = this.add
              .sprite(this.gemX(c), this.gemY(-1 - (writeRow - r)), "gem")
              .setTint(this.gemColors[type])
              .setData({ row: r, col: c, type, special: null })
              .setInteractive();
            this.gems[r][c] = sprite;
            this.tweens.add({
              targets: sprite,
              y: this.gemY(r),
              duration: 280 + (writeRow - r) * 40,
              ease: "Cubic.easeIn"
            });
          }
        }
        this.time.delayedCall(420, () => {
          const more = this.findMatches();
          if (more.cells.length > 0) {
            this.cascade(more, depth + 1);
          } else {
            this.busy = false;
          }
        });
      });
    };
    this.swap = (a, b) => {
      if (this.busy) return;
      this.busy = true;
      const sA = this.gems[a.row][a.col];
      const sB = this.gems[b.row][b.col];
      const tA = this.gemTypes[a.row][a.col];
      const tB = this.gemTypes[b.row][b.col];
      this.gemTypes[a.row][a.col] = tB;
      this.gemTypes[b.row][b.col] = tA;
      this.gems[a.row][a.col] = sB;
      this.gems[b.row][b.col] = sA;
      sA.setData({ row: b.row, col: b.col, type: tA, special: sA.getData("special") });
      sB.setData({ row: a.row, col: a.col, type: tB, special: sB.getData("special") });
      this.tweens.add({
        targets: sA,
        x: this.gemX(b.col),
        y: this.gemY(b.row),
        duration: 200,
        ease: "Cubic.easeOut"
      });
      this.tweens.add({
        targets: sB,
        x: this.gemX(a.col),
        y: this.gemY(a.row),
        duration: 200,
        ease: "Cubic.easeOut",
        onComplete: () => {
          const matches = this.findMatches();
          if (matches.cells.length === 0) {
            this.gemTypes[a.row][a.col] = tA;
            this.gemTypes[b.row][b.col] = tB;
            this.gems[a.row][a.col] = sA;
            this.gems[b.row][b.col] = sB;
            sA.setData({ row: a.row, col: a.col, type: tA, special: sA.getData("special") });
            sB.setData({ row: b.row, col: b.col, type: tB, special: sB.getData("special") });
            this.tweens.add({
              targets: sA,
              x: this.gemX(a.col),
              y: this.gemY(a.row),
              duration: 200,
              ease: "Cubic.easeOut"
            });
            this.tweens.add({
              targets: sB,
              x: this.gemX(b.col),
              y: this.gemY(b.row),
              duration: 200,
              ease: "Cubic.easeOut",
              onComplete: () => { this.busy = false; }
            });
          } else {
            this.cascade(matches, 0);
          }
        }
      });
    };
    this.scene.launch("UIScene");
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super("UIScene");
  }

  create() {
    this.scoreText = this.add
      .text(20, 16, "Score: 0", {
        fontSize: "22px",
        color: "#ffffff",
        fontStyle: "bold"
      })
      .setScrollFactor(0);
    const main = this.scene.get("MainScene");
    main.events.on("scoreUpdate", (score) => {
      this.scoreText.setText("Score: " + score);
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
