---
id: 66faa80000000000000000ac
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Define `this.swap` - the closure that animates two adjacent gems swapping positions. The state swap is immediate (we update `gemTypes` and `gems` arrays and the `setData` row/col tags) so subsequent logic sees the post-swap world. The tween only handles the visual interpolation.

After `/* create-methods */`, add:

```js
    this.swap = (a, b) => {
      if (this.busy) return;
      this.busy = true;
      const sA = this.gems[a.row][a.col];
      const sB = this.gems[b.row][b.col];
      const tA = this.gemTypes[a.row][a.col];
      const tB = this.gemTypes[b.row][b.col];
      this.gemTypes[a.row][a.col] = tB;
      this.gemTypes[b.row][b.col] = tA;
      this.gems[a.row][a.col] = sB;
      this.gems[b.row][b.col] = sA;
      sA.setData({ row: b.row, col: b.col, type: tA, special: sA.getData("special") });
      sB.setData({ row: a.row, col: a.col, type: tB, special: sB.getData("special") });
      this.tweens.add({
        targets: sA,
        x: this.gemX(b.col),
        y: this.gemY(b.row),
        duration: 200,
        ease: "Cubic.easeOut"
      });
      this.tweens.add({
        targets: sB,
        x: this.gemX(a.col),
        y: this.gemY(a.row),
        duration: 200,
        ease: "Cubic.easeOut",
        onComplete: () => {
          this.busy = false;
        }
      });
    };
```

`if (this.busy) return;` is the input lock - while a tween is running, additional swap requests are dropped. The lock releases in `onComplete`. Two parallel tweens (one per gem) are simpler than a chain because both stages are identical durations.

# --hints--
Define `this.swap = (a, b) => { ... };` on the scene.

```js
assert.match(code, /this\.swap\s*=\s*\(\s*a\s*,\s*b\s*\)\s*=>/);
```

Bail with `if (this.busy) return;` before locking.

```js
assert.match(code, /if\s*\(\s*this\.busy\s*\)\s*return\s*;\s*this\.busy\s*=\s*true/);
```

Swap `this.gemTypes[a.row][a.col]` with `this.gemTypes[b.row][b.col]`.

```js
assert.match(code, /this\.gemTypes\[a\.row\]\[a\.col\]\s*=\s*tB[\s\S]*this\.gemTypes\[b\.row\]\[b\.col\]\s*=\s*tA/);
```

Tween sA toward `this.gemX(b.col), this.gemY(b.row)` over 200ms.

```js
assert.match(code, /targets:\s*sA[\s\S]*x:\s*this\.gemX\(b\.col\)[\s\S]*y:\s*this\.gemY\(b\.row\)[\s\S]*duration:\s*200/);
```

Release the input lock from sB's `onComplete`.

```js
assert.match(code, /onComplete:\s*\(\s*\)\s*=>\s*\{\s*this\.busy\s*=\s*false\s*;/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillRoundedRect(0, 0, 48, 48, 10);
    g.lineStyle(2, 0xffffff, 0.5);
    g.strokeRoundedRect(0, 0, 48, 48, 10);
    g.generateTexture("gem", 48, 48);
    g.destroy();
    this.gems = [];
    this.gemTypes = [];
    for (let r = 0; r < this.rows; r++) {
      this.gems[r] = [];
      this.gemTypes[r] = [];
      for (let c = 0; c < this.cols; c++) {
        let type;
        do {
          type = Phaser.Math.Between(0, this.gemColors.length - 1);
        } while (
          (c >= 2 &&
            this.gemTypes[r][c - 1] === type &&
            this.gemTypes[r][c - 2] === type) ||
          (r >= 2 &&
            this.gemTypes[r - 1][c] === type &&
            this.gemTypes[r - 2][c] === type)
        );
        this.gemTypes[r][c] = type;
        const sprite = this.add
          .sprite(this.gemX(c), this.gemY(r), "gem")
          .setTint(this.gemColors[type])
          .setData({ row: r, col: c, type, special: null });
        this.gems[r][c] = sprite;
      }
    }
    this.busy = false;
    this.selected = null;
    for (let r = 0; r < this.rows; r++) {
      for (let c = 0; c < this.cols; c++) {
        this.gems[r][c].setInteractive();
      }
    }
    this.input.on("gameobjectdown", (pointer, gem) => {
      if (this.busy) return;
      if (!this.selected) {
        this.selected = gem;
        this.tweens.add({
          targets: gem,
          scale: 1.15,
          duration: 120,
          ease: "Cubic.easeOut"
        });
        return;
      }
      this.tweens.add({ targets: this.selected, scale: 1, duration: 120 });
      const prev = this.selected;
      this.selected = null;
      if (gem === prev) return;
      const dr = Math.abs(gem.getData("row") - prev.getData("row"));
      const dc = Math.abs(gem.getData("col") - prev.getData("col"));
      if (dr + dc !== 1) return;
      const a = { row: prev.getData("row"), col: prev.getData("col") };
      const b = { row: gem.getData("row"), col: gem.getData("col") };
      this.swap(a, b);
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene]
};

new Phaser.Game(config);
```
