---
id: 66faa80000000000000000aa
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
If the random fill creates a 3-in-a-row, the player sees matches vanish before they can move - boring start. Reroll the type when placing it would complete a 3-in-a-row to the left or above.

Replace the random type pick:

```js
let type;
do {
  type = Phaser.Math.Between(0, this.gemColors.length - 1);
} while (
  (c >= 2 &&
    this.gemTypes[r][c - 1] === type &&
    this.gemTypes[r][c - 2] === type) ||
  (r >= 2 &&
    this.gemTypes[r - 1][c] === type &&
    this.gemTypes[r - 2][c] === type)
);
this.gemTypes[r][c] = type;
```

The two `(c >= 2 ...)` and `(r >= 2 ...)` clauses check the two previous gems in this row/column. If both already match the candidate type, we'd be the third in a row - reroll. With 5 colors, the average expected reroll count is ~0.04, so this loop terminates fast.

# --hints--
Replace `const type = ...` with `let type; do { ... } while (...);`.

```js
assert.match(code, /do\s*\{\s*type\s*=\s*Phaser\.Math\.Between\(\s*0\s*,\s*this\.gemColors\.length\s*-\s*1\s*\)/);
```

Guard horizontal: `c >= 2 && this.gemTypes[r][c - 1] === type && this.gemTypes[r][c - 2] === type`.

```js
assert.match(code, /c\s*>=\s*2\s*&&\s*this\.gemTypes\[r\]\[c\s*-\s*1\]\s*===\s*type\s*&&\s*this\.gemTypes\[r\]\[c\s*-\s*2\]\s*===\s*type/);
```

Guard vertical: `r >= 2 && this.gemTypes[r - 1][c] === type && this.gemTypes[r - 2][c] === type`.

```js
assert.match(code, /r\s*>=\s*2\s*&&\s*this\.gemTypes\[r\s*-\s*1\]\[c\]\s*===\s*type\s*&&\s*this\.gemTypes\[r\s*-\s*2\]\[c\]\s*===\s*type/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillRoundedRect(0, 0, 48, 48, 10);
    g.lineStyle(2, 0xffffff, 0.5);
    g.strokeRoundedRect(0, 0, 48, 48, 10);
    g.generateTexture("gem", 48, 48);
    g.destroy();
    this.gems = [];
    this.gemTypes = [];
    for (let r = 0; r < this.rows; r++) {
      this.gems[r] = [];
      this.gemTypes[r] = [];
      for (let c = 0; c < this.cols; c++) {
--fcc-editable-region--
        const type = Phaser.Math.Between(0, this.gemColors.length - 1);
        this.gemTypes[r][c] = type;
--fcc-editable-region--
        const sprite = this.add
          .sprite(this.gemX(c), this.gemY(r), "gem")
          .setTint(this.gemColors[type])
          .setData({ row: r, col: c, type, special: null });
        this.gems[r][c] = sprite;
      }
    }
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene]
};

new Phaser.Game(config);
```
