---
id: 66faa80000000000000000a9
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Render the gems. First, draw a single white rounded-rectangle into a generated texture called `gem` - we tint it five different ways instead of shipping five sprite sheets. Then spawn an 8x8 grid of `gem` sprites, each with a random type and tint.

After `/* create-grid */`, add:

```js
const g = this.add.graphics();
g.fillStyle(0xffffff, 1);
g.fillRoundedRect(0, 0, 48, 48, 10);
g.lineStyle(2, 0xffffff, 0.5);
g.strokeRoundedRect(0, 0, 48, 48, 10);
g.generateTexture("gem", 48, 48);
g.destroy();
this.gems = [];
this.gemTypes = [];
for (let r = 0; r < this.rows; r++) {
  this.gems[r] = [];
  this.gemTypes[r] = [];
  for (let c = 0; c < this.cols; c++) {
    const type = Phaser.Math.Between(0, this.gemColors.length - 1);
    this.gemTypes[r][c] = type;
    const sprite = this.add
      .sprite(this.gemX(c), this.gemY(r), "gem")
      .setTint(this.gemColors[type])
      .setData({ row: r, col: c, type, special: null });
    this.gems[r][c] = sprite;
  }
}
```

Two side-effects worth noting: `setData({ row, col, type })` stashes grid coordinates on the sprite for later input handling, and `special: null` is a placeholder for stripe/color-bomb/wrapped tagging in step 17.

# --hints--
Generate a `"gem"` texture with `g.generateTexture("gem", 48, 48)`.

```js
assert.match(code, /g\.generateTexture\(\s*["\']gem["\']\s*,\s*48\s*,\s*48\s*\)/);
```

Initialize `this.gems = []` and `this.gemTypes = []`.

```js
assert.match(code, /this\.gems\s*=\s*\[\s*\][\s\S]*this\.gemTypes\s*=\s*\[\s*\]/);
```

Loop `r` from 0 to `this.rows`.

```js
assert.match(code, /for\s*\(\s*let\s+r\s*=\s*0\s*;\s*r\s*<\s*this\.rows\s*;\s*r\+\+\s*\)/);
```

Random type via `Phaser.Math.Between(0, this.gemColors.length - 1)`.

```js
assert.match(code, /Phaser\.Math\.Between\(\s*0\s*,\s*this\.gemColors\.length\s*-\s*1\s*\)/);
```

Add a `"gem"` sprite at `this.gemX(c), this.gemY(r)` and `setTint(this.gemColors[type])`.

```js
assert.match(code, /this\.add\s*\.\s*sprite\(\s*this\.gemX\(c\)\s*,\s*this\.gemY\(r\)\s*,\s*["\']gem["\']\s*\)[\s\S]*setTint\(\s*this\.gemColors\[type\]\s*\)/);
```

Stash grid coords with `.setData({ row: r, col: c, type, special: null })`.

```js
assert.match(code, /setData\(\s*\{\s*row:\s*r\s*,\s*col:\s*c\s*,\s*type\s*,\s*special:\s*null\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene]
};

new Phaser.Game(config);
```
