---
id: 66faa80000000000000000a8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
Match-3 needs a coordinate system. Inside `create`, define grid constants and two helpers for converting `(col, row)` pairs into screen pixels. The grid is `8x8` with `56`-pixel cells, centered horizontally and pushed `30px` down to leave room for HUD.

After the `/* create-config */` comment, add:

```js
this.cols = 8;
this.rows = 8;
this.cell = 56;
this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
this.gemX = (col) => this.gridX + col * this.cell;
this.gemY = (row) => this.gridY + row * this.cell;
this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
```

The five hex colors are the gem palette - red, blue, green, yellow, purple. We'll tint a single white `gem` texture five ways next step so we don't ship a separate sprite sheet.

# --hints--
Set `this.cols = 8`, `this.rows = 8`, and `this.cell = 56`.

```js
assert.match(code, /this\.cols\s*=\s*8[\s\S]*this\.rows\s*=\s*8[\s\S]*this\.cell\s*=\s*56/);
```

Center the grid: `this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2`.

```js
assert.match(code, /this\.gridX\s*=\s*400\s*-\s*\(\s*this\.cols\s*\*\s*this\.cell\s*\)\s*\/\s*2\s*\+\s*this\.cell\s*\/\s*2/);
```

Define `this.gemX = (col) => this.gridX + col * this.cell`.

```js
assert.match(code, /this\.gemX\s*=\s*\(\s*col\s*\)\s*=>\s*this\.gridX\s*\+\s*col\s*\*\s*this\.cell/);
```

Define `this.gemY = (row) => this.gridY + row * this.cell`.

```js
assert.match(code, /this\.gemY\s*=\s*\(\s*row\s*\)\s*=>\s*this\.gridY\s*\+\s*row\s*\*\s*this\.cell/);
```

Define `this.gemColors` array with five hex colors.

```js
assert.match(code, /this\.gemColors\s*=\s*\[\s*0x[0-9a-f]{6}\s*,\s*0x[0-9a-f]{6}\s*,\s*0x[0-9a-f]{6}\s*,\s*0x[0-9a-f]{6}\s*,\s*0x[0-9a-f]{6}\s*\]/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene]
};

new Phaser.Game(config);
```
