---
id: 66faa80000000000000000ca
title: Step 36
challengeType: 0
dashedName: step-36
---

# --description--
If `findHint()` returns `null`, the player is stuck on a dead board with no possible matches. Add a `reshuffle()` helper that scrambles all gem types in place using `Phaser.Utils.Array.Shuffle`, then validates by running `findMatches()` - if the shuffle accidentally created matches, kick straight into a cascade.

```js
this.reshuffle = () => {
  this.busy = true;
  const all = [];
  for (let r = 0; r < this.rows; r++) {
    for (let c = 0; c < this.cols; c++) {
      if (this.gems[r][c]) all.push(this.gemTypes[r][c]);
    }
  }
  Phaser.Utils.Array.Shuffle(all);
  let idx = 0;
  for (let r = 0; r < this.rows; r++) {
    for (let c = 0; c < this.cols; c++) {
      if (!this.gems[r][c]) continue;
      const newType = all[idx++];
      this.gemTypes[r][c] = newType;
      this.gems[r][c]
        .setTint(this.gemColors[newType])
        .setData({ row: r, col: c, type: newType, special: null });
    }
  }
  this.cameras.main.shake(220, 0.005);
  const matches = this.findMatches();
  if (matches.cells.length > 0) {
    this.cascade(matches, 0);
  } else {
    this.busy = false;
  }
};
```

After cascade settles, call `findHint()` first - if `null`, trigger the reshuffle instead of starting the idle timer. The `cameras.main.shake` gives haptic feedback that something structural just happened. Reshuffles preserve gem identity (no destroyed/respawned objects), so combo/score state is untouched.

# --hints--
Add `this.reshuffle = () => { ... }` as a new helper.

```js
assert.match(code, /this\.reshuffle\s*=\s*\(\s*\)\s*=>/);
```

Use `Phaser.Utils.Array.Shuffle(all)` to randomize the type pool.

```js
assert.match(code, /Phaser\.Utils\.Array\.Shuffle\(/);
```

Call `this.cameras.main.shake` for haptic feedback.

```js
assert.match(code, /this\.cameras\.main\.shake\(/);
```

After cascade settles, `if (!this.findHint()) this.reshuffle();` instead of starting the hint timer.

```js
assert.match(code, /if\s*\(\s*!\s*pair\s*\)\s*\{\s*\n\s*this\.reshuffle\(\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  init(data) {
    this.startLevel = (data && data.level) || 0;
  }


  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillRoundedRect(0, 0, 48, 48, 10);
    g.lineStyle(2, 0xffffff, 0.5);
    g.strokeRoundedRect(0, 0, 48, 48, 10);
    g.generateTexture("gem", 48, 48);
    g.destroy();
    this.gems = [];
    this.gemTypes = [];
    for (let r = 0; r < this.rows; r++) {
      this.gems[r] = [];
      this.gemTypes[r] = [];
      for (let c = 0; c < this.cols; c++) {
        let type;
        do {
          type = Phaser.Math.Between(0, this.gemColors.length - 1);
        } while (
          (c >= 2 &&
            this.gemTypes[r][c - 1] === type &&
            this.gemTypes[r][c - 2] === type) ||
          (r >= 2 &&
            this.gemTypes[r - 1][c] === type &&
            this.gemTypes[r - 2][c] === type)
        );
        this.gemTypes[r][c] = type;
        const sprite = this.add
          .sprite(this.gemX(c), this.gemY(r), "gem")
          .setTint(this.gemColors[type])
          .setData({ row: r, col: c, type, special: null });
        this.gems[r][c] = sprite;
      }
    }
    this.busy = false;
    this.selected = null;
    this.dragStart = null;
    for (let r = 0; r < this.rows; r++) {
      for (let c = 0; c < this.cols; c++) {
        this.gems[r][c].setInteractive();
      }
    }
    this.input.on("gameobjectdown", (pointer, gem) => {
      if (this.busy) return;
      this.dragStart = { x: pointer.x, y: pointer.y, gem };
      if (!this.selected) {
        this.selected = gem;
        this.tweens.add({
          targets: gem,
          scale: 1.15,
          duration: 120,
          ease: "Cubic.easeOut"
        });
        return;
      }
      this.tweens.add({ targets: this.selected, scale: 1, duration: 120 });
      const prev = this.selected;
      this.selected = null;
      if (gem === prev) return;
      const dr = Math.abs(gem.getData("row") - prev.getData("row"));
      const dc = Math.abs(gem.getData("col") - prev.getData("col"));
      if (dr + dc !== 1) return;
      const a = { row: prev.getData("row"), col: prev.getData("col") };
      const b = { row: gem.getData("row"), col: gem.getData("col") };
      this.swap(a, b);
    });
    this.input.on("pointerup", (pointer) => {
      if (!this.dragStart) return;
      const start = this.dragStart;
      this.dragStart = null;
      const dx = pointer.x - start.x;
      const dy = pointer.y - start.y;
      const adx = Math.abs(dx);
      const ady = Math.abs(dy);
      if (Math.max(adx, ady) < 20) return;
      if (this.busy) return;
      const r = start.gem.getData("row");
      const c = start.gem.getData("col");
      let target = null;
      if (adx > ady) {
        const tc = c + (dx > 0 ? 1 : -1);
        if (tc >= 0 && tc < this.cols) target = this.gems[r][tc];
      } else {
        const tr = r + (dy > 0 ? 1 : -1);
        if (tr >= 0 && tr < this.rows) target = this.gems[tr][c];
      }
      if (!target) return;
      if (this.selected) {
        this.tweens.add({ targets: this.selected, scale: 1, duration: 120 });
        this.selected = null;
      }
      this.swap(
        { row: r, col: c },
        { row: target.getData("row"), col: target.getData("col") }
      );
    });
--fcc-editable-region--
    this.findMatches = () => {
      const cells = new Map();
      const runs = [];
      for (let r = 0; r < this.rows; r++) {
        let runStart = 0;
        let runType = this.gemTypes[r][0];
        for (let c = 1; c <= this.cols; c++) {
          const t = c < this.cols ? this.gemTypes[r][c] : -2;
          if (t === runType && runType >= 0) continue;
          if (c - runStart >= 3 && runType >= 0) {
            runs.push({ orient: "h", row: r, col: runStart, len: c - runStart, type: runType });
            for (let k = runStart; k < c; k++) {
              cells.set(`${r},${k}`, { row: r, col: k, type: runType });
            }
          }
          runStart = c;
          runType = t;
        }
      }
      for (let c = 0; c < this.cols; c++) {
        let runStart = 0;
        let runType = this.gemTypes[0][c];
        for (let r = 1; r <= this.rows; r++) {
          const t = r < this.rows ? this.gemTypes[r][c] : -2;
          if (t === runType && runType >= 0) continue;
          if (r - runStart >= 3 && runType >= 0) {
            runs.push({ orient: "v", row: runStart, col: c, len: r - runStart, type: runType });
            for (let k = runStart; k < r; k++) {
              cells.set(`${k},${c}`, { row: k, col: c, type: runType });
            }
          }
          runStart = r;
          runType = t;
        }
      }
      return { cells: [...cells.values()], runs };
    };
    this.activateSpecial = (gem, allCells) => {
      const special = gem.getData("special");
      if (!special) return;
      gem.setData({
        row: gem.getData("row"),
        col: gem.getData("col"),
        type: gem.getData("type"),
        special: null
      });
      const sr = gem.getData("row");
      const sc = gem.getData("col");
      const targets = [];
      if (special === "stripe-h") {
        for (let c = 0; c < this.cols; c++) {
          if (this.gems[sr][c]) {
            targets.push({ row: sr, col: c, type: this.gemTypes[sr][c] });
          }
        }
      } else if (special === "stripe-v") {
        for (let r = 0; r < this.rows; r++) {
          if (this.gems[r][sc]) {
            targets.push({ row: r, col: sc, type: this.gemTypes[r][sc] });
          }
        }
      } else if (special === "color-bomb") {
        const colorTarget = gem.getData("type");
        for (let r = 0; r < this.rows; r++) {
          for (let c = 0; c < this.cols; c++) {
            if (this.gemTypes[r][c] === colorTarget && this.gems[r][c]) {
              targets.push({ row: r, col: c, type: colorTarget });
            }
          }
        }
      } else if (special === "wrapped") {
        for (let dr = -1; dr <= 1; dr++) {
          for (let dc = -1; dc <= 1; dc++) {
            const r = sr + dr;
            const c = sc + dc;
            if (r < 0 || r >= this.rows) continue;
            if (c < 0 || c >= this.cols) continue;
            if (!this.gems[r][c]) continue;
            targets.push({ row: r, col: c, type: this.gemTypes[r][c] });
          }
        }
      }
      targets.forEach((cell) => {
        const key = `${cell.row},${cell.col}`;
        if (allCells.has(key)) return;
        allCells.set(key, cell);
        const chainGem = this.gems[cell.row][cell.col];
        if (chainGem && chainGem.getData("special")) {
          this.activateSpecial(chainGem, allCells);
        }
      });
    };
    this.checkLevelEnd = () => {
      if (typeof this.movesLeft !== "number") return;
      const bestKey = "match3-best-" + this.level;
      const stored = parseInt(
        window.localStorage?.getItem(bestKey) || "0",
        10
      );
      this.bestScore = Math.max(stored, this.score);
      if (this.bestScore > stored) {
        window.localStorage?.setItem(bestKey, String(this.bestScore));
      }
      if (this.matchedTarget >= this.targetCount) {
        this.events.emit("levelComplete", { level: this.level });
        this.cameras.main.flash(400, 100, 220, 180);
        const nextLevel = this.level + 1;
        this.time.delayedCall(700, () => {
          if (nextLevel < this.LEVELS.length) {
            this.scene.restart({ level: nextLevel });
          } else {
            this.events.emit("gameWon", {
              score: this.score,
              best: this.bestScore
            });
          }
        });
      } else if (this.movesLeft <= 0) {
        this.events.emit("levelFailed", { level: this.level });
      }
    };
    this.cascade = (matches, depth) => {
      this.reduceMotion ??=
        window.matchMedia?.("(prefers-reduced-motion: reduce)").matches || false;
      const dur = (ms) =>
        Math.round(ms * (this.reduceMotion ? 0.45 : 1));
      this.score ??= 0;
      this.combo ??= 0;
      this.combo = depth + 1;
      this.LEVELS ??= [
        { color: 0, target: 25, moves: 20 },
        { color: 1, target: 30, moves: 22 },
        { color: 2, target: 35, moves: 24 }
      ];
      this.level ??= this.startLevel ?? 0;
      this.movesLeft ??= this.LEVELS[this.level].moves;
      this.targetColor ??= this.LEVELS[this.level].color;
      this.targetCount ??= this.LEVELS[this.level].target;
      this.matchedTarget ??= 0;
      const targetMatched = matches.cells.filter(
        (c) => c.type === this.targetColor
      ).length;
      if (targetMatched > 0) {
        this.matchedTarget += targetMatched;
        this.events.emit("targetUpdate", {
          matched: this.matchedTarget,
          target: this.targetCount
        });
      }
      this.events.emit("comboUpdate", this.combo);
      const allCells = new Map(
        matches.cells.map((c) => [`${c.row},${c.col}`, c])
      );
      [...allCells.values()].forEach((cell) => {
        const g2 = this.gems[cell.row][cell.col];
        if (g2 && g2.getData("special")) this.activateSpecial(g2, allCells);
      });
      this.score += allCells.size * 10 * this.combo;
      this.events.emit("scoreUpdate", this.score);
      const spared = new Set();
      matches.runs.forEach((run) => {
        if (run.len < 4) return;
        const key = `${run.row},${run.col}`;
        if (spared.has(key)) return;
        const existingGem = this.gems[run.row][run.col];
        if (existingGem && existingGem.getData("special")) return;
        spared.add(key);
        const gem = existingGem;
        if (!gem) return;
        let kind;
        if (run.len >= 5) {
          kind = "color-bomb";
        } else {
          kind = run.orient === "h" ? "stripe-h" : "stripe-v";
        }
        gem.setData({
          row: run.row,
          col: run.col,
          type: run.type,
          special: kind
        });
        const bloom = gem.postFX.addBloom(0xffffff, 1, 1, 0.5, 1.6, 6);
        this.tweens.add({
          targets: bloom,
          strength: { from: 1.6, to: 0 },
          duration: 700,
          onComplete: () => gem.postFX.remove(bloom)
        });
        this.tweens.chain({
          targets: gem,
          tweens: [
            { scale: 1.5, duration: 120, ease: "Cubic.easeOut" },
            { scale: 0.85, duration: 80, ease: "Cubic.easeIn" },
            { scale: 1, duration: 140, ease: "Back.easeOut" }
          ]
        });
      });
      const hRuns = matches.runs.filter((r) => r.orient === "h");
      const vRuns = matches.runs.filter((r) => r.orient === "v");
      hRuns.forEach((h) => {
        vRuns.forEach((v) => {
          if (h.type !== v.type) return;
          const hr = h.row;
          const vc = v.col;
          const inH = vc >= h.col && vc < h.col + h.len;
          const inV = hr >= v.row && hr < v.row + v.len;
          if (!(inH && inV)) return;
          const key = `${hr},${vc}`;
          if (spared.has(key)) return;
          const gem = this.gems[hr][vc];
          if (!gem) return;
          if (gem.getData("special")) return;
          spared.add(key);
          gem.setData({
            row: hr,
            col: vc,
            type: h.type,
            special: "wrapped"
          });
          const bloom = gem.postFX.addBloom(0xff8be6, 1, 1, 0.6, 2, 8);
          this.tweens.add({
            targets: bloom,
            strength: { from: 2, to: 0 },
            duration: 800,
            onComplete: () => gem.postFX.remove(bloom)
          });
          this.tweens.chain({
            targets: gem,
            tweens: [
              { scale: 1.7, duration: 140, ease: "Cubic.easeOut" },
              { scale: 0.9, duration: 80, ease: "Cubic.easeIn" },
              { scale: 1, duration: 140, ease: "Back.easeOut" }
            ]
          });
        });
      });
      [...allCells.values()].forEach(({ row, col }) => {
        if (spared.has(`${row},${col}`)) return;
        const gem = this.gems[row][col];
        this.gems[row][col] = null;
        this.gemTypes[row][col] = -1;
        if (!gem) return;
        this.tweens.add({
          targets: gem,
          scale: 0,
          duration: dur(200),
          ease: "Cubic.easeIn",
          onComplete: () => gem.destroy()
        });
      });
      this.time.delayedCall(220, () => {
        for (let c = 0; c < this.cols; c++) {
          let writeRow = this.rows - 1;
          for (let r = this.rows - 1; r >= 0; r--) {
            if (this.gemTypes[r][c] === -1) continue;
            if (writeRow !== r) {
              this.gemTypes[writeRow][c] = this.gemTypes[r][c];
              this.gems[writeRow][c] = this.gems[r][c];
              this.gemTypes[r][c] = -1;
              this.gems[r][c] = null;
              const moved = this.gems[writeRow][c];
              moved.setData({
                row: writeRow,
                col: c,
                type: this.gemTypes[writeRow][c],
                special: moved.getData("special")
              });
              this.tweens.add({
                targets: moved,
                y: this.gemY(writeRow),
                duration: dur(280),
                ease: "Cubic.easeIn"
              });
            }
            writeRow--;
          }
          for (let r = writeRow; r >= 0; r--) {
            const type = Phaser.Math.Between(0, this.gemColors.length - 1);
            this.gemTypes[r][c] = type;
            const sprite = this.add
              .sprite(this.gemX(c), this.gemY(-1 - (writeRow - r)), "gem")
              .setTint(this.gemColors[type])
              .setData({ row: r, col: c, type, special: null })
              .setInteractive();
            this.gems[r][c] = sprite;
            this.tweens.add({
              targets: sprite,
              y: this.gemY(r),
              duration: dur(280 + (writeRow - r) * 40),
              ease: "Cubic.easeIn"
            });
          }
        }
        this.time.delayedCall(420, () => {
          const more = this.findMatches();
          if (more.cells.length > 0) {
            this.cascade(more, depth + 1);
          } else {
            this.busy = false;
            this.checkLevelEnd();
            this.hintTimer?.remove();
            this.hintTimer = this.time.delayedCall(5000, () => {
              if (this.busy) return;
              const pair = this.findHint();
              if (!pair) return;
              this.tweens.add({
                targets: pair,
                scale: 1.15,
                duration: 320,
                yoyo: true,
                repeat: 2,
                ease: "Sine.easeInOut"
              });
            });
          }
        });
      });
    };
    this.swap = (a, b) => {
      if (this.busy) return;
      this.busy = true;
      this.hintTimer?.remove();
      const sA = this.gems[a.row][a.col];
      const sB = this.gems[b.row][b.col];
      const tA = this.gemTypes[a.row][a.col];
      const tB = this.gemTypes[b.row][b.col];
      this.gemTypes[a.row][a.col] = tB;
      this.gemTypes[b.row][b.col] = tA;
      this.gems[a.row][a.col] = sB;
      this.gems[b.row][b.col] = sA;
      sA.setData({ row: b.row, col: b.col, type: tA, special: sA.getData("special") });
      sB.setData({ row: a.row, col: a.col, type: tB, special: sB.getData("special") });
      this.tweens.add({
        targets: sA,
        x: this.gemX(b.col),
        y: this.gemY(b.row),
        duration: 200,
        ease: "Cubic.easeOut"
      });
      this.tweens.add({
        targets: sB,
        x: this.gemX(a.col),
        y: this.gemY(a.row),
        duration: 200,
        ease: "Cubic.easeOut",
        onComplete: () => {
          const matches = this.findMatches();
          if (matches.cells.length === 0) {
            this.gemTypes[a.row][a.col] = tA;
            this.gemTypes[b.row][b.col] = tB;
            this.gems[a.row][a.col] = sA;
            this.gems[b.row][b.col] = sB;
            sA.setData({ row: a.row, col: a.col, type: tA, special: sA.getData("special") });
            sB.setData({ row: b.row, col: b.col, type: tB, special: sB.getData("special") });
            this.tweens.add({
              targets: sA,
              x: this.gemX(a.col),
              y: this.gemY(a.row),
              duration: 200,
              ease: "Cubic.easeOut"
            });
            this.tweens.add({
              targets: sB,
              x: this.gemX(b.col),
              y: this.gemY(b.row),
              duration: 200,
              ease: "Cubic.easeOut",
              onComplete: () => { this.busy = false; }
            });
          } else {
            if (typeof this.movesLeft === "number") {
              this.movesLeft--;
              this.events.emit("movesUpdate", this.movesLeft);
            }
            this.cascade(matches, 0);
          }
        }
      });
    };
    this.findHint = () => {
      for (let r = 0; r < this.rows; r++) {
        for (let c = 0; c < this.cols; c++) {
          if (c + 1 < this.cols) {
            [
              this.gemTypes[r][c],
              this.gemTypes[r][c + 1]
            ] = [
              this.gemTypes[r][c + 1],
              this.gemTypes[r][c]
            ];
            const ok = this.findMatches().cells.length > 0;
            [
              this.gemTypes[r][c],
              this.gemTypes[r][c + 1]
            ] = [
              this.gemTypes[r][c + 1],
              this.gemTypes[r][c]
            ];
            if (ok) return [this.gems[r][c], this.gems[r][c + 1]];
          }
          if (r + 1 < this.rows) {
            [
              this.gemTypes[r][c],
              this.gemTypes[r + 1][c]
            ] = [
              this.gemTypes[r + 1][c],
              this.gemTypes[r][c]
            ];
            const ok = this.findMatches().cells.length > 0;
            [
              this.gemTypes[r][c],
              this.gemTypes[r + 1][c]
            ] = [
              this.gemTypes[r + 1][c],
              this.gemTypes[r][c]
            ];
            if (ok) return [this.gems[r][c], this.gems[r + 1][c]];
          }
        }
      }
      return null;
    };
--fcc-editable-region--
    this.cameras.main.postFX.addVignette(0.5, 0.5, 0.7, 0.42);
    this.scene.launch("UIScene");
    if (!window.localStorage?.getItem("match3-tutorial-seen")) {
      const dim = this.add
        .rectangle(width / 2, height / 2, width, height, 0x000000, 0.7)
        .setDepth(100);
      const card = this.add
        .rectangle(width / 2, height / 2, 360, 220, 0x12121f, 1)
        .setStrokeStyle(2, 0x55c8ff)
        .setDepth(101);
      const title = this.add
        .text(width / 2, height / 2 - 70, "Match 3", {
          fontFamily: "monospace",
          fontSize: "24px",
          color: "#fff"
        })
        .setOrigin(0.5)
        .setDepth(102);
      const body = this.add
        .text(
          width / 2,
          height / 2 - 10,
          "Drag or tap two adjacent gems\nto swap. Match 3+ to clear.",
          {
            fontFamily: "monospace",
            fontSize: "14px",
            color: "#cdd",
            align: "center"
          }
        )
        .setOrigin(0.5)
        .setDepth(102);
      const cta = this.add
        .text(width / 2, height / 2 + 60, "Tap to start", {
          fontFamily: "monospace",
          fontSize: "16px",
          color: "#55c8ff"
        })
        .setOrigin(0.5)
        .setDepth(102);
      this.tweens.add({
        targets: cta,
        alpha: { from: 1, to: 0.3 },
        duration: 700,
        yoyo: true,
        repeat: -1
      });
      const dismiss = () => {
        window.localStorage?.setItem("match3-tutorial-seen", "1");
        [dim, card, title, body, cta].forEach((o) => o.destroy());
      };
      dim.setInteractive().on("pointerdown", dismiss);
      this.input.keyboard?.once("keydown", dismiss);
    }
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super("UIScene");
  }

  create() {
    this.scoreText = this.add
      .text(20, 16, "Score: 0", {
        fontSize: "22px",
        color: "#ffffff",
        fontStyle: "bold"
      })
      .setScrollFactor(0);
    this.movesText = this.add
      .text(400, 16, "Moves: --", {
        fontSize: "20px",
        color: "#ffd34d",
        fontStyle: "bold"
      })
      .setScrollFactor(0)
      .setOrigin(0.5, 0);
    this.targetText = this.add
      .text(780, 16, "0 / 0", {
        fontSize: "20px",
        color: "#ffffff"
      })
      .setScrollFactor(0)
      .setOrigin(1, 0);
    const main = this.scene.get("MainScene");
    main.events.on("scoreUpdate", (score) => {
      this.scoreText.setText("Score: " + score);
    });
    main.events.on("movesUpdate", (moves) => {
      this.movesText.setText("Moves: " + moves);
    });
    main.events.on("targetUpdate", ({ matched, target }) => {
      const clamped = Math.min(matched, target);
      this.targetText.setText(clamped + " / " + target);
    });
    main.events.on("levelFailed", () => {
      const payload = {
        win: false,
        level: main.level,
        score: main.score,
        best: main.bestScore || 0
      };
      this.scene.stop("MainScene");
      this.time.delayedCall(200, () => {
        this.scene.launch("GameOverScene", payload);
      });
    });
  }
}


class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  init(data) {
    this.payload = data || {};
  }

  create() {
    const { width, height } = this.scale;
    this.add
      .rectangle(width / 2, height / 2, width, height, 0x000000, 0.75)
      .setScrollFactor(0);
    const title = this.payload.win ? "Level cleared!" : "Out of moves";
    const titleColor = this.payload.win ? "#6dff4d" : "#ff4d6d";
    this.add
      .text(width / 2, height / 2 - 80, title, {
        fontSize: "32px",
        color: titleColor,
        fontStyle: "bold"
      })
      .setOrigin(0.5);
    this.add
      .text(
        width / 2,
        height / 2 - 30,
        "Score " + (this.payload.score || 0),
        { fontSize: "22px", color: "#ffffff" }
      )
      .setOrigin(0.5);
    this.add
      .text(
        width / 2,
        height / 2 + 4,
        "Best  " + (this.payload.best || 0),
        { fontSize: "18px", color: "#ffd34d" }
      )
      .setOrigin(0.5);
    const retry = this.add
      .text(width / 2, height / 2 + 80, "Tap to retry", {
        fontSize: "20px",
        color: "#ffd34d",
        fontStyle: "bold"
      })
      .setOrigin(0.5)
      .setInteractive();
    retry.on("pointerdown", () => {
      this.scene.stop("UIScene");
      this.scene.stop("GameOverScene");
      const target = this.payload.win
        ? (this.payload.level || 0)
        : (this.payload.level || 0);
      this.scene.start("MainScene", { level: target });
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene, UIScene, GameOverScene]
};

new Phaser.Game(config);
```
