---
id: 66faa80000000000000000c1
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--
On win, advance to the next level. Use `this.scene.restart` with init data so MainScene re-runs `create` with a fresh grid and fresh level state.

Three changes:

1. **Add a `init(data)` method** to MainScene that captures `data.level` as `this.startLevel`.
1. **Lazy-init `this.level`** from `this.startLevel ?? 0` so restarts pick up the new level.
1. **Update `checkLevelEnd`** to flash blue, then `this.scene.restart({ level: nextLevel })` after a 700ms beat.

If the player clears the last level, emit `gameWon` instead. We'll wire a victory screen later.

# --hints--
Define `init(data)` on MainScene capturing `data.level`.

```js
assert.match(code, /init\s*\(\s*data\s*\)\s*\{[\s\S]*this\.startLevel/);
```

Lazy-init `this.level ??= this.startLevel ?? 0`.

```js
assert.match(code, /this\.level\s*\?\?=\s*this\.startLevel\s*\?\?\s*0/);
```

Flash the camera blue with `cameras.main.flash(400, 100, 220, 180)`.

```js
assert.match(code, /this\.cameras\.main\.flash\(\s*400\s*,\s*100\s*,\s*220\s*,\s*180\s*\)/);
```

Restart MainScene with `this.scene.restart({ level: nextLevel })`.

```js
assert.match(code, /this\.scene\.restart\(\s*\{\s*level:\s*nextLevel\s*\}\s*\)/);
```

Emit `"gameWon"` when there is no next level.

```js
assert.match(code, /this\.events\.emit\(\s*["\']gameWon["\']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Match-3 Puzzle</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #07071a;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 30px rgba(140, 90, 255, 0.4);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

--fcc-editable-region--
  preload() {
  }
--fcc-editable-region--

  create() {
    this.cols = 8;
    this.rows = 8;
    this.cell = 56;
    this.gridX = 400 - (this.cols * this.cell) / 2 + this.cell / 2;
    this.gridY = 300 - (this.rows * this.cell) / 2 + this.cell / 2 + 30;
    this.gemX = (col) => this.gridX + col * this.cell;
    this.gemY = (row) => this.gridY + row * this.cell;
    this.gemColors = [0xff4d6d, 0x4dadff, 0x6dff4d, 0xffd34d, 0xc14dff];
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillRoundedRect(0, 0, 48, 48, 10);
    g.lineStyle(2, 0xffffff, 0.5);
    g.strokeRoundedRect(0, 0, 48, 48, 10);
    g.generateTexture("gem", 48, 48);
    g.destroy();
    this.gems = [];
    this.gemTypes = [];
    for (let r = 0; r < this.rows; r++) {
      this.gems[r] = [];
      this.gemTypes[r] = [];
      for (let c = 0; c < this.cols; c++) {
        let type;
        do {
          type = Phaser.Math.Between(0, this.gemColors.length - 1);
        } while (
          (c >= 2 &&
            this.gemTypes[r][c - 1] === type &&
            this.gemTypes[r][c - 2] === type) ||
          (r >= 2 &&
            this.gemTypes[r - 1][c] === type &&
            this.gemTypes[r - 2][c] === type)
        );
        this.gemTypes[r][c] = type;
        const sprite = this.add
          .sprite(this.gemX(c), this.gemY(r), "gem")
          .setTint(this.gemColors[type])
          .setData({ row: r, col: c, type, special: null });
        this.gems[r][c] = sprite;
      }
    }
    this.busy = false;
    this.selected = null;
    for (let r = 0; r < this.rows; r++) {
      for (let c = 0; c < this.cols; c++) {
        this.gems[r][c].setInteractive();
      }
    }
    this.input.on("gameobjectdown", (pointer, gem) => {
      if (this.busy) return;
      if (!this.selected) {
        this.selected = gem;
        this.tweens.add({
          targets: gem,
          scale: 1.15,
          duration: 120,
          ease: "Cubic.easeOut"
        });
        return;
      }
      this.tweens.add({ targets: this.selected, scale: 1, duration: 120 });
      const prev = this.selected;
      this.selected = null;
      if (gem === prev) return;
      const dr = Math.abs(gem.getData("row") - prev.getData("row"));
      const dc = Math.abs(gem.getData("col") - prev.getData("col"));
      if (dr + dc !== 1) return;
      const a = { row: prev.getData("row"), col: prev.getData("col") };
      const b = { row: gem.getData("row"), col: gem.getData("col") };
      this.swap(a, b);
    });
    this.findMatches = () => {
      const cells = new Map();
      const runs = [];
      for (let r = 0; r < this.rows; r++) {
        let runStart = 0;
        let runType = this.gemTypes[r][0];
        for (let c = 1; c <= this.cols; c++) {
          const t = c < this.cols ? this.gemTypes[r][c] : -2;
          if (t === runType && runType >= 0) continue;
          if (c - runStart >= 3 && runType >= 0) {
            runs.push({ orient: "h", row: r, col: runStart, len: c - runStart, type: runType });
            for (let k = runStart; k < c; k++) {
              cells.set(`${r},${k}`, { row: r, col: k, type: runType });
            }
          }
          runStart = c;
          runType = t;
        }
      }
      for (let c = 0; c < this.cols; c++) {
        let runStart = 0;
        let runType = this.gemTypes[0][c];
        for (let r = 1; r <= this.rows; r++) {
          const t = r < this.rows ? this.gemTypes[r][c] : -2;
          if (t === runType && runType >= 0) continue;
          if (r - runStart >= 3 && runType >= 0) {
            runs.push({ orient: "v", row: runStart, col: c, len: r - runStart, type: runType });
            for (let k = runStart; k < r; k++) {
              cells.set(`${k},${c}`, { row: k, col: c, type: runType });
            }
          }
          runStart = r;
          runType = t;
        }
      }
      return { cells: [...cells.values()], runs };
    };
    this.activateSpecial = (gem, allCells) => {
      const special = gem.getData("special");
      if (!special) return;
      gem.setData({
        row: gem.getData("row"),
        col: gem.getData("col"),
        type: gem.getData("type"),
        special: null
      });
      const sr = gem.getData("row");
      const sc = gem.getData("col");
      const targets = [];
      if (special === "stripe-h") {
        for (let c = 0; c < this.cols; c++) {
          if (this.gems[sr][c]) {
            targets.push({ row: sr, col: c, type: this.gemTypes[sr][c] });
          }
        }
      } else if (special === "stripe-v") {
        for (let r = 0; r < this.rows; r++) {
          if (this.gems[r][sc]) {
            targets.push({ row: r, col: sc, type: this.gemTypes[r][sc] });
          }
        }
      } else if (special === "color-bomb") {
        const colorTarget = gem.getData("type");
        for (let r = 0; r < this.rows; r++) {
          for (let c = 0; c < this.cols; c++) {
            if (this.gemTypes[r][c] === colorTarget && this.gems[r][c]) {
              targets.push({ row: r, col: c, type: colorTarget });
            }
          }
        }
      } else if (special === "wrapped") {
        for (let dr = -1; dr <= 1; dr++) {
          for (let dc = -1; dc <= 1; dc++) {
            const r = sr + dr;
            const c = sc + dc;
            if (r < 0 || r >= this.rows) continue;
            if (c < 0 || c >= this.cols) continue;
            if (!this.gems[r][c]) continue;
            targets.push({ row: r, col: c, type: this.gemTypes[r][c] });
          }
        }
      }
      targets.forEach((cell) => {
        const key = `${cell.row},${cell.col}`;
        if (allCells.has(key)) return;
        allCells.set(key, cell);
        const chainGem = this.gems[cell.row][cell.col];
        if (chainGem && chainGem.getData("special")) {
          this.activateSpecial(chainGem, allCells);
        }
      });
    };
    this.checkLevelEnd = () => {
      if (typeof this.movesLeft !== "number") return;
      if (this.matchedTarget >= this.targetCount) {
        this.events.emit("levelComplete", { level: this.level });
      } else if (this.movesLeft <= 0) {
        this.events.emit("levelFailed", { level: this.level });
      }
    };
    this.cascade = (matches, depth) => {
      this.score ??= 0;
      this.combo ??= 0;
      this.combo = depth + 1;
      this.LEVELS ??= [
        { color: 0, target: 25, moves: 20 },
        { color: 1, target: 30, moves: 22 },
        { color: 2, target: 35, moves: 24 }
      ];
      this.level ??= 0;
      this.movesLeft ??= this.LEVELS[this.level].moves;
      this.targetColor ??= this.LEVELS[this.level].color;
      this.targetCount ??= this.LEVELS[this.level].target;
      this.matchedTarget ??= 0;
      const targetMatched = matches.cells.filter(
        (c) => c.type === this.targetColor
      ).length;
      if (targetMatched > 0) {
        this.matchedTarget += targetMatched;
        this.events.emit("targetUpdate", {
          matched: this.matchedTarget,
          target: this.targetCount
        });
      }
      this.events.emit("comboUpdate", this.combo);
      const allCells = new Map(
        matches.cells.map((c) => [`${c.row},${c.col}`, c])
      );
      [...allCells.values()].forEach((cell) => {
        const g2 = this.gems[cell.row][cell.col];
        if (g2 && g2.getData("special")) this.activateSpecial(g2, allCells);
      });
      this.score += allCells.size * 10 * this.combo;
      this.events.emit("scoreUpdate", this.score);
      const spared = new Set();
      matches.runs.forEach((run) => {
        if (run.len < 4) return;
        const key = `${run.row},${run.col}`;
        if (spared.has(key)) return;
        const existingGem = this.gems[run.row][run.col];
        if (existingGem && existingGem.getData("special")) return;
        spared.add(key);
        const gem = existingGem;
        if (!gem) return;
        let kind;
        if (run.len >= 5) {
          kind = "color-bomb";
        } else {
          kind = run.orient === "h" ? "stripe-h" : "stripe-v";
        }
        gem.setData({
          row: run.row,
          col: run.col,
          type: run.type,
          special: kind
        });
        this.tweens.chain({
          targets: gem,
          tweens: [
            { scale: 1.5, duration: 120, ease: "Cubic.easeOut" },
            { scale: 0.85, duration: 80, ease: "Cubic.easeIn" },
            { scale: 1, duration: 140, ease: "Back.easeOut" }
          ]
        });
      });
      const hRuns = matches.runs.filter((r) => r.orient === "h");
      const vRuns = matches.runs.filter((r) => r.orient === "v");
      hRuns.forEach((h) => {
        vRuns.forEach((v) => {
          if (h.type !== v.type) return;
          const hr = h.row;
          const vc = v.col;
          const inH = vc >= h.col && vc < h.col + h.len;
          const inV = hr >= v.row && hr < v.row + v.len;
          if (!(inH && inV)) return;
          const key = `${hr},${vc}`;
          if (spared.has(key)) return;
          const gem = this.gems[hr][vc];
          if (!gem) return;
          if (gem.getData("special")) return;
          spared.add(key);
          gem.setData({
            row: hr,
            col: vc,
            type: h.type,
            special: "wrapped"
          });
          this.tweens.chain({
            targets: gem,
            tweens: [
              { scale: 1.7, duration: 140, ease: "Cubic.easeOut" },
              { scale: 0.9, duration: 80, ease: "Cubic.easeIn" },
              { scale: 1, duration: 140, ease: "Back.easeOut" }
            ]
          });
        });
      });
      [...allCells.values()].forEach(({ row, col }) => {
        if (spared.has(`${row},${col}`)) return;
        const gem = this.gems[row][col];
        this.gems[row][col] = null;
        this.gemTypes[row][col] = -1;
        if (!gem) return;
        this.tweens.add({
          targets: gem,
          scale: 0,
          duration: 200,
          ease: "Cubic.easeIn",
          onComplete: () => gem.destroy()
        });
      });
      this.time.delayedCall(220, () => {
        for (let c = 0; c < this.cols; c++) {
          let writeRow = this.rows - 1;
          for (let r = this.rows - 1; r >= 0; r--) {
            if (this.gemTypes[r][c] === -1) continue;
            if (writeRow !== r) {
              this.gemTypes[writeRow][c] = this.gemTypes[r][c];
              this.gems[writeRow][c] = this.gems[r][c];
              this.gemTypes[r][c] = -1;
              this.gems[r][c] = null;
              const moved = this.gems[writeRow][c];
              moved.setData({
                row: writeRow,
                col: c,
                type: this.gemTypes[writeRow][c],
                special: moved.getData("special")
              });
              this.tweens.add({
                targets: moved,
                y: this.gemY(writeRow),
                duration: 280,
                ease: "Cubic.easeIn"
              });
            }
            writeRow--;
          }
          for (let r = writeRow; r >= 0; r--) {
            const type = Phaser.Math.Between(0, this.gemColors.length - 1);
            this.gemTypes[r][c] = type;
            const sprite = this.add
              .sprite(this.gemX(c), this.gemY(-1 - (writeRow - r)), "gem")
              .setTint(this.gemColors[type])
              .setData({ row: r, col: c, type, special: null })
              .setInteractive();
            this.gems[r][c] = sprite;
            this.tweens.add({
              targets: sprite,
              y: this.gemY(r),
              duration: 280 + (writeRow - r) * 40,
              ease: "Cubic.easeIn"
            });
          }
        }
        this.time.delayedCall(420, () => {
          const more = this.findMatches();
          if (more.cells.length > 0) {
            this.cascade(more, depth + 1);
          } else {
            this.busy = false;
            this.checkLevelEnd();
          }
        });
      });
    };
    this.swap = (a, b) => {
      if (this.busy) return;
      this.busy = true;
      const sA = this.gems[a.row][a.col];
      const sB = this.gems[b.row][b.col];
      const tA = this.gemTypes[a.row][a.col];
      const tB = this.gemTypes[b.row][b.col];
      this.gemTypes[a.row][a.col] = tB;
      this.gemTypes[b.row][b.col] = tA;
      this.gems[a.row][a.col] = sB;
      this.gems[b.row][b.col] = sA;
      sA.setData({ row: b.row, col: b.col, type: tA, special: sA.getData("special") });
      sB.setData({ row: a.row, col: a.col, type: tB, special: sB.getData("special") });
      this.tweens.add({
        targets: sA,
        x: this.gemX(b.col),
        y: this.gemY(b.row),
        duration: 200,
        ease: "Cubic.easeOut"
      });
      this.tweens.add({
        targets: sB,
        x: this.gemX(a.col),
        y: this.gemY(a.row),
        duration: 200,
        ease: "Cubic.easeOut",
        onComplete: () => {
          const matches = this.findMatches();
          if (matches.cells.length === 0) {
            this.gemTypes[a.row][a.col] = tA;
            this.gemTypes[b.row][b.col] = tB;
            this.gems[a.row][a.col] = sA;
            this.gems[b.row][b.col] = sB;
            sA.setData({ row: a.row, col: a.col, type: tA, special: sA.getData("special") });
            sB.setData({ row: b.row, col: b.col, type: tB, special: sB.getData("special") });
            this.tweens.add({
              targets: sA,
              x: this.gemX(a.col),
              y: this.gemY(a.row),
              duration: 200,
              ease: "Cubic.easeOut"
            });
            this.tweens.add({
              targets: sB,
              x: this.gemX(b.col),
              y: this.gemY(b.row),
              duration: 200,
              ease: "Cubic.easeOut",
              onComplete: () => { this.busy = false; }
            });
          } else {
            if (typeof this.movesLeft === "number") {
              this.movesLeft--;
              this.events.emit("movesUpdate", this.movesLeft);
            }
            this.cascade(matches, 0);
          }
        }
      });
    };
    this.scene.launch("UIScene");
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super("UIScene");
  }

  create() {
    this.scoreText = this.add
      .text(20, 16, "Score: 0", {
        fontSize: "22px",
        color: "#ffffff",
        fontStyle: "bold"
      })
      .setScrollFactor(0);
    this.movesText = this.add
      .text(400, 16, "Moves: --", {
        fontSize: "20px",
        color: "#ffd34d",
        fontStyle: "bold"
      })
      .setScrollFactor(0)
      .setOrigin(0.5, 0);
    this.targetText = this.add
      .text(780, 16, "0 / 0", {
        fontSize: "20px",
        color: "#ffffff"
      })
      .setScrollFactor(0)
      .setOrigin(1, 0);
    const main = this.scene.get("MainScene");
    main.events.on("scoreUpdate", (score) => {
      this.scoreText.setText("Score: " + score);
    });
    main.events.on("movesUpdate", (moves) => {
      this.movesText.setText("Moves: " + moves);
    });
    main.events.on("targetUpdate", ({ matched, target }) => {
      const clamped = Math.min(matched, target);
      this.targetText.setText(clamped + " / " + target);
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#0a0a1a",
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
