---
id: 67bf4350777ac6ffdc027745
title: Build a Color Picker App
challengeType: 25
dashedName: build-a-color-picker-app
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should use the `useState` hook.
2. You should have a `#color-picker-container` element with a white background.
3. You should have a `#color-input` element which should be a `color` input.
4. When `#color-input` is changed, `#color-picker-container` should have its background set to that new value.

# --hints--

You should use the `useState` hook.

```js
assert.match(code, /useState/);
```

You should use the return value of the `useState` hook.

```js
assert.match(code, /=\s*useState\(/);
```

You should render an element with `color-picker-container` as the `id`.

```js
assert.exists(document.getElementById("color-picker-container"));
```

Your `#color-picker-container` should have a white background initially.

```js
const el = document.getElementById("color-picker-container");
const style = window.getComputedStyle(el);
assert.equal(style.backgroundColor, "rgb(255, 255, 255)");
```

You should render an `input` element with an `id` of `color-input`.

```js
const inputMaybe = document.getElementById("color-input");
assert.exists(inputMaybe);
assert.equal(inputMaybe?.tagName, "INPUT");
```

Your `#color-input` should be a `color` input.

```js
const inputDefinitely = document.getElementById("color-input");
assert.equal(inputDefinitely?.type, "color");
```

When your `#color-input` value is changed, your `#color-picker-container` should change its background color to reflect that value.

```js
const inputForSure = document.getElementById("color-input");
assert.exists(inputForSure);
// This entire approach is very hacky, but mounting the app with Enzyme
// was throwing Rule of Hooks errors. 😒
const key = Object.keys(inputForSure).find(key => key.startsWith("__reactProps"));
const reactInstance = inputForSure[key]
if (reactInstance && reactInstance.onChange) {
  reactInstance.onChange({ target: { value: "#000000" } });
}
setTimeout(() => {
  const container = document.getElementById("color-picker-container");
  assert.exists(container);
  const style = window.getComputedStyle(container);
  console.log(style.backgroundColor);
  assert.equal(style.backgroundColor, "rgb(0, 0, 0)");
}, 1000);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MyApp } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<MyApp />);
    </script>
</body>

</html>
```

```css

```

```jsx
export const MyApp = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MyApp } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<MyApp />);
    </script>
</body>

</html>
```

```css
body,
html {
    margin: 0;
    padding: 0;
    height: 100%;
    font-family: Arial, sans-serif;
}

#color-picker-container {
    height: 100vh;
    display: flex;
    justify-content: center;
    align-items: center;
    flex-direction: column;
    background-color: #ffffff;
}

#pick-color-btn {
    padding: 15px 30px;
    font-size: 18px;
    color: white;
    background-color: #14183c;
    border: none;
    border-radius: 5px;
    cursor: pointer;
    margin-top: 20px;
}

#pick-color-btn:hover {
    background-color: #08131f;
}

input[type="color"] {
    position: absolute;
    margin-top: 50px;
    height: 40px;
}

.stylep{
    color: #14183c;
    font-size: 15px;
    font-weight: bold;
    margin-bottom: 50px;
    
}
```

```jsx
const { useState } = React;

export const MyApp = () => {
    const [bgColor, setBgColor] = useState('#ffffff'); 
    
    const handleColorChange = (event) => {
        setBgColor(event.target.value);
    };

    return (
            <div id="color-picker-container" style={{ backgroundColor: bgColor }}>
            <p className="stylep">
                Choose a color using the color input below:
            </p>

            <input
                id="color-input"
                type="color"
                value={bgColor}
                onChange={handleColorChange}
            />
        </div>
    );
}
```
