---
id: 67f77e09071e523297138424
title: Build a Multimedia Player
challengeType: 14
dashedName: build-a-multimedia-player
# demoType can either be 'onClick' or 'onLoad'. If the project or lab doesn't have a preview, delete the property
demoType: onClick
---

# --description--

You will practice working with the <code>audio</code> and <code>video</code> elements, the <code>controls</code> attribute, and the <code>aria-label</code> attribute.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have an `h1` element for the main title of the page.
2. You should have three `section` elements.
3. Inside the first `section` element, you should have an `h2` element for the title of song playing.
4. Below the `h2` element, you should have an `audio` element with `controls` attribute and an `aria-label` attribute.
5. Inside the `audio` element, you should have a `source` element with a `src` attribute pointing to an audio file and a `type` attribute.
6. Inside the second `section` element, you should have an `h2` element for the title of the video playing.
7. Below the `h2` element, you should have a `video` element with `controls`, `width` attributes and an `aria-label` attribute.
8. Inside the `video` element, you should have a `source` element with a `src` attribute pointing to a video file and a `type` attribute.
9. Below the `source` element, you should have a `track` element with a `src` attribute pointing to a subtitles file and a `kind` attribute, a `srclang` attribute and a `label` attribute.
10. Inside the third `section` element, you should have an `h2` element for the title of the section eg. "Transcript".
11. Below the `h2` element, you should have a `p` element with the transcript of the video.

# --hints--

You should have an `h1` element.

```js
assert.isNotNull(document.querySelector("h1"));
```

You should have three `section` elements.

```js
assert.equal(document.querySelectorAll("section").length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Multimedia Player</title>
</head>
<body>
  <h1>Multimedia Player</h1>
  <section>
      <h2>Now Playing: Sailing Away</h2>
      <audio controls aria-label="Audio Player for Sailing Away">
          <source src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" type="audio/mpeg">
      </audio>
  </section>

  <section>
      <h2>What is a map method and how does it work?</h2>
      <video controls width="600" aria-label="Video Player for Is AI coming for developer jobs?">
          <source src="https://cdn.freecodecamp.org/curriculum/labs/What is the map method and how does it work.mp4" type="video/mp4">

          <track 
              src="what-is-a-map-method.vtt" 
              kind="captions" 
              srclang="en" 
              label="English" 
              default>
      </video>
  </section>

  <section>
      <h2>Transcript</h2>
      <p>
          What is a map method and how does it work?
          The map method is a powerful and widely used function in JavaScript that operates on arrays.
          It's designed to create a new array by applying a given function to each element of the original array.
          This method does not modify the original array, but instead returns the new array containing the results of the function applied to each element. 
      </p>
  </section>

</body>
</html> 
```
