---
id: 6675df027b8c9574fa7bbffa
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Now that we have created a function that returns `"Hi"`, let's create a second function to pass into 
our `sayHello` function. Create a function that returns your name called `sayMyName`. 

# --hints--

You need to define `sayMyName` as a function. 

```js
assert.isFunction(sayMyName); 
```

The function `sayMyName` should return a string. 

```js
assert.isString(sayMyName()); 
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--


--fcc-editable-region--

function sayHello()
{
  return "Hi"; 
}
```
