---
id: 6675b238e128702f011e5395
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

Callbacks are among the most powerful tools in Javascript. Callbacks are used to pass a function into another function. The passed in function is then executed inside of the other function. The output then influences the other function.

Let's start building out an example. Begin by creating a function called `sayHello` that returns `"Hi"`. 

# --hints--

You need to define `sayHello` as a function. 

```js
assert.isFunction(sayHello); 
```

You need to make `sayHello` return `"Hi"`, with no extra characters. 

```js
assert.equal(sayHello(),"Hi"); 
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
