---
id: 6675df027b8c9574fa7bbffb
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

For this lesson, we are going to make it possible to actually pass a function as a variable
by adding an parameter to our `sayHello` function called `retrievedName`. 

# --hints--

There should be one parameter listed in the `sayHello` function. 

```js
const params = __helpers.getFunctionParams(sayHello.toString());
assert.equal(params.length,1); 
```

The parameter should be called `retrievedName`.

```js
const params = __helpers.getFunctionParams(sayHello.toString());
assert.equal(params[0].name,"retrievedName"); 
```

# --seed--

## --seed-contents--

```js

function sayMyName()
{
  return "Quincy"; 
}

--fcc-editable-region--

function sayHello()
{
  return "Hi"; 
}
--fcc-editable-region--
```
