---
id: 615f522dea4f776f64dc3e91
title: Hatua ya 31
challengeType: 0
dashedName: step-31
---

# --description--

Kitengo cha `rem` kinasimama `root em`, na kinahusiana na saizi ya fonti ya kipengele cha `html`.

Unda kichaguzi cha `.small-text` na uweke `font-size` kuwa `0.85rem`, ambayo itakokotoa hadi takriban `13.6px` (kumbuka kwamba uliweka `html` yako kuwa na `font-size` ya `16px`).

# --hints--

Unapaswa kuwa na kichaguzi cha `.small-text`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text'));
```

Kichaguzi chako cha `.small-text` kinafaa kuwa na sifa ya `font-size` iliyowekwa kuwa `0.85rem`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text')?.fontSize === '0.85rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
      </div>
    </div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
