---
id: 615f2d4150fe0d4cbd0f2628
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Chini ya kipengele chako cha `h1`, ongeza kipengele cha `p` chenye maandishi `8 servings per container`.

# --hints--

Unapaswa kuongeza kipengele kipya cha `p`.

```js
assert.exists(document.querySelector('p'));
```

Kipengele chako cha `p` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

Kipengele chako cha `p` kinapaswa kuja baada ya kipengele chako cha `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
