---
id: 615f41c979787462e76dab90
title: Hatua ya 20
challengeType: 0
dashedName: step-20
---

# --description--

Darasa lako jipya bado halina mtindo wowote. Unda kichaguzi cha `.bold` na ukipe `font-weight` iliyowekwa kuwa `800` ili kufanya maandishi kuwa ya mazito.

Endelea na uondoe kipengele cha `font-weight` kutoka kwa kichaguzi chako cha `h1` pia.

# --hints--

Unapaswa kuwa na kichaguzi cha `.bold`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

Kichaguzi chako cha `.bold` kinafaa kuwa na sifa ya `font-weight` iliyowekwa kuwa `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

Kichaguzi chako cha `h1` hakipaswi kuwa na sifa ya `font-weight`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

Hupaswi kuondoa kichaguzi chako cha `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
