---
id: 615f34ecc1091b4fd5a8a484
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Ndani ya kipengele chako cha `head`, ongeza kipengele cha `link` chenye sifa ya `rel` iliyowekwa kuwa `stylesheet` na sifa ya `href` iliyowekwa kuwa `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Hii italeta familia ya fonti ya `Open Sans`, yenye thamani za uzito wa fonti `400`, `700` na `800`.

Pia ongeza kipengele cha `link` ili kuunganisha faili yako ya `styles.css`.

# --hints--

Your code should have two `link` elements.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Your `link` elements should be void elements, they should not have a end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Your two `link` elements should be inside the `head` element.

```js
const headContentRegex = /(?<=<head\s*>)(?:.|\s*)*?(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Vipengele vyako vyote viwili vya `link` vinapaswa kuwa na sifa ya `rel` iliyowekwa kuwa `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

One of your `link` elements should have an `href` attribute set to `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Moja ya vipengele vyako vya `link` lazima kiwe na sifa ya `href` yenye thamani `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
