---
id: 615f3ed16592445e57941aec
title: Hatua ya 16
challengeType: 0
dashedName: step-16
---

# --description--

Mistari inaweza kusaidia kutenganisha na kupanga maudhui muhimu, hasa wakati nafasi ni chache.

Unda kipengele cha `div` chini ya kipengele chako `h1` na ukipe sifa ya `class` iliyowekwa kuwa `divider`.

# --hints--

Unapaswa kuunda `div` mpya.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

`div` yako mpya inafaa kuwa na sifa ya `class` iliyowekwa kuwa `divider`.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

Kipengele chako cha `.divider` kinapaswa kuwa ndani ya kipengele chako cha `.label`.

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

Kipengele chako cha `.divider` kinapaswa kuja baada ya kipengele chako cha `h1`.

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
