---
id: 615f34ecc1091b4fd5a8a484
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Ndani ya kipengele chako cha `head`, ongeza kipengele cha `link` chenye sifa ya `rel` iliyowekwa kuwa `stylesheet` na sifa ya `href` iliyowekwa kuwa `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Hii italeta familia ya fonti ya `Open Sans`, yenye thamani za uzito wa fonti `400`, `700` na `800`.

Pia ongeza kipengele cha `link` ili kuunganisha faili yako ya `styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na vipengele viwili vya kujifunga vya `link`.

```js
assert(document.querySelectorAll('link').length === 2);
```

Vipengele vyako vyote viwili vya `link` vinapaswa kuwa na sifa ya `rel` iliyowekwa kuwa `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Moja ya vipengele vyako vya `link` lazima kiwe na sifa ya `href` yenye thamani `./styles.css`.

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Moja ya vipengele vyako vya `link` lazima kiwe na sifa ya `href` yenye thamani `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
