---
id: 615f2abbe7d18d49a1e0e1c8
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Tumekuandalia mfano wa msingi wa HTML.

Unda kipengele cha `h1` ndani ya kipengele chako cha `body` na ukipe maandishi `Nutrition Facts`.

# --hints--

Unapaswa kuongeza kipengele kipya cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
