---
id: 615f34948891834dd77655a6
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Ongeza kipengele cha pili cha `p` chenye maandishi `Serving size 2/3 cup (55g)`.

# --hints--

Unapaswa kuongeza kipengele cha pili cha `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Kipengele chako cha pili cha `p` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Kipengele chako cha pili cha `p` kinapaswa kuja baada ya kipengele chako cha `p` kilichopo.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Kipengele chako cha pili cha `p` kinapaswa kuwa na maandishi `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
