---
id: 65ef1be2b958cabf712296cd
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер ви будете постійно звужувати інтервал, знаходячи середину поточного інтервалу та порівнюючи квадрат середини з шуканим значенням.

Для цього створіть цикл `for` в межах блоку `else`, який виконується не більше `max_iterations` разів.

# --hints--

Створіть цикл `for`, щоб ітерувати над `range(max_iterations)`. Використайте `_` як змінну циклу.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].is_equivalent("for _ in range(max_iterations):\\n    pass")`));
    }
})

```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None

--fcc-editable-region--

--fcc-editable-region--
```
