---
id: 65ef190c6b51e9b5a5f7ed29
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Надайте функції `square_root_bisection` такі параметри:

- `square_target`: число, для якого ви хочете знайти квадратний корінь;
- `tolerance` (необов’язковий параметр): припустима різниця між квадратом приблизного значення кореня та фактичним шуканим значенням (за замовчуванням `1e-7`).  Допуск `1e-7` позначає, що розв’язок буде точним у межах `0.0000001` від істинного значення і є хорошим вибором за замовчуванням, що збалансовує точність та продуктивність;
- `max_iterations` (необов’язковий параметр): максимальна кількість ітерацій для виконання (за замовчуванням `100`). Якщо метод не збігається з цим обмеженням, то припускають, що розв’язок не знайдено.

# --hints--

Функція повинна мати такі параметри: `square_target`, `tolerance = 1e-7` та `max_iterations = 100`. Порядок має значення.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("square_root_bisection").has_args("square_target, tolerance=1e-7, max_iterations=100")`))
})
```

# --seed--

## --seed-contents--

```py

--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

