---
id: 65ef198fde24dfb7ff675b42
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Інструкція `raise` дозволяє видати конкретний виняток. Вона складається з ключового слова `raise`, після якого вказано тип винятку, та дозволяє надати власне повідомлення для користувача:

```py
raise ValueError("Invalid value")
```

Коли код вище запускається, виникає помилка `ValueError`, а користувачу відтворюється повідомлення `"Invalid value"`.

Якщо `square_target` менше за `0`, неможливо обчислити дійснозначний корінь. Тому має виникнути помилка `ValueError` з повідомленням `'Square root of negative number is not defined in real numbers'`. Не забудьте видалити ключове слово `pass`.

# --hints--

Видаліть ключове слово `pass`.

```js
({
    test: () => {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].has_pass()`))
    }
})
```

В межах тіла `if` має виникнути помилка `ValueError` з повідомленням `'Square root of negative number is not defined in real numbers'`.

```js
({
    test: () => {

        assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[0].find_bodies()[0].is_equivalent("raise ValueError('Square root of negative number is not defined in real numbers')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        pass

--fcc-editable-region--
```
