---
id: 65ef1cb031de30c2cde64c12
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Якщо різниця знаходиться в межах заданих значень `tolerance`, то встановіть значення `root` на `mid` та вийдіть з циклу.

# --hints--

Видаліть наявну інструкцію `pass`.

```js
({
    test: () => 
    {
        assert.isFalse(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[0].has_pass()`))
    }

})

```

Призначте значення `mid` до `root` та вийдіть з циклу.

```js
({
    test: () => 
    {
          assert(runPython(`_Node(_code).find_function("square_root_bisection").find_ifs()[1].find_bodies()[2].find_for_loops()[0].find_ifs()[0].find_bodies()[0].is_equivalent("root = mid\\nbreak")`))
    }

})
```

# --seed--

## --seed-contents--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = 0
        high = max(1, square_target)
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

--fcc-editable-region--       
            if abs(square_mid - square_target) < tolerance:
                pass

--fcc-editable-region--
```
