---
id: 65ef181dad3d96b14a21a03e
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви знайдете приблизний квадратний корінь заданого числа за допомогою методу бісекції.

Метод бісекції — це техніка для знаходження коренів дійснозначної функції. Він працює шляхом звуження інтервалу, де знаходиться корінь, поки не збіжиться до значення в межах вказаного допуску.

Для початку створіть функцію під назвою `square_root_bisection`. Наразі залиште параметри порожніми.

# --hints--

Визначте функцію `square_root_bisection` без параметрів.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
