---
id: 66da326c02141df538f29ba5
title: Build a Checkout Page
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have an `h1` element with the text `Checkout`.
1. You should have two `section` elements immediately after the `h1` element.
1. You should have an `h2` element with the text `Your Cart` within the first section.
1. You should have an image of an item in the first section with appropriate alternate text. You can use this image if you would like: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. You should have an `h2` element with the text `Payment Information` within the second section.
1. You should have a `form` element within the second section.
1. You should have an input with an `id` and `name` of `card-name` within your form and a `label` associated with it.
1. You should have an input with an `id` and `name` of `card-number` within your form and a `label` associated with it.
1. At least two of your input elements should be `required`.
1. Add a span within the `label` element for the `card-name` input, within the span, it should include an asterisk (*) and an aria-hidden attribute of true.
1. Add a span within the `label` element for the `card-number` input, within the span, it should include an asterisk (*) and an aria-hidden attribute of true.
1. Add an input element with an id of `password` and a `type` of `password`. With a `aria-describedby` attribute of `passwordHelp`.
1. Add a `label` element for the `password` input with an id of `passwordHelp`.
1. Add a `p` tag with the id of `passwordHelp` with the following text: `Your password must be at least 8 characters long`.


# --hints--

You should have an `h1` element with the text `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

You should only have one `h1` element on your page.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

You should have at least two `section` elements after the `h1` element.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

You should have an `h2` element with the text `Your Cart` within the first section.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

You should have at least one image with alternate text in your first section.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

You should have an `h2` element with the text `Payment Information` within the second section.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

You should have a `form` element within the second `section` element.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

You should have an `input` element with an `id` and `name` of `card-name` within your form.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

You should have an `input` element with an `id` and `name` of `card-number` within your form.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

All of your `input` elements that aren't a `type` of `submit` should have a `label` element associated with them.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

At least two of your `input` elements should be required.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

At least two of your `span` elements should have an `aria-hidden` attribute set to `true`.

```js
const spans = document.querySelectorAll('span[aria-hidden="true"]');
assert.isAtLeast(spans.length, 2, 'At least two spans should have aria-hidden="true"');
```

At least one of your `input` elements should have an `aria-describedby` attribute.

```js
const passwordInput = document.querySelector('input#password');
assert.exists(passwordInput?.getAttribute('aria-describedby'));
```

At least one of the `p` elements should have the text `Your password must be at least 8 characters long.`.

```js
const passwordHelp = document.querySelector('p#passwordHelp');
assert.exists(passwordHelp);
assert.equal(passwordHelp?.innerText, 'Your password must be at least 8 characters long.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <h2>Login</h2>
                    <label for="username">Username:</label>
                    <input id="username" type="text" aria-describedby="usernameHelp" />
                    <p id="usernameHelp">Your username must be at least 8 characters long.</p>
                </div>
                <div>
                    <label for="password">Password:</label>
                    <input id="password" type="password" aria-describedby="passwordHelp" />
                    <p id="passwordHelp">Your password must be at least 8 characters long.</p>
                </div>
                <div>
                    <label for="card-name">
                        Cardholder Name
                        <span aria-hidden="true">*</span>
                    </label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">
                        Card Number
                        <span aria-hidden="true">*</span>
                    </label>
                    <input type="number" id="card-number" name="card-number" required>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="number" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
