---
id: 5dc17dc8f86c76b9248c6eb5
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Commenting allows you to leave messages without affecting the browser display. It also allows you to make code inactive. A comment in HTML starts with `<!--`, contains any number of lines of text, and ends with `-->`. 

Here is an example of a comment with the `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

In this step, add the string "Everyone loves cute cats online!" inside the paragraph (`p`) element.

# --hints--

Your paragraph should contain the string "Everyone loves cute cats online!".

```js
assert(code.match(/<p>\s*Everyone\s+loves\s+cute\s+cats\s+online!\s*<\/p>/i));
```

The `p` element should be inside the editable region.

```js
assert(code.match(/--fcc-editable-region--\s*<p>Everyone\s+loves\s+cute\s+cats\s+online!\s*<\/p>\s*--fcc-editable-region--/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```

