---
id: 5dc23991f86c76b9248c6eb8
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In the previous step, you placed an `h2`, comment, and `p` within the opening and closing `main` tags. These three things are all nested within the `main` element. For readability purposes, nested elements like this should have two additional spaces in front of them than the element they are nested in. In the example below, the `body` element is two spaces over from the `html` element and the `h1` element is two spaces over from the `body` element.

```html
<html>
  <body> 
    <h1>Hello World</h1>
  </body>
</html>
```

Add spaces in front of your `h2` element, comment, and `p` element so they have two spaces more than are in front of the `main` tags.

# --hints--

Your code should still have a `main` element.  Make sure you did not move or delete it.

```js
assert(document.querySelector('main') && code.match(/<\/main>/));
```

Your `main` element should not be moved. It should be located `--fcc-expected--` spaces over from the start of the line but it is currently `--fcc-actual--` spaces over.

```js
const mainMatch = code.toLowerCase().match(/( *)<main>/);
const mainLocation = mainMatch ? mainMatch[1].length : 0;
assert.equal(mainLocation, 4);
```

Your code should have an `h2` element with text `Cat Photos`. You may have accidentally deleted it, it is missing both opening and closing tags, or the text has changed.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

The `h2` element should be on the line below the `main` element's opening tag.

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s*<h2>/));
```

The `h2` element should be located `--fcc-expected--` spaces over from the start of the line but it is currently `--fcc-actual--` spaces over.

```js
const h2Match = code.toLowerCase().match(/( *)<h2>/);
const h2Location = h2Match[1].length;
assert.equal(h2Location, 6);
```

Your code should have a comment. You removed the comment from an earlier step.

```js
assert(code.match(/<!--.*-->/));
```

The comment's text should be `TODO: Add link to cat photos`. Do not change the text or spacing of the comment.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

The comment should be below the `h2` element.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s*<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

The comment should be located `--fcc-expected--` spaces over from the start of the line but it is currently `--fcc-actual--` spaces over.

```js
const commentMatch = code.toLowerCase().match(/( *)<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/);
const commentLocation = commentMatch[1].length;
assert.equal(commentLocation, 6);
```

Your code should have a `p` element. You have removed the `p` element from an earlier step.

```js
assert(document.querySelector('p'));
```

The text of the `p` element should be `Click here to view more cat photos.` Do not change the text, spacing, or punctuation of the `p` element.

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

The `p` element should be below the comment.

```js
assert(code.toLowerCase().match(/-->\s*\n\s*<p>/));
```

The `p` element should be located `--fcc-expected--` spaces over from the start of the line but it is currently `--fcc-actual--` spaces over.

```js
const pMatch = code.toLowerCase().match(/( *)<p>/);
const pLocation = pMatch[1].length;
assert.equal(pLocation, 6);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
