---
id: 5dc23991f86c76b9248c6eb8
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In the previous step, you nested the `h2` element, comment, and `p` element within the `main` element. Indenting nested elements two more spaces than their parent element improves readability:

```html
<ul> 
  <li> This `li` element indented </li>
  <li> This `li` element is also indented </li>
</ul>
```

Add two more spaces in front of the `h2`, comment, and `p` elements so your HTML is more readable.

# --hints--

Your code should have an `h2` element with text `Cat Photos`. You may have accidentally deleted it, it is missing both opening and closing tags, or the text has changed.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

You should not add the `ul` or `li` elements from the example.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

Your `h2` element should be below the `main` element's opening tag and its opening tag should start 6 spaces over from the start of the line.

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s{6}<h2>/));
```

Your code should have a comment. You removed the comment from an earlier step.

```js
assert(code.match(/<!--.*-->/));
```

The comment's text should be `TODO: Add link to cat photos`. Do not change the text or spacing of the comment.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Your comment should be below the `h2` element and start 6 spaces over from the start of the line.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Your code should have a `p` element. You have removed the `p` element from an earlier step.

```js
assert(document.querySelector('p'));
```

The text of the `p` element should be `Click here to view more cat photos.` Do not change the text, spacing, or punctuation of the `p` element.

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

Your `p` element should be below the comment and its opening tag should start 6 spaces over from the start of the line.

```js
assert(code.toLowerCase().match(/-->\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

