---
id: 5dfa22d1b521be39a3de7be0
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

In the previous steps you turned the words `cat photos` into a link by surrounding them with an opening `<a>` tag and closing `</a>` tag. In the preview you can see the words `cat photos` in blue and underlined next to the image. That's your link; feel free to give it a click.

You can turn words inside of a `p` element into a link by surrounding them with an opening `<a>` tag and closing `</a>` tag. Inside the `p` element, turn the words `cat photos` into a link to `https://freecatphotoapp.com`.

Do not forget to remove the `<a>` to left of the cat image. We do not need to links.

# --hints--

Your code should only contain one anchor (`a`) element. Remove any extra anchor elements.

```js
assert(document.querySelectorAll('a').length === 1);
```

Your anchor (`a`) element should be nested within the `p` element.

```js
assert($('p > a').length);
```

The link's href value should be `https://freecatphotoapp.com`. You have either omitted the href value or have a typo.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

The link's text should be `cat photos`. You have either omitted the text or have a typo.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

After nesting the anchor (`a`) element, the only `p` element content visible in the browser should be `Click here to view more cat photos.` Double check the text, spacing, or punctuation of both the `p` and nested anchor element.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

