---
id: 5dfa2407b521be39a3de7be1
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

To enhance the user experience, you will now turn the text "cute cats" into a clickable link that directs users to a different picture. You will achieve this by wrapping the text "cute cats" in an anchor (`a`) element with the appropriate `href` attribute pointing to the desired image URL.

Here’s the basic syntax for creating a link:

```html
<a href="URL">text</a>
```

In this case, the URL should be set to the address of the picture you want to link to, such as `https://cdn.freecodecamp.org/curriculum/cat-photo-app/cute-cat.jpg`.

# --hints--

Your `p` element should have a nested anchor (`a`) element with the text `cat photos`. You may have deleted it or have a typo.

```js
const anchor = document.querySelectorAll('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Your anchor (`a`) element does not have a `target` attribute. Check that there is a space after the opening tag's name and/or there are spaces before all attribute names.

```js
assert(document.querySelector('a').hasAttribute('target'));
```

The value of the `target` attribute should be `_blank`. You have either omitted the value or have a typo. Remember that attribute values should be surrounded with quotation marks.

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
--fcc-editable-region--
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cute-cat.jpg" >cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

