---
id: 5dfa22d1b521be39a3de7be0
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

You can see the words `cat photos` underlined next to the image in the preview now. That's your link; feel free to give it a click. In the text of your `p` element, turn the words `cat photos` into a link to `https://freecatphotoapp.com`. When you are done, remove the old anchor tag and text below the paragraph.

# --hints--

Your anchor (`a`) element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Your anchor (`a`) element should have a closing tag. Closing tags have a `/` just after the `<` character.

```js
assert(code.match(/<\/a\>/));
```

Your anchor (`a`) element should be nested within the `p` element.

```js
assert(document.querySelectorAll('p > a').length);
```

Although you have set the anchor (`a`) element's `href` attribute to the correct link, it is recommended to always surround the value of an attribute with quotation marks.

```js
assert(!/\<a\s+href\s*=\s*https:\/\/freecatphotoapp\.com/.test(code));
```

The link's href value should be `https://freecatphotoapp.com`. You have either omitted the href value or have a typo.

```js
const nestedAnchor = document.querySelectorAll('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

The link's text should be `cat photos`. You have either omitted the text or have a typo.

```js
const nestedAnchor = document.querySelectorAll('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

After nesting the anchor (`a`) element, the only `p` element content visible in the browser should be `Click here to view more cat photos.` Double check the text, spacing, or punctuation of both the `p` and nested anchor element.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/^click here to view more cat photos\.?$/));
```

Your code should only contain one anchor (`a`) element. Remove any extra anchor elements.

```js
assert(document.querySelectorAll('a').length === 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

