---
id: 5dfa22d1b521be39a3de7be0
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Add the words `"See more "` before the anchor element and `" in our gallery."` after the anchor element.


# --hints--

The link's `href` value should be `https://freecatphotoapp.com`. You have either omitted the `href` value or have a typo.

```js
assert(
  document.querySelector('a').href === 'https://freecatphotoapp.com/'
);
```

The link's text should be `cat photos`. You have either omitted the text or have a typo.

```js
const nestedAnchor = document.querySelector('a');
const innerContent = nestedAnchor.innerHTML;
assert.isTrue(innerContent.trim() === 'cat photos');
```

The text inside your anchor element has extra leading or trailing whitespace. The only space in the anchor text should be between the word `cat` and the word `photos`.

```js
const nestedAnchor = document.querySelector('a');
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

After nesting the anchor (`a`) element, the only content visible in the browser should be `See more cat photos in our gallery.` Double check the text, spacing, or punctuation of both the text and nested anchor element.

```js
assert.match(code, /see more <a[^>]*>cat photos<\/a> in our gallery\.?/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

