---
id: 62dabe2ef403a12d5d295273
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Add `p` tags to turn the text with the anchor element into a paragraph.

# --hints--

Your code should have two paragraph (`p`) elements.

```js
assert.isTrue(document.querySelectorAll('p').length === 2);
```

Your code should only have one anchor (`a`) element.

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

Your `p` element's text and anchor (`a`) element should be `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.` You have either omitted the text or have a typo.

```js
assert(/\<p>See more <a href="https:\/\/freecatphotoapp.com">cat photos<\/a> in our gallery./.test(code));
```

There should not be `cat photos` text below the `p` element.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
