---
id: 668bae23be54d01f57738a11
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

You can turn any text into a link, such as the text inside of a `p` element.

```html
<p>I think <a href="https://www.freecodecamp.org">freeCodeCamp</a> is great.</p>
```

In the text of your first `p` element, turn the words `cute cats` into a link by adding opening and closing anchor (`a`) tags around these words. Then set the `href` attribute to `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg`

# --hints--

Your code should have two paragraph (`p`) elements.

```js
assert.isTrue(document.querySelectorAll('p').length === 2);
```

Your code should have two anchor (`a`) elements.

```js
assert.isTrue(document.querySelectorAll('a').length === 2);
```

Do not make any changes to the elements you created in the previous step.

```js
assert(/\<p>See\s+more\s+\<a href="https:\/\/freecatphotoapp.com">cat\s+photos<\/a>\s+in\s+our\s+gallery./.test(code));
```

The link's `href` value should be `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg`. You have either omitted the `href` value or have a typo.

```js
assert(
  document.querySelector('a').href === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
