---
id: 66faa70000000000000000a8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
A soundboard is nothing without sounds. Load all six SFX plus the ambient music track in `preload`, with both `.ogg` and `.mp3` fallbacks so the asset works in every browser.

Declare a local `SFX_KEYS` array, iterate it, and call `this.load.audio(key, [...])` for each in array form so the browser can pick its preferred format. Then load the `"music-ambient"` track too.

```js
const SFX_KEYS = [
  "sfx-jump",
  "sfx-coin",
  "sfx-hit",
  "sfx-explosion",
  "sfx-ui-click",
  "sfx-ping"
];
SFX_KEYS.forEach((key) => {
  this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
});
this.load.audio("music-ambient", [
  "audio/music-ambient.ogg",
  "audio/music-ambient.mp3"
]);
```

Listing `.ogg` first lets Chrome and Firefox win on file size while Safari falls through to the `.mp3` fallback.

# --hints--
Inside `preload`, declare a local `SFX_KEYS` array with the six sfx keys.

```js
assert.match(code, /const\s+SFX_KEYS\s*=\s*\[\s*["\']sfx-jump["\']\s*,\s*["\']sfx-coin["\']\s*,\s*["\']sfx-hit["\']\s*,\s*["\']sfx-explosion["\']\s*,\s*["\']sfx-ui-click["\']\s*,\s*["\']sfx-ping["\']\s*\]/);
```

Call `this.load.audio(key, [...ogg, mp3])` for each SFX key in a `forEach`.

```js
assert.match(code, /SFX_KEYS\.forEach\([\s\S]*this\.load\.audio\(\s*key\s*,\s*\[\s*`audio\/\$\{key\}\.ogg`\s*,\s*`audio\/\$\{key\}\.mp3`\s*\]\s*\)/);
```

Load the music track with `this.load.audio("music-ambient", [...ogg, mp3])`.

```js
assert.match(code, /this\.load\.audio\(\s*["\']music-ambient["\']\s*,\s*\[\s*["\']audio\/music-ambient\.ogg["\']\s*,\s*["\']audio\/music-ambient\.mp3["\']\s*\]\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
