---
id: 66faa70000000000000000ac
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Add a one-shot `pointerdown` listener that calls `this.sound.unlock()` if the manager is locked, then destroys both the overlay rectangle and its text.

After the overlay block:

```js
this.input.once("pointerdown", () => {
  if (this.sound.locked) this.sound.unlock();
  this.overlay.destroy();
  this.overlayText.destroy();
});
```

`this.sound.locked` is `true` until the first user gesture; `unlock()` resumes the suspended `AudioContext`. After this fires, every subsequent `sound.play` works normally.

# --hints--
Use `this.input.once("pointerdown", ...)` so the listener fires exactly once.

```js
assert.match(code, /this\.input\.once\(\s*["\']pointerdown["\']/);
```

Inside the handler, call `this.sound.unlock()` only when `this.sound.locked` is true.

```js
assert.match(code, /if\s*\(\s*this\.sound\.locked\s*\)\s*this\.sound\.unlock\(\s*\)/);
```

Destroy both `this.overlay` and `this.overlayText`.

```js
assert.match(code, /this\.overlay\.destroy\(\s*\)[\s\S]*this\.overlayText\.destroy\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.sfxVolume = 0.6;
    this.musicVolume = 0.5;
    this.playSfx = (key) => this.sound.play(key);
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
    this.overlay = this.add
      .rectangle(0, 0, 600, 480, 0x000000, 0.85)
      .setOrigin(0)
      .setDepth(1000);
    this.overlayText = this.add
      .text(300, 240, "TAP TO ENABLE AUDIO", {
        fontSize: "24px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(1001);
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
