---
id: 66faa70000000000000000b0
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Add an SFX volume slider. We'll need a slider helper for several controls, so define `this.makeSlider` first, then add the slider and finally extend `this.playSfx` to pass the live volume.

Replace the line `this.playSfx = (key) => this.sound.play(key);` with the entire block below. The four parts: a `this.makeSlider(x, y, w, h, initial, onChange)` helper that draws a track and clamped knob, an `"SFX Vol"` label, the slider itself, and a new `this.playSfx` that passes `{ volume: this.sfxVolume }` to `this.sound.play`.

```js
this.makeSlider = (x, y, w, h, initial, onChange) => {
  const track = this.add
    .rectangle(x, y, w, h, 0x222244)
    .setOrigin(0)
    .setStrokeStyle(1, 0x6c8aff);
  const knob = this.add
    .rectangle(x + initial * w - 4, y - 2, 8, h + 4, 0xffd166)
    .setOrigin(0);
  track.setInteractive();
  track.on("pointerdown", (pointer) => {
    const v = Phaser.Math.Clamp((pointer.x - x) / w, 0, 1);
    knob.x = x + v * w - 4;
    onChange(v);
  });
  return { track, knob };
};
this.add.text(8, 108, "SFX Vol", {
  fontSize: "10px",
  color: "#ffffff"
});
this.makeSlider(80, 102, 300, 16, this.sfxVolume, (v) => {
  this.sfxVolume = v;
});
this.playSfx = (key) =>
  this.sound.play(key, { volume: this.sfxVolume });
```

Click anywhere on the track and the knob jumps to your click (no drag yet) and the SFX get louder or quieter from the next press onward.

# --hints--
Define `this.makeSlider` as an arrow function taking `(x, y, w, h, initial, onChange)`.

```js
assert.match(code, /this\.makeSlider\s*=\s*\(\s*x\s*,\s*y\s*,\s*w\s*,\s*h\s*,\s*initial\s*,\s*onChange\s*\)\s*=>/);
```

The track rectangle should use `(x, y, w, h, 0x222244)` with `setStrokeStyle(1, 0x6c8aff)` and a `0xffd166` knob 8px wide.

```js
assert.match(code, /\.rectangle\(\s*x\s*,\s*y\s*,\s*w\s*,\s*h\s*,\s*0x222244\s*\)[\s\S]*\.setStrokeStyle\(\s*1\s*,\s*0x6c8aff\s*\)[\s\S]*\.rectangle\(\s*x\s*\+\s*initial\s*\*\s*w\s*-\s*4\s*,\s*y\s*-\s*2\s*,\s*8\s*,\s*h\s*\+\s*4\s*,\s*0xffd166\s*\)/);
```

On `pointerdown`, clamp `(pointer.x - x) / w` to `0..1` with `Phaser.Math.Clamp` and call `onChange(v)`.

```js
assert.match(code, /Phaser\.Math\.Clamp\(\s*\(\s*pointer\.x\s*-\s*x\s*\)\s*\/\s*w\s*,\s*0\s*,\s*1\s*\)[\s\S]*onChange\(\s*v\s*\)/);
```

Build a slider at `(80, 102, 300, 16, this.sfxVolume, (v) => { this.sfxVolume = v; })` and an `"SFX Vol"` label at `(8, 108)`.

```js
assert.match(code, /this\.add\.text\(\s*8\s*,\s*108\s*,\s*["\']SFX Vol["\'][\s\S]*this\.makeSlider\(\s*80\s*,\s*102\s*,\s*300\s*,\s*16\s*,\s*this\.sfxVolume/);
```

Replace `this.playSfx` to pass `{ volume: this.sfxVolume }` to `this.sound.play`.

```js
assert.match(code, /this\.playSfx\s*=\s*\(\s*key\s*\)\s*=>\s*this\.sound\.play\(\s*key\s*,\s*\{\s*volume:\s*this\.sfxVolume\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.sfxVolume = 0.6;
    this.musicVolume = 0.5;
--fcc-editable-region--
    this.playSfx = (key) => this.sound.play(key);
--fcc-editable-region--
    this.makeButton = (x, y, w, h, label, onClick) => {
      const bg = this.add
        .rectangle(x, y, w, h, 0x222244)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      const text = this.add
        .text(x + w / 2, y + h / 2, label, {
          fontSize: "12px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", onClick);
      return { bg, label: text };
    };
    this.music = this.sound.add("music-ambient", {
      loop: true,
      volume: this.musicVolume
    });
    this.musicBtn = this.makeButton(
      8,
      56,
      140,
      32,
      "Music: Play",
      () => {
        if (this.music.isPlaying) {
          this.music.pause();
          this.musicBtn.label.setText("Music: Resume");
        } else if (this.music.isPaused) {
          this.music.resume();
          this.musicBtn.label.setText("Music: Pause");
        } else {
          this.music.play();
          this.musicBtn.label.setText("Music: Pause");
        }
      }
    );
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
    this.overlay = this.add
      .rectangle(0, 0, 600, 480, 0x000000, 0.85)
      .setOrigin(0)
      .setDepth(1000);
    this.overlayText = this.add
      .text(300, 240, "TAP TO ENABLE AUDIO", {
        fontSize: "24px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(1001);
    this.input.once("pointerdown", () => {
      if (this.sound.locked) this.sound.unlock();
      this.overlay.destroy();
      this.overlayText.destroy();
    });
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
