---
id: 66faa70000000000000000a9
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Now build the visible row of six sound-effect buttons. Each is a 90x36 rectangle with the SFX name centered on top, laid out in a horizontal row across the top of the canvas.

In `create`, redeclare `SFX_KEYS` (we'll de-duplicate later) and `forEach` it. For each key, build a `rectangle` at `(8 + i * 95, 8)` and a centered `text` showing the key with its `sfx-` prefix stripped.

```js
const SFX_KEYS = [
  "sfx-jump",
  "sfx-coin",
  "sfx-hit",
  "sfx-explosion",
  "sfx-ui-click",
  "sfx-ping"
];
SFX_KEYS.forEach((key, i) => {
  const x = 8 + i * 95;
  const bg = this.add
    .rectangle(x, 8, 90, 36, 0x2a3045)
    .setOrigin(0)
    .setStrokeStyle(2, 0x6c8aff);
  this.add
    .text(x + 45, 26, key.replace("sfx-", ""), {
      fontSize: "10px",
      color: "#ffffff"
    })
    .setOrigin(0.5);
});
```

The row spans `8 + 95 * 5 + 90 = 573` pixels - well within the 600px canvas.

# --hints--
Inside `create`, declare a local `SFX_KEYS` array with the six sfx keys.

```js
assert.match(code, /create\s*\(\s*\)\s*\{[\s\S]*const\s+SFX_KEYS\s*=\s*\[\s*["\']sfx-jump["\']/);
```

Iterate `SFX_KEYS` with `forEach((key, i) => ...)` and build a 90x36 rectangle at `(8 + i * 95, 8)`.

```js
assert.match(code, /SFX_KEYS\.forEach\(\s*\(\s*key\s*,\s*i\s*\)[\s\S]*const\s+x\s*=\s*8\s*\+\s*i\s*\*\s*95[\s\S]*\.rectangle\(\s*x\s*,\s*8\s*,\s*90\s*,\s*36\s*,\s*0x2a3045\s*\)/);
```

Add a centered `text` at `(x + 45, 26)` showing `key.replace("sfx-", "")`.

```js
assert.match(code, /\.text\(\s*x\s*\+\s*45\s*,\s*26\s*,\s*key\.replace\(\s*["\']sfx-["\']\s*,\s*["\']{2}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
