---
id: 66faa70000000000000000af
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--
Right now clicking the music button after it's already playing creates a *new* `Sound` instance every time. Replace the click handler with a tri-state toggle: pause when playing, resume when paused, otherwise play. Update the button label each time.

Replace the click handler line `() => this.music.play()` with:

```js
() => {
  if (this.music.isPlaying) {
    this.music.pause();
    this.musicBtn.label.setText("Music: Resume");
  } else if (this.music.isPaused) {
    this.music.resume();
    this.musicBtn.label.setText("Music: Pause");
  } else {
    this.music.play();
    this.musicBtn.label.setText("Music: Pause");
  }
}
```

# --hints--
When `this.music.isPlaying`, call `this.music.pause()` and set the label to `"Music: Resume"`.

```js
assert.match(code, /if\s*\(\s*this\.music\.isPlaying\s*\)[\s\S]*this\.music\.pause\(\s*\)[\s\S]*this\.musicBtn\.label\.setText\(\s*[\"']Music: Resume[\"']\s*\)/);
```

When `this.music.isPaused`, call `this.music.resume()` and set the label to `"Music: Pause"`.

```js
assert.match(code, /else\s+if\s*\(\s*this\.music\.isPaused\s*\)[\s\S]*this\.music\.resume\(\s*\)[\s\S]*this\.musicBtn\.label\.setText\(\s*[\"']Music: Pause[\"']\s*\)/);
```

Otherwise, call `this.music.play()` and set the label to `"Music: Pause"`.

```js
assert.match(code, /else\s*\{\s*this\.music\.play\(\s*\)\s*;?\s*this\.musicBtn\.label\.setText\(\s*[\"']Music: Pause[\"']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.sfxVolume = 0.6;
    this.musicVolume = 0.5;
    this.playSfx = (key) => this.sound.play(key);
    this.makeButton = (x, y, w, h, label, onClick) => {
      const bg = this.add
        .rectangle(x, y, w, h, 0x222244)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      const text = this.add
        .text(x + w / 2, y + h / 2, label, {
          fontSize: "12px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", onClick);
      return { bg, label: text };
    };
    this.music = this.sound.add("music-ambient", {
      loop: true,
      volume: this.musicVolume
    });
    this.musicBtn = this.makeButton(
      8,
      56,
      140,
      32,
      "Music: Play",
--fcc-editable-region--
      () => this.music.play()
--fcc-editable-region--
    );
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
    this.overlay = this.add
      .rectangle(0, 0, 600, 480, 0x000000, 0.85)
      .setOrigin(0)
      .setDepth(1000);
    this.overlayText = this.add
      .text(300, 240, "TAP TO ENABLE AUDIO", {
        fontSize: "24px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(1001);
    this.input.once("pointerdown", () => {
      if (this.sound.locked) this.sound.unlock();
      this.overlay.destroy();
      this.overlayText.destroy();
    });
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
