---
id: 66faa70000000000000000b6
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--
Add a rate slider that drives `pingRate`, then rebuild the ping with both `volume` AND `rate`.

Replace the entire pan block from step 15 with the version below. The diff: a `this.pingRate = 1` initialiser, `rate: this.pingRate` added to the `sound.add` config, and a brand-new rate slider after the pan slider:

```js
this.pingPan = 0;
this.pingRate = 1;
this.playPing = () => {
  const ping = this.sound.add("sfx-ping", {
    volume: this.sfxVolume,
    rate: this.pingRate
  });
  ping.setPan(this.pingPan);
  ping.play();
  ping.once("complete", () => ping.destroy());
};
this.add.text(8, 188, "Pan", {
  fontSize: "10px",
  color: "#ffffff"
});
this.makeSlider(80, 182, 300, 16, 0.5, (v) => {
  this.pingPan = (v - 0.5) * 2;
  this.playPing();
});
this.add.text(8, 228, "Rate", {
  fontSize: "10px",
  color: "#ffffff"
});
this.makeSlider(80, 222, 300, 16, 1 / 3, (v) => {
  this.pingRate = 0.5 + v * 1.5;
  this.playPing();
});
```

The rate slider's `1/3` initial position maps to `pingRate = 1` (`0.5 + 1/3 * 1.5 = 1.0`). Drag left for slow + low, drag right for fast + high.

# --hints--
`this.playPing` should pass `rate: this.pingRate` in the `sound.add` config.

```js
assert.match(code, /this\.sound\.add\(\s*[\"']sfx-ping[\"']\s*,\s*\{\s*volume:\s*this\.sfxVolume\s*,\s*rate:\s*this\.pingRate/);
```

Initialise `this.pingRate = 1`.

```js
assert.match(code, /this\.pingRate\s*=\s*1\b/);
```

Add a `"Rate"` slider at `(80, 222, 300, 16, 1 / 3, ...)` whose handler sets `this.pingRate = 0.5 + v * 1.5` and calls `this.playPing()`.

```js
assert.match(code, /this\.makeSlider\(\s*80\s*,\s*222\s*,\s*300\s*,\s*16\s*,\s*1\s*\/\s*3[\s\S]*this\.pingRate\s*=\s*0\.5\s*\+\s*v\s*\*\s*1\.5[\s\S]*this\.playPing\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.sfxVolume = parseFloat(
      localStorage.getItem("sb:sfxVol") ?? "0.6"
    );
    this.musicVolume = parseFloat(
      localStorage.getItem("sb:musicVol") ?? "0.5"
    );
    this.makeSlider = (x, y, w, h, initial, onChange) => {
      const track = this.add
        .rectangle(x, y, w, h, 0x222244)
        .setOrigin(0)
        .setStrokeStyle(1, 0x6c8aff);
      const knob = this.add
        .rectangle(x + initial * w - 4, y - 2, 8, h + 4, 0xffd166)
        .setOrigin(0);
      track.setInteractive();
      track.on("pointerdown", (pointer) => {
        const v = Phaser.Math.Clamp((pointer.x - x) / w, 0, 1);
        knob.x = x + v * w - 4;
        onChange(v);
      });
      return { track, knob };
    };
    this.add.text(8, 108, "SFX Vol", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.makeSlider(80, 102, 300, 16, this.sfxVolume, (v) => {
      this.sfxVolume = v;
      localStorage.setItem("sb:sfxVol", String(v));
    });
    this.playSfx = (key) =>
      this.sound.play(key, { volume: this.sfxVolume });
    this.add.text(8, 148, "Music Vol", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.makeSlider(80, 142, 300, 16, this.musicVolume, (v) => {
      this.musicVolume = v;
      this.music.setVolume(v);
      localStorage.setItem("sb:musicVol", String(v));
    });
    this.makeButton = (x, y, w, h, label, onClick) => {
      const bg = this.add
        .rectangle(x, y, w, h, 0x222244)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      const text = this.add
        .text(x + w / 2, y + h / 2, label, {
          fontSize: "12px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", onClick);
      return { bg, label: text };
    };
    this.music = this.sound.add("music-ambient", {
      loop: true,
      volume: this.musicVolume
    });
    this.musicBtn = this.makeButton(
      8,
      56,
      140,
      32,
      "Music: Play",
      () => {
        if (this.music.isPlaying) {
          this.music.pause();
          this.musicBtn.label.setText("Music: Resume");
        } else if (this.music.isPaused) {
          this.music.resume();
          this.musicBtn.label.setText("Music: Pause");
        } else {
          this.music.play();
          this.musicBtn.label.setText("Music: Pause");
        }
      }
    );
    this.muteBtn = this.makeButton(
      158,
      56,
      140,
      32,
      "Mute: Off",
      () => {
        this.sound.mute = !this.sound.mute;
        this.muteBtn.label.setText(
          this.sound.mute ? "Mute: On" : "Mute: Off"
        );
      }
    );
--fcc-editable-region--
    this.pingPan = 0;
    this.playPing = () => {
      const ping = this.sound.add("sfx-ping", {
        volume: this.sfxVolume
      });
      ping.setPan(this.pingPan);
      ping.play();
      ping.once("complete", () => ping.destroy());
    };
    this.add.text(8, 188, "Pan", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.makeSlider(80, 182, 300, 16, 0.5, (v) => {
      this.pingPan = (v - 0.5) * 2;
      this.playPing();
    });
--fcc-editable-region--
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
    this.overlay = this.add
      .rectangle(0, 0, 600, 480, 0x000000, 0.85)
      .setOrigin(0)
      .setDepth(1000);
    this.overlayText = this.add
      .text(300, 240, "TAP TO ENABLE AUDIO", {
        fontSize: "24px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(1001);
    this.input.once("pointerdown", () => {
      if (this.sound.locked) this.sound.unlock();
      this.overlay.destroy();
      this.overlayText.destroy();
    });
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
