---
id: 66faa70000000000000000b3
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--
Persisting only matters if we read the saved value back. Replace the two volume initialisers at the top of `create` to read from `localStorage` with sane defaults.

Replace:

```js
this.sfxVolume = 0.6;
this.musicVolume = 0.5;
```

with:

```js
this.sfxVolume = parseFloat(
  localStorage.getItem("sb:sfxVol") ?? "0.6"
);
this.musicVolume = parseFloat(
  localStorage.getItem("sb:musicVol") ?? "0.5"
);
```

Reload the preview after touching the sliders and the saved values come back.

# --hints--
Initialise `this.sfxVolume` from `parseFloat(localStorage.getItem("sb:sfxVol") ?? "0.6")`.

```js
assert.match(code, /this\.sfxVolume\s*=\s*parseFloat\(\s*localStorage\.getItem\(\s*["\']sb:sfxVol["\']\s*\)\s*\?\?\s*["\']0\.6["\']\s*\)/);
```

Initialise `this.musicVolume` from `parseFloat(localStorage.getItem("sb:musicVol") ?? "0.5")`.

```js
assert.match(code, /this\.musicVolume\s*=\s*parseFloat\(\s*localStorage\.getItem\(\s*["\']sb:musicVol["\']\s*\)\s*\?\?\s*["\']0\.5["\']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
--fcc-editable-region--
    this.sfxVolume = 0.6;
    this.musicVolume = 0.5;
--fcc-editable-region--
    this.makeSlider = (x, y, w, h, initial, onChange) => {
      const track = this.add
        .rectangle(x, y, w, h, 0x222244)
        .setOrigin(0)
        .setStrokeStyle(1, 0x6c8aff);
      const knob = this.add
        .rectangle(x + initial * w - 4, y - 2, 8, h + 4, 0xffd166)
        .setOrigin(0);
      track.setInteractive();
      track.on("pointerdown", (pointer) => {
        const v = Phaser.Math.Clamp((pointer.x - x) / w, 0, 1);
        knob.x = x + v * w - 4;
        onChange(v);
      });
      return { track, knob };
    };
    this.add.text(8, 108, "SFX Vol", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.makeSlider(80, 102, 300, 16, this.sfxVolume, (v) => {
      this.sfxVolume = v;
      localStorage.setItem("sb:sfxVol", String(v));
    });
    this.playSfx = (key) =>
      this.sound.play(key, { volume: this.sfxVolume });
    this.add.text(8, 148, "Music Vol", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.makeSlider(80, 142, 300, 16, this.musicVolume, (v) => {
      this.musicVolume = v;
      this.music.setVolume(v);
      localStorage.setItem("sb:musicVol", String(v));
    });
    this.makeButton = (x, y, w, h, label, onClick) => {
      const bg = this.add
        .rectangle(x, y, w, h, 0x222244)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      const text = this.add
        .text(x + w / 2, y + h / 2, label, {
          fontSize: "12px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", onClick);
      return { bg, label: text };
    };
    this.music = this.sound.add("music-ambient", {
      loop: true,
      volume: this.musicVolume
    });
    this.musicBtn = this.makeButton(
      8,
      56,
      140,
      32,
      "Music: Play",
      () => {
        if (this.music.isPlaying) {
          this.music.pause();
          this.musicBtn.label.setText("Music: Resume");
        } else if (this.music.isPaused) {
          this.music.resume();
          this.musicBtn.label.setText("Music: Pause");
        } else {
          this.music.play();
          this.musicBtn.label.setText("Music: Pause");
        }
      }
    );
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
    this.overlay = this.add
      .rectangle(0, 0, 600, 480, 0x000000, 0.85)
      .setOrigin(0)
      .setDepth(1000);
    this.overlayText = this.add
      .text(300, 240, "TAP TO ENABLE AUDIO", {
        fontSize: "24px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(1001);
    this.input.once("pointerdown", () => {
      if (this.sound.locked) this.sound.unlock();
      this.overlay.destroy();
      this.overlayText.destroy();
    });
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
