---
id: 66faa70000000000000000ab
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
Mobile browsers refuse to start audio that wasn't triggered by a user gesture, so right now your SFX buttons may be silent on iOS. Add a full-canvas overlay with a `TAP TO ENABLE AUDIO` label.

After the SFX `forEach`:

```js
this.overlay = this.add
  .rectangle(0, 0, 600, 480, 0x000000, 0.85)
  .setOrigin(0)
  .setDepth(1000);
this.overlayText = this.add
  .text(300, 240, "TAP TO ENABLE AUDIO", {
    fontSize: "24px",
    color: "#ffffff"
  })
  .setOrigin(0.5)
  .setDepth(1001);
```

Depth `1000`/`1001` keeps the overlay on top of every game object you'll add later.

# --hints--
Add `this.overlay` as a 600x480 rectangle at `(0, 0)` with `0x000000, 0.85` and depth `1000`.

```js
assert.match(code, /this\.overlay\s*=\s*this\.add\s*\.rectangle\(\s*0\s*,\s*0\s*,\s*600\s*,\s*480\s*,\s*0x000000\s*,\s*0\.85\s*\)[\s\S]*\.setDepth\(\s*1000\s*\)/);
```

Add `this.overlayText` at `(300, 240)` reading `"TAP TO ENABLE AUDIO"` with depth `1001`.

```js
assert.match(code, /this\.overlayText\s*=\s*this\.add\s*\.text\(\s*300\s*,\s*240\s*,\s*["\']TAP TO ENABLE AUDIO["\'][\s\S]*\.setDepth\(\s*1001\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.sfxVolume = 0.6;
    this.musicVolume = 0.5;
    this.playSfx = (key) => this.sound.play(key);
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
