---
id: 66faa70000000000000000ad
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--
Time to add a control button for the music. We'll need several rectangular buttons, so extract the pattern into a `this.makeButton(x, y, w, h, label, onClick)` helper that returns `{ bg, label }` so callers can update the label later.

Right after `this.playSfx = ...`:

```js
this.makeButton = (x, y, w, h, label, onClick) => {
  const bg = this.add
    .rectangle(x, y, w, h, 0x222244)
    .setOrigin(0)
    .setStrokeStyle(2, 0x6c8aff);
  const text = this.add
    .text(x + w / 2, y + h / 2, label, {
      fontSize: "12px",
      color: "#ffffff"
    })
    .setOrigin(0.5);
  bg.setInteractive();
  bg.on("pointerdown", onClick);
  return { bg, label: text };
};
this.musicBtn = this.makeButton(
  8,
  56,
  140,
  32,
  "Music: Play",
  () => {}
);
```

The button does nothing yet - we wire actual playback in step 8. Defining the helper next to `this.playSfx` keeps every later button (mute in step 14, crossfade in step 17) below its own creator.

# --hints--
Define `this.makeButton` as an arrow function `(x, y, w, h, label, onClick) => ...` that returns `{ bg, label }`.

```js
assert.match(code, /this\.makeButton\s*=\s*\(\s*x\s*,\s*y\s*,\s*w\s*,\s*h\s*,\s*label\s*,\s*onClick\s*\)\s*=>\s*\{[\s\S]*return\s*\{\s*bg\s*,\s*label\s*:\s*text\s*\}/);
```

The bg rectangle should be `(x, y, w, h, 0x222244)` with `setStrokeStyle(2, 0x6c8aff)`.

```js
assert.match(code, /\.rectangle\(\s*x\s*,\s*y\s*,\s*w\s*,\s*h\s*,\s*0x222244\s*\)[\s\S]*\.setStrokeStyle\(\s*2\s*,\s*0x6c8aff\s*\)/);
```

Wire the bg with `setInteractive()` and `on("pointerdown", onClick)`.

```js
assert.match(code, /bg\.setInteractive\(\s*\)[\s\S]*bg\.on\(\s*[\"']pointerdown[\"']\s*,\s*onClick\s*\)/);
```

Use the helper to build `this.musicBtn` at `(8, 56, 140, 32)` with label `"Music: Play"` and an empty handler.

```js
assert.match(code, /this\.musicBtn\s*=\s*this\.makeButton\(\s*8\s*,\s*56\s*,\s*140\s*,\s*32\s*,\s*["\']Music: Play["\']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.sfxVolume = 0.6;
    this.musicVolume = 0.5;
    this.playSfx = (key) => this.sound.play(key);
--fcc-editable-region--

--fcc-editable-region--
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
      bg.setInteractive();
      bg.on("pointerdown", () => this.playSfx(key));
    });
    this.overlay = this.add
      .rectangle(0, 0, 600, 480, 0x000000, 0.85)
      .setOrigin(0)
      .setDepth(1000);
    this.overlayText = this.add
      .text(300, 240, "TAP TO ENABLE AUDIO", {
        fontSize: "24px",
        color: "#ffffff"
      })
      .setOrigin(0.5)
      .setDepth(1001);
    this.input.once("pointerdown", () => {
      if (this.sound.locked) this.sound.unlock();
      this.overlay.destroy();
      this.overlayText.destroy();
    });
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
