---
id: 66faa70000000000000000aa
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
Buttons that don't react are decorations. Wire each SFX button to play its sound on `pointerdown` - and route the play through a `this.playSfx(key)` helper now, so we can tune volume and add anti-double-tap guards in later steps without rewriting every handler.

Replace the entire `SFX_KEYS` block in `create` with the version below. The new lines at the top declare the SFX and music volume defaults plus the `this.playSfx` helper, then the `forEach` calls `setInteractive()` on the button background and wires `pointerdown` to `this.playSfx(key)`.

```js
this.sfxVolume = 0.6;
this.musicVolume = 0.5;
this.playSfx = (key) => this.sound.play(key);
const SFX_KEYS = [
  "sfx-jump",
  "sfx-coin",
  "sfx-hit",
  "sfx-explosion",
  "sfx-ui-click",
  "sfx-ping"
];
SFX_KEYS.forEach((key, i) => {
  const x = 8 + i * 95;
  const bg = this.add
    .rectangle(x, 8, 90, 36, 0x2a3045)
    .setOrigin(0)
    .setStrokeStyle(2, 0x6c8aff);
  this.add
    .text(x + 45, 26, key.replace("sfx-", ""), {
      fontSize: "10px",
      color: "#ffffff"
    })
    .setOrigin(0.5);
  bg.setInteractive();
  bg.on("pointerdown", () => this.playSfx(key));
});
```

Click any button - you'll see a click flash but won't hear anything yet. iOS and Safari block audio until the user has interacted with the page; we'll fix that in step 6.

# --hints--
Initialise `this.sfxVolume = 0.6` at the top of `create`.

```js
assert.match(code, /this\.sfxVolume\s*=\s*0\.6/);
```

Initialise `this.musicVolume = 0.5` at the top of `create`.

```js
assert.match(code, /this\.musicVolume\s*=\s*0\.5/);
```

Define `this.playSfx = (key) => this.sound.play(key)` so later steps can extend it.

```js
assert.match(code, /this\.playSfx\s*=\s*\(\s*key\s*\)\s*=>\s*this\.sound\.play\(\s*key\s*\)/);
```

Call `bg.setInteractive()` inside the `forEach`.

```js
assert.match(code, /bg\.setInteractive\(\s*\)/);
```

Wire `bg.on("pointerdown", () => this.playSfx(key))`.

```js
assert.match(code, /bg\.on\(\s*["\']pointerdown["\']\s*,\s*\(\s*\)\s*=>\s*this\.playSfx\(\s*key\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
--fcc-editable-region--
    const SFX_KEYS = [
      "sfx-jump",
      "sfx-coin",
      "sfx-hit",
      "sfx-explosion",
      "sfx-ui-click",
      "sfx-ping"
    ];
    SFX_KEYS.forEach((key, i) => {
      const x = 8 + i * 95;
      const bg = this.add
        .rectangle(x, 8, 90, 36, 0x2a3045)
        .setOrigin(0)
        .setStrokeStyle(2, 0x6c8aff);
      this.add
        .text(x + 45, 26, key.replace("sfx-", ""), {
          fontSize: "10px",
          color: "#ffffff"
        })
        .setOrigin(0.5);
    });
--fcc-editable-region--
  }

  update(time, delta) {}
}

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
