---
id: 66faa70000000000000000a7
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the soundboard workshop. Across 20 steps you'll wire 6 sound-effect buttons, a music play/pause toggle, volume sliders that persist to `localStorage`, a tap-to-enable audio overlay for mobile, a stereo pan and rate slider, a music crossfade, a HUD list of currently playing sounds, and a `prefers-reduced-motion`-aware accessibility default.

Start by declaring `MainScene` extending `Phaser.Scene` with a constructor that calls `super("MainScene")`, an empty `preload` method, an empty `create` method, and an empty `update(time, delta)` method.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {}

  create() {}

  update(time, delta) {}
}
```

The `config` object below is already filled in - a 600x480 canvas with no physics.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["\']MainScene["\']\s*\)/);
```

You should declare a `preload` method.

```js
assert.match(code, /\bpreload\s*\(\s*\)\s*\{/);
```

You should declare a `create` method.

```js
assert.match(code, /\bcreate\s*\(\s*\)\s*\{/);
```

You should declare an `update` method that takes `time` and `delta`.

```js
assert.match(code, /\bupdate\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Soundboard</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

const config = {
  type: Phaser.AUTO,
  width: 600,
  height: 480,
  parent: "game-container",
  backgroundColor: "#101020",
  scene: [MainScene]
};

new Phaser.Game(config);
```
