---
id: 6823e036ea4b71553558c01b
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Notice that each heading looks a little different in the preview. Same with the paragraph.

Next, add another `p` element that displays the following text:

```md
HTML represents the content and structure of a webpage
```

# --hints--

You should have a second `p` element.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Your `p` element should have the text `HTML represents the content and structure of a webpage`. You can copy the text from the instructions.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
