---
id: 6823ac607bfdbc46331b2559
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML stands for HyperText Markup Language. It's the code that defines the structure and content of a webpage. This is your code editor, where you'll write HTML.

Find line 1 in the editor and type this text:

```md
Welcome to freeCodeCamp
```

When you are done, click the "check your code" button to see if it's correct.

# --hints--

You should have the text `Welcome to freeCodeCamp` in your editor. Double-check for spelling.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
