---
id: 682cd20b883fc7b25eb539c6
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Below the other two lines of text, add:

```md
Learn the skills to become a full stack developer
```

# --hints--

You should have the text `Learn the skills to become a full stack developer` in your editor. Double-check for spelling.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--

--fcc-editable-region--
```
