---
id: 656485a8a3496d4a36b1496a
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

These three variables constitute the possible characters to choose from when generating the password.

Just before them, add a comment saying `Define the possible characters for the password`.

# --hints--

You should add the comment just above the `letters` variable.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
