---
id: 619665c9abd72906f3ad30f9
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

你將創建一隻快樂飛揚的企鵝，並在此過程中進一步學習 CSS transforms 和動畫的知識。

先從基本的 HTML 模板開始。 需要包含 `DOCTYPE` 聲明，語言設置爲英文的 `html` 元素，相應的 `meta` 標籤，`head`、`body` 以及 `title` 元素。 此外，還需要把樣式錶鏈接到頁面裏。

# --hints--

代碼應該有一個 `<!DOCTYPE html>` 聲明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

代碼應該有一個 `html` 元素。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

你的 `html` 元素應該有一個帶有值爲 `en` 的 `lang` 屬性的開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的代碼應該在 `html` 元素中包含一個 `head` 元素。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

你的代碼應該在 `html` 元素中包含一個 `body` 元素。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

你的 `head` 元素應該在你的 `body` 元素之前。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

一個 `meta` 元素應將 `name` 設置爲 `viewport`，並將 `content` 設置爲 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一個 `meta` 元素應將 `charset` 屬性設置爲 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 應該有一些文本。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

你的代碼應該有一個 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
