---
id: 619691693bc14b0e528f5a20
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

選擇 `.ground` 元素，設置其 `width` 爲視口的寬度。 然後設置其 `height` 爲 `400px`。

# --hints--

應該用到 `.ground` 選擇器。

```js
assert.match(code, /\.ground\s*\{/);
```

元素 `.ground` 的 `width` 的值應該爲 `100vw`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

元素 `.ground` 的 `height` 的值應該爲 `400px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
