---
id: 61969aa6acef5b12200f672e
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

來添加兩座冰山，讓背景更豐富。

在 `.penguin` 元素上面，添加一個 `div`，並設置其 `class` 爲 `left-mountain`。

# --hints--

應該在 `body` 內添加一個新的 `div`。 應該是 `--fcc-expected--` `div` 元素，而不是 `--fcc-actual--`。

```js
assert.equal(document.querySelectorAll('body > div')?.length, 3);
```

`div` 的 `class` 應該爲 `left-mountain`。

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

`.left-mountain` 應該在 `.penguin` 前面。

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
