---
id: 6197f40a16afea068c7e60c8
title: 步驟 29
challengeType: 0
dashedName: step-29
---

# --description--

修改 `.penguin` 元素的堆疊層級，使其顯示在`.ground` 元素的前面，並設置其 `position` 屬性值爲 `relative`。

# --hints--

應該使用 `z-index` 屬性來設置堆疊層級。

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('.penguin')?.zIndex);
```

應該將 `.penguin` 元素的 `z-index` 設置爲 `4`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.zIndex, '4');
```

應該設置 `.penguin` 的 `position` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.penguin')?.position);
```

`.penguin` 的 `position` 屬性值應該爲 `relative`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin')?.position, 'relative');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin">
      <div class="penguin-head"></div>
      <div class="penguin-body"></div>
    </div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

--fcc-editable-region--
.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;

}
--fcc-editable-region--

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
