---
id: 6143869bb45bd85e3b1926aa
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro de `.heading`, crie um elemento `header` com o atributo `class` definido como `hero`.

Neste elemento, crie um elemento `img` com o `src` definido como `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, o `alt` definido como `freecodecamp logo`, e a `class` definida como `hero-img`.

O atributo `loading` em um elemento `img` pode ser definido como `lazy` para dizer ao navegador que não carregue a imagem até que seja necessário (por exemplo, quando o usuário rola a imagem para visualização). Como um benefício adicional, os elementos carregados com atraso não serão carregados até que os elementos sem atrasdo sejam carregados - isso significa que os usuários com conexões de internet lentas podem visualizar o conteúdo da sua página sem ter que esperar o carregamento das imagens.

Dê ao novo elemento `img` um atributo `loading` definido como `lazy`.

# --hints--

Você deve criar um elemento `header`.

```js
assert.exists(document.querySelector('header'));
```

O elemento `header` deve estar dentro do elemento `.heading`.

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

O elemento `header` deve ter a `class` definida como `hero`.

```js
assert(document.querySelector('header')?.className === 'hero');
```

Você deve criar um elemento `img`.

```js
assert.exists(document.querySelector('img'));
```

O elemento `img` deve estar dentro do elemento `header`.

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

O elemento `img` deve ter o `src` definido como `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

O elemento `img` deve ter o atributo `alt` definido como `freecodecamp logo`.

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

O elemento `img` deve ter o atributo `loading` definido como `lazy`.

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

O elemento `img` deve ter a `class` definida como `hero-img`.

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
