---
id: 61438b5b66d76a6264430f2a
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Dentro de cada um dos novos elementos `a`, adicione um elemento `i` e dê a eles as seguintes classes:

- O primeiro elemento `i` deve ter a classe `fab fa-facebook-f`
- O segundo elemento `i` deve ter a classe `fab fa-twitter`
- O terceiro elemento `i` deve ter a classe `fab fa-instagram`
- O quarto elemento `i` deve ter a classe `fab fa-linkedin-in`
- O quinto elemento `i` deve ter a classe `fab fa-youtube`

# --hints--

Você deve ter cinco elementos `i`.

```js
assert(document.querySelectorAll('i')?.length === 5);
```

Cada elemento `a` deve ter apenas um elemento `i`.

```js
const aaaaa = [...document.querySelectorAll('.social-icons a')];
assert(aaaaa?.every(a => a?.children?.length === 1 && a?.children?.[0]?.localName === 'i'));
```

Cada elemento `i` deve ter um atributo `class`, que inclui `fab`.

```js
const iiiii = [...document.querySelectorAll('i')];
assert(iiiii?.every(i => i?.classList?.contains('fab')));
```

O primeiro elemento `i` deve ter um atributo `class` que inclui `fa-facebook-f`.

```js
assert(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

O segundo elemento `i` deve ter um atributo `class` que inclui `fa-twitter`.

```js
assert(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

O terceiro elemento `i` deve ter um atributo `class` que inclui `fa-instagram`.

```js
assert(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

O quarto elemento `i` deve ter um atributo `class` que inclui `fa-linkedin-in`.

```js
assert(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

O quinto elemento `i` deve ter um atributo `class` que inclui `fa-youtube`.

```js
assert(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
