---
id: 61437d575fb98f57fa8f7f36
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Comece com o boilerplate padrão do HTML. Adicione a declaração `DOCTYPE` e os elementos `html` (especificando que esta página está em inglês), `head` e `body`.

Adicione a tag `<meta>` com o `charset` apropriado e uma tag `<meta>` para responsividade móvel dentro do elemento `head`.

# --hints--

O código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento para ser `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Você deve ter uma tag de abertura `<html>` e deve ter o atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Você deve acrescentar uma tag de abertura `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

Você deve acrescentar uma tag de fechamento `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

Você deve acrescentar uma tag de abertura `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

Você deve acrescentar uma tag de fechamento `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

O elemento `head` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Você deve ter dois elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Um elemento `meta` deve ter o atributo `name` definido como `viewport` e o atributo `content` definido como `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

O outro elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
