---
id: 6143862a5e54455d262c212e
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro do `body`, crie um elemento `main`. Em seguida, nesse elemento, crie uma `section` com uma `class` chamada `heading`.

# --hints--

Você deve ter apenas um elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

O elemento `main` deve estar dentro do elemento `body`.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

Você deve ter apenas um elemento `section`.

```js
assert.exists(document.querySelector('section'));
```

O elemento `section` deve estar dentro do elemento `main`.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

O elemento `section` deve ter o atributo `class` definido como `heading`.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
