---
id: 614385513d91ae5c251c2052
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione um elemento `title` com o texto `Magazine`, um elemento `link` para a folha de estilo da fonte `https://fonts.googleapis.com/css?family=Anton|Baskervville|Raleway&display=swap`, um `link` para a folha de estilo do FontAwesome `https://use.fontawesome.com/releases/v5.8.2/css/all.css` e um `link` para sua folha de estilo `./styles.css`.

A folha de estilos de fonte carregará três fontes separadas: `Anton`, `Baskervville`, e `Raleway`.

# --hints--

O código deve ter três elementos `link`.

```js
assert(code.match(/<link/g)?.length === 3);
```

Os elementos `link` devem ser de fechamento automático.

```js
assert(code.match(/<\/link>/i) === null);
```

Os elementos `link` devem estar dentro do elemento `head`.

```js
const head = code.match(/<head>(.|\r|\n)*<\/head>/i)?.[0];
assert(head.match(/<link/g)?.length === 3)
```

Os três elementos `link` devem ter o atributo `rel` com o valor `stylesheet`.

```js
assert(code.match(/<link[\s\S]*?rel=('|"|`)stylesheet\1/gi)?.length === 3);
```

Um dos seus elementos de link deve ter o `href` definido como `https://fonts.googleapis.com/css?family=Anton|Baskervville|Raleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton|Baskervville|Raleway&display=swap'));
```

Um dos elementos link deve ter o atributo `href` definido como `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Um dos elementos `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O projeto deve ter o título `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
