---
id: 614387cbefeeba5f3654a291
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Após o elemento `header`, crie uma `div` com a `class` definida como `author`.

Dentro dessa `div`, crie um elemento `p` com a `class` definida como `author-name` e dê a ele o texto `By freeCodeCamp`. Coloque o elemento `freeCodeCamp` dentro de um elemento `a` com o `href` definido como `https://freecodecamp.org` e o `target` definido como `_blank`.


Abaixo disso, adicione um segundo elemento `p` com a classe `publish-date` e o texto `March 7, 2019`.

# --hints--

Crie um novo elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

O elemento `div` deve vir depois do elemento `header`.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

O elemento `div` deve ter a `class` definida como `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

Você deve criar dois novos elementos `p`.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

Os dois novos elementos `p` devem estar dentro do elemento `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

O primeiro elemento `p` deve ter uma `class` definida como `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

O primeiro elemento `p` deve ter o texto `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

O segundo elemento `p` deve ter uma `class` definida como `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

O segundo elemento `p` deve ter o texto `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

Você deve criar um elemento `a`.

```js
assert.exists(document.querySelector('a'));
```

O elemento `a` deve estar dentro do primeiro elemento `p`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

O elemento `a` deve ter o `href` definido como `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

O elemento `a` deve ter o `target` definido como `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

O elemento `a` deve envolver o texto `freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.textContent === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
