---
id: 614389f601bb4f611db98563
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Abaixo do elemento `.author`, crie um elemento `div` com a classe `social-icons`.

Adicione cinco elementos `a` dentro dessa nova `div` e dê a eles os seguintes atributos `href`.

- O primeiro elemento `a` deve ter o `href` definido como `https://www.facebook.com/freecodecamp`.
- O segundo elemento `a` deve ter o `href` definido como `https://twitter.com/freecodecamp`.
- O terceiro elemento `a` deve ter o `href` definido como `https://instagram.com/freecodecamp`.
- O quarto elemento `a` deve ter o `href` definido como `https://www.linkedin.com/school/free-code-camp`.
- O quinto elemento `a` deve ter o `href` definido como `https://www.youtube.com/freecodecamp`.

# --hints--

Crie um elemento `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

O novo elemento `div` deve estar depois do elemento `.author`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

O novo elemento `div` deve ter a classe `social-icons`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

O elemento `.social-icons` deve ter cinco elementos `a`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

O primeiro elemento `a` deve ter o `href` definido como `https://www.facebook.com/freecodecamp`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

O segundo elemento `a` deve ter o `href` definido como `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

O terceiro elemento `a` deve ter o `href` definido como `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

O quarto elemento `a` deve ter o `href` definido como `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

O quinto elemento `a` deve ter o `href` definido como `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
