---
id: 65f569725359e10d345bc52a
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

You are going to build a program that calculates and draws the trajectory of a projectile given the angle, speed and height of the throw.

Start by importing `math`, you will use it a lot in this project as it has useful methods like `math.radians`, `math.cos`, `math.sin` and others.

# --hints--

Use `import math` to import the `math` module.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
