---
id: 65f569725359e10d345bc52a
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

You are going to build a program that calculates and draws the trajectory of a projectile given the starting angle, speed and height.

Start by importing `math`, it's going to be used a lot.

# --hints--

Use `import math` to import the `math` module.

```js
assert(code.match(/import\s+math/))
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
