---
id: 65fd4bd84561a14d3e720061
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

You should use the `__slots__` variable to define which properties the class has: assign a tuple to `__slots__` containing the strings `"__height"`, `"__speed"` and `"__angle"`, these will be the names of the properties.

# --hints--

You should assign to `__slots__` the tuple `("__height", "__speed", "__angle")`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().find_variable("__slots__").is_equivalent('__slots__ = ("__height", "__speed", "__angle")')`))})

```

# --seed--

## --seed-contents--

```py
import math

--fcc-editable-region--
class Projectile:
    pass
--fcc-editable-region--
```
