---
id: 65fd4c95b0444b4f86a8203c
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Create an `__init__` method that accepts, other than `self`, the three parameters `height`, `speed` and `angle`.

# --hints--

You should have an `__init__` method inside the `Projectile` class.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__") `))})
```

Your `__init__` method should have the arguments `self`, `height`, `speed` and `angle`, in this order.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().find_function("__init__").has_args("self, height, speed, angle")`))})
```

# --seed--

## --seed-contents--

```py
import math

--fcc-editable-region--
class Projectile:
    __slots__ = ('__height', '__speed', '__angle')


--fcc-editable-region--
```
