---
id: 65fd4b3bedee044b4b957d33
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Create a `Projectile` class, and its `__init__` method.

The class should be instantiated with three arguments: the starting speed, the starting height and the starting angle of the throw of the projectile in this order. These three values should be saved internally as attributes `__speed`, `__height` and `__angle`.

Two underscores in front of attribute or method names activates name mangling, meaning that from outside a class (`ball = Projectile(12, 12, 12)`) the attribute name is obscured like this: attribute `__height` has to be called from outside as `ball._Projectile__height`. It's a way to indicate that a method or attribute is used only inside the class and should not be used on instances.

For the value of the angle, the argument should bee a degrees value, but it should be saved internally in radians, so you need to use `math.radians()` to convert it.

# --hints--

You should use `class Projectile:` to create the class.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

You should have an `__init__` method inside the `Projectile` class.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

A new `Projectile` class instance should have the values of the three attributes propperly initialised.

```js
({test: () => assert(runPython(`
import math
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import math
--fcc-editable-region--
```
