---
id: 65fd4b3bedee044b4b957d33
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Create a `Projectile` class.

# --hints--

You should use `class Projectile:` to create the class.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import math
--fcc-editable-region--
```
