---
id: 65fd4cff221e7450e347f749
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Inside the `__init__` method, assign `height` to the `__height` property, and `speed` to the `__speed` property.

For `angle`, convert the value using `math.radians`, then assign it to the `__angle` property.

# --hints--

Assign `height` to the `__height` property with `self.__height = height`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__init__").find_variable("self.__height").is_equivalent("self.__height = height")`))})
```

Assign `speed` to the `__speed` property with `self.__speed = speed`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__init__").find_variable("self.__speed").is_equivalent("self.__speed = speed")`))})
```

Assign `angle` to the `__angle` property with `self.__angle = math.radians(angle)`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__init__").find_variable("self.__angle").is_equivalent("self.__angle = math.radians(angle)")`))})
```

# --seed--

## --seed-contents--

```py
import math

--fcc-editable-region--
class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        pass
--fcc-editable-region--
```
