---
id: 660400f02cf3a76795b82d4e
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now let's set up the calculation for the displacement of the projectile, which is the horizontal space traveled from the throw to when it touches the ground.

Create a function `__calculate_displacement`, with no additional arguments. Two underscrores in front of the function name trigger name mangling, the method is accessible from outside the class with the name `_Projectile__calculate_displacement`.

# --hints--

You should declare a function called `__calculate_displacement` with `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

The `__calculate_displacement` function should have only the `self` argument.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

# --seed--

## --seed-contents--

```py
import math

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
--fcc-editable-region--

--fcc-editable-region--
```
