---
id: 660bcfaee1f1fb1114369937
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Inside the `__calculate_displacement` function declare a variable `horizontal_component` and assign to it the multiplication between the speed and the cosine of the angle at throw.

# --hints--

You should declare a variable `horizontal_component`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('horizontal_component')`))})
```

`horizontal_component` should have value of `self.__speed * math.cos(self.__angle)`.

```js
({test: () => assert(runPython(`any([_Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('horizontal_component').is_equivalent(x) for x in ['horizontal_component = self.__speed * math.cos(self.__angle)', 'horizontal_component = math.cos(self.__angle) * self.__speed'])`))})
```

# --seed--

## --seed-contents--

```py
import math

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    def __calculate_displacement(self):
        pass

--fcc-editable-region--
```
