---
id: 660c470f0f6e8f2a2bc3f523
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Now you need to create the next component of the calculation. Create a variable called `vertical_component` and assign to it the multiplication of the speed and sine of the angle of the projectile at throw.

# --hints--

You should declare a variable `vertical_component`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('vertical_component')`))})
```

`horizontal_component` should have value of `self.__speed * math.cos(self.__angle)`.

```js
({test: () => assert(runPython(`
node = _Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('vertical_component')
any([node.is_equivalent(x) for x in ['vertical_component = self.__speed * math.sin(self.__angle)', 'vertical_component = math.sin(self.__angle) * self.__speed']])`))})
```

# --seed--

## --seed-contents--

```py
import math

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)

--fcc-editable-region--
```
