---
id: 660d956033806673ba8c48cc
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Declare a variable `gh_component` and assign to it the multiplication between the number `2`, the `GRAVITATIONAL_ACCELERATION` and the height of the projectile at throw.

# --hints--

You should have a variable `gh_component`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('gh_component')`))})
```

Your variable `gh_component` should have value of `2 * GRAVITATIONAL_ACCELERATION * self.__height`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('gh_component').is_equivalent('gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height')`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)

--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        
--fcc-editable-region--
        
```
