---
id: 660d877ba97710440b731144
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Before you continue with the displacement, you need to define a global variable.
<!-- TODO: WHat is gravitational_acceleration -->

Create the variable `GRAVITATIONAL_ACCELERATION` and give it a value of `9.81`.

# --hints--

You should create a variable `GRAVITATIONAL_ACCELERATION`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('GRAVITATIONAL_ACCELERATION')`))})
```

The variable `GRAVITATIONAL_ACCELERATION` should have a value of `9.81`.

```js
({test: () => assert(runPython(`_Node(_code).get_variable('GRAVITATIONAL_ACCELERATION') == 9.81`))})
```

# --seed--

## --seed-contents--

```py
import math

--fcc-editable-region--

--fcc-editable-region--

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        
```
