---
id: 660d89bd20c05e4ca73962bd
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Create a variable `square_component`, its value should be the square of the vertical component. You can exponentiate a value by using the `**` operator, for example you would calculate \\( 2^4 \\) by writing `2 ** 4`.

# --hints--

You should declare a variable `squared_component`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('squared_component')`))})
```

The variable `squared_component` should have value of `vertical_component**2`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('squared_component').is_equivalent('squared_component = vertical_component**2')`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)

--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        
--fcc-editable-region--
        
```
