---
id: 660d89bd20c05e4ca73962bd
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Create a variable `square_component`, its value should be the square of the multiplication between the speed and the sine of the angle. You can exponentiate a value by using the `**` operator, for example `2 ** 4` is \\( 2^4 \\).

# --hints--

You should declare a variable `squared_component`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('squared_component')`))})
```

The variable `squared_component` should have value of `(self.__speed * math.cos(self.__angle))**2`.

```js
({test: () => assert(runPython(`
node = _Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('squared_component')
any([node.is_equivalent(x) for x in ['squared_component = (self.__speed * math.sin(self.__angle))**2', 'squared_component = (math.sin(self.__angle) * self.__speed)**2']])`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)

--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        
--fcc-editable-region--
        
```
