---
id: 660e9039407bc743d7b5f18e
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

Create a variable `sqrt_component` and give it the value of the square root (`math.sqrt()`) of the sum between `squared_component` and `gh_component`. 

# --hints--

The variable `sqrt_component` should exist.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('sqrt_component')`))})
```

The variable `sqrt_component` should have value of `math.sqrt(squared_component + gh_component)`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('sqrt_component').is_equivalent('sqrt_component = math.sqrt(squared_component + gh_component)')`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)

--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        
--fcc-editable-region--
        
```
