---
id: 660fcf3e1b9bb056b2edb567
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

At this point you are ready to create the string rapresentation. Let's start by creating an instance of the Projectile. Create a `ball` variable and assign it a call to `Projectile` using `45, 45, 45` as arguments. Then print the `ball` variable.

# --hints--

You should have a `ball` variable.

```js
{test: () => assert(runPython(`_Node(_code).has_variable('ball')`))}
```

The `ball` variable should have a value of `Projectile(45, 45, 45)`.

```js
{test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(45, 45, 45)')`))}
```

You should print the `ball` variable.

```js
{test: () => assert(runPython(`_Node(_code).has_call('print(ball)')`))}
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
--fcc-editable-region--

--fcc-editable-region--
        
```
