---
id: 660fcdae4653184d49e22742
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

The `displacement` calculation is not finished, now multiply `horizontal_component` by the sum of `vertical_component` and `sqrt_component`. Remember to wrap the sum in parenthesis, so that the sum happens before the multiplication.

# --hints--

You should have a `displacement` variable.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('displacement')`))})
```

The `displacement` variable should have a value of `horizontal_component * (vertical_component + sqrt_component)`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('displacement').is_equivalent('displacement = horizontal_component * (vertical_component + sqrt_component)')`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)

--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component
        
--fcc-editable-region--
        
```
