---
id: 660fce8e32639f52d4157545
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now you need to finish the `displacement` declaration by dividing the current value by `GRAVITATIONAL_ACCELERATION`. With this you have finished the calculation of the `displacement`.

Finish the function by returning the `displacement`.

# --hints--

You should have a `displacement` variable.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_variable('displacement')`))})
```

The `displacement` variable should have a value of `horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").find_variable('displacement').is_equivalent('displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION')`))})
```

The function should return the `displacement` variable.

```js
// ({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_return('displacement')`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)

--fcc-editable-region--
    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component)
        
--fcc-editable-region--
        
```
