---
id: 661019a8586cf63cb35bcbd1
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

The print function is printing the pointer of an object. You will need to define the string rapresentation of the object to have an interesting print output.

Create a `__str__` function inside the `Projectile` class, and return an empty f-string inside this new function.

# --hints--

You should have a function `__str__`.

```js
({test: () => assert(runPython(`_Node(_code).find_class('Projectile').has_function('__str__')`))})
```

The function `__str__` should have only `self` as arguments.

```js
({test: () => assert(runPython(`_Node(_code).find_class('Projectile').find_function('__str__').has_args('self')`))})
```

The function `__str__` should return `f''`.

```js
//({test: () => assert(runPython(`_Node(_code).find_class('Projectile').find_function('__str__').has_return("f''")`))
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
ball = Projectile(45, 45, 45)
print(ball)
   
```
