---
id: 6613e31b7091ba34343a68dc
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

Add to the f string the words `Projectile details` followed by a colon, and a new line character.

# --hints--

The function should return `f'Projectile details:\n'`.

```js
//({test: () => assert(runPython(`_Node(_code).find_class('Projectile').find_function('__str__').has_return("f'Projectile details:\\\\n'")`))
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    def __str__(self):
        return f''
--fcc-editable-region--

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
ball = Projectile(45, 45, 45)
print(ball)
   
```
