---
id: 6616df314fdad508d9d28274
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Add to the f-string the word `height` followed by a colon, then a space and the value of the `__height` attribute, then a space, and `'m'`, and finish with a new line character.

# --hints--

The function should return `f'Projectile details:\nheight: {self.__height} m\n'`

```js
\\({test: () => assert(runPython(`_Node(_code).find_class('Projectile').find_function('__str__').has_return("f'Projectile details:\\\\nheight: {self._height} m\\\\n'")`))
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    def __str__(self):
        return f'Projectile details:\n'
--fcc-editable-region--

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
ball = Projectile(45, 45, 45)
print(ball)
   
```
