---
id: 6616ec0bf9ee8548ce6b0f08
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Now let's start writing the function to calculate the coordinates of the trajectory! Create a method of the `Projectile` class called `__calculate_y_coordinate`, it should have, other than `self`, an `x` argument.

# --hints--

You should have a function called `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

The `__calculate_y_coordinate` function should have `self, x` as arguments.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
height: {self.__height} m
speed: {self.__speed} m/s
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    --fcc-editable-region--

    --fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```
