---
id: 661ab452d755931aafdeb178
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

\\( y = y_0 + x \tan(\theta) - \frac{g x^2}{2 v_0^2 \cos^2(\theta)} \\)

This formula is the formula to calculate the vertical position ( \\( y \\) ) for any given horizontal position ( \\(x\\) ) having the starting angle ( \\(\theta\\) ), speed ( \\(v_0\\) ) and height ( \\(y_0\\) ).

You will need to use `math.tan()` and `math.cos()` and remember that `x ** y` is the way to write \\( x^y \\), and that the value of \\( g \\) is in the variable `GRAVITATIONAL_ACCELERATION`.

Implement the method so that it returns the y coordinate.

# --hints--

The method should return the correct value. For `Projectile(12, 12, 12)`, given the 4 x coordinates of 0, 1, 2 and 3, the y coordinate should be approssimately 12, 12.18, 12.28, 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

The method should return the correct value. For `Projectile(45, 12, 22)`, given the x coordinates of 6, 9, 9.5 and 12 the y coordinate should be approssimately 46.0, 45.43, 45.26 and 44.14.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 46.0), (9, 45.43), (9.5, 45.26), (12, 44.14)]])
    `))
}})
```

The `__calculate_y_coordinate` method should not round the result.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
height: {self.__height} m
speed: {self.__speed} m/s
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
--fcc-editable-region--
    def __calculate_y_coordinate(self, x):
        pass
--fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```
