---
id: 661ab452d755931aafdeb178
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

\\( y = y_0 + x \tan(\theta) - \frac{g x^2}{2 v_0^2 \cos^2(\theta)} \\)

This formula is the formula to calculate the vertical position ( \\( y \\) ) for any given horizontal position ( \\(x\\) ) having the starting angle ( \\(\theta\\) ), speed ( \\(v_0\\) ) and height ( \\(y_0\\) ).

This function, like `__calculate_displacement`, will have the calculation broken down into smaller pieces: start creating a variable `height_component` and assign to it the starting height; and also a `angle_component` variable and assign to it the multiplication between the tangent of the angle (use `math.tan()`) and `x`.

\\( height component = y_0, angle component = x \tan(\theta) \\)

# --hints--

You should have a variable `height_component` and a variable `angle_component` inside the `__calculate_y_coordinate` function.

```js
({test: () => {
    assert(runPython(`
    func_node = _Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate')
    func_node.has_variable('height_component') and func_node.has_variable('angle_component')
    `))
}})
```

The `height_component` variable should have value of `self.__height`.

```js
({test: () => assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').find_variable('height_component').is_equivalent('height_component = self.__height')`))})
```

The `angle_component` variable should have value of `math.tan(self.__angle) * x`.

```js
({test: () => assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').find_variable('angle_component').is_equivalent('angle_component = math.tan(self.__angle) * x')`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, height, speed, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
height: {self.__height} m
speed: {self.__speed} m/s
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
--fcc-editable-region--
    def __calculate_y_coordinate(self, x):
        pass
--fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```
