---
id: 662fa453ec8033322b3ab2e3
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Let's talk of encapsulation and getters and setters. You have written the three attributes of the class to be private using the two underscores in front. The attributes are hidden and can't be accessed by the outside.

But it may be that you want to see or update the values of the attributes of an already existing instance: in this case you use getters to see the values of the attributes and getters to set the new values.

In this step you will write the getters of the three attributes.

An example of how to write a getter is:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

This `number_of_eggs` getter returns the value of the private property `__number_of_eggs`.

To define a getter, you define a method and give it a `@property` decorator. The decorator makes so that the method is not called like a method, but used like a property.

```py
n = Nest()
print(n.number_of_eggs)
```

You can access the number of eggs in the nest simply with `n.number_of_eggs`.

Create the three getters for the three attributes. Remember that the angle is stored internally in radians, but getter should return a degrees value (`math.degrees()`, together with the `round` function could be useful).

Give the getters the same name as the private property but without the two underscores in front.

# --hints--

You should define a `speed` method.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

Your `speed` method should have a single parameter, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

Your `speed` method should have a `@property` decorator.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

Your `speed` method should return the value of the private `__speed` property.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

You should define a `height` method.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('height')
`))})
```

Your `height` method should have a single parameter, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('height').has_args('self')
`))})
```

Your `height` method should have a `@property` decorator.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('height').has_decorators('property')
`))})
```

Your `height` method should return the value of the private `__height` property.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.height == 23
`))})
```

You should define an `angle` method.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('angle')
`))})
```

Your `angle` method should have a `@property` decorator.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('angle').has_decorators('property')
`))})
```

Your `angle` method should have a single parameter, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('angle').has_args('self')
`))})
```

Your `angle` method should return the value of the private `__angle` property converted to degrees.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.angle == 24
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__height', '__speed', '__angle')

    def __init__(self, speed, height, angle):
        self.__height = height
        self.__speed = speed
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
height: {self.__height} m
speed: {self.__speed} m/s
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--

--fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```
