---
id: 663353465bfb14259717da93
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Create a method `create_coordinates_table`, then at the end of the code print `graph.create_coordinates_table())`.

The method shoud use the `__coordinates` property and return a string containing all the coordinates, where x and y have always the same position, the values for x are always integers, and the values for y are always rounded to two decimal places:

```py
  x      y
  1   0.10
  2   0.20
  3   0.30
  4   0.40  
```

The table should start with a newline character and end with a new line character.

# --hints--

You should have a `create_coordinates_table()` method.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('create_coordinates_table')`
    ))
})
```

The method should return the correct output.

```js
({
    test: () => assert(runPython(
        `str = """
  x      y
  0  45.00
  1  46.00
  2  46.98
  3  47.96
  4  48.92
  5  49.88
  6  50.83
  7  51.76
  8  52.69
  9  53.61
 10  54.52
 11  55.41
 12  56.30
 13  57.18
 14  58.05
 15  58.91
 16  59.76
 17  60.60
 18  61.43
 19  62.25
 20  63.06
 21  63.86
 22  64.66
 23  65.44
 24  66.21
 25  66.97
 26  67.73
 27  68.47
 28  69.20
 29  69.93
 30  70.64
 31  71.34
 32  72.04
 33  72.72
 34  73.40
 35  74.07
 36  74.72
 37  75.37
 38  76.00
 39  76.63
 40  77.25
 41  77.86
 42  78.45
 43  79.04
 44  79.62
 45  80.19
 46  80.75
 47  81.30
 48  81.84
 49  82.37
 50  82.89
 51  83.40
 52  83.90
 53  84.39
 54  84.87
 55  85.35
 56  85.81
 57  86.26
 58  86.70
 59  87.14
 60  87.56
 61  87.97
 62  88.38
 63  88.77
 64  89.16
 65  89.53
 66  89.90
 67  90.25
 68  90.60
 69  90.94
 70  91.26
 71  91.58
 72  91.89
 73  92.18
 74  92.47
 75  92.75
 76  93.02
 77  93.28
 78  93.53
 79  93.77
 80  94.00
 81  94.22
 82  94.43
 83  94.63
 84  94.82
 85  95.00
 86  95.17
 87  95.33
 88  95.48
 89  95.63
 90  95.76
 91  95.88
 92  96.00
 93  96.10
 94  96.19
 95  96.28
 96  96.35
 97  96.42
 98  96.47
 99  96.52
100  96.56
101  96.58
102  96.60
103  96.61
104  96.60
105  96.59
106  96.57
107  96.54
108  96.49
109  96.44
110  96.38
111  96.31
112  96.23
113  96.14
114  96.04
115  95.93
116  95.81
117  95.68
118  95.55
119  95.40
120  95.24
121  95.07
122  94.90
123  94.71
124  94.51
125  94.31
126  94.09
127  93.86
128  93.63
129  93.38
130  93.13
131  92.86
132  92.59
133  92.31
134  92.01
135  91.71
136  91.40
137  91.07
138  90.74
139  90.40
140  90.05
141  89.69
142  89.32
143  88.94
144  88.55
145  88.15
146  87.74
147  87.32
148  86.89
149  86.45
150  86.00
151  85.54
152  85.07
153  84.60
154  84.11
155  83.61
156  83.11
157  82.59
158  82.06
159  81.53
160  80.98
161  80.43
162  79.86
163  79.29
164  78.70
165  78.11
166  77.51
167  76.89
168  76.27
169  75.64
170  75.00
171  74.34
172  73.68
173  73.01
174  72.33
175  71.64
176  70.94
177  70.23
178  69.51
179  68.78
180  68.04
181  67.29
182  66.53
183  65.76
184  64.99
185  64.20
186  63.40
187  62.59
188  61.78
189  60.95
190  60.12
191  59.27
192  58.41
193  57.55
194  56.67
195  55.79
196  54.90
197  53.99
198  53.08
199  52.16
200  51.22
201  50.28
202  49.33
203  48.37
204  47.39
205  46.41
206  45.42
207  44.42
208  43.41
209  42.39
210  41.36
211  40.32
212  39.27
213  38.21
214  37.14
215  36.07
216  34.98
217  33.88
218  32.77
219  31.66
220  30.53
221  29.39
222  28.25
223  27.09
224  25.93
225  24.75
226  23.57
227  22.37
228  21.17
229  19.95
230  18.73
231  17.50
232  16.25
233  15.00
234  13.74
235  12.47
236  11.18
237   9.89
238   8.59
239   7.28
240   5.96
241   4.63
242   3.29
243   1.94
244   0.58
"""
ball = Projectile(10, 3, 45)
g = Graph(ball.calculate_all_coordinates())
g.create_coordinates_table() == str`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s
    
    def __repr__(self):
        return f'{self.__class__}({self.speed}, {self.height}, {self.angle})'

class Graph:
    __slots__ = ('__coordinates')

    def __init__(self, coord):
        self.__coordinates = coord
    
    def __repr__(self):
        return f'Graph({self.__coordinates})'

--fcc-editable-region--
    


ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

graph = Graph(coordinates)

--fcc-editable-region--
```
