---
id: 6633c06601c081735063b528
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

You have built a projectile trajectory calculator.

Now to conclude, create a little utility function that takes in the desired values for speed, height and angle and print to the terminal in sequence, the details of the projectile, the table of coordinates and the graph of the trajectory.

The function name should be `projectile_helper`, and call it once with values of your choice. 

# --hints--

Test 1

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('projectile_helper')    
    `)
})
```

You should print the requested things in order

```js
({
    test: () => runPython(`
    prints = _Node(_code).find_function('projectile_helper').find_calls('print')
    # 1. print the instance

    # 2. print the table
    assert len(prints[1].find_calls('create_coordinates_table')) == 1
    # 3. print the graph
    assert len(prints[2].find_calls('create_trajectory')) == 1
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = '∙'
x_axe_tick = '⊤'
y_axe_tick = '⊣'

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s
    
    def __repr__(self):
        return f'{self.__class__}({self.speed}, {self.height}, {self.angle})'

class Graph:
    __slots__ = ('__coordinates')

    def __init__(self, coord):
        self.__coordinates = coord

    def __repr__(self):
        return f"Graph({self.__coordinates})"

    def create_coordinates_table(self):
        table = '\n  x      y\n'
        for x, y in self.__coordinates:
            table += f'{x:>3}{round(y, 2):>7.2f}\n'

        return table

    def create_trajectory(self):
        rounded_coord = [(x, round(y)) for x, y in self.__coordinates]
        rows = max([y for x, y in rounded_coord]) + 1
        columns = max([x for x, y in rounded_coord]) + 1
        graph_matrix = [
            [PROJECTILE if (x, y) in rounded_coord else " " for x in range(columns)]
            for y in range(rows, -1, -1)
        ]
        graph_rows = ["⊣" + "".join(row) for row in graph_matrix] + [
            " " + "⊤" * columns
        ]
        graph = "\n" + "\n".join(graph_rows) + "\n"
        return graph

--fcc-editable-region--

--fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```

# --solutions--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = '∙'
x_axe_tick = '⊤'
y_axe_tick = '⊣'

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s
    
    def __repr__(self):
        return f'{self.__class__}({self.speed}, {self.height}, {self.angle})'

class Graph:
    __slots__ = ('__coordinates')

    def __init__(self, coord):
        self.__coordinates = coord

    def __repr__(self):
        return f"Graph({self.__coordinates})"
        
    def __str__(self):
    	return self.create_trajectory()

    def create_coordinates_table(self):
        table = '\n  x      y\n'
        for x, y in self.__coordinates:
            table += f'{x:>3}{round(y, 2):>7.2f}\n'

        return table

    def create_trajectory(self):
        rounded_coord = [(x, round(y)) for x, y in self.__coordinates]
        rows = max([y for x, y in rounded_coord]) + 1
        columns = max([x for x, y in rounded_coord]) + 1
        graph_matrix = [
            [PROJECTILE if (x, y) in rounded_coord else " " for x in range(columns)]
            for y in range(rows, -1, -1)
        ]
        graph_rows = ["⊣" + "".join(row) for row in graph_matrix] + [
            " " + "⊤" * columns
        ]
        graph = "\n" + "\n".join(graph_rows) + "\n"
        return graph

ball = Projectile(45, 45, 45)
print(ball)

def projectile_helper(speed, height, angle):
    p = Projectile(speed, height, angle)
    print(p)
    coord = p.calculate_all_coordinates()
    g = Graph(coord)
    print(g.create_coordinates_table())
    print(g.create_trajectory())
    
projectile_helper(12, 43, 1)
    
```
