---
id: 6633b4d3271d5b3aea1c1195
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now it's time for the graph.
Create a function `__create_trajectory`.

The function should return a list of strings, where each string is a row of the final graph, you should add a `∙` for each point rapresenting the projectile coordinates.

Important notes: round each coordinate to the nearest integer to have the points to write in the graph, and for the symbol you can use the `PROJECTILE` variable.

For example, for the coordinates from `Projectile(12, 3, 45)`, the output of the method should be as follow:

```py
[
    '                  ',
    '      ∙∙∙         ',
    '    ∙∙   ∙∙∙      ',
    '  ∙∙        ∙     ',
    ' ∙           ∙∙   ',
    '∙              ∙  ',
    '                ∙ ',
    '                  ',
    '                 ∙',
]
```

# --hints--

You should have the method.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__create_trajectory')`
    ))
})
```

The `__create_trajectory` method should return the correct output.

```js
({
    test: () => assert(runPython(
        `
gr = [
    "                  ",
    "      ∙∙∙         ",
    "    ∙∙   ∙∙∙      ",
    "  ∙∙        ∙     ",
    " ∙           ∙∙   ",
    "∙              ∙  ",
    "                ∙ ",
    "                  ",
    "                 ∙",
]
ball = Projectile(12, 3, 45)
coords = ball.calculate_all_coordinates()
g = Graph(coords)
g._Graph__create_trajectory() == gr`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = '∙'

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s
    
    def __repr__(self):
        return f'{self.__class__}({self.speed}, {self.height}, {self.angle})'

class Graph:
    __slots__ = ('__coordinates')

    def __init__(self, coord):
        self.__coordinates = coord

    def __repr__(self):
        return f"Graph({self.__coordinates})"

    def __create_coordinates_table(self):
        table = '\n  x      y\n'
        for x, y in self.__coordinates:
            table += f'{x:>3}{round(y, 2):>7.2f}\n'

        return table

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```
