---
id: 6633bdc3ca38c068c3735984
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Add the x and y axis to the graph so to obtain the following output:

```py
[
    "⊣                  ",
    "⊣      ∙∙∙         ",
    "⊣    ∙∙   ∙∙∙      ",
    "⊣  ∙∙        ∙     ",
    "⊣ ∙           ∙∙   ",
    "⊣∙              ∙  ",
    "⊣                ∙ ",
    "⊣                  ",
    "⊣                 ∙",
    " ⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤",
]
```

The symbols are available as `x_axe_tick` = '⊤' and `y_axe_tick`.

# --hints--

The method should output the correct output.

```js
({
    test: () => assert(runPython(
        `
gr = [
    "⊣                  ",
    "⊣      ∙∙∙         ",
    "⊣    ∙∙   ∙∙∙      ",
    "⊣  ∙∙        ∙     ",
    "⊣ ∙           ∙∙   ",
    "⊣∙              ∙  ",
    "⊣                ∙ ",
    "⊣                  ",
    "⊣                 ∙",
    " ⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤",
]
ball = Projectile(12, 3, 45)
coords = ball.calculate_all_coordinates()
g = Graph(coords)
g._Graph__create_trajectory() == gr`
    ))
})
```

The method should output the correct output.

```js
({
    test: () => assert(runPython(
        `
gr = [
    "⊣                    ",
    "⊣     ∙∙             ",
    "⊣ ∙∙∙∙  ∙∙∙          ",
    "⊣∙         ∙∙        ",
    "⊣            ∙       ",
    "⊣             ∙      ",
    "⊣              ∙     ",
    "⊣               ∙    ",
    "⊣                ∙   ",
    "⊣                 ∙  ",
    "⊣                  ∙ ",
    "⊣                    ",
    "⊣                   ∙",
    " ⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤⊤",
]
ball = Projectile(11, 9, 30)
coords = ball.calculate_all_coordinates()
g = Graph(coords)
g._Graph__create_trajectory() == gr`
    ))
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = '∙'
x_axe_tick = '⊤'
y_axe_tick = '⊣'

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        displacement = horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        return displacement
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s
    
    def __repr__(self):
        return f'{self.__class__}({self.speed}, {self.height}, {self.angle})'

class Graph:
    __slots__ = ('__coordinates')

    def __init__(self, coord):
        self.__coordinates = coord

    def __repr__(self):
        return f"Graph({self.__coordinates})"

    def __create_coordinates_table(self):
        table = '\n  x      y\n'
        for x, y in self.__coordinates:
            table += f'{x:>3}{round(y, 2):>7.2f}\n'

        return table

--fcc-editable-region--
    def __create_trajectory(self):
        rounded_coord = [(x, round(y)) for x, y in self.__coordinates]
        rows = max([y for x, y in rounded_coord]) + 1
        columns = max([x for x, y in rounded_coord]) + 1
        graph_matrix = [
            [PROJECTILE if (x, y) in rounded_coord else " " for x in range(columns)]
            for y in range(rows, -1, -1)
        ]
        return ["".join(x) for x in graph_matrix]
--fcc-editable-region--

ball = Projectile(45, 45, 45)
print(ball)
   
```
