---
id: 660f44f10ea40f300b896a5e
title: Step 75
challengeType: 1
dashedName: step-75
---

# --description--

The text is gone again! Empty strings evaluate to `false`, making them a <dfn>falsy value</dfn>. You will learn more about truthy and falsy values in future projects.

For now, remove your `if` statement entirely. Use `let` to declare a `continueLoop` variable and assign it the boolean `false`. Then use `let` to declare a `done` variable and assign it the value `0`.

# --hints--

You should not have an `if` statement.

```js
assert.notMatch(code, /if\s*\(/);
```

You should use `let` to declare a `continueLoop` variable.

```js
assert.match(code, /let\s+continueLoop/);
```

Your `continueLoop` variable should have the value `false`.

```js
assert.isFalse(continueLoop);
```

You should use `let` to declare a `done` variable.

```js
assert.match(code, /let\s+done/);
```

Your `done` variable should have the value `0`.

```js
assert.equal(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
