---
id: 660f09a2694b59c3a10ee304
title: Step 22
challengeType: 1
dashedName: step-22
---

# --description--

For now, remove your first console log and your `rows[rows.length - 1]` assignment. Leave the second `rows` log statement for later.

# --hints--

You should remove your `console.log(rows[0])` statement.

```js
assert.notMatch(code, /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

You should remove your `rows[rows.length - 1]` reassignment.

```js
assert.notMatch(code, /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

You should not remove your `console.log(rows)` statement.

```js
assert.match(code, /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
