---
id: 646cfc2b8e6fe95c20a819d5
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

As orelhas devem sempre ser colocadas acima da parte da cabeça a que se sobrepõem. Você pode fazer isso com a propriedade `z-index`.

`z-index` é uma propriedade que você pode usar para definir a ordem de sobreposição dos elementos do HTML. Qualquer elemento com um `z-index` maior sempre será posicionado sobre um elemento com um `z-index` menor.

Para ver o `z-index` em ação, defina a propriedade `z-index` da orelha esquerda como `-1`.

# --hints--

O seletor `.cat-left-ear` deve ter a propriedade `z-index` definida como `-1`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.zIndex === '-1')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
