---
id: 646c5d7057c45f432fcdd46c
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Usando um seletor de classe, dê ao elemento `.cat-head` uma largura de `205px` e uma altura de `180px`. Além disso, dê a ele uma borda de `1px solid #000` e um `border-radius` de `46%`.

# --hints--

Você deve ter um seletor `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

O seletor `.cat-head` deve ter a propriedade `width` definida como `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

O seletor `.cat-head` deve ter a propriedade `height` definida como `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

O seletor `.cat-head` deve ter a propriedade `border` definida para `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

O seletor `.cat-head` deve ter a propriedade `border-radius` definida como `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
