---
id: 646ceae4d57c214e6b03576c
title: Passo 22
challengeType: 0
dashedName: step-22
---

# --description--

Dentro do elemento `.cat-right-ear`, crie outro elemento `div` com a classe `cat-right-inner-ear`.

# --hints--

Você não deve alterar o elemento `div` existente com a classe `cat-right-ear`.

```js
assert(document.querySelectorAll('div.cat-right-ear').length === 1);
```

Você deve ter um elemento `div` dentro do elemento `.cat-right-ear`.

```js
assert(document.querySelectorAll('.cat-right-ear div').length === 1);
```

O elemento `div` deve ter a classe `cat-right-inner-ear`.

```js
assert(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          --fcc-editable-region--
          <div class="cat-right-ear">

          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
