---
id: 646c59e942f35541923104bf
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

De volta ao HTML, crie um elemento `main`. Dentro do elemento `main`, adicione um elemento `div` com a classe `cat-head`.

# --hints--

Você deve ter apenas um elemento `main`.

```js
assert(document.querySelectorAll('main').length === 1);
```

Você deve ter um elemento `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

O elemento `div` deve ter a classe `cat-head`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

O elemento `div` deve estar dentro das tags do elemento `main`.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
