---
id: 646c48df8674cf2b91020ecc
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Adicione um elemento `link` dentro do elemento `head`. Para esse elemento `link`, defina o atributo `rel` como `stylesheet` e `href` como `./styles.css`.

# --hints--

Você deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

O elemento `link` deve ter os atributos `rel` e `href` definidos como `stylesheet` e `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
