---
id: 646c48df8674cf2b91020ecc
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Adicione um elemento `link` dentro do elemento `head`. Para esse elemento `link`, defina o atributo `rel` como `stylesheet` e `href` como `./styles.css`.

# --hints--

Você deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

O elemento `link` deve ter um atributo `rel` definido como `stylesheet`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1/)
```

O elemento `link` deve ter um atributo `href` com o valor `./styles.css`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1\s+href=\1\.\/styles\.css\1\s*\/?>/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
