---
id: 646c48df8674cf2b91020ecb
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do elemento `head`, adicione uma tag `meta` com o atributo `charset` de `utf-8`. Adicione também um elemento `title` com o texto `fCC Cat Painting`.

# --hints--

Você deve adicionar um único elemento `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

O elemento `meta` deve ter o atributo `charset` com o valor `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Você deve adicionar um único elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

O elemento `title` deve ter o texto `fCC Cat Painting`. Atente-se a ortografia e letras maiúsculas.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
