---
id: 646c5ace05e4be4211407935
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Para ver o elemento `.cat-head`, dê a ele um plano de fundo com gradiente linear com os valores `#5e5e5e` em `85%` e `#45454f` em `100%`.

Você pode não notar a diferença entre estas duas cores, mas elas estão lá.

# --hints--

O seletor `.cat-head` deve ter uma propriedade `background`.

```js
assert.match(code, /background:/)
```

A propriedade `background` deve ser a função `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

A função `linear-gradient` deve definir a primeira cor como `#5e5e5e` em `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

A função `linear-gradient` deve definir a segunda cor como `#45454f` em `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
