---
id: 646c5e727e56e743c9aed4a1
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

O posicionamento do CSS permite que você defina como você deseja que um elemento seja posicionado no navegador. Ele possui uma propriedade `position` que você pode definir como `static`, `absolute`, `relative`, `sticky` ou `fixed`.

Uma vez que você definir a propriedade `position` do elemento, você pode mover o elemento definindo um pixel ou um valor percentual para um ou mais dentre as propriedades `top`, `right`, `left` ou `bottom`.

`static` é o posicionamento padrão para todos os elementos. Se você o atribuir a um elemento, não poderá movê-lo usando `top`, `right`, `left` ou `bottom`.

Dê a `.cat-head` uma propriedade `position` `static`, definindo as propriedades `top` e `left` como `100px`.


# --hints--

O seletor `.cat-head` deve ter a propriedade `position` definida como `static`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

O seletor `.cat-head` deve ter a propriedade `top` definida como `100px`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

O seletor `.cat-head` deve ter a propriedade `left` definida como `100px`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
