---
id: 646ceb843412c74edee27a79
title: Passo 22
challengeType: 0
dashedName: step-22
---

# --description--

Agora, você aprenderá um truque do CSS para desenhar triângulos.

Ele será usado para desenhar as orelhas do gato.

Usando um seletor de classe, dê ao elemento `.cat-right-ear` as propriedades `height` (altura) e `width` (largura) definido as duas como `100px`. Defina a`background-color` como `white`. Defina as bordas da direita e da esquerda como `35px solid blue`. Defina as bordas superior e inferior como `35px solid red`.

# --hints--

Você deve ter um seletor `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

O seletor `.cat-right-ear` deve ter a propriedade `height` definida como `100px`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

O seletor `.cat-right-ear` deve ter a propriedade `width` definida como `100px`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

O seletor `.cat-right-ear` deve ter a propriedade `background-color` definida como `white`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

O seletor `.cat-right-ear` deve ter a propriedade `border-left` definida como `35px solid blue`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

O seletor `.cat-right-ear` deve ter a propriedade `border-right` definida como `35px solid blue`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

O seletor `.cat-right-ear` deve ter a propriedade `border-top` definida como `35px solid red`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

O seletor `.cat-right-ear` deve ter a propriedade `border-bottom` definida como `35px solid red`. Não se esqueça de adicionar um ponto-e-vírgula.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
