---
id: 646c47867800472a4ed5d2ea
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Comece com a estrutura HTML básica. Adicione uma referência ao `DOCTYPE` do `html` e um elemento `html` com um atributo `lang` definido como `en`. Adicione, também, os elementos `head` e `body` dentro do elemento `html`.

# --hints--

Você deve criar uma declaração de `DOCTYPE` para `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

A declaração `DOCTYPE` deve ser no início do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Você deve ter uma tag de abertura `html` com um atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Você deve acrescentar uma tag de fechamento para `html`.

```js
assert(code.match(/<\/html>/i));
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert(code.match(/<head>/i));
```

Você deve ter uma tag de fechamento para `head`.

```js
assert(code.match(/<\/head>/i));
```

Você deve ter uma tag de abertura para `body`.

```js
assert(code.match(/<body>/i));
```

Você deve ter uma tag de fechamento para `body`.

```js
assert(code.match(/<\/body>/i));
```

O elemento `body` deve estar depois do elemento `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Os elementos `head` e `body` devem estar dentro do elemento `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
