---
id: 6476f5c17f99146071ee884c
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Podemos ver que nada acontece. O elemento `.cat-head` não se moveu apesar de definirmos um `top` e um `left` de `100px`. Isso, porém, não ocorre com o posicionamento `relative`.

Quando você usa o valor `relative`, o elemento ainda está posicionado de acordo com o fluxo normal do documento, mas os valores de `top`, `left`, `bottom` e `right` tornam-se ativos.

Em vez de `static`, dê a `.cat-head` uma propriedade `position` com o valor `relative` e deixe as propriedades `top` e `left` como estão.

# --hints--

O seletor `.cat-head` deve ter a propriedade `position` definida como `relative`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

O seletor `.cat-head` deve ter a propriedade `top` definida como `100px`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

O seletor `.cat-head` deve ter a propriedade `left` definida como `100px`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
