---
id: 6476fc5cf14b276e6d04e82a
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Use um seletor de classe para dar ao elemento `.box` uma largura de `200px`, uma altura de `600px` e uma cor de fundo de `#000`. Dê a ele, também, uma `position` `absolute`, um valor de `top` de `800px` e um valor de `left` de `650px`.

# --hints--

Você deve ter um seletor `.box`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

O seletor `.box` deve ter a propriedade `width` definida como `200px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

O seletor `.box` deve ter a propriedade `height` definida como `600px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

O seletor `.box` deve ter a propriedade `background-color` definida como `#000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

O seletor `.box` deve ter a propriedade `position` definida como `absolute`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

O seletor `.box` deve ter a propriedade `left` definida como `650px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

O seletor `.box` deve ter a propriedade `top` definida como `800px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

