---
id: 6476fb7cbaafa36d65e9cf35
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` é um valor de propriedade `position` que permite deixar um elemento fixo na página, não importando para onde o usuário role a página.

Você terá que criar mais marcações do html para ver como o posicionamento `fixed` funciona. No HTML, crie um elemento `div` com a classe `box`.

# --hints--

Você deve ter um elemento `div`.

```js
assert(document.querySelectorAll('div').length === 2);
```

O elemento `div` deve ter a classe `box`.

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
