---
id: 6476f7a4827bcc61682f2347
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

A propriedade `position` seguinte é a `absolute`. Quando você usa o valor `absolute` para a propriedade `position`, o elemento é retirado do fluxo normal do documento. Então, a posição é determinada pelas propriedades `top`, `right`, `bottom` e `left`.

Defina a propriedade `position` do elemento `.cat-head` como `absolute`. Em seguida, defina as propriedades `top` e `left` com qualquer valor em pixels.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

O seletor `.cat-head` deve ter a propriedade `position` definida como `absolute`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

O seletor `.cat-head` deve ter a propriedade `top` definida como um número qualquer em `px`. Não se esqueça de adicionar um ponto e vírgula.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

O seletor `.cat-head` deve ter a propriedade `left` definida como um número qualquer em `px`. Não se esqueça de adicionar um ponto e vírgula.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
