---
id: 64770351e8586671ec0911f0
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

O valor da última propriedade `position` é `sticky`. O posicionamento `sticky` é um híbrido dos posicionamentos `relative` e `fixed`. Ele permite que um elemento "**se grude**" a uma posição específica dentro do elemento que o contém ou da viewport, com base na posição de rolagem.

Altere o valor da propriedade `position` de `.cat-head` para `sticky`, defina `top` como `0` e remova `left` e seu valor.

**Observação**: para ver como `sticky` funciona, você tem de colocar alguns textos antes e depois do elemento `div` de `.cat-head`. Se você rolar para baixo depois disso, verá que `.cat-head` fica preso na parte superior e permanece lá.

# --hints--

O seletor `.cat-head` deve ter a propriedade `position` definida como `sticky`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

O seletor `.cat-head` deve ter a propriedade `top` definida como `0`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Você não deve ter a propriedade `left` e seu valor no código.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
