---
id: 6477062778c85972eb648030
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Você, agora, deve centralizar a cabeça do gato.

Dê ao elemento `.cat-head` uma propriedade `position` definida como `absolute`. Defina um valor de `0` para as propriedades `right`, `left`, `top` e `bottom`. Em seguida, defina a propriedade `margin` de todos os lados como `auto`. Essa é uma maneira de centralizar um elemento verticalmente e horizontalmente usando posicionamento em CSS.

# --hints--

O seletor `.cat-head` deve ter a propriedade `position` definida como `absolute`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

O seletor `.cat-head` deve ter a propriedade `top` definida como `0`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

O seletor `.cat-head` deve ter a propriedade `left` definida como `0`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

O seletor `.cat-head` deve ter a propriedade `right` definida como `0`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

O seletor `.cat-head` deve ter a propriedade `bottom` definida como `0`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

O seletor `.cat-head` deve ter a propriedade `margin` definida como `auto`. Não se esqueça de adicionar um ponto e vírgula.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
