---
id: 647d855ac12cd436059acd39
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Use o seletor universal para adicionar `box-sizing: border-box;` ao CSS. Isso garante que os elementos incluirão preenchimento e borda na largura e altura especificadas.

# --hints--

Você deve ter o seletor universal (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

O seletor universal deve ter a propriedade `box-sizing` definida como `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
