---
id: 6676a8a8d9861319ee901a90
title: Passo 24
challengeType: 0
dashedName: step-24
---

# --description--

Agora, você pode começar a trabalhar nas orelhas do gato.

Limpe seu código de revisão removendo o seletor `.cat-right-ear` e todas as suas propriedades.

Usando um seletor de classe, dê ao elemento `.cat-left-ear` uma borda esquerda e uma borda direita de `35px solid transparent`. Além disso, defina a borda inferior como `70px solid #5e5e5e`.

# --hints--


O seletor `.cat-right-ear` não deve ter a propriedade `height`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

O seletor `.cat-right-ear` não deve ter a propriedade `width`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


O seletor `.cat-right-ear` não deve ter uma propriedade `border-left`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

O seletor `.cat-right-ear` não deve ter uma propriedade `border-right`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

O seletor `.cat-right-ear` não deve ter uma propriedade `border-top`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

O seletor `.cat-right-ear` não deve ter uma propriedade `border-bottom`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

Você não deve ter um seletor `.cat-right-ear`.

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Você deve ter um seletor `.cat-left-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

O seletor `.cat-left-ear` deve ter a propriedade `border-left` definida como `35px solid transparent`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

O seletor `.cat-left-ear` deve ter a propriedade `border-right` definida como `35px solid transparent`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

O seletor `.cat-left-ear` deve ter a propriedade `border-bottom` definida como `70px solid #5e5e5e`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
