---
id: 6690e10ebe2181212abc9652
title: Design a Business Card
challengeType: 14
dashedName: lab-business-card
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. Your page background should be `rosybrown`. The overall font should be `Arial`, with a fallback of `sans-serif`.   
2. All of the `p` elements should have a top and bottom margin of `5px` and a left and right margin of `0px`.
3. All links on the page should have a default of no underline. When you hover over the links, there should be an underline.
4. You should have a `div` with the class `business-card`. 
5. The `business-card` class should be `300px` wide and it should have a background color of `lavenderblush`. There should be a padding of `20px` all around and a top margin of `100px`. The text should be center aligned and the font size should be `16px`. The left and right margins should be set to `auto`.
6. Inside the `buisness-card` div, there should be an `img` element with the class `profile-image`. You can set the image source to `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` if you like. There should be an `alt` with  a meaningful description. The image size should be `150px by 150px`. The bottom margin should be `10px` and there should be a `2px` solid `lightgray` border around the image.
7. Inside the `business-card` `div`, after the `img` element, you should have three `p` elements with the class `full-name`, `designation`, and `company`, respectively.
8. The first `p` element should contain your name. Its font size should be `24px` and the `color` should be `darkslategray`.
9.  The second `p` element should contain your designation. Its `font-size` property should be `18px` and the `color` should be set to `gray`.
10. The third `p` element should have your company name. Its font size should be `16px` and the `color` should be set to `darkgray`.
11. There should be an `hr` element below the `p` element containing your company name. The `hr` element in general should have a `1px` dashed top border in `slategray` color. There should be a top and bottom `margin` of `20px` and a left and right `margin` of `0px`.
12. After the `hr` element, there should be two `p` elements. In the first `p` element, you should have an email address as the text. In the second `p` element, the text should be a phone number. After the two `p` elements, there should be an `a` element with the text `Portfolio` pointing to any valid link. Add an `hr` element after the `a` element containing the portfolio link.
13. You should have another `div` with the class `social-media`. This div should have a top margin of `10px` and a font size of `14px`. The links inside this div should be `maroon` and should have a margin of `10px` all around. Within this class, there should be an `h2` element with the text `Connect with me`. 
14. Inside the `social-media` div, there should be three `a` elements. The `a` elements should have the text `Twitter`, `LinkedIn` and `GitHub` respectively with links to valid websites.

# --hints--

Your page background color should be `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body').getPropVal('background-color'), 'rosybrown');
```

Your page font should be set to `Arial` with a fallback of `sans-serif`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body').getPropVal('font-family'), 'Arial, sans-serif');
```

Your `p` elements should have a top and bottom margin of `5px` and a left and right margin of `0px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('p').getPropVal('margin'), '5px 0px');
```

The links of the page should have no underline by default.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a').getPropVal('text-decoration'), 'none');
```

The page links should have an underline when hovered over.

```js
const gs = (s) => new __helpers.CSSHelp(document).getStyle(s);
const a_hover = gs('a:hover')?.textDecoration;
assert.equal(a_hover, 'underline');
```

You should have one main `div` with the class `business-card`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'business-card');
```

Your `.business-card` class should have a width of `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('width'), '300px');
```

Your `.business-card` class should have a background color of `lavenderblush`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('background-color'), 'lavenderblush');
```

Your `.business-card` class should have a padding of `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('padding'), '20px');
```

Your `.business-card` class should have a margin of `100px` from the top.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('margin-top'), '100px');
```

The text inside the `.business-card` class should be center-aligned.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('text-align'), 'center');
```

The font size of the text inside the `.business-card` class should be `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('font-size'), '16px');
```

The left and right margins of the `.business-card` class should be set to `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card').getPropVal('margin-left'), 'auto');
```

Inside the `.business-card` class, there should be an `img` element with the class `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

The image source should be set to a valid image.

```js
assert.match(document.querySelector('.profile-image').getAttribute('src'),/.+/);
```

(Always Passing) The `alt` attribute of the image should be set to a meaningful description.

```js

assert.match(document.querySelector('.img')?.getAttribute('alt'), /^.+$/);

```

The image size should be `150px by 150px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image').getPropVal('width'), '150px');
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image').getPropVal('height'), '150px');
```

The bottom margin of the image should be `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image').getPropVal('margin-bottom'), '10px');
```

There should be a `2px` solid `lightgray` border around the image.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image').getPropVal('border'), '2px solid lightgray');
```

Inside the `.business-card` class, after the `img` element, there should be a `p` element with the class `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

After the `full-name` `p` element, there should be a `p` element with the class `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

After the `designation` `p` element, there should be a `p` element with the class `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

(Always Passing, to test) The first `p` element with the class `full-name` should contain your name.

```js
assert.match(document.querySelector('p.full-name')?.textContent, /\S+/);
```

The first `p` element with the class `full-name` should have a font size of `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.full-name').getPropVal('font-size'), '24px');
```

The font color of the first `p` element with the class `full-name` should be `darkslategray`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.full-name').getPropVal('color'), 'darkslategray');
```

(Always Passing) The second `p` element with the class `designation` should contain your designation.

```js
assert.match(document.querySelector('p.designation')?.textContent, /.+/);
```

The second `p` element with the class `designation` should have a font size of `18px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.designation').getPropVal('font-size'), '18px');
```

The font color of the second `p` element with the class `designation` should be `gray`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.designation').getPropVal('color'), 'gray');
```

(Always Passing) The third `p` element with the class `company` should contain your company name.

```js
assert.match(document.querySelector('p.company')?.textContent, /.+/);
```

The third `p` element with the class `company` should have a font size of `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.company').getPropVal('font-size'), '16px');
```

The font color of the third `p` element with the class `company` should be `darkgray`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.company').getPropVal('color'), 'darkgray');
```

After the third `p` element with the class `company`, there should be an `hr` element.

```js
assert.exists(document.querySelector('p.company + hr'));
```

The `hr` elements should have a `1px` dashed top border in `slategray` color.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('hr').getPropVal('border-top'), '1px dashed slategray');
```

The `hr` elements should have a top and bottom margin of `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('hr').getPropVal('margin'), '20px 0px');
```

The `hr` elements should have a top and bottom margin of `20px` and a left and right margin of `0px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('hr').getPropVal('margin'), '20px 0px');
```

(Always Passing) After the first `hr` element, there should be a `p` element with an email address as your text.

```js
assert.match(document.querySelector('hr + p')?.textContent, /.+/);
```

After the email `p` element, there should be another `p` element with a phone number as your text.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

After the phone number `p` element, there should be an `a` element with the text `Portfolio` 

```js
assert.equal(document.querySelector('.business-card a').textContent, 'Portfolio');
```

(Always Passing) Your portfolio `a` element should have a valid `href`.

```js
assert.match(document.querySelector('.business-card a')?.href, /.+/);
```

After the `a` element containing your portfolio, there should be another `hr` element as a divider.

```js
assert.exists(document.querySelector('a + hr'));
```

After the second `hr` element, there should be a `div` element with the class `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

The `.social-media` class should have a top margin of `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.social-media').getPropVal('margin-top'), '10px');
```

The `.social-media` class should have a font size of `14px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.social-media').getPropVal('font-size'), '14px');
```

The links inside the `.social-media` class should be `maroon`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.social-media a').getPropVal('color'), 'maroon');
```

The links inside the `.social-media` class should have a margin of `10px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.social-media a').getPropVal('margin'), '10px');
```

Inside the `.social-media` class, there should be an `h2` element with the text `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2').textContent, 'Connect with me');
```

Inside the `.social-media` class, there should be an `a` element with the text `Twitter` 

```js
assert.equal(document.querySelectorAll(".social-media a")[0].textContent, 'Twitter');
```

(Always Passing) The Twitter `href` should point to a valid link.

```js
const paragraphs = document.querySelectorAll('a');
const twitterP = Array.from(paragraphs).find(a => a.textContent.includes('Twitter'));
assert.match(twitterP?.href, /.+/);
```

Inside the `.social-media` class, there should be an `a` element with the text `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1].textContent, 'LinkedIn');
```

The LinkedIn `href` should point to a valid link.

```js
const anchors = document.querySelectorAll('a');
const linkedInA = Array.from(anchors).find(a => a.textContent.includes('LinkedIn'));
assert.match(linkedInA?.href, /.+/);
```

Inside the `.social-media` class, there should be an `a` element with the text `GitHub`

```js
assert.equal(document.querySelectorAll(".social-media a")[2].textContent, 'GitHub');
```

(Always Passing) The GitHub `href` should point to a valid link.

```js
const anchors = document.querySelectorAll('a');
const GitHubA = Array.from(anchors).find(a => a.textContent.includes('GitHub'));
assert.match(GitHubA?.href, /.+/);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class = "company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

a:hover {
    text-decoration: underline;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    border-top: 1px dashed slategray;
    margin: 20px 0;
}
```
