---
id: 65c4ef8463f869d0eaf87c67
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside the `TreeNode` class, replace `pass` with an `__init__` method so that you can initialize the attributes of the object. Don't add any parameters and use the `pass` keyword to avoid an error.

# --hints--

You should define an `__init__` method inside your `TreeNode` class. Remember to use the `pass` keyword.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\([^(]*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```

