---
id: 65c4ef8463f869d0eaf87c67
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside the `TreeNode` class, replace `pass` with an `__init__` method so that you can initialize the attributes of the object. Don't add any parameters for now.

# --hints--

You should remove `pass` keyword from the `TreeNode` class.

```js
({ test: () => {
    assert.isFalse((runPython(`_Node(_code).find_class("TreeNode").has_pass()`)))
}})
```

You should define an `__init__` method inside your `TreeNode` class. Remember to use the `pass` keyword inside.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`))
}})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
