---
id: 65ca06475e9b8aff8ea5a5f1
title: Step 34
challengeType: 20
dashedName: step-34
---

# --description--

After your existing `if`, write another one that checks if the target key is less than the current node key. Remember to use `pass`.

# --hints--

You should have an `if` condition that checks if `key` is less than `node.key`. Remember to use `pass`.

```js
const after_split = code.split('def _delete(self, node, key):')[1];
assert.match(after_split, /if\s+key\s+<\s+node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node

--fcc-editable-region--

```
