---
id: 65cb40ee75b4f787b21066a8
title: Step 48
challengeType: 20
dashedName: step-48
---

# --description--

Assign the result of the `_delete` method to `self.root` in the `delete` method.

# --hints--

You should assign the result of the `_delete` method to `self.root` in the `delete` method.

```js
const after_split = code.split('def delete(self, key):')[1];
assert.match(after_split, /self\.root\s*=\s*self\._delete\(self\.root,\s*key\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key
        
--fcc-editable-region--
    def delete(self, key):
        self._delete(self.root, key)

--fcc-editable-region--

```
