---
id: 61408f155e798909b6908712
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Per abilitare la navigazione sulla pagina, aggiungi una lista puntata con i tre seguenti elementi:

- `INFO`
- `HTML`
- `CSS`

Il testo degli elementi della lista dovrebbe essere racchiuso in tag di ancoraggio.

# --hints--

Dovresti annidare un elemento `ul` all'interno dell'elemento `nav`.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

Dovresti annidare tre elementi `li` dentro l'elemento `ul`.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

Dovresti annidare un elemento `a` dentro ogni elemento `li`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

Dovresti dare al primo elemento `a` il testo `INFO`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

Dovresti dare al secondo elemento `a` il testo `HTML`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

Dovresti dare al terzo elemento `a` il testo `CSS`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
