---
id: 5a24c314108439a4d403614d
title: Визначте дію Redux
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

Оскільки Redux — це фреймворк керування станами, оновлення стану є одним з його основних завдань. Всі оновлення стану в Redux запускаються через відправлення дій. Дія — це звичайний об’єкт JavaScript, який містить інформацію про подію, що відбулася. Сховище Redux отримує ці об’єкти дій, а потім відповідно оновлює стан. Інколи дія Redux також містить певні дані. Наприклад, дія може містити ім’я користувача після його входу. Хоча дані необов’язкові, дії повинні містити властивість `type`, яка вказує «тип» дії, що відбулася.

Вважайте дії Redux месенджерами, які доставляють інформацію про події, що відбуваються у застосунку, до сховища Redux. Після цього сховище здійснює оновлення стану на основі дії, що відбулася.

# --instructions--

Написати дію Redux так само просто, як і оголосити об’єкт з властивістю типу. Оголосіть об’єкт `action` та надайте йому властивість `type` зі значенням рядка `'LOGIN'`.

# --hints--

Об’єкт `action` має існувати.

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

Об’єкт `action` повинен мати властивість `type` зі значенням `LOGIN`.

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
