---
id: 646cf1206cac5f51804f49cf
title: 步骤 25
challengeType: 0
dashedName: step-25
---

# --description--

要正确定位左耳，你可以使用 CSS transform 将其旋转一定的角度。

`transform` 属性允许你在不改变布局或影响周围元素的情况下修改元素的形状、位置和大小。 它具有 `translate()`、`rotate()`、`scale()`、`skew()` 和 `matrix()` 等功能。

将 `transform` 属性设置为 `rotate(-45deg)`，看看会发生什么。

# --hints--

你应该将你的 `.cat-left-ear` 元素的 `transform` 属性设置为 `rotate(-45deg)`。 Don't forget to add a semicolon.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform === 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
