---
id: 646ce8bb4b35544d501c7184
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

在你的 `.cat-ears` 元素中创建两个 `div` 元素，分别使用类 `cat-left-ear` 和 `cat-right-ear` 。

# --hints--

You should not change the existing `div` element with the class `cat-ears`.

```js
assert(document.querySelectorAll('div.cat-ears').length === 1);
```

你应该在你的 `.cat-ears` 元素中创建两个 `div` 元素。

```js
assert(document.querySelectorAll('.cat-ears div').length === 2);
```

你的第一个 `div` 元素应该有 `cat-left-ear` 类。

```js
assert(document.querySelectorAll('.cat-ears div')[0]?.classList.contains('cat-left-ear'));
```

你的第二个 `div` 元素应该有 `cat-right-ear` 类。

```js
assert(document.querySelectorAll('.cat-ears div')[1]?.classList.contains('cat-right-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        --fcc-editable-region--
        <div class="cat-ears">

        </div>
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
