---
id: 646c47867800472a4ed5d2ea
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

从基本的 HTML 结构开始。添加`html`的`DOCTYPE`引用和`html`元素，其`lang`属性设置为`en`。此外，在`html`元素中添加`head`和`body`元素。

# --hints--

您应该具有`html`的`DOCTYPE`声明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你的 `DOCTYPE` 声明应该位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你的 `html` 元素应该有一个带有值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html>/i));
```

你需要一个开始标签 `<head>`

```js
assert(code.match(/<head>/i));
```

你需要一个结束标签 `</head>`

```js
assert(code.match(/<\/head>/i));
```

你需要一个开始标签 `<body>`

```js
assert(code.match(/<body>/i));
```

你需要一个结束标签 `</body>`

```js
assert(code.match(/<\/body>/i));
```
你的 `body` 元素需要 在 `head` 元素之后

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

你的 `head` 和 `body` 元素应该在 `html` 元素中。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
