---
id: 646c59e942f35541923104bf
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

回到你的 HTML 中，创建一个 `main` 元素。 在那个 `main` 元素内部，添加一个带有类 `cat-head` 的 `div` 元素。

# --hints--

你应该有一个 `main` 元素。

```js
assert(document.querySelectorAll('main').length === 1);
```

你应该有一个 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 1);
```

你的 `div` 元素应该有类 `cat-head`。

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

你的 `div` 元素应该在 `main` 标签内。

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
