---
id: 646c47867800472a4ed5d2ea
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

从基本的 HTML 结构开始。 添加一个 `html` 的 `DOCTYPE` 声明，再添加一个 `html` 元素并设定它的属性 `lang` 为 `en`。 然后，在 `html` 元素中添加一个 `head` 和 `body` 元素。

# --hints--

你应该有 `DOCTYPE` 声明为 `html`。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你应该有一个属性 `lang` 设定为 `en` 的 `html` 开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你应该有一个 `html` 结束标签。

```js
assert(code.match(/<\/html>/i));
```

你应该有一个 `head` 开始标签。

```js
assert(code.match(/<head>/i));
```

你应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head>/i));
```

你应该有一个 `body` 开始标签。

```js
assert(code.match(/<body>/i));
```

你应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body>/i));
```

你的 `body` 元素应该在 `head` 元素之后。

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

你的 `head` 和 `body` 元素应该在 `html` 元素里面。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
