---
id: 646c5e727e56e743c9aed4a1
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

CSS定位使您可以设置您希望在浏览器中如何定位一个元素。 它有一个 `position` 属性，你可以设置值为 `static`, `absolute`, `relative`, `sticky` 或 `fixed`.

一旦您设置了元素的 `position` 属性，您可以通过为`top`, `right`, `left`, 或 `bottom` 属性中的一个或者多个设置像素值或百分比值来移动该元素。

`static` 是所有元素的默认定位。 If you assign it to an element, you won't be able to move it around with `top`, `right`, `left`, or `bottom`.

Give `.cat-head` a `position` property of `static`, then set the `top` and `left` properties to `100px` each.


# --hints--

Your `.cat-head` selector should have a `position` property set to `static`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Your `.cat-head` selector should have a `top` property set to `100px`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Your `.cat-head` selector should have a `left` property set to `100px`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
