---
id: 646ceb843412c74edee27a79
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

现在你将学习一个 CSS 技巧来绘制三角形。

这将用于绘制猫耳朵。

使用类选择器，将 `.cat-right-ear` 元素的 `height` 和 `width` 属性设置为 `100px`。 将 `background-color` 设置为 `white`。 将左侧和右侧的边框设置为 `35px solid blue`。 将顶部和底部的边框设置为 `35px solid red`。

# --hints--

你应该有一个 `.cat-right-ear` 选择器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

你的 `.cat-right-ear` 选择器应将 `height` 属性设置为 `100px`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

你的 `.cat-right-ear` 选择器应将 `width` 属性设置为 `100px`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

你的 `.cat-right-ear` 选择器应将 `background-color` 属性设置为 `white`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

你的 `.cat-right-ear` 选择器应设置一个 `border-left` 属性为 `35px solid blue`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

你的 `.cat-right-ear` 选择器应设置一个 `border-right` 属性为 `35px solid blue`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

你的 `.cat-right-ear` 选择器应设置一个 `border-top` 属性为 `35px solid red`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

你的 `.cat-right-ear` 选择器应设置一个 `border-bottom` 属性为 `35px solid red`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
