---
id: 646c5d7057c45f432fcdd46c
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

使用类选择器，设置 `.cat-head` 元素的宽度为 `205px`、高度为 `180px`。 同时，设置它的边框为 `1px solid #000`、`border-radius` 为 `46%`。

# --hints--

你应该有一个 `.cat-head` 选择器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

你的 `.cat-head` 选择器应将 `width` 设置为 `205px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

你的 `.cat-head` 选择器应将 `height` 设置为 `180px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

你的 `.cat-head` 选择器应将 `border` 设置为 `1px solid #000`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

你的 `.cat-head` 选择器应将 `border-radius` 设置为 `46%`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
