---
id: 646c48df8674cf2b91020ecb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一个 `meta` 标签，其 `charset` 属性为 `utf-8`。还要添加一个带有文本 `fcc Cat painting` 的 `title` 元素。

# --hints--

你需要添加一个 `meta` 元素。

```js
assert(document.querySelectorAll('meta').length === 1);
```

你的 `meta` 元素应该有一个 `charset` 属性，其值为 `utf-8`。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

你需要添加一个 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

你的 `title` 元素应该有一个文本 `fCC Cat Painting`。请注意，拼写和大小写是重要的。

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
