---
id: 646c5e727e56e743c9aed4a1
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

CSS 定位可让你设定元素在浏览器中的位置。 它有一个 `position` 属性，你可以设置其值为 `static`、`absolute`、`relative`、`sticky` 或 `fixed`。

当你设置了元素的 `position` 属性，你可以通过为 `top`、`right`、`left`、 `bottom` 属性中的一个或者多个设置像素值或百分比值来移动该元素。

`static` 是所有元素的默认定位。 如果你将元素设置为这个定位，你将无法用 `top`、`right`、`left` 或 `bottom` 移动元素。

设置 `.cat-head` 元素的 `position` 属性为 `static`，分别设置 `top` 和 `left` 属性为 `100px`。


# --hints--

你的 `.cat-head` 选择器应将 `position` 属性设置为 `static`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

你的 `.cat-head` 选择器应将 `top` 属性设置为 `100px`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

你的 `.cat-head` 选择器应将 `left` 属性设置为 `100px`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
