---
id: 646c5ffef5598d449b52ec12
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

现在你应该绘制猫的耳朵。 猫有左耳和右耳，每只耳朵里都有内耳。

在你的 `.cat-head` 元素中，创建一个具有 `cat-ears` 类的 `div` 元素。

# --hints--

你不应该更改现有的 `div` 元素及其 `cat-head` 类。

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

你应该在你的 `.cat-head` 元素内创建一个 `div` 元素。

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

你的 `div` 元素应具有类 `cat-ears`。

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">

      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
