---
id: 646c48df8674cf2b91020ecc
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `head` 元素中添加一个 `link` 元素。 对于那个 `link` 元素，将 `rel` 属性设置为 `stylesheet`，并将 `href` 设置为 `./styles.css`。

# --hints--

你应该有一个 `link` 元素。

```js
assert.match(code, /<link/)
```

你的 `link` 元素应该有 `rel` 和 `href` 属性，分别设置为 `stylesheet` 和 `./styles.css`。

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
