---
id: 646cf2249f02ca5233d9af7c
title: 步骤 28
challengeType: 0
dashedName: step-28
---

# --description--

现在你可以开始处理猫的右耳了。 你已经有它的 HTML 了。

使用类选择器，为 `.cat-right-ear` 元素的左右边框分别设置 `35px solid transparent`。 另外，将底边设置为 `70px solid #5e5e5e`。

# --hints--

你应该有一个 `.cat-right-ear` 选择器。

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear'))
```

你的 `.cat-right-ear` 选择器应设置一个 `border-left` 属性为 `35px solid transparent`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

你的 `.cat-right-ear` 选择器应设置一个 `border-right` 属性为 `35px solid transparent`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

你的 `.cat-right-ear` 选择器应设置一个 `border-bottom` 属性为 `70px solid #5e5e5e`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
