---
id: 646c48df8674cf2b91020ecb
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一个 `meta` 标签，设置它的 `charset` 属性为 `utf-8`。 同时添加一个包含文本 `fCC Cat Painting` 的 `title` 元素。

# --hints--

你应该只添加一个 `meta` 元素。

```js
assert(document.querySelectorAll('meta').length === 1);
```

你的 `meta` 元素应该有一个 `charset` 属性值为 `utf-8`。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

你应该只添加一个 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

你的 `title` 元素应包含文本 `fCC Cat Painting`。 请注意拼写和大小写的问题。

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
