---
id: 6476f5c17f99146071ee884c
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

你可以看到什么也没有发生。 尽管将 `top` 和 `left` 都设置为 `100px`，但 `.cat-head` 元素并未移动。 但 `relative` 定位的情况并非如此。

当你使用 `relative` 值时，元素仍然按照文档的正常流程定位，但是 `top`、`left`、`bottom` 和 `right` 值变为活动状态。

不要使用 `static`，而是将 `.cat-head` 设置为 `relative` 定位，并将 `top` 和 `left` 属性保持原样。

# --hints--

你的 `.cat-head` 选择器应将 `position` 属性设置为 `relative`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

你的 `.cat-head` 选择器应将 `top` 属性设置为 `100px`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

你的 `.cat-head` 选择器应将 `left` 属性设置为 `100px`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
