---
id: 6476f7a4827bcc61682f2347
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

下一个位置属性是 `absolute`。 当你为 `position` 属性使用 `absolute` 值时，该元素将从文档的正常流程中取出，然后其位置由 `top`、`right`、`bottom` 和 `left` 属性决定。

将 `.cat-head` 元素的 position 属性设置为 `absolute`，然后将 `top` 和 `left` 属性设置为任意像素值。

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

你的 `.cat-head` 选择器应将 `position` 属性设置为 `absolute`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

你的 `.cat-head` 选择器应将 `top` 属性设置为以 `px` 结尾的任意数字。 确保你添加了分号。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

你的 `.cat-head` 选择器应将 `left` 属性设置为以 `px` 结尾的任意数字。 确保你添加了分号。

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
