---
id: 6476fd4213318f6ee211028a
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

现在将 `.cat-head` 的 `position` 属性值替换为 `fixed`。 保留 `top` 和 `left` 不变。

之后，上下滚动查看 `fixed` 值的工作原理。

# --hints--

你的 `.cat-head` 选择器应将 `position` 属性设置为 `fixed`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

你的 `.cat-head` 选择器应将 `top` 属性设置为 `100px`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

你的 `.cat-head` 选择器应将 `left` 属性设置为 `100px`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
