---
id: 6476fc5cf14b276e6d04e82a
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

使用类选择器为你的 `.box` 元素指定宽度为 `200px`、高度为 `600px` 以及背景颜色为 `#000`。 另外，为其设置 `position` 为 `absolute`、`top` 为 `800px` 和 `left` > 为 `650px`。

# --hints--

你应该有一个 `.box` 选择器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

你的 `.box` 选择器应将 `width` 属性设置为 `200px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

你的 `.box` 选择器应将 `height` 属性设置为 `600px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

你的 `.box` 选择器应将 `background-color` 属性设置为 `#000`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

你的 `.box` 选择器应将 `position` 属性设置为 `absolute`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

你的 `.box` 选择器应将 `left` 属性设置为 `650px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

你的 `.box` 选择器应将 `top` 属性设置为 `800px`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

