---
id: 6476fb7cbaafa36d65e9cf35
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` 是一个 `position` 属性值，可让你将元素固定到页面，无论用户滚动到页面上的哪个位置。

你必须创建更多标记才能了解 `fixed` 定位的工作原理。 在 HTML 中，创建一个带有 `box` 类的 `div` 元素。

# --hints--

你应该有一个 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 2);
```

你的 `div` 元素应该具有 `box` 类。

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
