---
id: 64770351e8586671ec0911f0
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

最后一个位置属性值为 `sticky`。 `sticky` 定位是 `relative` 定位和 `fixed` 定位的混合体。 它允许元素根据滚动位置**粘**到其包含元素或视口中的特定位置。

将 `.cat-head` 的 `position` 属性值更改为 `sticky`，将 `top` 设置为 `0`，然后删除 `left` 及其值。

**注意**：要了解 `sticky` 的工作原理，你必须在`.cat-head` `div` 之前和之后放置一些文本。 如果你在此之后向下滚动，你会看到 `.cat-head` 卡在顶部并保留在那里。

# --hints--

你的 `.cat-head` 选择器应将 `position` 属性设置为 `sticky`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

你的 `.cat-head` 选择器应将 `top` 属性设置为 `0`。 确保你添加了分号。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

你的代码中不应包含 `left` 属性及其值。

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
