---
id: 647d855ac12cd436059acd39
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

使用通用选择器将 `box-sizing: border-box;` 添加到你的 CSS 中。 这可确保元素在其指定的宽度和高度内包含内边距和边框。

# --hints--

你应该有通用选择器（`*`）。

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

你的通用选择器应将 `box-sizing` 属性设置为 `border-box`。

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
