---
id: 6676a8b01e56ec1a1e07c202
title: 步骤-23
challengeType: 0
dashedName: step-23
---

# --description--

注意你现在有一个厚边框的白色方块。 边框有对角线，可用于创建三角形。

在同一个类选择器中，将 `height` 和 `width` 调整为 `0`。 将左侧、右侧和顶部边框设置为 `transparent`。

删除 `background-color` 属性，你应该看到一个三角形。

# --hints--


你应该有 `.cat-right-ear` 选择器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

你的 `.cat-right-ear` 选择器应将 `height` 属性设置为 `0`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

你的 `.cat-right-ear` 选择器应将 `width` 属性设置为 `0`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

你的 `.cat-right-ear` 选择器应将 `border-left` 属性设置为 `35px solid transparent`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

你的 `.cat-right-ear` 选择器应将 `border-right` 属性为 `35px solid transparent`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

你的 `.cat-right-ear` 选择器应将 `border-top` 属性设置为 `35px solid transparent`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


你的 `.cat-right-ear` 选择器应将 `border-bottom` 属性为 `35px solid red`。 不要忘记加上分号。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

你的 `.cat-right-ear` 选择器不应该有 `background-color` 属性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
