---
id: 66faa50000000000000000cb
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the dungeon explorer. Across 50 steps you will load a 4-room Tiled map, drive a player with cursor keys, talk to an NPC, pick up a key, unlock a door, fight a slime, save your progress, and reach the victory stairs.

A `UIScene` shell is already defined below — that's the HUD layer for dialogue, notices, HP, and minimap. You'll fill it in later. For now, declare `MainScene` extending `Phaser.Scene` with a constructor that calls `super("MainScene")` and empty `preload`, `create`, and `update(time, delta)` methods.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {}

  create() {}

  update(time, delta) {}
}
```

The `config` object below already references both scenes.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["\']MainScene["\']\s*\)/);
```

You should declare a `preload` method.

```js
assert.match(code, /\bpreload\s*\(\s*\)\s*\{/);
```

You should declare a `create` method.

```js
assert.match(code, /\bcreate\s*\(\s*\)\s*\{/);
```

You should declare an `update` method that takes `time` and `delta`.

```js
assert.match(code, /\bupdate\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
