---
id: 66faa50000000000000000cc
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
A Tiled map needs two assets in the loader: the tileset PNG and the map JSON. Start with the tileset image. Inside `preload`, load `/curriculum-assets/phaser/tilemaps/tileset-dungeon.png` under the cache key `"tiles"`:

```js
this.load.image(
  "tiles",
  "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
);
```

# --hints--
You should load `tileset-dungeon.png` under the key `"tiles"` in `preload`.

```js
assert.match(code, /this\.load\.image\(\s*["\']tiles["\']\s*,\s*["\']\/curriculum-assets\/phaser\/tilemaps\/tileset-dungeon\.png["\']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
