---
id: 66faa50000000000000000ce
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
With both assets cached, the next job is to construct the runtime `Tilemap` object. `this.make.tilemap({ key })` reads the cached JSON; `map.addTilesetImage("dungeon", "tiles")` connects the map's logical tileset name (the one Tiled stored, `"dungeon"`) to the loader cache key for the PNG (`"tiles"`).

Inside `create`, add:

```js
const map = this.make.tilemap({ key: "dungeon" });
const tileset = map.addTilesetImage("dungeon", "tiles");
this.map = map;
```

# --hints--
You should call `this.make.tilemap({ key: "dungeon" })` and store the result in a `map` constant.

```js
assert.match(code, /const\s+map\s*=\s*this\.make\.tilemap\(\s*\{\s*key:\s*["\']dungeon["\']\s*\}\s*\)/);
```

You should call `map.addTilesetImage("dungeon", "tiles")` and store the result in a `tileset` constant.

```js
assert.match(code, /const\s+tileset\s*=\s*map\.addTilesetImage\(\s*["\']dungeon["\']\s*,\s*["\']tiles["\']\s*\)/);
```

You should keep a reference on the scene with `this.map = map`.

```js
assert.match(code, /this\.map\s*=\s*map/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
--fcc-editable-region--
--fcc-editable-region--
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
