---
id: 66faa50000000000000000d4
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Top-down movement with 4 cursor keys lets you press two at once for diagonals — but a naive implementation makes diagonals 1.41× faster. Normalize the direction vector before scaling by speed so all 8 directions feel the same. Also remember the latest direction in `this.facing` so the next steps (attacks, dialogue glyphs) know which way the player is facing.

Inside `update`:

```js
this.facing = this.facing || { x: 0, y: 1 };
const speed = 100;
const body = this.player.body;
body.setVelocity(0);
let vx = 0;
let vy = 0;
if (this.cursors.left.isDown) vx = -1;
else if (this.cursors.right.isDown) vx = 1;
if (this.cursors.up.isDown) vy = -1;
else if (this.cursors.down.isDown) vy = 1;
if (vx !== 0 || vy !== 0) {
  const len = Math.hypot(vx, vy);
  body.setVelocity((vx / len) * speed, (vy / len) * speed);
  this.facing = { x: vx, y: vy };
}
```

# --hints--
You should reset the player's velocity to zero each frame.

```js
assert.match(code, /body\.setVelocity\(\s*0\s*\)/);
```

You should detect the four cursor key presses to set `vx` and `vy`.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*this\.cursors\.right\.isDown[\s\S]*this\.cursors\.up\.isDown[\s\S]*this\.cursors\.down\.isDown/);
```

You should normalize the direction vector with `Math.hypot(vx, vy)` before scaling by `speed`.

```js
assert.match(code, /const\s+len\s*=\s*Math\.hypot\(\s*vx\s*,\s*vy\s*\)[\s\S]*body\.setVelocity\(\s*\(\s*vx\s*\/\s*len\s*\)\s*\*\s*speed\s*,\s*\(\s*vy\s*\/\s*len\s*\)\s*\*\s*speed\s*\)/);
```

You should track the current facing direction in `this.facing`.

```js
assert.match(code, /this\.facing\s*=\s*\{\s*x:\s*vx\s*,\s*y:\s*vy\s*\}/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
    walls.setCollisionByProperty({ collides: true });
    const spawn = map.findObject("objects", (o) => o.type === "spawn");
    this.player = this.add.rectangle(
      spawn.x + 8,
      spawn.y + 8,
      12,
      12,
      0x66ddff
    );
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.player.setDepth(20);
    this.physics.add.collider(this.player, walls);
    this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.startFollow(this.player);
    this.cameras.main.setLerp(0.12, 0.12);
    this.cursors = this.input.keyboard.createCursorKeys();
  }

  update(time, delta) {
--fcc-editable-region--
--fcc-editable-region--
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
