---
id: 66faa50000000000000000d3
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--
The map is bigger than the viewport. Bound the physics world to the map, bound the camera to the map, follow the player with a small lerp for smoothness, and grab the cursor keys.

```js
this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
this.cameras.main.startFollow(this.player);
this.cameras.main.setLerp(0.12, 0.12);
this.cursors = this.input.keyboard.createCursorKeys();
```

# --hints--
You should set the physics world bounds to the map size.

```js
assert.match(code, /this\.physics\.world\.setBounds\(\s*0\s*,\s*0\s*,\s*map\.widthInPixels\s*,\s*map\.heightInPixels\s*\)/);
```

You should set the main camera bounds to the map size.

```js
assert.match(code, /this\.cameras\.main\.setBounds\(\s*0\s*,\s*0\s*,\s*map\.widthInPixels\s*,\s*map\.heightInPixels\s*\)/);
```

You should start the main camera following `this.player`.

```js
assert.match(code, /this\.cameras\.main\.startFollow\(\s*this\.player\s*\)/);
```

You should set the camera lerp to `(0.12, 0.12)`.

```js
assert.match(code, /this\.cameras\.main\.setLerp\(\s*0\.12\s*,\s*0\.12\s*\)/);
```

You should create cursor keys with `this.input.keyboard.createCursorKeys()`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
    walls.setCollisionByProperty({ collides: true });
    const spawn = map.findObject("objects", (o) => o.type === "spawn");
    this.player = this.add.rectangle(
      spawn.x + 8,
      spawn.y + 8,
      12,
      12,
      0x66ddff
    );
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.player.setDepth(20);
    this.physics.add.collider(this.player, walls);
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
