---
id: 66faa50000000000000000cf
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
A tilemap is data; layers are what render. The `floor` layer carries the ground tiles, the `walls` layer carries the solid wall tiles. Render order matters — walls must be drawn on top of floor. Use `setDepth` to lock that order so future objects (player, props) can sit cleanly between them.

```js
const floor = map.createLayer("floor", tileset, 0, 0);
const walls = map.createLayer("walls", tileset, 0, 0);
floor.setDepth(0);
walls.setDepth(10);
this.walls = walls;
```

# --hints--
You should call `map.createLayer("floor", tileset, 0, 0)` and store it in `floor`.

```js
assert.match(code, /const\s+floor\s*=\s*map\.createLayer\(\s*["\']floor["\']\s*,\s*tileset\s*,\s*0\s*,\s*0\s*\)/);
```

You should call `map.createLayer("walls", tileset, 0, 0)` and store it in `walls`.

```js
assert.match(code, /const\s+walls\s*=\s*map\.createLayer\(\s*["\']walls["\']\s*,\s*tileset\s*,\s*0\s*,\s*0\s*\)/);
```

You should call `floor.setDepth(0)`.

```js
assert.match(code, /floor\.setDepth\(\s*0\s*\)/);
```

You should call `walls.setDepth(10)`.

```js
assert.match(code, /walls\.setDepth\(\s*10\s*\)/);
```

You should expose the walls layer on the scene with `this.walls = walls`.

```js
assert.match(code, /this\.walls\s*=\s*walls/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
