---
id: 66faa50000000000000000f4
title: Step 42
challengeType: 0
dashedName: step-42
---

# --description--
Add a transient room-label HUD: when MainScene emits `"ui:room"`, show the room name centred at the top of the screen for 2 seconds.

After the victory block in `UIScene.create`:

```js
this.roomLabel = this.add
  .text(160, 36, "", {
    fontSize: "12px",
    color: "#ffffff",
    backgroundColor: "rgba(0,0,0,0.6)",
    padding: { x: 6, y: 2 }
  })
  .setOrigin(0.5)
  .setVisible(false);
main.events.on("ui:room", (name) => {
  this.roomLabel.setText(name);
  this.roomLabel.setVisible(true);
  this.time.delayedCall(2000, () => this.roomLabel.setVisible(false));
});
```

# --hints--
You should add a centred `roomLabel` text element at `(160, 36)`.

```js
assert.match(code, /this\.roomLabel\s*=\s*this\.add\s*\.text\(\s*160\s*,\s*36\s*,/);
```

You should listen for `"ui:room"` and set the text on the label.

```js
assert.match(code, /main\.events\.on\(\s*["\']ui:room["\']\s*,\s*\(\s*name\s*\)\s*=>\s*\{[\s\S]*?this\.roomLabel\.setText\(\s*name\s*\)/);
```

You should hide the label after 2000 ms.

```js
assert.match(code, /this\.time\.delayedCall\(\s*2000\s*,\s*\(\s*\)\s*=>\s*this\.roomLabel\.setVisible\(\s*false\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
    walls.setCollisionByProperty({ collides: true });
    const spawn = map.findObject("objects", (o) => o.type === "spawn");
    this.player = this.add.rectangle(
      spawn.x + 8,
      spawn.y + 8,
      12,
      12,
      0x66ddff
    );
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.player.setDepth(20);
    this.physics.add.collider(this.player, walls);
    this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.startFollow(this.player);
    this.cameras.main.setLerp(0.12, 0.12);
    this.cursors = this.input.keyboard.createCursorKeys();
    const npcObj = map.findObject("objects", (o) => o.type === "npc");
    this.npc = this.add.rectangle(
      npcObj.x + 8,
      npcObj.y + 8,
      12,
      14,
      0xffaa66
    );
    this.npc.setDepth(20);
    this.physics.add.existing(this.npc, true);
    this.physics.add.collider(this.player, this.npc);
    this.npcLines = npcObj.properties.find((p) => p.name === "lines").value.split("|");
    this.dialogueIndex = 0;
    this.eKey = this.input.keyboard.addKey("E");
    this.lastTalkAt = 0;
    this.flags = new Set();
    this.scene.launch("UIScene");
    this.inventory = new Set();
    const keyObj = map.findObject("objects", (o) => o.type === "key");
    this.keyItem = this.add.rectangle(
      keyObj.x + 8,
      keyObj.y + 8,
      8,
      8,
      0xffe066
    );
    this.physics.add.existing(this.keyItem, true);
    this.keyItem.setDepth(15);
    this.physics.add.overlap(this.player, this.keyItem, () => {
      if (!this.keyItem.active) return;
      this.keyItem.destroy();
      this.inventory.add("bronze_key");
      this.events.emit("ui:notice", "Bronze Key obtained");
    });
    const lockedDoor = map.findObject("objects", (o) => {
      if (o.type !== "door" || !o.properties) return false;
      return o.properties.some((p) => p.name === "locked" && p.value === true);
    });
    this.lockedDoor = lockedDoor;
    this.lockedDoorTx = lockedDoor.properties.find((p) => p.name === "tx").value;
    this.lockedDoorTy = lockedDoor.properties.find((p) => p.name === "ty").value;
    this.lastDoorAt = 0;
    this.score = 0;
    const chestObj = map.findObject("objects", (o) => o.type === "chest");
    this.chest = this.add
      .sprite(chestObj.x + 8, chestObj.y + 8, "tiles", 27)
      .setDepth(15);
    this.chestOpened = false;
    const saveObj = map.findObject("objects", (o) => o.type === "save");
    this.savePoint = this.add.rectangle(
      saveObj.x + 8,
      saveObj.y + 8,
      12,
      12,
      0xffff66
    );
    this.savePoint.setDepth(15);
    this.saveGame = () => {
      const save = {
        x: this.player.x,
        y: this.player.y,
        hp: this.hp,
        score: this.score,
        inventory: Array.from(this.inventory),
        flags: Array.from(this.flags),
        chestOpened: this.chestOpened
      };
      localStorage.setItem(
        "phaser-dungeon-explorer-save",
        JSON.stringify(save)
      );
    };
    this.loadGame = () => {
      const raw = localStorage.getItem("phaser-dungeon-explorer-save");
      if (!raw) return;
      try {
        const save = JSON.parse(raw);
        this.player.x = save.x;
        this.player.y = save.y;
        this.hp = save.hp;
        this.score = save.score;
        this.inventory = new Set(save.inventory);
        this.flags = new Set(save.flags);
        if (save.chestOpened) {
          this.chestOpened = true;
          this.chest.setFrame(28);
        }
      } catch (e) {
        return;
      }
    };
    this.loadGame();
    const slimeObj = map.findObject("objects", (o) => o.type === "slime");
    this.slime = this.add.rectangle(
      slimeObj.x,
      slimeObj.y,
      14,
      12,
      0x88ee66
    );
    this.physics.add.existing(this.slime);
    this.slime.body.setCollideWorldBounds(true);
    this.slime.setDepth(18);
    this.physics.add.collider(this.slime, walls);
    this.slimePoly = slimeObj.polyline.map((p) => ({
      x: slimeObj.x + p.x,
      y: slimeObj.y + p.y
    }));
    this.slimeTarget = 1;
    this.slimeHp = 2;
    this.hp = 3;
    this.maxHp = 3;
    this.iframesUntil = 0;
    this.physics.add.overlap(this.player, this.slime, () => {
      if (this.time.now < this.iframesUntil || !this.slime.active) return;
      this.iframesUntil = this.time.now + 800;
      this.hp = Math.max(0, this.hp - 1);
      this.events.emit("ui:hp", this.hp);
      this.cameras.main.shake(150, 0.005);
    });
    this.spaceKey = this.input.keyboard.addKey("SPACE");
    this.lastAttackAt = 0;
    this.swing = this.add
      .rectangle(0, 0, 16, 16, 0xffffff, 0.6)
      .setDepth(25)
      .setVisible(false);
    const goalObj = map.findObject("objects", (o) => o.type === "goal");
    this.goal = this.add
      .sprite(goalObj.x + 8, goalObj.y + 8, "tiles", 31)
      .setDepth(15);
    this.physics.add.existing(this.goal, true);
    this.physics.add.overlap(this.player, this.goal, () => {
      if (this.flags.has("victory")) return;
      this.flags.add("victory");
      this.events.emit("ui:victory");
      this.cameras.main.fade(800, 0, 0, 0);
    });
    map.filterObjects("objects", (o) => o.type === "torch").forEach(
      (t) => {
        const torch = this.add
          .sprite(t.x + 8, t.y + 8, "tiles", 25)
          .setDepth(15);
        this.tweens.add({
          targets: torch,
          alpha: { from: 1, to: 0.6 },
          duration: 500,
          yoyo: true,
          repeat: -1
        });
      }
    );
    this.currentRoom = "NW";
    this.visitedRooms = new Set(["NW"]);
    this.transitionLockedUntil = 0;
    this.detectRoom = (px, py) => {
      const tx = px / 16;
      const ty = py / 16;
      if (tx < 22 && ty < 17) return "NW";
      if (tx >= 22 && ty < 17) return "NE";
      if (tx < 22 && ty >= 17) return "SW";
      return "SE";
    };
  }

  update(time, delta) {
    this.facing = this.facing || { x: 0, y: 1 };
    const speed = 100;
    const body = this.player.body;
    body.setVelocity(0);
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown) vx = -1;
    else if (this.cursors.right.isDown) vx = 1;
    if (this.cursors.up.isDown) vy = -1;
    else if (this.cursors.down.isDown) vy = 1;
    if (vx !== 0 || vy !== 0) {
      const len = Math.hypot(vx, vy);
      body.setVelocity((vx / len) * speed, (vy / len) * speed);
      this.facing = { x: vx, y: vy };
    }
    const distToNpc = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.npc.x,
      this.npc.y
    );
    if (
      distToNpc < 28 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      time > this.lastTalkAt + 250
    ) {
      this.lastTalkAt = time;
      const line = this.npcLines[Math.min(this.dialogueIndex, this.npcLines.length - 1)];
      this.events.emit("ui:dialogue", line);
      this.dialogueIndex++;
      this.flags.add("talked-to-npc");
    }
    const distToDoor = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.lockedDoor.x + 8,
      this.lockedDoor.y + 8
    );
    if (
      distToDoor < 24 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      time > this.lastDoorAt + 500
    ) {
      this.lastDoorAt = time;
      if (this.inventory.has("bronze_key")) {
        this.walls.removeTileAt(this.lockedDoorTx, this.lockedDoorTy);
        this.events.emit("ui:notice", "Bronze door unlocks");
        this.flags.add("door-unlocked");
      } else {
        this.events.emit("ui:notice", "It is locked. You need a bronze key.");
      }
    }
    const distToChest = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.chest.x,
      this.chest.y
    );
    if (
      distToChest < 24 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      !this.chestOpened
    ) {
      this.chestOpened = true;
      this.chest.setFrame(28);
      this.score += 50;
      this.events.emit("ui:score", this.score);
      this.events.emit("ui:notice", "Chest opens. +50 gold");
    }
    const distToSave = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.savePoint.x,
      this.savePoint.y
    );
    if (
      distToSave < 24 &&
      Phaser.Input.Keyboard.JustDown(this.eKey)
    ) {
      this.saveGame();
      this.events.emit("ui:notice", "Game saved");
    }
    if (this.slime.active && this.slimeHp > 0) {
      const target = this.slimePoly[this.slimeTarget];
      const dx = target.x - this.slime.x;
      const dy = target.y - this.slime.y;
      const d = Math.hypot(dx, dy);
      if (d < 4) {
        this.slimeTarget = (this.slimeTarget + 1) % this.slimePoly.length;
      } else {
        const sp = 30;
        this.slime.body.setVelocity((dx / d) * sp, (dy / d) * sp);
      }
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.spaceKey) &&
      time > this.lastAttackAt + 350
    ) {
      this.lastAttackAt = time;
      const ax = this.player.x + this.facing.x * 14;
      const ay = this.player.y + this.facing.y * 14;
      this.swing.setPosition(ax, ay).setVisible(true).setAlpha(0.6);
      this.tweens.add({
        targets: this.swing,
        alpha: { from: 0.6, to: 0 },
        duration: 200,
        onComplete: () => this.swing.setVisible(false)
      });
      if (
        this.slime.active &&
        Phaser.Math.Distance.Between(ax, ay, this.slime.x, this.slime.y) < 14
      ) {
        this.slimeHp--;
        this.slime.body.setVelocity(
          this.facing.x * 200,
          this.facing.y * 200
        );
        this.cameras.main.shake(80, 0.003);
        if (this.slimeHp <= 0) {
          this.slime.destroy();
          this.score += 30;
          this.events.emit("ui:score", this.score);
          this.events.emit("ui:notice", "Slime defeated. +30 gold");
        }
      }
    }
    const room = this.detectRoom(this.player.x, this.player.y);
    if (
      room &&
      room !== this.currentRoom &&
      time > this.transitionLockedUntil
    ) {
      this.currentRoom = room;
      this.transitionLockedUntil = time + 600;
      this.cameras.main.fadeOut(150, 0, 0, 0);
      this.cameras.main.once("camerafadeoutcomplete", () =>
        this.cameras.main.fadeIn(150)
      );
      this.visitedRooms.add(room);
      this.events.emit("ui:room", room);
    }
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
    const main = this.scene.get("MainScene");
    this.box = this.add
      .rectangle(160, 210, 304, 44, 0x000000, 0.7)
      .setVisible(false);
    this.box.setStrokeStyle(1, 0xffffff, 0.4);
    this.text = this.add.text(16, 192, "", {
      fontSize: "10px",
      color: "#ffffff",
      wordWrap: { width: 288 }
    });
    this.text.setVisible(false);
    main.events.on("ui:dialogue", (line) => {
      this.fullLine = line;
      this.shown = 0;
      this.typing = true;
      this.text.setText("").setVisible(true);
      this.box.setVisible(true);
    });
    this.fullLine = "";
    this.shown = 0;
    this.typing = false;
    this.notice = this.add.text(160, 24, "", {
      fontSize: "10px",
      color: "#ffe066",
      backgroundColor: "rgba(0,0,0,0.7)",
      padding: { x: 4, y: 2 }
    });
    this.notice.setOrigin(0.5, 0.5).setVisible(false);
    main.events.on("ui:notice", (msg) => {
      this.notice.setText(msg);
      this.notice.setVisible(true);
      this.time.delayedCall(2000, () => this.notice.setVisible(false));
    });
    this.hpBar = this.add.graphics();
    this.hpBar.setScrollFactor(0).setDepth(60);
    this.drawHp = (hp) => {
      this.hpBar.clear();
      this.hpBar.fillStyle(0x222222, 1).fillRect(8, 8, 60, 8);
      this.hpBar.fillStyle(0xff5566, 1).fillRect(8, 8, (hp / 3) * 60, 8);
      this.hpBar.lineStyle(1, 0xffffff, 1).strokeRect(8, 8, 60, 8);
    };
    this.drawHp(3);
    main.events.on("ui:hp", (hp) => this.drawHp(hp));
    main.events.on("ui:victory", () => {
      this.add
        .rectangle(160, 120, 320, 240, 0x000000, 0.7)
        .setDepth(80);
      this.add
        .text(160, 120, "Victory!", {
          fontSize: "32px",
          color: "#ffe066"
        })
        .setOrigin(0.5)
        .setDepth(81);
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    if (this.typing) {
      this.shown += delta * 0.06;
      const slice = this.fullLine.slice(0, Math.floor(this.shown));
      this.text.setText(slice);
      if (slice.length >= this.fullLine.length) this.typing = false;
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
