---
id: 66faa50000000000000000d1
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--
The map's object layer carries a single `spawn` object. Read it with `map.findObject` and place a 12×12 cyan rectangle there as the player. Object coordinates are in pixels and refer to the top-left corner; add `8` to centre on a 16×16 tile.

```js
const spawn = map.findObject("objects", (o) => o.type === "spawn");
this.player = this.add.rectangle(
  spawn.x + 8,
  spawn.y + 8,
  12,
  12,
  0x66ddff
);
```

# --hints--
You should read the spawn object from the `"objects"` layer.

```js
assert.match(code, /const\s+spawn\s*=\s*map\.findObject\(\s*["\']objects["\']\s*,\s*\(\s*\w+\s*\)\s*=>\s*\w+\.type\s*===\s*["\']spawn["\']\s*\)/);
```

You should create the player as a 12×12 rectangle with colour `0x66ddff`, centred on the spawn tile.

```js
assert.match(code, /this\.player\s*=\s*this\.add\.rectangle\(\s*spawn\.x\s*\+\s*8\s*,\s*spawn\.y\s*\+\s*8\s*,\s*12\s*,\s*12\s*,\s*0x66ddff\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
    walls.setCollisionByProperty({ collides: true });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
