---
id: 66faa50000000000000000e5
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--
Restore on scene start. After the save/load closures in `MainScene.create`, call `this.loadGame()`. If a save exists, the player teleports to where they were; otherwise it's a no-op.

```js
this.loadGame();
```

# --hints--
You should call `this.loadGame()` from `MainScene.create`.

```js
assert.match(code, /this\.loadGame\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
    walls.setCollisionByProperty({ collides: true });
    const spawn = map.findObject("objects", (o) => o.type === "spawn");
    this.player = this.add.rectangle(
      spawn.x + 8,
      spawn.y + 8,
      12,
      12,
      0x66ddff
    );
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.player.setDepth(20);
    this.physics.add.collider(this.player, walls);
    this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.startFollow(this.player);
    this.cameras.main.setLerp(0.12, 0.12);
    this.cursors = this.input.keyboard.createCursorKeys();
    const npcObj = map.findObject("objects", (o) => o.type === "npc");
    this.npc = this.add.rectangle(
      npcObj.x + 8,
      npcObj.y + 8,
      12,
      14,
      0xffaa66
    );
    this.npc.setDepth(20);
    this.physics.add.existing(this.npc, true);
    this.physics.add.collider(this.player, this.npc);
    this.npcLines = npcObj.properties.find((p) => p.name === "lines").value.split("|");
    this.dialogueIndex = 0;
    this.eKey = this.input.keyboard.addKey("E");
    this.lastTalkAt = 0;
    this.flags = new Set();
    this.scene.launch("UIScene");
    this.inventory = new Set();
    const keyObj = map.findObject("objects", (o) => o.type === "key");
    this.keyItem = this.add.rectangle(
      keyObj.x + 8,
      keyObj.y + 8,
      8,
      8,
      0xffe066
    );
    this.physics.add.existing(this.keyItem, true);
    this.keyItem.setDepth(15);
    this.physics.add.overlap(this.player, this.keyItem, () => {
      if (!this.keyItem.active) return;
      this.keyItem.destroy();
      this.inventory.add("bronze_key");
      this.events.emit("ui:notice", "Bronze Key obtained");
    });
    const lockedDoor = map.findObject("objects", (o) => {
      if (o.type !== "door" || !o.properties) return false;
      return o.properties.some((p) => p.name === "locked" && p.value === true);
    });
    this.lockedDoor = lockedDoor;
    this.lockedDoorTx = lockedDoor.properties.find((p) => p.name === "tx").value;
    this.lockedDoorTy = lockedDoor.properties.find((p) => p.name === "ty").value;
    this.lastDoorAt = 0;
    this.score = 0;
    const chestObj = map.findObject("objects", (o) => o.type === "chest");
    this.chest = this.add
      .sprite(chestObj.x + 8, chestObj.y + 8, "tiles", 27)
      .setDepth(15);
    this.chestOpened = false;
    const saveObj = map.findObject("objects", (o) => o.type === "save");
    this.savePoint = this.add.rectangle(
      saveObj.x + 8,
      saveObj.y + 8,
      12,
      12,
      0xffff66
    );
    this.savePoint.setDepth(15);
    this.saveGame = () => {
      const save = {
        x: this.player.x,
        y: this.player.y,
        hp: this.hp,
        score: this.score,
        inventory: Array.from(this.inventory),
        flags: Array.from(this.flags),
        chestOpened: this.chestOpened
      };
      localStorage.setItem(
        "phaser-dungeon-explorer-save",
        JSON.stringify(save)
      );
    };
    this.loadGame = () => {
      const raw = localStorage.getItem("phaser-dungeon-explorer-save");
      if (!raw) return;
      try {
        const save = JSON.parse(raw);
        this.player.x = save.x;
        this.player.y = save.y;
        this.hp = save.hp;
        this.score = save.score;
        this.inventory = new Set(save.inventory);
        this.flags = new Set(save.flags);
        if (save.chestOpened) {
          this.chestOpened = true;
          this.chest.setFrame(28);
        }
      } catch (e) {
        return;
      }
    };
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    this.facing = this.facing || { x: 0, y: 1 };
    const speed = 100;
    const body = this.player.body;
    body.setVelocity(0);
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown) vx = -1;
    else if (this.cursors.right.isDown) vx = 1;
    if (this.cursors.up.isDown) vy = -1;
    else if (this.cursors.down.isDown) vy = 1;
    if (vx !== 0 || vy !== 0) {
      const len = Math.hypot(vx, vy);
      body.setVelocity((vx / len) * speed, (vy / len) * speed);
      this.facing = { x: vx, y: vy };
    }
    const distToNpc = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.npc.x,
      this.npc.y
    );
    if (
      distToNpc < 28 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      time > this.lastTalkAt + 250
    ) {
      this.lastTalkAt = time;
      const line = this.npcLines[Math.min(this.dialogueIndex, this.npcLines.length - 1)];
      this.events.emit("ui:dialogue", line);
      this.dialogueIndex++;
      this.flags.add("talked-to-npc");
    }
    const distToDoor = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.lockedDoor.x + 8,
      this.lockedDoor.y + 8
    );
    if (
      distToDoor < 24 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      time > this.lastDoorAt + 500
    ) {
      this.lastDoorAt = time;
      if (this.inventory.has("bronze_key")) {
        this.walls.removeTileAt(this.lockedDoorTx, this.lockedDoorTy);
        this.events.emit("ui:notice", "Bronze door unlocks");
        this.flags.add("door-unlocked");
      } else {
        this.events.emit("ui:notice", "It is locked. You need a bronze key.");
      }
    }
    const distToChest = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.chest.x,
      this.chest.y
    );
    if (
      distToChest < 24 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      !this.chestOpened
    ) {
      this.chestOpened = true;
      this.chest.setFrame(28);
      this.score += 50;
      this.events.emit("ui:score", this.score);
      this.events.emit("ui:notice", "Chest opens. +50 gold");
    }
    const distToSave = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.savePoint.x,
      this.savePoint.y
    );
    if (
      distToSave < 24 &&
      Phaser.Input.Keyboard.JustDown(this.eKey)
    ) {
      this.saveGame();
      this.events.emit("ui:notice", "Game saved");
    }
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
    const main = this.scene.get("MainScene");
    this.box = this.add
      .rectangle(160, 210, 304, 44, 0x000000, 0.7)
      .setVisible(false);
    this.box.setStrokeStyle(1, 0xffffff, 0.4);
    this.text = this.add.text(16, 192, "", {
      fontSize: "10px",
      color: "#ffffff",
      wordWrap: { width: 288 }
    });
    this.text.setVisible(false);
    main.events.on("ui:dialogue", (line) => {
      this.fullLine = line;
      this.shown = 0;
      this.typing = true;
      this.text.setText("").setVisible(true);
      this.box.setVisible(true);
    });
    this.fullLine = "";
    this.shown = 0;
    this.typing = false;
    this.notice = this.add.text(160, 24, "", {
      fontSize: "10px",
      color: "#ffe066",
      backgroundColor: "rgba(0,0,0,0.7)",
      padding: { x: 4, y: 2 }
    });
    this.notice.setOrigin(0.5, 0.5).setVisible(false);
    main.events.on("ui:notice", (msg) => {
      this.notice.setText(msg);
      this.notice.setVisible(true);
      this.time.delayedCall(2000, () => this.notice.setVisible(false));
    });
  }

  update(time, delta) {
    if (this.typing) {
      this.shown += delta * 0.06;
      const slice = this.fullLine.slice(0, Math.floor(this.shown));
      this.text.setText(slice);
      if (slice.length >= this.fullLine.length) this.typing = false;
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
