---
id: 66faa50000000000000000dd
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--
Time to populate the dungeon with collectibles. Initialize an inventory `Set`, read the `key` object from the map, render it as a small yellow square with a static body, and on overlap with the player destroy it, add `"bronze_key"` to inventory, and emit a notice.

After `LAUNCH_UI` (the `this.scene.launch("UIScene")` line) in `MainScene.create`:

```js
this.inventory = new Set();
const keyObj = map.findObject("objects", (o) => o.type === "key");
this.keyItem = this.add.rectangle(
  keyObj.x + 8,
  keyObj.y + 8,
  8,
  8,
  0xffe066
);
this.physics.add.existing(this.keyItem, true);
this.keyItem.setDepth(15);
this.physics.add.overlap(this.player, this.keyItem, () => {
  if (!this.keyItem.active) return;
  this.keyItem.destroy();
  this.inventory.add("bronze_key");
  this.events.emit("ui:notice", "Bronze Key obtained");
});
```

# --hints--
You should initialize `this.inventory` as a new Set.

```js
assert.match(code, /this\.inventory\s*=\s*new\s+Set\(\s*\)/);
```

You should read the key object from the `"objects"` layer.

```js
assert.match(code, /const\s+keyObj\s*=\s*map\.findObject\(\s*["\']objects["\']\s*,\s*\(\s*\w+\s*\)\s*=>\s*\w+\.type\s*===\s*["\']key["\']\s*\)/);
```

You should render the key item as an 8×8 yellow rectangle.

```js
assert.match(code, /this\.keyItem\s*=\s*this\.add\.rectangle\(\s*keyObj\.x\s*\+\s*8\s*,\s*keyObj\.y\s*\+\s*8\s*,\s*8\s*,\s*8\s*,\s*0xffe066\s*\)/);
```

You should add an overlap between the player and the key that destroys the key and adds `"bronze_key"` to inventory.

```js
assert.match(code, /physics\.add\.overlap\(\s*this\.player\s*,\s*this\.keyItem[\s\S]*this\.keyItem\.destroy\(\s*\)[\s\S]*this\.inventory\.add\(\s*["\']bronze_key["\']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
    walls.setCollisionByProperty({ collides: true });
    const spawn = map.findObject("objects", (o) => o.type === "spawn");
    this.player = this.add.rectangle(
      spawn.x + 8,
      spawn.y + 8,
      12,
      12,
      0x66ddff
    );
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.player.setDepth(20);
    this.physics.add.collider(this.player, walls);
    this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.startFollow(this.player);
    this.cameras.main.setLerp(0.12, 0.12);
    this.cursors = this.input.keyboard.createCursorKeys();
    const npcObj = map.findObject("objects", (o) => o.type === "npc");
    this.npc = this.add.rectangle(
      npcObj.x + 8,
      npcObj.y + 8,
      12,
      14,
      0xffaa66
    );
    this.npc.setDepth(20);
    this.physics.add.existing(this.npc, true);
    this.physics.add.collider(this.player, this.npc);
    this.npcLines = npcObj.properties.find((p) => p.name === "lines").value.split("|");
    this.dialogueIndex = 0;
    this.eKey = this.input.keyboard.addKey("E");
    this.lastTalkAt = 0;
    this.flags = new Set();
    this.scene.launch("UIScene");
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    this.facing = this.facing || { x: 0, y: 1 };
    const speed = 100;
    const body = this.player.body;
    body.setVelocity(0);
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown) vx = -1;
    else if (this.cursors.right.isDown) vx = 1;
    if (this.cursors.up.isDown) vy = -1;
    else if (this.cursors.down.isDown) vy = 1;
    if (vx !== 0 || vy !== 0) {
      const len = Math.hypot(vx, vy);
      body.setVelocity((vx / len) * speed, (vy / len) * speed);
      this.facing = { x: vx, y: vy };
    }
    const distToNpc = Phaser.Math.Distance.Between(
      this.player.x,
      this.player.y,
      this.npc.x,
      this.npc.y
    );
    if (
      distToNpc < 28 &&
      Phaser.Input.Keyboard.JustDown(this.eKey) &&
      time > this.lastTalkAt + 250
    ) {
      this.lastTalkAt = time;
      const line = this.npcLines[Math.min(this.dialogueIndex, this.npcLines.length - 1)];
      this.events.emit("ui:dialogue", line);
      this.dialogueIndex++;
      this.flags.add("talked-to-npc");
    }
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
    const main = this.scene.get("MainScene");
    this.box = this.add
      .rectangle(160, 210, 304, 44, 0x000000, 0.7)
      .setVisible(false);
    this.box.setStrokeStyle(1, 0xffffff, 0.4);
    this.text = this.add.text(16, 192, "", {
      fontSize: "10px",
      color: "#ffffff",
      wordWrap: { width: 288 }
    });
    this.text.setVisible(false);
    main.events.on("ui:dialogue", (line) => {
      this.fullLine = line;
      this.shown = 0;
      this.typing = true;
      this.text.setText("").setVisible(true);
      this.box.setVisible(true);
    });
    this.fullLine = "";
    this.shown = 0;
    this.typing = false;
    this.notice = this.add.text(160, 24, "", {
      fontSize: "10px",
      color: "#ffe066",
      backgroundColor: "rgba(0,0,0,0.7)",
      padding: { x: 4, y: 2 }
    });
    this.notice.setOrigin(0.5, 0.5).setVisible(false);
    main.events.on("ui:notice", (msg) => {
      this.notice.setText(msg);
      this.notice.setVisible(true);
      this.time.delayedCall(2000, () => this.notice.setVisible(false));
    });
  }

  update(time, delta) {
    if (this.typing) {
      this.shown += delta * 0.06;
      const slice = this.fullLine.slice(0, Math.floor(this.shown));
      this.text.setText(slice);
      if (slice.length >= this.fullLine.length) this.typing = false;
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
