---
id: 66faa50000000000000000cd
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Now load the Tiled JSON map. `tilemapTiledJSON` parses Tiled's export format — layers, tilesets, objects — and caches it under the key you provide. Inside `preload`, after the tileset load, add:

```js
this.load.tilemapTiledJSON(
  "dungeon",
  "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
);
```

Keep the keys distinct: `"tiles"` for the image, `"dungeon"` for the map.

# --hints--
You should call `this.load.tilemapTiledJSON("dungeon", "/curriculum-assets/phaser/tilemaps/dungeon-01.json")` in `preload`.

```js
assert.match(code, /this\.load\.tilemapTiledJSON\(\s*["\']dungeon["\']\s*,\s*["\']\/curriculum-assets\/phaser\/tilemaps\/dungeon-01\.json["\']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
