---
id: 66faa50000000000000000d0
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Tiles in the `walls` layer carry a `collides: true` custom property. Tell arcade physics to honour it — every tile with `collides: true` becomes a solid body for collision checks.

```js
walls.setCollisionByProperty({ collides: true });
```

This is a one-shot setup call: Phaser walks the layer once and caches the collision flags.

# --hints--
You should call `walls.setCollisionByProperty({ collides: true })`.

```js
assert.match(code, /walls\.setCollisionByProperty\(\s*\{\s*collides:\s*true\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Top-Down Dungeon Explorer</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "dungeon",
      "/curriculum-assets/phaser/tilemaps/dungeon-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "dungeon" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    this.map = map;
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.walls = walls;
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: false });
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 320,
  height: 240,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
