---
id: 69eefea5a952ae85871df0dc
title: Implement Data Sanitizer Sweep Process
challengeType: 26
dashedName: lab-data-sanitizer-sweep
---

# --description--

In this lab, you will clean semi-structured log files by stripping blacklisted characters and tracking repair statistics.

Here is an example of an array of log entries:

```log
  "DEBUG: Server started on port 3000",
  "INFO: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%",
  "/INFO: Cache cleared successfully",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
```

For example, if blacklisted characters are `[e, o, r]` and the replacement character is `#`, sanitizing the first line changes:

`DEBUG: Server started on port 3000` to `DEBUG: S##v## sta#t#d #n p##t 3000`

Encountering a fatal marker should stop the sanitizing process immediately. Looking at the example log entries using `$` as the fatal marker. Please see second-to-last log entry in the example log entries.

If the process fails, then the output summary should return a reason as below:

`Sanitizing stopped Failure Reason: Fatal character occurred.`

An array of log entries can also include informational entries. This is indicated by a specific leading character.

Assuming any log entry that starts with `/` is considered an informational log entry then this entry should be skipped and count as a successfully sanitized entry. For example, if there are 4 log entries in total and each entry starts with `/` then the output summary should return as follows.

`Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.`

**The blacklisted characters will be case sensitive. For example if blacklisted character is `e` then `Debug` will become `D#bug` whilst `DEBUG` will still be `DEBUG`.**

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Story:** 

1. Create a function called `sanitizeLog(entry, blacklist)`
   - `entry` is a single string of characters
   - `blacklist` is an array of single letter characters and each character is case sensitive.
   - The function should take the entry log and replace any character that appears in `blacklist` with `#`
2. `sanitizeLog` should return the actual modified log entry.
3. Create a function called `cleanLogs(logs, blacklist)`
   - `logs` is an array of strings, essentially an array of what was input into the `sanitizeLog` function as `entry` parameter.
   - `blacklist` is the same array of single-letter, case-sensitive characters used in `sanitizeLog`.
   - `cleanLogs` should iteratively call the `sanitizeLog` function as part of the solution.
4. `cleanLogs` should iteratively call `sanitizeLog` function to process an array of log entries.
   - The function should return a customized string containing either an error message or a string with the total number of log entries and the number of sanitized log entries.
   - The fatal marker will be hardcoded as `$`
   - Informational indicator will be hardcoded as `/`.
      - All log entries that start with this character are to be skipped.
   - If `cleanLogs` does not encounter any fatal marker, then the function should return the below line. A full stop is required as part of the return, and the sanitized entry count should equal the total entry count.
      - `Total Entry: X Sanitized Entry Count: Y finished sanitize logs process.`
   - If `cleanLogs` encounters a fatal marker, then the function should return the line below and exit before finishing the entire log. A full stop is required as part of the return.
      - `Sanitizing stopped Failure Reason: Fatal character occurred.`
5. If `cleanLogs` encounters `/` at the start of the log entry, then the function should continue processing the remaining log entries, and this entry counts as a successfully processed log entry.

# --before-each--

```js
let logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

let logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

let logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --hints--

You should have a function called `sanitizeLog`.

```js
assert.isFunction(sanitizeLog)
```

You should have a function called `cleanLogs`.

```js
assert.isFunction(cleanLogs)
```

`sanitizeLog("DEBUG: Server started on port 3000",["a", "A", "c", "T"])` should return string `DEBUG: Server st#rted on port 3000`

```js
assert.strictEqual(sanitizeLog("DEBUG: Server started on port 3000",["a", "A", "c", "T"]),"DEBUG: Server st#rted on port 3000");
```

`sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "c", "T"])` should return string `W#RN: Disk us#ge #t 92%, threshold=90%`

```js
assert.strictEqual(sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "c", "T"]),'W#RN: Disk us#ge #t 92%, threshold=90%');
```

`cleanLogs(logs_1,["a", "A", "c", "T"])` should process all log entries without error and return string `Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process.`

```js
assert.strictEqual(cleanLogs(logs_1,["a", "A", "c", "T"]),'Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process.');
```

`cleanLogs(logs_2,["a", "A", "c", "T"])` should result in error message due to fatal character and return string `Sanitizing stopped Failure Reason: Fatal character occurred.`.

```js
assert.strictEqual(cleanLogs(logs_2,["a", "A", "c", "T"]),'Sanitizing stopped Failure Reason: Fatal character occurred.');
```

`cleanLogs(logs_3,["a", "A", "c", "T"])` should process all log entries without error and return string `Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.`.

```js
assert.strictEqual(cleanLogs(logs_3,["a", "A", "c", "T"]),'Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.');
```

# --seed--

## --seed-contents--

```js
let logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

let logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

let logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --solutions--

```js
function sanitizeLog(entry, blacklist) {
  var sanitizedEntry = [];

  for (let i = 0; i < entry.length; i++) {
    if (blacklist.includes(entry[i])) {
      sanitizedEntry.push("#");
    } else {
      sanitizedEntry.push(entry[i]);
    }
  }

  return sanitizedEntry.join("");
};

function cleanLogs(logs, blacklist) {
  let accumulator = 0;
  let fatalError = "$";
  let informationalIndicator = "/";

  let sanitizedLogs = [];

  for (let i = 0; i < logs.length; i++) {
    if (logs[i].includes(fatalError)) {
      return("Sanitizing stopped Failure Reason: Fatal character occurred.");
    } else if (logs[i].startsWith(informationalIndicator)) {
      accumulator += 1;
      continue;
    } else {
      let tmpSanitized = sanitizeLog(logs[i], blacklist);
      sanitizedLogs.push(tmpSanitized);

      accumulator += 1;
    }
  }

  return("Total Entry: " +
        logs.length +
        " Sanitized Entry Count: " +
        accumulator +
        " finished sanitize logs process.");
};
```
