---
id: 69eefea5a952ae85871df0dc
title: Implement Data Sanitizer Sweep Process
challengeType: 26
dashedName: lab-data-sanitizer-sweep
---

# --description--

In this lab, you will clean semi-structured log files by stripping disallowed tokens and tracking repair statistics.

Here are some examples log entries in a log file:

```log
  "DEBUG: Server started on port 3000",
  "INFO: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%",
  "/INFO: Cache cleared successfully",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
```

For example, if the disallowed tokens are `e, o, r` and the replacement character is `#`, sanitizing the first line changes:

`DEBUG: Server started on port 3000` to `DEBUG: S##v## sta#t#d #n p##t 3000`

Fatal markers are characters that should stop parsing immediately. Looking at the example, `$` is a fatal marker. If it appears (like in the second-to-last log line), the program should stop processing additional logs.

After processing, the program should print a repair summary such as:

`Total Entry: 6 Sanitized Entry Count: 4 finished sanitize logs process.`

If sanitization fails, the output summary should also include a reason, such as:

`Sanitizing stopped Failure Reason: Fatal character occurred.`

For this lab, if a log string starts with `/` then in `cleanLogs` function this string should be skipped and count as a successfully sanitized entry count. For example, if the log array have 4 logs strings in total and all strings start with `/` then this means the result will be 

`Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.`

**The disallowed characters will be case sensitive. So for example if disallowed character is `e` then `Debug` will become `D#bug` whilst `DEBUG` will still be `DEBUG`.**

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Story:** 

1. Create a function called `sanitizeLog(entry, blacklist)`
   - Entry is a single string of characters
   - Blacklist is an array of single letter characters and each character is case sensitive.
   - The function should take the entry log and modify it by replacing any character within the `blacklist` with `#`
   - The function should return the actual modified entry log.
2. `sanitizeLog` should return a string.
3. Create a function called `cleanLogs(logs, blacklist)`
   - `logs` is an array of strings, essentially an array of what was inputted into `sanitizeLog` function as `entry` parameter.
   - `blacklist` is the same array of single string characters from the `sanitizeLog` function.
4. `cleanLogs` should return a customised string returning either an error message or a string containing both the total number of log string and number of sanitized log string.
   - The fatal marker will be hardcoded as `$`
   - All logs lines to be skipped will start with `/`
   - If `cleanLogs` function did not encounter any fatal marker, then the funtion should return below line
      - `Total Entry: X Sanitized Entry Count: Y finished sanitize logs process.` (Full stop is needed as part of the return)
   - If `cleanLogs` did encounter a fatal marker, then the function should return below line and exit before finishing the entire log.
      - `Sanitizing stopped Failure Reason: Fatal character occurred.` (Full stop is needed as part of the return)
   - If `cleanLogs` did encounter `/` at the start of the log entry, then the function should continue processing the remaining log files and this line counts as a successfully process log string essentially skipping the log string.

# --before-each--

```js
let logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

let logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

let logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --hints--

You should have a function called `sanitizeLog`.

```js
assert.isFunction(sanitizeLog)
```

You should have a function called `cleanLogs`.

```js
assert.isFunction(cleanLogs)
```

`sanitizeLog("DEBUG: Server started on port 3000",["a", "A", "c", "T"])` should return string `DEBUG: Server st#rted on port 3000`

```js
assert.strictEqual(sanitizeLog("DEBUG: Server started on port 3000",["a", "A", "c", "T"]),"DEBUG: Server st#rted on port 3000");
```

`sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "c", "T"])` should return string `W#RN: Disk us#ge #t 92%, threshold=90%`

```js
assert.strictEqual(sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "c", "T"]),'W#RN: Disk us#ge #t 92%, threshold=90%');
```

`cleanLogs(logs_1,["a", "A", "c", "T"])` should process all log strings without error and return string `Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process.`

```js
assert.strictEqual(cleanLogs(logs_1,["a", "A", "c", "T"]),'Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process.');
```

`cleanLogs(logs_2,["a", "A", "c", "T"])` should result in error message due to fatal character and return string `Sanitizing stopped Failure Reason: Fatal character occurred.`.

```js
assert.strictEqual(cleanLogs(logs_2,["a", "A", "c", "T"]),'Sanitizing stopped Failure Reason: Fatal character occurred.');
```

`cleanLogs(logs_3,["a", "A", "c", "T"])` should process all log strings without error and return string `Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.`.

```js
assert.strictEqual(cleanLogs(logs_3,["a", "A", "c", "T"]),'Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.');
```

# --seed--

## --seed-contents--

```js
var logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

var logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

var logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --solutions--

```js
function sanitizeLog(entry, blacklist) {
  var sanitizedEntry = [];

  for (let i = 0; i < entry.length; i++) {
    if (blacklist.includes(entry[i])) {
      sanitizedEntry.push("#");
    } else {
      sanitizedEntry.push(entry[i]);
    }
  }

  return sanitizedEntry.join("");
};

function cleanLogs(logs, blacklist) {
  let accumulator = 0;
  let fatalError = "$";
  let informationalIndicator = "/";

  let sanitizedLogs = [];

  for (let i = 0; i < logs.length; i++) {
    if (logs[i].includes(fatalError)) {
      return("Sanitizing stopped Failure Reason: Fatal character occurred.");
    } else if (logs[i].includes(informationalIndicator)) {
      accumulator += 1;
      continue;
    } else {
      let tmpSanitized = sanitizeLog(logs[i], blacklist);
      sanitizedLogs.push(tmpSanitized);

      accumulator += 1;
    }
  }

  return("Total Entry: " +
        logs.length +
        " Sanitized Entry Count: " +
        accumulator +
        " finished sanitize logs process.");
};
```
