---
id: 69eefea5a952ae85871df0dc
title: Implement Data Sanitizer Sweep Process
challengeType: 26
dashedName: lab-data-sanitizer-sweep
---

# --description--

In this lab, you will clean semi-structured log files by replacing blacklisted characters with `#` and tracking how many entries were processed. 

If blacklisted characters are `[e, E, o, r]` then `D#BUG: Server started on port 3000` becomes `D#BUG: S##v## sta#t#d #n p##t 3000`. Blacklisted characters are case sensitive - for example if `e` is blacklisted, `Debug` becomes `D#ebug`, but `DEBUG` is unaffected.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a function called `sanitizeLog(entry, blacklist)`, where `entry` is a string representing the entry log, and `blacklist` is an array of single character strings, each representing a disallowed character.
1. The function should return a modified `entry` string with each blacklisted character replaced with `#`.
1. You should create a function called `cleanLogs(logs, blacklist)`, where `logs` is an array of strings, each representing a log entry before data sanitization, and `blacklist` is the same array of disallowed characters used in `sanitizeLog`.
1. `cleanLogs` should iteratively call the `sanitizeLog` function to process the given array of log entries.
1. The function should return a string containing either an error message or the total number of log entries and the number of sanitized log entries.
1. When a log entry from the `logs` array contains the fatal marker `$`, `cleanLogs` should return `Sanitizing stopped Failure Reason: Fatal character occurred.`
1. When a log entry from the `logs` array starts with the `/` character, that entry should be considered an informational entry, and the `informationalLine` variable should increment by 1.
1. In `cleanLogs` function the sum of informational entry and successfully sanitized entry should equal to the total number of entries if sanitize process completed successfully and return `Total Entry: X Sanitized Entry Count: Y finished sanitize logs process.`, where `X` and `Y` should be replaced by the sum of informational entry plus sanitized entry count and total entry count, respectively.

**Below are the log entry array variables for tests**

```js
let logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

let logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

let logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --before-each--

```js
let logs_1 = [
  "INFO: Server started on port 3000",
  "DEBUG: user=admin token=abc123 connected",
  "WARN: High memory usage detected at 87%"
];

let logs_2 = [
  "FATAL: System overload detected, shutting down",
  "DEBUG: token=def456$ refreshed for session",
  "ERROR: Null pointer exception in module auth",
];

let logs_3 = [
  "ERROR: Failed to connect to db, retry=3",
  "DEBUG: api_key=XYZ789 used in request",
  "/INFO: Cache cleared successfully",
  "WARN: Disk usage at 92%, threshold=90%"
];
```

# --hints--

You should have a function called `sanitizeLog`.

```js
assert.isFunction(sanitizeLog)
```

You should have a function called `cleanLogs`.

```js
assert.isFunction(cleanLogs)
```

`sanitizeLog("DEBUG: Server started on port 3000",["a", "S", "r", "t"])` should return string `DEBUG: #e#ve# s####ed on po## 3000`

```js
assert.strictEqual(sanitizeLog("DEBUG: Server started on port 3000",["a", "S", "r", "t"]),"DEBUG: #e#ve# s####ed on po## 3000");

// Extra test to guard against hardcode
let hardCodeTestCase = "DEBUG: Running new Script to confirm result";
let newBlackList = ["R","c","e","n"];

assert.strictEqual(sanitizeLog(hardCodeTestCase,newBlackList),"DEBUG: #u##i#g ##w S#ript to #o#firm r#sult");
```

`sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "d", "e"])` should return string `W#RN: Disk us#g# #t 92%, thr#shol#=90%`

```js
assert.strictEqual(sanitizeLog("WARN: Disk usage at 92%, threshold=90%",["a", "A", "d", "e"]),'W#RN: Disk us#g# #t 92%, thr#shol#=90%');

// Extra test to guard against hardcode to check that this string is not sanitized.
let hardCodeTestCaseNoSanitize = "TESTING: No Blacklisted Characters";
let newBlackListNoSanitize = ["A","b","z","Y"];

assert.strictEqual(sanitizeLog(hardCodeTestCaseNoSanitize,newBlackListNoSanitize),"TESTING: No Blacklisted Characters");
```

`cleanLogs(logs_1,["a", "A", "c", "T"])` should process all log entries without error and return string `Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process.`

```js
assert.strictEqual(cleanLogs(logs_1,["a", "A", "c", "T"]),'Total Entry: 3 Sanitized Entry Count: 3 finished sanitize logs process.');

// Extra test to guard against hardcode
let test3logs = [
  "BEGIN: Running Test cases",
  "TEST: Doing first test case",
  "TEST: Doing second test case",
  "STOP: Completed all test cases"
];

let test3Blacklist = ["s","u","a","s"];

assert.strictEqual(cleanLogs(test3logs,test3Blacklist),'Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.');
```

`cleanLogs(logs_2,["a", "A", "c", "T"])` should result in error message due to fatal character and return string `Sanitizing stopped Failure Reason: Fatal character occurred.`.

```js
assert.strictEqual(cleanLogs(logs_2,["a", "A", "c", "T"]),'Sanitizing stopped Failure Reason: Fatal character occurred.');

// Extra test to guard against hardcode
let test4logs = [
  "/INFO: Info Log1",
  "/INFO: Info Log2",
  "/INFO: Info Log3",
  "/INFO: Info Log4",
  "/INFO: Info Log5",
  "/INFO: Info Log6"
];

let test4Blacklist = ["I","o"];

assert.strictEqual(cleanLogs(test4logs,test4Blacklist),'Total Entry: 6 Sanitized Entry Count: 6 finished sanitize logs process.');
```

`cleanLogs(logs_3,["a", "A", "c", "T"])` should process all log entries without error and return string `Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.`.

```js
assert.strictEqual(cleanLogs(logs_3,["a", "A", "c", "T"]),'Total Entry: 4 Sanitized Entry Count: 4 finished sanitize logs process.');

// Extra test to guard against hardcode
let test5logs = [
  "ERROR: Error ch$aracter",
  "LOG: Fatal character found"
];

let test5Blacklist = ["a","F","f","j"];

assert.strictEqual(cleanLogs(test5logs,test5Blacklist),'Sanitizing stopped Failure Reason: Fatal character occurred.');
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sanitizeLog(entry, blacklist) {
  let sanitizedEntry = [];

  for (let i = 0; i < entry.length; i++) {
    if (blacklist.includes(entry[i])) {
      sanitizedEntry.push("#");
    } else {
      sanitizedEntry.push(entry[i]);
    }
  }

  return sanitizedEntry.join("");
};

function cleanLogs(logs, blacklist) {
  let accumulator = 0;
  let informationalLine = 0;
  let fatalError = "$";
  let informationalIndicator = "/";

  let sanitizedLogs = [];

  for (let i = 0; i < logs.length; i++) {
    if (logs[i].includes(fatalError)) {
      console.log("Total Entry: " + logs.length + " Sanitized Entry Count: " + (accumulator + informationalLine) + " finished sanitize logs process.");
      return("Sanitizing stopped Failure Reason: Fatal character occurred.");
    } else if (logs[i].startsWith(informationalIndicator)) {
      informationalLine += 1;
      continue;
    } else {
      let tmpSanitized = sanitizeLog(logs[i], blacklist);
      sanitizedLogs.push(tmpSanitized);

      accumulator += 1;
    }
  }

  return("Total Entry: " +
        logs.length +
        " Sanitized Entry Count: " +
        (accumulator + informationalLine) +
        " finished sanitize logs process.");
};
```
