---
id: 658225d560369446ece5518b
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

Add another key-value pair to the dictionary you are appending to the `expense` list. Use the string `'category'` as the key, and the `category` parameter as the value.

# --hints--

You should have `expenses.append({'amount': amount, 'category': category})` in your code.

```js
assert.match(code, /^\s+expenses\.append\s*\(\s*\{\s*("|')amount\1\s*:\s*amount\s*,\s*("|')category\2\s*:\s*category\s*\}\s*\)/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
