---
id: 5c6c06847491271903d37cfd
title: Use the value attribute with Radio Buttons and Checkboxes
challengeType: 0
forumTopicId: 301099
dashedName: use-the-value-attribute-with-radio-buttons-and-checkboxes
---

# --description--

When a form gets submitted, the data is sent to the server and includes entries for the options selected. Inputs of type `radio` and `checkbox` report their values from the `value` attribute.

For example:

```html
<label for="indoor">
  <input id="indoor" value="indoor" type="radio" name="indoor-outdoor">Indoor
</label>
<label for="outdoor">
  <input id="outdoor" value="outdoor" type="radio" name="indoor-outdoor">Outdoor
</label>
```

Here, you have two `radio` inputs. When the user submits the form with the `indoor` option selected, the form data will include the line: `indoor-outdoor=indoor`. This comes from the `name` and `value` attributes of the "indoor" input.

If you omit the `value` attribute, the submitted form data uses the default value, which is `on`. In this scenario, if the user clicked the "indoor" option and submitted the form, the resulting form data would be `indoor-outdoor=on`, which is not useful. So the `value` attribute needs to be set to something to identify the option.

**How it works:**

- The `name` attribute groups related inputs together (like radio buttons where only one can be selected)
- The `value` attribute specifies what data gets sent when that option is selected
- When submitted, the format is: `name=value`
- Without a `value` attribute, inputs default to sending `on`, which doesn't tell you which option was chosen

**Practical example:**

If you have a form asking "What's your favorite color?" with radio buttons for "red", "blue", and "green", you'd set:

- All three inputs: `name="favorite-color"`
- First input: `value="red"`
- Second input: `value="blue"`
- Third input: `value="green"`

When submitted, you'd receive `favorite-color=blue` if the user selected blue, making it clear which option they chose.

# --instructions--

Give each of the existing `radio` and `checkbox` inputs the `value` attribute. Do not create any new elements. Use the input label text, in lowercase, as the value for the attribute.

# --hints--

One of your radio buttons should have the `value` attribute of `indoor`.

```js
assert.match($('label:contains("Indoor") > input[type="radio"]').attr('value'), /^indoor$/);
```

One of your radio buttons should have the `value` attribute of `outdoor`.

```js
assert.match($('label:contains("Outdoor") > input[type="radio"]').attr('value'), /^outdoor$/);
```

One of your checkboxes should have the `value` attribute of `loving`.

```js
assert.match($('label:contains("Loving") > input[type="checkbox"]').attr('value'), /^loving$/);

```

One of your checkboxes should have the `value` attribute of `lazy`.

```js
assert.match($('label:contains("Lazy") > input[type="checkbox"]').attr('value'), /^lazy$/);
```

One of your checkboxes should have the `value` attribute of `energetic`.

```js
assert.match($('label:contains("Energetic") > input[type="checkbox"]').attr('value'), /^energetic$/);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

