---
id: 694cc1b515976889afd4cded
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Now you need to store some string values about the movie ticket. You can store strings in Python this way:

```py
name = 'Alice'
```

Create a variable named `seat_type` to store the type of seat the user has selected and set its value to the string `Gold`. Create another variable named `show_time` to store the show time of the movie and set its value to the string `Evening`.

Remember to surround both values with either single or double quotes.

# --hints--

You should declare a variable named `seat_type`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

You should assign the string `Gold` to `seat_type`. Remember to surround the value with either single or double quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

You should declare a variable named `show_time`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

You should assign the string `Evening` to `show_time`. Remember to surround the value with either single or double quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
