---
id: 6958dd85587e10c99162dd81
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

As you learned in previous lessons, in Python, an `if` statement can be used to run code depending on if a condition is true.

An `if` statement consists of the `if` keyword, followed by a condition and a colon. The code to run when the condition is true, which must be indented, is called the body of the `if` statement.

```py
if condition:
    # Code to execute if condition is True
```

In this step, you will check if the user is eligible to book a movie ticket based on their age.

Create an `if` statement to check if `age` is greater than `17`. Inside the body of the `if` statement, print `User is eligible to book a ticket`. This will print the message only when the user's age is greater than `17`.

Remember to indent the body of the `if` statement and surround the message with single or double quotes inside the `print()` call.

# --hints--

You should have an `if` statement.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

The `if` statement condition should be `age > 17`.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

You should have a `print()` call inside your `if` statement. Remember to indent the body of the `if` statement.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

You should print `User is eligible to book a ticket` inside your `if` statement. Remember to surround the message with single or double quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
