---
id: 6958e2377cacbc46a7f7383a
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Now, add an `else` clause to your `if` statement and print `No extra charges will be applied` inside the body of `else` statement. This will print the message only when the extra charges condition is false.

Then, below the `else` clause, use the `print()` function to display a message that shows `Extra charges:` followed by the updated value of `extra_charges` and check the output in the terminal.

# --hints--

You should have an `else` clause.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1] == _Node()`)) })
```

You should print `No extra charges will be applied` inside your `else` clause. Remember to surround the message with single or double quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

You should print the value of `extra_charges` variable below your `if...else` statement.

```js
({ test: () => (runPython(`assert _Node(_code).find_calls("print")[1]`)) })
```

You should use the label `Extra charges:` as first argument. Remember to surround the message with single or double quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_calls("print")[1].find_call_args()[0].is_equivalent("'Extra charges:'")
`)) })
```

You should use the `extra_charges` variable as second argument.

```js
({ test: () => (runPython(`
assert _Node(_code).find_calls("print")[1].find_call_args()[1].is_equivalent("extra_charges")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 18:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualify for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
