---
id: 6958e59845777391b86d7331
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

The `if...elif...else` statement is used to check multiple conditions in order.

```py
if condition1:
   # Code to execute if condition1 is True
elif condition2:
   # Code to execute if condition2 is True
else:
   # Code to execute if all conditions are False
```

In this step, you will calculate service charges based on the type of seat the user has selected.

First, create a variable named `service_charges` and set it to `0`.

Then, you should use the `if...elif...else` statement to calculate the service charges based on `seat_type`:

- If `seat_type` is `Premium`, assign `5` to `service_charges`
- Else if `seat_type` is `Gold`, assign `3` to `service_charges`
- Otherwise, assign `1` to `service_charges`

Finally, print a message that shows `Service charges:` followed by the value of `service_charges`.

# --hints--

You should declare a variable named `service_charges`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

You should assign the integer value `0` to `service_charges`.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

You should have an `if` statement.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

The if statement condition should be `seat_type == 'Premium'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'") or _cond.is_equivalent("'Premium' == seat_type")
`)) })
```

You should assign the integer value `5` to `service_charges` variable inside your `if` statement.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

The elif statement condition should be `seat_type == 'Gold'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[1] 
assert _cond.is_equivalent("seat_type == 'Gold'") or _cond.is_equivalent("'Gold' == seat_type")
`)) })
```

You should assign the integer value `3` to `service_charges` variable inside your `elif` statement.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[1].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 3")
`)) })
```

You should have an `else` clause.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[2] == _Node() 
`)) })
```

You should assign the integer value `1` to `service_charges` variable inside your `else` clause.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[2].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 1")
`)) })
```

You should print the value of `service_charges` variable below your `if...elif...if` statement.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_bodies()[0].find_calls("print")[1]`)) })
```

You should use the label `Service charges:` as first argument.

```js
({ test: () => (runPython(`
_arg = _Node(_code).find_ifs()[4].find_bodies()[0].find_calls("print")[1].find_call_args()[0] 
assert _arg.is_equivalent("'Service charges:'")
`)) })
```

You should use the `service_charges` variable as second argument.

```js
({ test: () => (runPython(`
_arg = _Node(_code).find_ifs()[4].find_bodies()[0].find_calls("print")[1].find_call_args()[1] 
assert _arg.is_equivalent("service_charges")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 18:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualify for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age > 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
