---
id: 6958dfa0c74b237e1b56ce04
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

The user qualifies for a membership discount if they are a member.

Create a variable named `discount` and set its value to `0`. This will store the discount the user gets on the movie ticket.

# --hints--

You should declare a variable named `discount`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

You should assign the integer value `0` to `discount`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
