---
id: 6958dfa0c74b237e1b56ce04
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

The user qualify for a membership discount if they are a member.

First, create a variable named `discount` and set its value to `0`. This will store the discount the user gets on the movie ticket.

Then, create an `if` statement to check if `is_member` is `True`.

Inside the body of `if` statement, update the `discount` value to `3` and print `User qualify for membership discount` to the terminal.

# --hints--

You should declare a variable named `discount`.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

You should assign the integer value `0` to `discount`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

You should have an `if` statement.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

The if statement condition should be `is_member`.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

You should assign `3` to `discount` inside your `if` statement.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

You should print `User qualify for membership discount` inside your `if` statement.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualify for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 18:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
