---
id: 6958e0ecc121b0ae3dd9c5e0
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

You should also handle the case when the user does not qualify for a membership discount.

Add an `else` clause to your `if` statement and print `User does not qualify for membership discount`.

You can print a message followed by a variable like this:

```py
print('Total:', total)
```

You also want to display the updated value of `discount`. Below the `if...else` statement, use the `print()` function to display the message that shows `Discount:` followed by the value of `discount`. Then, check the output in the terminal.

# --hints--

You should have an `else` clause.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1] == _Node()`)) })
```

You should print `User does not qualify for membership discount` inside your `else` statement.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

You should print the value of `discount` variable below your `if...else` statement.

```js
({ test: () => (runPython(`assert _Node(_code).find_calls("print")[0]`)) })
```

You should use the label `Discount:` as first argument. Remember to surround the message with single or double quotes.

```js
({ test: () => (runPython(`
assert _Node(_code).find_calls("print")[0].find_call_args()[0].is_equivalent("'Discount:'")
`)) })
```

You should use the `discount` variable as second argument.

```js
({ test: () => (runPython(`
assert _Node(_code).find_calls("print")[0].find_call_args()[1].is_equivalent("discount")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 18:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualify for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
