---
id: 6960b8632b9544d1f57295f4
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Users between `18` and `21` can book tickets, but only when the `show_time` is not `Evening`.

Use the `or` operator to add other conditions to your existing `if` statement. Then, use the `and` operator to combine both the conditions checking if `age` is greater than `18` and `show_time` is not `Evening`. 

# --hints--

The if statment condition should be `age >= 21 or age > 18 and show_time != 'Evening'`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert  _cond.is_equivalent("age >= 21 or age > 18 and show_time != 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 18:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualify for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21:
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
