---
id: 6960a5165f7efe52a0c0aa39
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

In Python, the `and` operator is used to check if multiple conditions are true. Here's how you can use it to combine two conditions in an `if` statement:

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

The membership discount should only apply to members if their `age` is greater than or equal to `21`.

Use the `and` operator to combine the existing condition of your `if` statement with another condition checking if `age` is greater than or equal to `21`.

# --hints--

Your third `if` statement condition should be `is_member and age >= 21`.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
