---
id: 587d78a8367417b2b2512ae5
title: Elemente mit variablen Geschwindigkeiten animieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cZ89WA4'
forumTopicId: 301040
dashedName: animate-elements-at-variable-rates
---

# --description--

Es gibt eine Vielzahl von Möglichkeiten, die Geschwindigkeit von ähnlich animierten Elementen zu variieren. Bisher wurde dies durch die Verwendung der Eigenschaft `animation-iteration-count` (dt. Anzahl der Animationsdurchläufe) und das Setzen von `@keyframes`-Regeln (dt. Schlüsselbilder) erreicht.

Als Beispiel dient uns die Animation auf der rechten Seite, welche aus zwei Sternen besteht, welche je bei der 20 %-Marke in der `@keyframes`-Regel kleiner und transparent werden und die Funkel-Animation erzeugen. Du kannst die `@keyframes`-Regel für eines der Elemente ändern, sodass die Sterne mit unterschiedlicher Geschwindigkeit funkeln.

# --instructions--

Ändere die Animationsgeschwindigkeit für das Element mit dem Klassennamen `star-1`, indem du die `@keyframes`-Regel auf 50 % änderst.

# --hints--

Die `@keyframes`-Regel der `star-1`-Klasse sollte bei 50 % liegen.

```js
assert(code.match(/twinkle-1\s*?{\s*?50%/g));
```

# --seed--

## --seed-contents--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>

<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```

# --solutions--

```html
<style>
  .stars {
    background-color: white;
    height: 30px;
    width: 30px;
    border-radius: 50%;
    animation-iteration-count: infinite;
  }

  .star-1 {
    margin-top: 15%;
    margin-left: 60%;
    animation-name: twinkle-1;
    animation-duration: 1s;
  }

  .star-2 {
    margin-top: 25%;
    margin-left: 25%;
    animation-name: twinkle-2;
    animation-duration: 1s;
  }

  @keyframes twinkle-1 {
    50% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  @keyframes twinkle-2 {
    20% {
      transform: scale(0.5);
      opacity: 0.5;
    }
  }

  #back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: linear-gradient(black, #000099, #66c2ff, #ffcccc, #ffeee6);
  }
</style>
<div id="back"></div>
<div class="star-1 stars"></div>
<div class="star-2 stars"></div>
```
