---
id: 587d78a8367417b2b2512ae4
title: Erstelle einen CSS "Herzschlag" mit unendlichen Animationsdurchläufen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cDZpDUr'
forumTopicId: 301062
dashedName: make-a-css-heartbeat-using-an-infinite-animation-count
---

# --description--

Hier ist ein weiteres Beispiel einer kontinuierlichen Animation mit der `animation-iteration-count`-Eigenschaft, die das Herz verwendet, das du bei einer vorherigen Aufgabe entwickelt hast.

Die eine Sekunde lange Herzschlag-Animation besteht aus zwei Stücken. Die `heart` Elemente (einschließlich der `:before` und `:after` Teile) sind so animiert, dass sie ihre Größe mit der `transform`-Eigenschaft ändern. Das Hintergrund-`div` wechselt seine Farbe mit der `background`-Eigenschaft.

# --instructions--

Lass das Herz beständig weiter schlagen, indem du die `animation-iteration-count`-Eigenschaft sowohl für die `back` als auch die `heart`-Klasse hinzufügst und den Wert auf `infinite` setzt. Die `heart:before` und `heart:after` Selektoren benötigen keine Animationseigenschaften.

# --hints--

Die `animation-iteration-count`-Eigenschaft für die `heart`-Klasse sollte einen Wert von `infinite` besitzen.

```js
assert($('.heart').css('animation-iteration-count') == 'infinite');
```

Die Eigenschaft `animation-iteration-count` für die `back`-Klasse sollte einen Wert von `infinite` haben.

```js
assert($('.back').css('animation-iteration-count') == 'infinite');
```

# --seed--

## --seed-contents--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;

  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;

  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }

</style>
<div class="back"></div>
<div class="heart"></div>
```

# --solutions--

```html
<style>
  .back {
    position: fixed;
    padding: 0;
    margin: 0;
    top: 0;
    left: 0;
    width: 100%;
    height: 100%;
    background: white;
    animation-name: backdiv;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }

  .heart {
    position: absolute;
    margin: auto;
    top: 0;
    right: 0;
    bottom: 0;
    left: 0;
    background-color: pink;
    height: 50px;
    width: 50px;
    transform: rotate(-45deg);
    animation-name: beat;
    animation-duration: 1s;
    animation-iteration-count: infinite;
  }
  .heart:after {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: 0px;
    left: 25px;
  }
  .heart:before {
    background-color: pink;
    content: "";
    border-radius: 50%;
    position: absolute;
    width: 50px;
    height: 50px;
    top: -25px;
    left: 0px;
  }

  @keyframes backdiv {
    50% {
      background: #ffe6f2;
    }
  }

  @keyframes beat {
    0% {
      transform: scale(1) rotate(-45deg);
    }
    50% {
      transform: scale(0.6) rotate(-45deg);
    }
  }
</style>
<div class="back"></div>
<div class="heart"></div>
```
