---
id: 587d781d367417b2b2512ac5
title: Zeilenhöhe von Absätzen festlegen
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVWdcv'
forumTopicId: 301070
dashedName: set-the-line-height-of-paragraphs
---

# --description--

CSS bietet die Eigenschaft `line-height`, um die Höhe jeder Zeile in einem Textblock zu ändern. Wie der Name andeutet, ändert er die vertikale Ausdehnung der jeweiligen Zeilen.

# --instructions--

Füge eine `line-height`-Eigenschaft zum `p`-Tag hinzu und setze sie auf 25px.

# --hints--

Dein Code sollte die `line-height`-Eigenschaft für das `p`-Tag auf 25 Pixel setzen.

```js
assert($('p').css('line-height') == '25px');
```

# --seed--

## --seed-contents--

```html
<style>
  p {
    font-size: 16px;

  }
</style>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit, sed do eiusmod tempor incididunt ut labore et dolore magna aliqua. Ut enim ad minim veniam, quis nostrud exercitation ullamco laboris nisi ut aliquip ex ea commodo consequat. Duis aute irure dolor in reprehenderit in voluptate velit esse cillum dolore eu fugiat nulla pariatur.
</p>
```

# --solutions--

```html
<style>
  p {
    font-size: 16px;
    line-height: 25px;
  }
</style>
<p>
  Lorem ipsum dolor sit amet, consectetur adipiscing elit, sed do eiusmod tempor incididunt ut labore et dolore magna aliqua. Ut enim ad minim veniam, quis nostrud exercitation ullamco laboris nisi ut aliquip ex ea commodo consequat. Duis aute irure dolor in reprehenderit in voluptate velit esse cillum dolore eu fugiat nulla pariatur.
</p>
```
