---
id: 66cc0a9e06e00b75d6782be9
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

In previous lessons, you learned how to call (invoke) a function. Calling a function means to execute the code inside the function.

Here is a reminder of how to call a function:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Add a `console.log` statement, and inside it, call the `addTwoAndSeven` function.

You should now see the sum of `2` and `7` logged to the console.

# --hints--

You should call the `addTwoAndSeven` function inside a `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
