---
id: 66cc172af7159a7c67804544
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Now it is time to add the multiplication functionality to your calculator.

Declare a function called `calculateProduct` that takes two parameters, `num1` and `num2`.

Inside the function, return the product of the two parameters.

Then, add a `console.log` that calls the `calculateProduct` function with the arguments `13` and `5`.

# --hints--

You should have a function called `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

Your `calculateProduct` function should have two parameters called `num1` and `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

Your `calculateProduct` function should return the product of the two parameters.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

You should have a `console.log` that calls the `calculateProduct` function with the arguments `13` and `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
