---
id: 66cc1a3a39aef47d6473cb2f
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

The next step is to add the division functionality to your calculator.

Declare a function called `calculateQuotient` that takes two parameters, `num1` and `num2`.

Inside the function, return the quotient of the two parameters.

Then, add a `console.log` that calls the `calculateQuotient` function with the arguments `7` and `11`.

# --hints--

You should have a function called `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Your `calculateQuotient` function should have two parameters called `num1` and `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

Your `calculateQuotient` function should return the quotient of the two parameters.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

You should have a `console.log` that calls the `calculateQuotient` function with the arguments `7` and `11`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
