---
id: 66cc0ba5881acb7692cfc4de
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now, it is time to add another function.

Declare a function called `addThreeAndFour` that returns the sum of `3` and `4`.

Then call the `addThreeAndFour` function inside a `console.log` to see the result.

# --hints--

You should declare a function called `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Your `addThreeAndFour` function should return the sum of `3` and `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

Your `addThreeAndFour` function should return the sum using the `+` operator.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

You should call the `addThreeAndFour` function inside a `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
