---
id: 66cc0de2920bec775f610424
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Even though the functions work as expected, there is a lot of repetition in your code. 

Since you are building a calculator, you don't want to have to create a function for every possible combination of numbers you want to add together.

It would be better to create a single reusable function that can add any two numbers together.

Remove all of the code you have written so far.

In the next step, you will review how to work with parameters and arguments in functions.

# --hints--

You should not have a function called `addTwoAndSeven`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

You should not have a function call to `addTwoAndSeven`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

You should not have a function called `addThreeAndFour`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

You should not have a function call to `addThreeAndFour`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
