---
id: 66cc0f1ae40802781b2ea972
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

In previous lessons, you learned how to work with function parameters.

A function parameter is a variable that is defined in the function's declaration and acts as a placeholder.

Here is an example of a function that has a parameter:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Declare a function called `calculateSum` that takes two parameters, `num1` and `num2`.

# --hints--

You should declare a function called `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Your `calculateSum` function should take two parameters, `num1` and `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
