---
id: 66cc095dfe1682753d2ab030
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

One of concepts you learned in previous lessons was how to return values from a function.

Here is a reminder of how to return a value from a function:

```js
function myFunction() {
  return "Hello World";
}
```

Inside your `addTwoAndSeven` function, return the sum of `2` and `7`.

# --hints--

Your `addTwoAndSeven` function should return the sum of `2` and `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Your `addTwoAndSeven` function should return the sum using the `+` operator.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
const validExpressions = ["2 + 7", "7 + 2", "(2 + 7)", "(7 + 2)"];
assert.isTrue(validExpressions.some(expr => addTwoAndSeven?.hasReturn(expr)));
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
