---
id: 66cdf76685e4cb5a8726e27b
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

To further test out your `calculateSum` function, you will need to call it with different arguments.

Add another `console.log` that calls the `calculateSum` function with the arguments `10` and `10`.

Below that `console.log`, add another `console.log` that calls the `calculateSum` function with the arguments `5` and `5`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

You should have a `console.log` that calls the `calculateSum` function with the arguments `10` and `10`.

```js
assert.deepInclude(spy.calls, [calculateSum(10, 10)]);
```

You should have a `console.log` that calls the `calculateSum` function with the arguments `5` and `5`.

```js
assert.deepInclude(spy.calls, [calculateSum(5, 5)]);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
