---
id: 612e7d1c29fb872d6384379c
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Вкладіть другий `div` в свій наявний `div` та встановіть `class` на `keys`.

# --hints--

Ви повинні створити другий елемент `div`.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

Ваш новий елемент `div` повинен бути в межах наявного елемента `div`.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

Ваш новий елемент `div` повинен мати `class` зі значенням `keys`.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
