---
id: 612e8279827a28352ce83a72
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Тепер скопіюйте набір із семи елементів `.key` та вкладіть ще два набори в div `.keys`.

# --hints--

Ви повинні мати 21 елемент `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Ви повинні мати 15 елементів `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Всі елементи `.key` повинні бути в елементі `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Всі елементи `.black--key` повинні бути в елементі `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
