---
id: 612e78af05201622d4bab8aa
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Час почати працювати над піаніно. Cтворіть елемент `div` в елементі `body`, що має `id` зі значенням `piano`.

# --hints--

Ви повинні створити новий елемент `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Ваш `div` повинен бути в межах елемента `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

Ваш `div` повинен мати `id` зі значенням `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
