---
id: 6909ed80e1c195f49564edac
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you'll learn how to work with conditional statements and comparison operators.

Begin by declaring a variable named `truthyOrFalsy` and assign it the boolean value `true`.

# --hints--

You should have a variable named `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

You should assign `true` to your `truthyOrFalsy` variable.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
