---
id: 690a562fd611e8155599ce98
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Change the assignment of `truthyOrFalsy` by setting its value to `"freeCodeCamp"`. You will still see `true` logged to the console because a non-empty string is a truthy value.

# --hints--

Your `truthyOrFalsy` variable should have the value `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
