---
id: 690a51806fb743bc896c3af7
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

As you learned in a previous lesson, booleans are a data type with only `true` and `false` values.

A truthy value is a value that evaluates to `true` in a boolean context. Examples of truthy values are:

- Non-empty strings
- Non-zero numbers
- The boolean `true`

And other values you'll meet later on.

You can use the `Boolean()` function to check the truthiness of a value. For example, `Boolean("Hello World!")` will return `true` because `"Hello World!"` is truthy.

Now, call `Boolean(truthyOrFalsy)` and log the result to the console.

# --hints--

You should log `Boolean(truthyOrFalsy)` to the console.


```js 
const consoleSpy = __helpers.spyOn(console, 'log');
const calledWithBoolean = consoleSpy.calls.some(args => args.includes(false) || args.includes(true));
consoleSpy.restore();
assert.isFalse(calledWithBoolean);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
