---
id: 690a57fef66aac39dffc61a2
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Now that you know what truthy and falsy values are, remove both your variable declaration and `console.log` statement from your code.

# --hints--

You should not have a `truthyOrFalsy` variable.

```js
const explorer = await __helpers.Explorer(code);
assert.isFalse(explorer.variables.truthyOrFalsy);
```

You should not have `console.log(Boolean(truthyOrFalsy))` in your code.


```js
const consoleSpy = __helpers.spyOn(console,'log');
const calledWithFalse = consoleSpy.calls.some(args => args.includes(false) || args.includes(true));
consoleSpy.restore();
assert.isFalse(calledWithFalse);

```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
