---
id: 690a110a76d0496d9e7c28eb
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Declare another variable named `timmyAge` and set it to `18`.

# --hints--

You should have a variable named `timmyAge`.

```js
assert.exists(timmyAge);
```

You should assign `18` to your `timmyAge` variable.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}
--fcc-editable-region--

--fcc-editable-region--
```
