---
id: 690a57095c4afb27bbe751b3
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Falsy values are values that evaluate to `false` in a boolean context. Examples of falsy values are:

- `""` (empty string)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

You'll meet all of these going through this coursework.

For now, change the assignment of `truthyOrFalsy` by setting its value to an empty string. After that, you'll see `false` logged to the console.

# --hints--

You should assign an empty string to `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
