---
id: 62a3b41c9494f937560640ab
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

Declare two more variables named `monsterHealth` and `inventory`, but do not initialize them.

# --hints--

You should use `let` to declare your `monsterHealth` variable.

```js
assert.match(code, /let\s+monsterHealth/i);
```

You should use camelCase to name your `monsterHealth` variable.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` should not have a value.

```js
assert.isUndefined(monsterHealth);
```

You should not assign a value to your `monsterHealth` variable. Remember your semi-colon.

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

You should use `let` to declare your `inventory` variable.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` should not have a value.

```js
assert.isUndefined(inventory);
```

You should not assign a value to your `inventory` variable. Remember your semi-colon.

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
