---
id: 67f3c0952ebc82db347449e3
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Inside your `select` element, you need to create options for the user to choose their power source. The `select` should have seven options listed below:

| Power Source                  |
|-------------------------------|
| Select one                    |
| Bitten by a strange creature  |
| Radioactive exposure          |
| Science experiment            |
| Alien heritage                |
| Ancient artifact discovery    |
| Other                         |

# --hints--

Your should create the first `option` element with the text `Select one` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[0]?.textContent, "Select one");
```

Your should create a second `option` element with the text `"Bitten by a strange creature"` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[1]?.textContent, "Bitten by a strange creature");
```

Your should create a third `option` element with the text `"Radioactive exposure"` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[2]?.textContent, "Radioactive exposure");
```

Your should create a fourth `option` element with the text `"Science experiment"` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[3]?.textContent, "Science experiment");
```

Your should create a fifth `option` element with the text `"Alien heritage"` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[4]?.textContent, "Alien heritage");
```

Your should create a sixth `option` element with the text `"Ancient artifact discovery"` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[5]?.textContent, "Ancient artifact discovery");
```

Your should create a seventh `option` element with the text `"Other"` inside your `select` element.

```js
const optionEls = document?.querySelectorAll("select > option");
assert.equal(optionEls[6]?.textContent, "Other");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Superhero Application Form</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.0/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.0/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SuperheroForm } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<SuperheroForm />);
    </script>
  </body>
</html>
```

```css
body {
  margin: 0;
  padding: 0;
  min-height: 100vh;
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  background: linear-gradient(30deg, #ff9999 50%, #6699ff 50%)
}

.form-wrap {
  background-color: white;
  width: 400px;
  padding: 20px;
  border: 1px solid black;
  box-shadow: 5px 5px 10px black;
}

.form-wrap h2,
.form-wrap p {
  text-align: center;
}

.form-wrap p {
  position: relative;
  top: -18px;
}

.section {
  display: flex;
  margin-bottom: 30px;
}

.column {
  flex-direction: column;
}

.submit-wrap {
  text-align: center;
}

.submit-btn {
  display: block;
  margin: 0 auto;
  padding: 0.4rem 0.5rem;
  border: 1px solid black
}

.submit-btn:hover {
  cursor: pointer;
  background-color: #f3f3f3;
}

.submit-btn:disabled {
  cursor: not-allowed;
}
```

```jsx
const { useState } = React;

export const SuperheroForm = () => {

  const [heroName, setHeroName] = useState('');
  const [realName, setRealName] = useState('');
  const [powerSource, setPowerSource] = useState('');
  const [powers, setPowers] = useState([]);

  return (
    <div className='form-wrap'>
      <h2>Superhero Application Form</h2>
      <p>Please complete all fields</p>
      <form>
        <div className='section'>
          <label>
            Hero Name
            <input
              type='text'
              value={heroName}
              onChange={e => setHeroName(e.target.value)}
            />
          </label>
          <label>
            Real Name
            <input
              type='password'
              value={realName}
              onChange={e => setRealName(e.target.value)}
            />
          </label>
          <label className='section column'>
            How did you get your powers?
            --fcc-editable-region--
            <select>

            </select>
            --fcc-editable-region--
          </label>
        </div>
      </form>
    </div>
  )
};
```
