---
id: 672b4434dd6d6a1edc405419
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Add two `a` elements inside the `.card-links` element.

The first one should have a text of `Send Card`, a `class` of `send-link` and an `href` of `#send`.

The second one should have a text of `Share on Social Media`, a `class` of `share-link` and an `href` of `#share`.

# --hints--

The first `a` element inside `.card-links` should have a text of `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

The first `a` element should have a `class` with a value of `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

The first `a` element should have an `href` with a value of `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

The second `a` element inside `.card-links` should have a text of `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

The second `a` element should have a `class` with a value of `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

The second `a` element should have an `href` with a value of `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

--fcc-editable-region--

      <div class="card-links"></div>

--fcc-editable-region--

  	</div>
  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
