---
id: 672e44e811e610232fead333
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Add an `h2` to `#send` that contains the text `Sending your card...`, then add a `p` element with the text `Card successfully sent to your recipient!`.

# --hints--

The `#send` element should contain a `h2` element.

```js
assert.exists(document.querySelector("#send > h2"));
```

The `h2` element should contain the text `Sending your card...`.

```js
assert.strictEqual(document.querySelector("#send > h2").innerText.trim(), "Sending your card...");
```

The second element inside `#send` should be a `p` element.

```js
assert.strictEqual(document.querySelector("#send").children[1].tagName, "P");
```

The `p` element should have the text `Card successfully sent to your recipient!`.

```js
assert.strictEqual(document.querySelector("#send").children[1].innerText.trim(), "Card successfully sent to your recipient!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>

      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>

--fcc-editable-region--

  <section id="send"></section>

--fcc-editable-region--

  <section id="share"></section>

  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
