---
id: 66b205e6eacba4c4e54ea434
title: Build an Email Masker
challengeType: 14
dashedName: build-an-email-masker
---

# --description--

In this lab, you will mask the username part of an email address with asterisks. Masking is a term used to hide or replace sensitive information with asterisks or other characters.

For example, if the email address was `myEmail@email.com`, then the masked email address will be `m*******l@email.com`.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. Declare a variable named `email` to store the email address you want to mask. 
1. You should create a variable called `maskedName` and assign it the result of masking the `email`.  Remember that you can use methods like `slice`, `repeat`, `indexOf` `slice` or even `replace` to help you.
1. Log the `maskedName` variable to the console. 

# --hints--

You should have an `email` variable.

```js
assert.match(__helpers.removeJSComments(code), /email/);
```

You should find the index of `@` in `email` using the `indexOf()` method.

```js
assert.match(__helpers.removeJSComments(code), /email\.indexOf/);
```

You should pass `@` as an argument to the `indexOf()` method.

```js
assert.match(__helpers.removeJSComments(code), /indexOf\(['"]@['"]\)/);
```

You should store the index of `@` in a variable named `atIndex`.

```js
assert.match(__helpers.removeJSComments(code), /atIndex\s*=\s*email\.indexOf\("@"\)/);
```

You should split the username part of the email from the first index to the `atIndex` using the `slice()` method.

```js
assert.match(__helpers.removeJSComments(code), /email\.slice\(0,\s*atIndex\)/);
```

You should store the username part of the email in a variable named `userName`.

```js
assert.match(__helpers.removeJSComments(code), /userName\s*=\s*email\.slice\(0\,\s*atIndex\);/);
```

You should extract the domain part of the email from the `atIndex` to the end of the `email` using the `slice()` method.

```js
assert.match(__helpers.removeJSComments(code), /email\.slice\(atIndex\)/);
```

You should store the domain part of the email in a variable named `domain`.

```js
assert.match(__helpers.removeJSComments(code), /domain\s*=\s*email\.slice\(atIndex\);/);
```

You should mask the username by replacing the characters between the first and last letters with asterisks using the `repeat()` method.

```js
assert.match(__helpers.removeJSComments(code), /repeat\(\s*userName\.length\s*-\s*2\s*\)/);
```

You should create the full masked name by combining the first letter of the username, repeated asterisks, and the last letter of the username and store it in a variable named `maskedName`.

```js
assert.match(__helpers.removeJSComments(code), /maskedName\s*=\s*userName\[0\]\s*\+\s*"\*"\.repeat\(userName\.length\s*-\s*2\)\s*\+\s*userName\[userName\.length\s*-\s*1\]/);
```

You should append `domain` to the masked name to get the final masked email address and store it in a variable named `finalMaskedEmail`.

```js
assert.match(__helpers.removeJSComments(code), /finalMaskedEmail\s*=\s*maskedName\s*\+\s*domain/);
```

Print the `finalMaskedEmail` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(finalMaskedEmail\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const email = "apple.pie@example.com";

const atIndex = email.indexOf("@");

const userName = email.slice(0, atIndex);
const domain = email.slice(atIndex);

const maskedName =
  userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1];

let finalMaskedEmail = maskedName + domain;
console.log(finalMaskedEmail);
```
