---
id: 66b205e6eacba4c4e54ea434
title: Build an Email Masker
challengeType: 14
dashedName: build-an-email-masker
---

# --description--

In this lab, you will mask the username part of an email address with asterisks. Masking is a term used to hide or replace sensitive information with asterisks or other characters.

For example, if the email address was `myEmail@email.com`, then the masked email address will be `m*******l@email.com`.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. Create a function named `maskedEmail` that takes `email` as an argument.
2. Inside the function, you should create a variable called `maskedName` and assign it the result of masking the `email` along with concatenating the domain name.  Remember that you can use methods like `slice`, `repeat`, `indexOf` `slice` or even `replace` to help you. Return `maskedName` from the function.
3. Outside the function, declare a variable named `email` to store the email address you want to mask.
4. Call the `maskEmail` function with the `email` variable and output the result to the console. 
5. `maskEmail("apple.pie@example.com")` should return `"a*******e@example.com"`.
6. `maskEmail("freecodecamp@example.com")` should return `"f**********p@example.com"`.

# --hints--

You should define a funtion named `maskEmail`.

```js
assert(typeof maskEmail === 'function');
```

The `maskEmail` funtion should take `email` as an argument.

```js
const condition1 = /maskEmail\s*=\s*\(email\)\s*=>\s*{/.test(code);
const condition2 = /maskEmail\(\s*email\s*\)/.test(code);
assert(condition1 || condition2);
```

Outside the function, you should have an `email` variable.

```js
assert.isDefined(email);
```

You should assign a valid email address to your `email` variable.

```js
assert.match(email, /^[\w\.-]+@[a-zA-Z\d\.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` should return `"a*******e@example.com"`.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` should return `"f**********p@example.com"`.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

You should log the function call to `maskedEmail` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
