---
id: 655cd899f8de09431eabb40c
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, you'll learn data structures by building the merge sort algorithm. This is a sorting algotithm that uses the divide-and-conquer principle to sort collections of data. 

Create a function called `merge_sort`. This function will handle the task of sorting an list of numbers.

# --hints--

You should hava a function named `merge_sort` 

```js
({ test: () => {
		assert(__pyodide.runPython(`
        import inspect
        inspect.isfunction(__locals.get('merge_sort'))
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
