---
id: 6564aee9c077774ea49c3faf
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

The merge sort algorithm mainly performs three actions recursively: divide, conquer, and merge. Let's start with "divide".

First, replace the `pass` keyword with a variable `mid_point`. Use the integer division operator (`//`) to divide the length of the `array` list by 2 and assign the result to your new `mid_point` variable. Remember to indent your code.

# --hints--

You should not have `pass` in the body of the function.

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}pass/m);
  }
})
```

You should declare a `middle_point` variable and assign `len(array) // 2` to it.

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}middle_point\s*=\s*len\(\s*array\s*\)\s*\/\/\s*2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
