---
id: 6564a9fe51964c229d5b7f4c
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

We'll need a parameter that denotes the data collection to be sorted. Create a parameter called `array` in the `merge_sort` function. fsdffdfd

# --hints--

Function parameters are defined within the function's parenthesis

```js
({ test: () => assert(__pyodide.runPython(`
      import inspect
      str(inspect.signature(__locals.get('merge_sort'))) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    random_numbers
--fcc-editable-region--
```
