---
id: 656627b47bd2d2a4afbc945d
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

In the previous step, we got the mid point. You can use it to divide `array` into two and assign each part to new variables. 

You can slice and extract certain portions of a data structure using the slice notation (`:`). 

In the code below, a `left_part` variable has been created. We assigned `array[:middle_point]` to the variable which denotes every element from the beginning of `array` up to the element with index `middle_point`, but not including it.

Create a `right_part` variable and assign all the elements from the `mid_point` to the end of `array` to it.


# --hints--

You should have a variable named `right_part`

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*(?!=)/m);
  }
})
```

You should assign `array[middle_point:]` to the `right_part` variable

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /right_part\s*=\s*array\[middle_point:\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2

--fcc-editable-region--
    left_part = array[:middle_point]

--fcc-editable-region--
```
