---
id: 656639771fed09b5c0e8fe71
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Now that you've divided the `array` list into two separate lists, you'll do the same thing for each list until every element stands alone.

The idea is to keep dividing the array until individual elements are considered. A list with a single number is always sorted.

To do that, you have to recursively call the `merge_sort` function on the `left_part` and `right_part`.

Start by calling the `merge_sort()` function.
# --hints--

To call a function, you have to add parenthesis after the function's name. That is: `merge_sort()`

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
```

```py
--fcc-editable-region--

--fcc-editable-region--
```
