---
id: 656709e50ed928da35260276
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

In the previous step, we assigned the element in the `left_part` to the `array` list because it was smaller. But this will not always be the case. In some comparison cases, the element on the right could be smaller.

We've already created an `if` statement that checks if the `left_part` index is less than the `right_part` index. 

Using an `else` statement will execute code when `left_part` index is not less than the `right_part` index. So the `else` statement will only run when the `right_part` index is less than the `left_part` index.

Create an `else` statement. In its code block/suite, assign `right_part[right_array_index]` to `array[sorted_index]`.

# --hints--

You should have `array[sorted_index] = right_part[right_array_index]` in the body of your `else` statement.

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /else:\s*[^}]*array\[sorted_index\]\s*=\s*right_part\[right_array_index\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
