---
id: 656707e11e671ed54c96f7ec
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

When the `if` statement returns `True`, it means that the element in the `left_part` list is smaller than the element it is being compared to in the `right_part` list.

In that case, we can assign the `left_part` index to the sorted array. 

For the `if` statement's suite, assign `left_part[left_array_index]` to `array[sorted_index]`, and increment `left_array_index` by `1`.

# --hints--

You should have `array[sorted_index] = left_part[left_array_index]` in the body of the `if` statement. Remember to remove the `pass` keyword

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(left_part\)\s+and\s+right_array_index\s*<\s*len\(right_part\):\s*[^}]*if\s+left_part\[left_array_index\]\s*<\s*right_part\[right_array_index\]:\s*[^}]*array\[sorted_index\]\s*=\s*left_part\[left_array_index\]\s*(?!.*\bpass\b)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
--fcc-editable-region--
            pass

--fcc-editable-region--
```
