---
id: 656706afd65547d22bee0b68
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Inside of the `while` loop. Create an `if` statement that checks if the index of `left_part` is less than the index of `right_part`.

First, remove the `pass` keyword from the body of the `while` loop, and create the `if` statement in its place.

Then create the `if` statement and its conditions.

Lastly, use the `pass` keyword in the body of the `if` statement.

# --hints--

Remove the `pass` keyword and create an `if` statement with this condition: `left_part[left_array_index] < right_part[right_array_index]`. Add the `pass` keyword to the body of the `if` statement

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(left_part\)\s+and\s+right_array_index\s*<\s*len\(right_part\):\s*[^}]*if\s+left_part\[left_array_index\]\s*<\s*right_part\[right_array_index\]:\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
      pass

--fcc-editable-region--
```
