---
id: 6577254891048c8f2c19e961
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Call the `merge_sort()` function again. Do not pass in any parameter.

# --hints--

To call a function, you have to add parenthesis after the function's name. That is: `merge_sort()`

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
--fcc-editable-region--

--fcc-editable-region--
```
