---
id: 6577320da0d4c2e594d418e2
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Pass `right_part` as a parameter to the `merge_sort()` function you created in the last step.

# --hints--

`right_part` should be a parameter of the function. Paremeters are defined within the parenthesis

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*right_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
--fcc-editable-region--
    merge_sort()
--fcc-editable-region--
```
