---
id: 65772ef923f922cd720e5008
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Pass `left_part` as a parameter to the `merge_sort()` function you created in the last step.

# --hints--

`left_part` should be a parameter of the function. Paremeters are defined within the parenthesis

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*left_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
--fcc-editable-region--
    merge_sort()
--fcc-editable-region--
```
