---
id: 646c48df8674cf2b91020ecc
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Додайте елемент `link` в межах елемента `head`. Встановіть для цього елемента `link` атрибут `rel` на `stylesheet` та атрибут `href` на `./styles.css`.

# --hints--

Ви повинні мати елемент `link`.

```js
assert.match(code, /<link/)
```

Елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1/)
```

Елемент `link` повинен мати атрибут `href` зі значенням `./styles.css`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1\s+href=\1\.\/styles\.css\1\s*\/?>/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
