---
id: 646c5d7057c45f432fcdd46c
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Надайте елементу `.cat-head` ширину зі значенням `205px` та висоту зі значенням `180px` за допомогою селектора класу. Також надайте йому кордон зі значенням `1px solid #000` та `border-radius` зі значенням `46%`.

# --hints--

Ви повинні мати селектор `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Селектор `.cat-head` повинен мати `width` зі значенням `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

Селектор `.cat-head` повинен мати `height` зі значенням `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

Селектор `.cat-head` повинен мати `border` зі значенням `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

Селектор `.cat-head` повинен мати `border-radius` зі значенням `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
