---
id: 646c47867800472a4ed5d2ea
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Почніть із базової структури HTML. Додайте оголошення `DOCTYPE` `html` та елемент `html` з атрибутом `lang` зі значенням `en`. Також додайте `head` та `body` в межах елемента `html`.

# --hints--

Ви повинні мати оголошення `DOCTYPE` для `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Оголошення `DOCTYPE` має бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ви повинні мати початковий тег `html` з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ви повинні мати кінцевий тег `html`.

```js
assert(code.match(/<\/html>/i));
```

Ви повинні мати початковий тег `head`.

```js
assert(code.match(/<head>/i));
```

Ви повинні мати кінцевий тег `head`.

```js
assert(code.match(/<\/head>/i));
```

Ви повинні мати початковий тег `body`.

```js
assert(code.match(/<body>/i));
```

Ви повинні мати кінцевий тег `body`.

```js
assert(code.match(/<\/body>/i));
```

Елемент `body` повинен йти після елемента `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Елементи `head` та `body` мають знаходитись всередині елемента `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
