---
id: 646c5ace05e4be4211407935
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Щоб побачити елемент `cat-head`, надайте йому лінійний градієнт з `#5e5e5e` зі значенням `85%` та `#45454f` зі значенням `100%`.

Ви можете не помітити різницю між цими кольорами, але вона існує.

# --hints--

Селектор `.cat-head` повинен мати властивість `background`.

```js
assert.match(code, /background:/)
```

Властивість `background` має використати функцію `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

Функція `linear-gradient` має встановити перший колір `#5e5e5e` зі значенням `85%`.

```js
assert.match(code, /background:\s*linear-gradient\(#5e5e5e\s+85%/)
```

Функція `linear-gradient` має встановити другий колір `#45454f` зі значенням `100%`.

```js
assert.match(code, /background:\s*linear-gradient\(#5e5e5e\s+85%,\s+#45454f\s+100%\);?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
