---
id: 646c48df8674cf2b91020ecb
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

В межах елемента `head` додайте тег `meta` з атрибутом `charset` зі значенням `utf-8`. Також додайте елемент `title` з текстом `fCC Cat Painting`.

# --hints--

Додайте один елемент `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Додайте один елемент `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Елемент `title` повинен мати текст `fCC Cat Painting`. Зауважте, що написання та вживання великої літери мають значення.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
