---
id: 646c59e942f35541923104bf
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Повернемось до HTML: створіть елемент `main`. Всередині цього елемента `main` додайте елемент `div` з класом `cat-head`.

# --hints--

Ви повинні мати елемент `main`.

```js
assert(document.querySelectorAll('main').length === 1);
```

Ви повинні мати елемент `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

Елемент `div` повинен мати клас `cat-head`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

Елемент `div` має знаходитись всередині тегу `main`.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
