---
id: 646c5e727e56e743c9aed4a1
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Позиціювання CSS дозволяє розмістити елементи в браузері так, як ви хочете. Існує властивість `position`, значенням якої може бути `static`, `absolute`, `relative`, `sticky` або `fixed`.

Як тільки ви налаштуєте властивість `position`, ви зможете переміщати елемент, встановлюючи піксельні чи відсоткові значення властивостей `top`, `right`, `left` або `bottom`.

`static` — це позиціювання всіх елементів за замовчуванням. Якщо призначити його до елемента, ви не зможете переміщати його за допомогою `top`, `right`, `left` або `bottom`.

Надайте `.cat-head` властивість `position` зі значенням `static`, а потім встановіть властивості `top` та `left` на `100px`.


# --hints--

Селектор `.cat-head` повинен мати властивість `position` зі значенням `static`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Селектор `.cat-head` повинен мати властивість `top` зі значенням `100px`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Селектор `.cat-head` повинен мати властивість `left` зі значенням `100px`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
