---
id: 6476fb7cbaafa36d65e9cf35
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` — це властивість `position`, яка дозволяє прикріпити елемент до сторінки, незалежно від того, куди прогорне користувач.

Вам знадобиться зробити декілька розміток, щоб дізнатись, як працює позиціювання `fixed`. Створіть елемент `div` з класом `box` в HTML.

# --hints--

Ви повинні мати елемент `div`.

```js
assert(document.querySelectorAll('div').length === 2);
```

Елемент `div` повинен мати клас `box`.

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
