---
id: 6476f7a4827bcc61682f2347
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Наступна властивість позиціювання — це `absolute`. Коли ви використовуєте значення `absolute` для властивості `position`, то елемент виймається з нормального потоку документа, а його позиціювання визначається властивостями `top`, `right`, `bottom` та `left`.

Встановіть властивість позиціювання елемента `.cat-head` на `absolute`, а потім встановіть властивості `top` та `left` на будь-яке піксельне значення.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

Селектор `.cat-head` повинен мати властивість `position` зі значенням `absolute`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Селектор `.cat-head` повинен мати властивість `top` зі значенням будь-якого числа в `px`. Не забудьте про крапку з комою.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Селектор `.cat-head` повинен мати властивість `left` зі значенням будь-якого числа в `px`. Не забудьте про крапку з комою.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
