---
id: 6476f5c17f99146071ee884c
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Як бачите, нічого не відбувається. Елемент `.cat-head` не перемістився, хоча ви встановили `top` та `left` на `100px`. Але не у випадку з позиціюванням `relative`.

Коли ви використовуєте значення `relative`, то елемент розміщено відповідно до нормального потоку документа, але значення `top`, `left`, `bottom` та `right` стають активними.

Надайте елементу `.cat-head` позиціювання `relative`, а не `static`, і залиште властивості `top` та `left` якими вони є.

# --hints--

Селектор `.cat-head` повинен мати властивість `position` зі значенням `relative`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

Селектор `.cat-head` повинен мати властивість `top` зі значенням `100px`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Селектор `.cat-head` повинен мати властивість `left` зі значенням `100px`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
