---
id: 64770351e8586671ec0911f0
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Останнє значення властивості позиціювання — це `sticky`. Позиціювання `sticky` є гібридом `relative` та `fixed`. Воно дозволяє **прикріпити** елемент до певного місця в межах елемента чи демонстраційного вікна на основі прокручування.

Змініть значення властивості `position` селектора `.cat-head` на `sticky`, встановіть `top` на `0`, а потім видаліть `left` та її значення.

**Примітка:** щоб дізнатись, як працює `sticky`, потрібно розмістити певний текст до та після елемента `div` `.cat-head`. If you scroll down after that, you'll see that the `.cat-head` gets stuck to the top and remains there.

# --hints--

Селектор `.cat-head` повинен мати властивість `position` зі значенням `sticky`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

Селектор `.cat-head` повинен мати властивість `top` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Вилучіть властивість `left` та її значення зі свого коду.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
