---
id: 6477062778c85972eb648030
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Тепер потрібно відцентрувати голову кота.

Give the `.cat-head` element a `position` property set to `absolute`. Встановіть значення властивостей `right`, `left`, `top`, `bottom` на `0`, а потім встановіть властивість `margin` всіх сторін на `auto`. Це один зі способів відцентрувати елемент вертикально та горизонтально, використовуючи позиціювання CSS.

# --hints--

Селектор `.cat-head` повинен мати властивість `position` зі значенням `absolute`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Селектор `.cat-head` повинен мати властивість `top` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Селектор `.cat-head` повинен мати властивість `left` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

Селектор `.cat-head` повинен мати властивість `right` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

Селектор `.cat-head` повинен мати властивість `bottom` зі значенням `0`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

Селектор `.cat-head` повинен мати властивість `margin` зі значенням `auto`. Не забудьте про крапку з комою.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
