---
id: 5f3326b143638ee1a09ff1e3
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Per prepararti a creare del contenuto, aggiungi un elemento `body` sotto l'elemento `head`.

# --hints--

Dovresti avere un tag di apertura `<body>`.

```js
assert(code.match(/<body>/i));
```

Dovresti avere un tag di chiusura `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Non dovresti cambiare l'elemento `head`. Assicurati di non aver eliminato il tag di chiusura.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

