---
id: 5f33310c1851c6c4da013250
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

È ora di aggiungere qualche contenuto del menu. Aggiungi un elemento `main` al di sotto dell'elemento `header` esistente. Conterrà delle informazioni sui prezzi del caffè e dei dessert offerti dal bar.

# --hints--

Il codice dovrebbe avere un tag di apertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Il codice dovrebbe avere un tag di chiusura `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Non dovresti cambiare l'elemento `header`. Assicurati di non aver eliminato accidentalmente il tag di chiusura.

```js
assert($('header').length === 1);
```

Il tag `main` dovrebbe essere dopo il tag `header`.

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
