---
id: 5f332b23c2045fb843337579
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Dato che l'elemento `p` aggiunto nel passaggio precedente fornisce informazioni supplementari sul bar, annida entrambi gli elementi `h1` e `p` in un elemento `header`.

# --hints--

Dovresti avere un tag di apertura `<header>`.

```js
assert(code.match(/<header>/i));
```

Dovresti avere un tag di chiusura `</header>`.

```js
assert(code.match(/<\/header>/i));
```

L'elemento `h1` dovrebbe essere annidato nell'elemento `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

L'elemento `p` dovrebbe essere annidato nell'elemento `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
<html>
```

