---
id: 5f332a88dc25a0fd25c7687a
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Il nome del bar è `CAMPER CAFE`. Aggiungi un elemento `h1` all'interno dell'elemento `main`. Dagli il nome del bar in lettere maiuscole per farlo risaltare.

# --hints--

Dovresti avere un tag di apertura `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Dovresti avere un tag di chiusura `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Non dovresti cambiare l'elemento `main` esistente.

```js
assert($('main').length === 1);
```

L'elemento `h1` dovrebbe essere annidato nell'elemento `main`.

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

L'elemento `h1` dovrebbe avere il testo `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

