---
id: 5f331e55dfab7a896e53c3a1
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

L'elemento `title` è uno dei vari elementi che forniscono informazioni aggiuntive non visibili sulla pagina web, ma utili per i motori di ricerca e la visualizzazione della pagina.

All'interno dell'elemento `head`, annida un elemento `meta` con `utf-8` come valore dell'attributo `charset` per comunicare al browser come codificare i caratteri della pagina. Nota che gli elementi `meta` sono autoconcludenti.

# --hints--

Il codice dovrebbe avere un tag `meta`.

```js
assert(code.match(/<meta.*?\/?>/is));
```

Il tag `meta` dovrebbe avere un attributo `charset`.

```js
assert(code.match(/<meta charset=/i));
```

L'attributo `charset` dovrebbe avere valore `utf-8`.

```js
assert(code.match(/charset=('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
