---
id: 5f33294a6af5e9188dbdb8f3
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Il nome del bar è `CAMPER CAFE`. Aggiungi un elemento `h1` all'interno dell'elemento `body`. Dagli il nome del bar in lettere maiuscole per farlo risaltare.

# --hints--

Dovresti avere un tag di apertura `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Dovresti avere un tag di chiusura `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Non dovresti cambiare l'elemento `body` esistente.

```js
assert($('body').length === 1);
```

L'elemento `h1` dovrebbe essere annidato nell'elemento `body`.

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

L'elemento `h1` dovrebbe avere il testo `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
