---
id: 5f33071498eb2472b87ddee4
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Come hai imparato negli ultimi step di Cat Photo App, c'è una struttura di base necessaria per iniziare a costruire la tua pagina web.

Aggiungi il tag `<!DOCTYPE html>` e un elemento `html` con un attributo `lang` con il valore `en`.

# --hints--

Dovresti avere la dichiarazione `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Dovresti avere un tag `<html>` di apertura con un attributo `lang` con il valore `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dovresti avere un tag di chiusura `<html>`. Ricorda che i tag di chiusura hanno un `/` dopo la virgoletta angolata di apertura `<`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
