---
id: 5f33294a6af5e9188dbdb8f3
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

È ora di aggiungere qualche contenuto al menu. Aggiungi un elemento `main` all'interno dell'elemento `body` esistente. Alla fine conterrà delle informazioni sui prezzi del caffè e dei dessert offerti dal bar.

# --hints--

Il codice dovrebbe avere un tag di apertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Il codice dovrebbe avere un tag di chiusura `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Non dovresti cambiare l'elemento `body`. Assicurati di non aver eliminato accidentalmente il tag di chiusura.

```js
assert($('body').length === 1);
```

Il tag `main` dovrebbe essere all'interno del tag `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
