---
id: 5f331e55dfab7a896e53c3a1
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

L'elemento `title` è uno dei vari elementi che forniscono informazioni aggiuntive non visibili sulla pagina web, ma utili per i motori di ricerca e la visualizzazione della pagina.

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8` to tell the browser how to encode characters for the page.

# --hints--

Il codice dovrebbe avere un tag `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Il tag `meta` dovrebbe avere un attributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

L'attributo `charset` dovrebbe avere valore `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
