---
id: 5f3477cb2e27333b1ab2b955
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Ora è necessario collegare il file `styles.css` in modo che gli stili vengano applicati di nuovo. Annida un elemento `link` autoconcludente nell'elemento `head`. Dagli un attributo `rel` con valore di `stylesheet` e un attributo `href` con valore di `styles.css`.

# --hints--

Il codice dovrebbe avere un elemento `link`.

```js
// link is removed -> if exists, replaced with style
const link = document.querySelector('head > style');
assert(link);
```

Non dovresti cambiare l'elemento `head` esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
assert($('head').length === 1);
```

L'elemento `link` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<link[\w\W\s]+\/?>/i));
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert(code.match(/rel\s*=\s*('|")stylesheet\1/i));
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert(code.match(/href\s*=\s*('|")styles.css\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
    <main>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
<html>
```

```css
h1, h2, p {
  text-align: center;
}
```
