---
id: 5f344fad8bf01691e71a30eb
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Finora, sei stato limitato per quanto riguarda la presentazione e l'aspetto del contenuto che hai creato. Per curare questa parte, dovresti creare un elemento `style` dentro l'elemento `head`.

# --hints--

Il codice dovrebbe avere un tag di apertura `<style>`.

```js
assert(code.match(/<style\s*>/i));
```

Il codice dovrebbe avere un tag di chiusura `</style>`.

```js
assert(code.match(/<\/style\s*>/));
```

L'elemento `style` dovrebbe essere annidato nell'elemento `head`.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
