---
id: 5f3477ae8466a9a3d2cc953c
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Ora che hai il CSS nel file `styles.css`, rimuovi l'elemento `style` e tutto il suo contenuto. Una volta rimosso, il testo centrato tornerà a sinistra.

# --hints--

Non dovresti avere un elemento `style` nel codice.

```js
assert(!code.match(/style/i));
```

Non dovresti avere alcun selettore CSS nel file HTML.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

