---
id: 5f356ed63c7807a4f1e6d054
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Adesso, l'obiettivo è far sì che `div` non occupi l'intera larghezza della pagina. La proprietà CSS `width` è perfetta per questo.

Puoi utilizzare il selettore `id` per scegliere come target un elemento specifico con un attributo `id`. <dfn>selettore di id</dfn> viene definito posizionando il simbolo del cancelletto `#` direttamente davanti al valore `id` dell'elemento. Ad esempio, se un elemento ha l'ID `cat`, sceglieresti come target quell'elemento in questo modo:

```css
#cat {
  width: 250px;
}
```

Utilizza il selettore `#menu` per dare al tuo elemento una larghezza di `300px`.

# --hints--

Dovresti avere un selettore `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

You should set the `width` property to `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

Il `div` dovrebbe avere una larghezza di 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

