---
id: 5f356ed69db0a491745e2bb6
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

Per applicare lo stile tramite classe all'elemento `div`, rimuovi l'attributo `id` e aggiungi un attributo `class` al tag di apertura dell'elemento `div`. Assicurati di impostare il valore della classe su `menu`.

# --hints--

Il `div` dovrebbe ancora essere renderizzato. Assicurati di aver scritto correttamente il tag `<div>`.

```js
assert($('div').length === 1);
```

L'elemento `div` dovrebbe avere la classe `menu`.

```js
assert($('div').attr('class').includes('menu'));
```

L'elemento `div` non dovrebbe più avere l'id `menu`.

```js
assert(!$('div#menu').length);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

