---
id: 60f0286404aefb0562a4fdf9
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Aggiungi degli elementi `title` e `meta` all'elemento `head`. Dai al tuo progetto il titolo `Registration Form`, e dai un attributo `charset` con il valore `UTF-8` all'elemento `meta`.

# --hints--

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

L'elemento `title` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Il progetto dovrebbe avere il titolo `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Fai attenzione all'uso delle maiuscole/minuscole e all'ortografia nel titolo.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
