---
id: 60f027099a15b00485563dd2
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi i tag di apertura e chiusura `html` sotto la dichiarazione `DOCTYPE` in modo da avere un posto dove iniziare a inserire del codice.

# --hints--

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio del documento HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

L'elemento `html` dovrebbe avere un tag di apertura.

```js
assert(code.match(/<html\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

I tag `html` dovrebbero essere nell'ordine corretto.

```js
assert(code.match(/<html\s*>\s*<\/html\s*>/));
```

Dovresti avere un solo elemento `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
