---
id: 60f02e7361b68405e27b62a5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

All'interno dell'elemento `body`, aggiungi l'intestazione per il contenuto della pagina con un elemento un `h1` con il testo `Registration Form`.

# --hints--

Dovresti aggiungere l'elemento `h1` all'interno del `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Dovresti assegnare all'elemento `h1` il testo `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" type="text/css" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
