---
id: 60f030d388cb74067cf291c3
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Sotto l'intestazione, utilizza il seguente testo all'interno di un elemento paragrafo per permettere agli utenti di registrarsi:

```md
Please fill out this form with the required information
```

# --hints--

Dovresti aggiungere un elemento `p` all'elemento `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Dovresti aggiungere l'elemento `p` al di sotto dell'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Dovresti assegnare all'elemento `p` il testo `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
