---
id: 60f0286404aefb0562a4fdf9
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Aggiungi un elemento `title` all'elemento `head` e dai al tuo progetto il titolo `Registration Form`. Annida un elemento `link` autoconcludente nell'elemento `head`. Dagli un attributo `rel` con valore di `stylesheet` e un attributo `href` con valore di `styles.css`.

# --hints--

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

L'elemento `title` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Il progetto dovrebbe avere il titolo `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Fai attenzione all'uso delle maiuscole/minuscole e all'ortografia nel titolo.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Il codice dovrebbe avere un elemento `link`.

```js
assert.exists(document.querySelector('link'));
```

Non dovresti cambiare i tuoi tag `head` esistenti. Assicurati di non aver eliminato il tag di chiusura.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

L'elemento `link` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<link[\w\W\s]+\/>/i));
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  --fcc-editable-region--
  <head>

  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
