---
id: 60f1922fcbd2410527b3bd89
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Adesso, iniziamo ad abbellire il progetto aggiungendo un po' di CSS. Inizia assegnando al `body` una proprietà `width` del `100%`, e una proprietà `height` di `100vh`.

# --hints--

Dovresti usare un selettore `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Dovresti assegnare al `body` una proprietà `width` del `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Dovresti assegnare al `body` una proprietà `height` di `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" type="text/css" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
