---
id: 60f1922fcbd2410527b3bd89
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--
L'unità `vh` sta per viewport height, ed è relativa all'1% dell'altezza (`height`) del viewport.

È tempo di abbellire il progetto con un po' di CSS. Inizia assegnando al `body` una proprietà `width` del `100%` e una proprietà `height` di `100vh`.

# --hints--

Dovresti usare il selettore di elemento `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Dovresti assegnare al `body` una proprietà `width` del `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Dovresti assegnare al `body` una proprietà `height` di `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
