---
id: 60f5d2776c854e069560fbe6
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Il primo elemento `fieldset` conterrà i campi nome, email e password. Inizia aggiungendo quattro elementi `label` al primo `fieldset`.

# --hints--

Ci dovrebbero essere tre elementi `fieldset` in totale.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

Gli elementi `fieldset` dovrebbero essere tutti figli diretti dell'elemento `form`.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

I quattro elementi `label` dovrebbero essere tutti figli diretti del primo elemento `fieldset`. Assicurati di chiudere gli elementi `label`.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Gli ultimi due elementi `fieldset` dovrebbero essere vuoti.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Assicurati di chiudere gli elementi `label`.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
