---
id: 60fad99e09f9d30c1657e790
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--

Centra l'elemento `form`, assegnandogli una proprietà `margin` con il valore `0 auto`. Quindi, fissa la sua larghezza massima a `500px` e la larghezza minima a `300px`. In questo intervallo, assegnagli una proprietà `width` di `60vw`.

# --hints--

Dovresti usare un selettore `form` per lo stile dell'elemento `form`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('form'));
```

Dovresti assegnare all'elemento `form` una proprietà `margin` con il valore `0 auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('form')?.margin, '0px auto');
```

Dovresti assegnare all'elemento `form` una proprietà `max-width` con il valore `500px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('form')?.maxWidth, '500px');
```

Dovresti assegnare all'elemento `form` una proprietà `min-width` con il valore `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('form')?.minWidth, '300px');
```

Dovresti assegnare all'elemento `form` una proprietà `width` con il valore `60vw`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('form')?.width, '60vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name: <input type="text" name="first-name" required /></label>
        <label>Enter Your Last Name: <input type="text" name="last-name" required /></label>
        <label>Enter Your Email: <input type="email" name="email" required /></label>
        <label>Create a New Password: <input type="password" name="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label>
          <input type="checkbox" name="terms" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
      <fieldset>
        <label>Upload a profile picture: <input type="file" name="file" /></label>
        <label>Input your age (years): <input type="number" name="age" min="13" max="120" />
        </label>
        <label>How did you hear about us?
          <select name="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label>Provide a bio:
          <textarea name="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
  font-family: Tahoma;
  font-size: 16px;
}

h1, p {
  margin: 1em auto;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

label {
  display: block;
  margin: 0.5rem 0;
}

```
