---
id: 62ff8e998d3e7eae14d6ae3b
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

Devi confermare che l'utente ha letto i termini e le condizioni.

Aggiungi un elemento `label` dopo il terzo `fieldset` e un elemento `input` con un attributo`type` del valore di `checkbox` all'interno dell'elemento `label` appena aggiunto. Fai sì che questo elemento `input` sia `required` in modo che gli utenti non possano iscriversi senza leggere i termini e le condizioni di utilizzo.

Non dimenticare di aggiungere un attributo `id` del valore `terms-and-conditions` per l'accessibilità.

# --hints--

Dovresti aggiungere un `label` dopo il terzo elemento `fieldset`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

Dovresti aggiungere un `input` all'elemento `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Dovresti aggiungere un attributo `type` con valore `checkbox` all'elemento `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Dovresti aggiungere un attributo `required` all'elemento `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

L'elemento `input` dovrebbe avere un `id` con il valore `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

L'elemento `label` dovrebbe avere un attributo `for` con il valore `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <label><input type="checkbox" required /></label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
