---
id: 62ff8b9dab5ac88e4d3d43a3
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Per seguire le buone pratiche di accessibilità, collega gli elementi `input` e gli elementi `label` usando l'attributo `for`.

Usa `first-name`, `last-name`, `email` e `new-password` come valori dei rispettivi attributi `id`.

# --hints--

Il primo elemento `input` dovrebbe avere un `id` con il valore `first-name`.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Il secondo elemento `input` dovrebbe avere un `id` con il valore `last-name`.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Il terzo elemento `input` dovrebbe avere un `id` con il valore `email`.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Il quarto elemento `input` dovrebbe avere un `id` con il valore `new-password`.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Il primo elemento `label` dovrebbe avere un attributo `for` con il valore `first-name`.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Il secondo elemento `label` dovrebbe avere un attributo `for` con il valore `last-name`.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Il terzo elemento `label` dovrebbe avere un attributo `for` con il valore `email`.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Il quarto elemento `label` dovrebbe avere un attributo `for` con il valore `new-password`.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
