---
id: 68e420e2d4bc7ae590dc7953
title: Design a Pricing Plans Layout Page
challengeType: 0
dashedName: design-a-pricing-plans-layout-page
demoType: onClick
---

# --description--

In this lab, you will create a responsive pricing plans layout using advanced flexbox properties. This project will teach you how to use the `order`, `flex-grow`, and `align-content` properties to create professional layouts.

**Objective:** Build a pricing plans page with three pricing cards that demonstrate advanced flexbox layout techniques.

**User Stories:**

1. Your page should have an `h1` element with the text `Pricing Plans`.

2. Your page should have three `div` elements with the class `pricing-card` to represent the pricing plans.
   - One of your `div` elements should have the class `pro-plan` in addition to the `pricing-card` class.
   - One of your `div` elements should have the class `basic-plan` in addition to the `pricing-card` class.
   - One of your `div` elements should have the class `premium-plan` in addition to the `pricing-card` class.

3. Each of your `.pricing-card` elements should use Flexbox to stack its content in a column and justify the space between the children using `space-between`. Each element should also set the `flex` property to `0 0 200px` to give it a consistent width and prevent it from growing or shrinking in the layout.

4. The `basic-plan` element should appear first in the layout. You should use the `order` property for this.

5. The `pro-plan` element should appear second in the layout. You should use the `order` property and set its `flex-grow` property to `2` so it takes up more space than the other plans.

6. The `premium-plan` element should come last in the layout. You should use the `order` property for this.

# --hints--

You should have an `h1` element with the text `Pricing Plans`.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
assert.equal(h1.textContent.trim(), 'Pricing Plans');
```

You should have three `div` elements with the class `pricing-card`.

```js
const pricingCards = document.querySelectorAll('div.pricing-card');
assert.equal(pricingCards.length, 3);
```

You should have a `div` element with both `pricing-card` and `pro-plan` classes.

```js
const proCard = document.querySelector('div.pricing-card.pro-plan');
assert.exists(proCard);
```

You should have a `div` element with both `pricing-card` and `basic-plan` classes.

```js
const basicCard = document.querySelector('div.pricing-card.basic-plan');
assert.exists(basicCard);
```

You should have a `div` element with both `pricing-card` and `premium-plan` classes.

```js
const premiumCard = document.querySelector('div.pricing-card.premium-plan');
assert.exists(premiumCard);
```

Each `.pricing-card` element should have `display` set to `flex`.

```js
const pricingCards = document.querySelectorAll('.pricing-card');
pricingCards.forEach(card => {
  const styles = getComputedStyle(card);
  assert.equal(styles.display, 'flex');
});
```

Each `.pricing-card` element should have `flex-direction` set to `column`.

```js
const pricingCards = document.querySelectorAll('.pricing-card');
pricingCards.forEach(card => {
  const styles = getComputedStyle(card);
  assert.equal(styles.flexDirection, 'column');
});
```

Each `.pricing-card` element should have `justify-content` set to `space-between`.

```js
const pricingCards = document.querySelectorAll('.pricing-card');
pricingCards.forEach(card => {
  const styles = getComputedStyle(card);
  assert.equal(styles.justifyContent, 'space-between');
});
```

Each `.pricing-card` element should have the `flex` property set to `0 0 200px`.

```js
const pricingCards = document.querySelectorAll('.pricing-card');
pricingCards.forEach(card => {
  const styles = getComputedStyle(card);
  assert.match(styles.flex, /^0\s+0\s+200px$/);
});
```

The `.basic-plan` element should have `order` set to `0`.

```js
const basicCard = document.querySelector('.basic-plan');
const styles = getComputedStyle(basicCard);
assert.equal(styles.order, '0');
```

The `.pro-plan` element should have `order` set to `1`.

```js
const proCard = document.querySelector('.pro-plan');
const styles = getComputedStyle(proCard);
assert.equal(styles.order, '1');
```

The `.pro-plan` element should have `flex-grow` set to `2`.

```js
const proCard = document.querySelector('.pro-plan');
const styles = getComputedStyle(proCard);
assert.equal(styles.flexGrow, '2');
```

The `.premium-plan` element should have `order` set to `2`.

```js
const premiumCard = document.querySelector('.premium-plan');
const styles = getComputedStyle(premiumCard);
assert.equal(styles.order, '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Pricing Plans Layout</title>
    <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
</body>
</html>
```

```css
/* Add your CSS here */

```
