---
id: 660ef55dd468079679ee0092
title: 步驟 11
challengeType: 1
dashedName: step-11
---

# --description--

你現在可以聲明下一個變量了。 移除 `console.log` 語句和爲 `character` 重新賦值的這兩類語句。

同時移除變量 `secondCharacter`，但是保留 `character` 初始化語句。

# --hints--

現在你的代碼中不應該有任何 `console.log()` 語句。

```js
assert.notMatch(code, /console/);
```

你不應該將 `character` 重新賦值。

```js
assert.isAtMost(code.match(/character/).length, 1);
```

你的變量 `character` 的值應該爲 `"Hello"`。

```js
assert.equal(character, "Hello");
```

You should not have a `secondCharacter` variable.

```js
assert.notMatch(code, /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
