---
id: 660ef19b95d3308e7dd31bb6
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

<dfn>console</dfn>（控制檯指令）允許你打印並查看 JavaScript 的輸出。 你可以使用 `console.log()` 來向控制檯發送信息。 舉個例子，下面這串代碼將會在控制檯打印 `"Naomi"`。

```js
let developer = "Naomi";
console.log(developer);
```

上面的代碼通過在 `console.log()` 中的名稱讀取到了變量 `developer` 。 請注意，括號內的值是將要被打印的內容。

在控制檯打印變量 `character` 的值。 然後點擊 “Console”（控制檯）按鈕來查看 JavaScript 控制檯。

# --hints--

你應該在你的代碼中使用 `console` （控制檯指令）。

```js
assert.match(code, /console/);
```

你應該使用 `console` 的 `log` 屬性。

```js
assert.match(code, /console\.log/);
```

你應該使用括號來調用 `.log()` 方法。

```js
assert.match(code, /console\.log\(/);
```

你應該在控制檯打印變量 `character` 的值。

```js
assert.match(code, /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
