---
id: 660ef31a5be625914a0102cd
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

當一個變量使用關鍵字 `let` 來聲明時，你可以在之後<dfn>重新賦值</dfn>。 在此示例中，`programmer` 的值就從 `"Naomi"` 被重新賦值爲了 `"CamperChan"`。

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Note that when reassigning a variable, you do **not** use the `let` keyword again.

在 `console.log` 語句後，將 `"World"` 賦值給變量 `character`。

# --hints--

你應該在你的代碼中兩次使用 `character`。

```js
assert.lengthOf(code.match(/character/g), 3);
```

你應該使用賦值運算符來重新賦值變量 `character`。

```js
assert.lengthOf(code.match(/character\s*\=/g), 2);
```

你應該在你的代碼中使用字符串 `"World"`。

```js
assert.match(code, /("|')World\1/);
```

你的 `character` 變量的值應該爲 `"World"`。

```js
assert.equal(character, "World");
```

你的重新賦值語句不需要使用 `let`。

```js
assert.isBelow(code.match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
