---
id: 660ef5c1904955978a986a5c
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

使用 `let` 來聲明一個變量 `count`。 給它賦值<dfn>數字類型</dfn> `8`。 當你要使用數字類型作爲值的時候，你不再需要用引號去包裹它。 舉個例子：

```js
let money = 100;
```

# --hints--

你的代碼中應該有第二個 `let` 關鍵字。

```js
assert.lengthOf(code.match(/let/g), 2);
```

你應該用 `let` 來聲明一個變量 `count`。

```js
assert.match(code, /let\s+count/);
```

你應該把數字類型的 `8` 賦值給變量 `count`。

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
