---
id: 660ef0f7c4b8e68ccd1f0786
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript 有七種原始數據類型，`String`（字符串類型）就是其中之一。 在 JavaScript 中，<dfn>字符串</dfn>代表的是一串字符內容，被包裹在單引號 （`'`）或者雙引號（`"`）之中。

注意，字符串是 <dfn>不可變的</dfn>，這意味着一旦被創建，它們就無法被更改。 仍然可以給該變量重新分配另一個值。

試着來使用單引號包裹你的 `"Hello"` 字符串。

# --hints--

你的代碼中不應該使用雙引號。

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

你應該對你的 `"Hello"` 字符串使用單引號。

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

你仍然應該在你的代碼中使用 `let` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

你仍然應該聲明一個 `character` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

你的 `character` 變量的值應該仍是字符串 `"Hello"`。

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
