---
id: 660ef857f2806aa626d29d17
title: Step 17
challengeType: 1
dashedName: step-17
---

# --description--

同樣你也可以使用減法（`-`）、乘法（`*`）以及除法（`/`）。 請用這些運算符和數字在你的 `console.log` 中盡情的做實驗吧。 當你準備好進入下一節時，移除 `console.log`。

# --hints--

你的代碼中不應該有 `console.log()`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你的 `character` 變量的值應該仍是 `"Hello"`。

```js
assert.equal(character, "Hello");
```

你的 `count` 變量的值應該仍是 `8`。

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
