---
id: 660eebd83100d37862268781
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

目前你的 `character` 變量還沒有一個值。 你可以利用<dfn>賦值</dfn>運算符 `=` 來賦值。 舉個例子：

```js
let hello = "Hello";
```

在聲明變量的同時給它賦值，叫做<dfn>初始化</dfn>。

在聲明變量的階段，通過將 `"Hello"` 賦值給 `character` 變量來實現變量的初始化。

# --hints--

你應該使用賦值運算符 `=` 。

```js
assert.match(code, /let\s+character\s*=/);
```

你應該使用字符串 `"Hello"` 。

```js
assert.match(code, /('|")Hello\1/);
```

你應該對你的 `"Hello"` 字符串使用雙引號。

```js
assert.match(code, /"Hello"/);
```

你應該將 `"Hello"` 賦值給變量 `character` 。

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
