---
id: 660ee6e3a242da6bd579de69
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

變量是編程中最重要的概念之一。 一個<dfn>變量</dfn>指向一個用來存儲值的特定內存地址。 變量被賦予一個名稱，可以在整個代碼中用來訪問該值。

聲明一個變量意味着爲它命名。 在 JavaScript 中，通常使用 `let` 關鍵字來實現。 舉個例子，你會這樣聲明一個 `hello` 變量。

```js
let hello;
```

變量的命名遵循特定的規則：名稱能夠包含字母、數字、美元符號以及下劃線，但是不能包含空格，且一定不能以數字開頭。

Use the `let` keyword to declare a variable called `character`.

_注意_ ：通常我們會以分號來結束JavaScript中的語句， `;`

# --hints--

你應該在你的代碼中使用 `let` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

你應該在你的代碼中使用 `character` 。

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

你應該使用 `let` 關鍵字來聲明一個 `character` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Your declaration should end with a semicolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
