---
id: 660ef5105b8ba095307a0e50
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

現在，再次在控制檯打印你的變量 `character`。 你應該在控制檯先看到字符串 `"Hello"`，然後是 `"World"`。

# --hints--

你應該再次使用 `console.log` 指令。

```js
assert.lengthOf(code.match(/console\.log/g), 2);
```

你應該再次打印 `character` 到控制檯。

```js
assert.lengthOf(code.match(/console\.log\(\s*character\s*\)/g), 2);
```


新的 `console.log()` 指令應該在你重新賦值之後。

```js
const reassign = code.split(/\n+/).findIndex(l => l.match(/character\s+=\s+("|')World\1/));
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
