---
id: 660ef6355e8f5a9e67fe5f46
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

通過 `number` 數據類型，你可以進行數學運算，比如加法運算。 試着在控制檯打印 `count + 1`。

# --hints--

你應該在你的代碼中使用 `console`。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

你應該使用 `console` 的 `.log()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你應該打印 `count + 1` 到控制檯。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
