---
id: 660f4de78f775e480ba2e451
title: Step 107
challengeType: 1
dashedName: step-107
---

# --description--

數組的 `.unshift()` 方法允許你將值添加到數組的 **開頭**，這與 `.push()` 不同，後者將值添加到數組的末尾。 `.unshift()` returns the new length of the array it was called on.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

使用 `const` 聲明一個 `unshifted` 變量，並將在 `numbers` 數組上調用 `.unshift()` 的結果分配給它。 傳遞 `5` 作爲參數。 然後打印你的 `unshifted` 變量。

# --hints--

你應該聲明你的 `unshifted` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

你應該在 `numbers` 數組上調用 `.unshift()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

你應該將 `5` 作爲參數傳遞給 `.unshift()` 調用。

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

你應該將 `.unshift()` 調用的結果分配給 `unshifted` 變量。

```js
assert.equal(unshifted, 4);
```

You should log your `unshifted` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
