---
id: 660f38c34a4de6141c0c369f
title: Step 71
challengeType: 1
dashedName: step-71
---

# --description--

你的金字塔生成器現在可以按預期運行。 但這是一個進一步探索你所編寫的代碼的絕佳機會。

加法運算符不是向變量添加值的唯一方法。 <dfn>加法賦值</dfn>運算符可以簡寫爲“取變量的原始值，添加該值，然後將結果賦給變量”。 例如，這兩個語句將產生相同的結果：

```js
test = test + 1;
test += 1;
```

更新 `for` 循環中的迭代器語句以使用加法賦值。

# --hints--

你的 `for` 循環不應該使用 `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

你的 `for` 循環應該使用帶有 `i` 的加法賦值。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
