---
id: 660f1e3f047bf4e403268713
title: Step 39
challengeType: 1
dashedName: step-39
---

# --description--

不幸的是，現在你無法看到循環正在做什麼。

使用 `let` 聲明一個 `result` 變量，併爲其分配一個空字符串。 空字符串用引號表示，引號之間沒有任何內容，例如 `""`。

# --hints--

你應該聲明一個 `result` 變量。

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

你應該使用 `let` 來聲明你的 `result` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s*result/);
```

你的 `result` 變量應該是一個空字符串。

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
