---
id: 660f4a1472f8e63d76162ce5
title: Step 99
challengeType: 1
dashedName: step-99
---

# --description--

如果將金字塔反過來，或者說 <dfn>倒置</dfn>，會怎樣？ 是時候嘗試一下了！

首先創建一個新的 `for` 循環。 聲明你的迭代器 `i` 併爲其分配 `count` 的值，然後使用布爾值 `false` 作爲條件和迭代語句。

# --hints--

你的代碼應該有一個 `for` 循環。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(/);
```

你的 `for` 循環應該使用 `count` 的值初始化 `i`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count/);
```

你的 `for` 循環應使用 `false` 作爲條件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false/);
```


你的 `for` 循環應使用 `false` 作爲迭代。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
