---
id: 660f039ff313dbb696b007ca
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

當數組包含值或 <dfn>元素</dfn> 時，這些值以逗號分隔。 這是一個包含兩個字符串的數組：

```js
let array = ["first", "second"];
```

將 `rows` 聲明更改爲包含字符串 `"Naomi"`、`"Quincy"` 和 `"CamperChan"` 的數組。 數組中值的順序很重要，因此請遵循該順序。 請記住，字符串區分大小寫。

# --hints--

數組中的第一個元素應該是字符串 `"Naomi"`。

```js
assert.equal(rows[0], "Naomi");
```

數組中的第二個元素應該是字符串 `"Quincy"`。

```js
assert.equal(rows[1], "Quincy");
```

數組中的第三個元素應該是字符串 `"CamperChan"`。

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
