---
id: 660f0c34aad72dc712b97624
title: Step 27
challengeType: 1
dashedName: step-27
---

# --description--

在這個項目中另一個必不可少的方法是 `.pop()`。 這個方法會將一個數組中的最後一個元素移除，並且 <dfn>returns</dfn>（返回）該元素。

當一個方法返回一個值時，你可以把它看作是把這個值還給你，讓你能在代碼的其他部分使用它。

創建一個新變量叫作 `popped` ，並將 `rows.pop()` 的結果分配給它。 然後，在控制檯記錄 `popped`。

# --hints--

你應該聲明一個名爲 `popped` 的變量。

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

你應該用關鍵字 `let` 來聲明你的 `popped` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

你應該調用 `.pop()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

你應該在你的數組 `rows` 上，調用 `.pop()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

你應該記錄你的變量 `popped`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
