---
id: 660f4934fb48f63abd5ae371
title: Step 95
challengeType: 1
dashedName: step-95
---

# --description--

使用 `done` 來跟蹤已生成的行數是可行的，但實際上你可以進一步清理邏輯。

數組具有特殊的 `length` 屬性，可讓你查看數組中有多少個值或 <dfn>元素</dfn>。 你可以使用類似 `myArray.length` 的語法訪問此屬性。

更新你的條件以檢查 `rows.length` 是否小於或等於 `count`。

# --hints--

你的 `while` 循環應檢查 `rows.length` 是否小於或等於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
