---
id: 660f34e99571070d56d2f231
title: Step 67
challengeType: 1
dashedName: step-67
---

# --description--

你的 `padRow` 函數有兩個你定義的參數。 當調用函數時，會爲這些參數提供值。

你向函數調用提供的值叫作<dfn>參數</dfn>，並且你將<dfn>傳遞</dfn>參數給函數調用。 以下是將 `"Hello"` 作爲參數傳遞的函數調用：

```js
test("Hello");
```

將 `i + 1` 和 `count` 作爲參數傳遞給 `padRow` 調用。 與參數一樣，傳遞的參數也用逗號分隔。

# --hints--

你應該將 `i + 1` 傳遞給 `padRow()` 調用。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

`i + 1` 參數後面應該有一個逗號。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

你應該傳遞 `count` 作爲第二個參數。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
