---
id: 660f53ad3d39175c5d4335ac
title: Step 117
challengeType: 1
dashedName: step-117
---

# --description--

你的金字塔生成器現已完成，其功能比你最初計劃的還要多！ 下一步是清理你的代碼。

從代碼中刪除所有註釋，包括單行和多行註釋。

# --hints--

你的代碼中不應該有任何單行註釋。

```js
assert.notMatch(__helpers.removeJSComments(code), /\/\//);
```

你的代碼中不應該有任何多行註釋。

```js
assert.notMatch(__helpers.removeJSComments(code), /(?:\\\*|\/\*)/);
```

你的代碼中不應該有任何註釋。

```js
assert.equal(code, __helpers.removeJSComments(code));
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
