---
id: 660f2fbd45b520046cac68e8
title: Step 65
challengeType: 1
dashedName: step-65
---

# --description--

記住在之前的步驟中，你瞭解了返回值。 函數可以<dfn>返回</dfn>一個值，供你的應用程序單獨使用。

在函數中，`return` 關鍵字用於指定返回值。 例如，此函數將返回賦予第一個參數的值：

```js
function name(parameter) {
  return parameter;
}
```

使用 `return` 關鍵字返回 `character` 變量的值，重複 `rowNumber` 次。

# --hints--

你應該使用 `.repeat()` 方法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

你應該在 `character` 變量上使用 `.repeat()` 方法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

你應該將 `rowNumber` 傳遞給 `.repeat()` 調用。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

你應該使用 `return` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

你應該返回 `.repeat()` 調用的結果。

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
