---
id: 660f2b6fd54ac1fc142804dd
title: Step 47
challengeType: 1
dashedName: step-47
---

# --description--

格式化這個金字塔的邏輯可能會變得複雜，這意味着現在是將該代碼提取到函數中的好時機。

<dfn>函數</dfn> 是可以在整個應用程序中重複使用的代碼塊。 函數使用以下語法聲明：

```js
function name(parameter) {

}
```

`function` 關鍵字告訴 JavaScript `name` 變量將是一個函數。 `parameter` 是一個變量，表示在使用函數時傳遞給函數的值。 一個函數可以具有任意數量的 <dfn>參數</dfn>。 類似於 `for` 循環，花括號之間的空間是 <dfn>函數主體</dfn>。

聲明一個 `padRow` 函數。 暫時不要創建任何參數變量。 函數體應該是空的。 請記住，你需要使用駝峯式命名約定。

# --hints--

你應該使用 `function` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

你應該聲明一個 `padRow` 函數。

```js
assert.isFunction(padRow);
```

你的 `padRow()` 函數不應有任何參數。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

你的 `padRow()` 函數應該有一個空的主體。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
