---
id: 660f17294346b7d69e79db3d
title: Step 32
challengeType: 1
dashedName: step-32
---

# --description--

你現在可以開始構建金字塔生成器了。 你的 `character` 變量將作爲金字塔的構建塊。

`"Hello"` 可能不太適合這種情況。 將 `character` 的值更改爲井號 `"#"`。

# --hints--

你的 `character` 變量應該是一個井號。

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
