---
id: 660f1cedf3676fe26122ebf6
title: Step 38
challengeType: 1
dashedName: step-38
---

# --description--

你應該會看到控制檯中打印出數字 0 到 7，每行一個。 這將作爲生成金字塔的基礎。

將你的日誌語句替換爲把 `i` 推送到 `rows` 數組的語句。

# --hints--

你不應該有 `console.log` 調用。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你應該在 `rows` 數組上調用 `.push()`。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

你應該將 `i` 推送到你的 `rows` 數組。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

你的 `.push()` 應該發生在你的 `for` 循環中。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
