---
id: 660f3dd626be3a1ffe27e5d1
title: Step 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript 還支持多行註釋。 多行註釋以 `/*` 開頭，以 `*/` 結尾。

與單行註釋不同，多行註釋將封裝多行。

使用 `/*` 和 `*/` 將當前的 `for` 循環（包括主體）轉換爲多行註釋。

# --hints--

你應該以 `/*` 開始多行註釋。

```js
assert.match(code, /\/\*/);
```

你應該以 `*/` 結束多行註釋。

```js
assert.match(code, /\*\//);
```

你的整個 `for` 循環應該被註釋掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
