---
id: 660f4377a359972c521d3f4b
title: Step 79
challengeType: 1
dashedName: step-79
---

# --description--

你將看到控制檯中打印的字符串，因爲 `true` 實際上是真的。

將 `if` 語句的條件更改爲布爾值 `false`。

# --hints--

你的 `if` 條件應該以 `false` 作爲條件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
