---
id: 660f447efc0e722f016c1be0
title: Step 81
challengeType: 1
dashedName: step-81
---

# --description--

文字又出現了！ 這是因爲 `"false"` 是一個字符串，當將其計算爲布爾值時變爲 `true`。 這意味着 `"false"` 是一個真值。

<dfn>真值</dfn>是作爲布爾值評估時被視爲真的值。 你在 JavaScript 中遇到的大多數值都是真值。

<dfn>假值</dfn>則相反 - 當作爲布爾值評估時被視爲假的值。 JavaScript 有一個定義的假值列表。 其中包括 `false`、`0`、`""`、`null`、`undefined` 和 `NaN`。

嘗試將你的 `if` 條件更改爲空字符串 `""`，這是一個假值。

# --hints--

你的 `if` 語句應以 `""` 作爲條件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
