---
id: 660f3915b41a441537ec9f5e
title: Step 72
challengeType: 1
dashedName: step-72
---

# --description--

因爲你僅將 `i` 增加 `1`，所以你可以使用 <dfn>自增運算符</dfn> `++`。 該運算符將變量的值增加 1，更新該變量的賦值。 例如，`test` 在這裏將變成 `8`：

```js
let test = 7;
test++;
```

將循環迭代的加法賦值替換爲自增運算符。

# --hints--

你的 `for` 循環不應使用帶有 `i` 的加法賦值。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

你的 `for` 循環應該在 `i` 上使用自增運算符。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
