---
id: 660f4a83373de83ea101685f
title: Step 100
challengeType: 1
dashedName: step-100
---

# --description--

因爲你要朝相反方向循環，所以你的循環需要在 `i` 大於 `0` 時運行。 你可以爲此使用 <dfn>大於</dfn> 運算符 `>`。

將循環條件設置爲當 `i` 大於 `0` 時運行。

# --hints--

當 `i` 大於 `0` 時，你的 `for` 循環應該運行。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
