---
id: 660f415b76859a2736771607
title: Step 78
challengeType: 1
dashedName: step-78
---

# --description--

你的金字塔又消失了。 沒關係——這是意料之中的事。

在創建新循環之前，你需要了解 `if` 語句。 <dfn>`if` 語句</dfn>允許你僅在滿足條件時運行代碼塊。 它們使用以下語法：

```js
if (condition) {
  logic
}
```

創建一個 `if` 語句，以布爾值 `true` 作爲條件。 在主體中，打印字符串 `"Condition is true"`。

# --hints--

你應該創建一個 `if` 語句。

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

你的 `if` 語句應以 `true` 作爲條件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

你的 `if` 主體應記錄 `"Condition is true"`。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
