---
id: 660f505d02b2bd513a1c3468
title: Step 111
challengeType: 1
dashedName: step-111
---

# --description--

你的金字塔不再是倒置的。 這是因爲你正在向數組的 **末尾** 添加新行。

更新循環體以將新行添加到數組的開頭。

# --hints--

你應該使用 `rows` 的 `unshift` 方法。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

你應該將 `padRow()` 調用作爲 `.unshift()` 方法的參數傳遞。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

你應該將 `i` 作爲第一個參數傳遞給 `padRow()` 調用。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

你應該將 `count` 作爲第二個參數傳遞給 `padRow()` 調用。

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
