---
id: 660f09a2694b59c3a10ee304
title: Step 23
challengeType: 1
dashedName: step-23
---

# --description--

現在，刪除你的第一個控制檯日誌和 `rows[rows.length - 1]` 賦值。 將第二個 `rows` 日誌語句留待以後使用。

# --hints--

你應該刪除 `console.log(rows[0])` 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

你應該刪除 `rows[rows.length - 1]` 重新賦值。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

你不應刪除 `console.log(rows)` 語句。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
