---
id: 660f4990b1caa03b9dc97a43
title: Step 97
challengeType: 1
dashedName: step-97
---

# --description--

現在你不再需要 `done` 變量。 從循環中刪除自增運算以及 `done` 的變量聲明。

# --hints--

你不應該自增 `done` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /done\+\+/);
```

你應該不再有 `done` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /done/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length < count) {
  done++;
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
