---
id: 660f0a55847d6cc485f29ba5
title: Step 26
challengeType: 1
dashedName: step-26
---

# --description--

JavaScript 中的 <dfn>方法</dfn> 是一個與特定值或特定對象相關聯的函數。 你已經遇到的一個例子是 `.log()` 方法，它是 `console` 對象的一部分。

數組也有自己的方法，你將首先探索的是 `.push()` 方法。 這個方法允許你去“推入”一個值到一個數組的末尾。 這裏有一個將數字 `12` 添加到數組末尾的例子。

```js
array.push(12);
```

使用 `.push()` 添加字符串 `"freeCodeCamp"` 到你的 `rows` 數組末尾。 在 `console.log` 前添加此代碼，以便你可以看到你對數組所做的更改。

# --hints--

你應該在你的代碼中使用 `.push()`。

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

你應該對你的 `rows` 數組使用 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

你應該將字符串 `"freeCodeCamp"` 傳遞給 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
