---
id: 660f07d231941bc11719f664
title: Step 21
challengeType: 1
dashedName: step-21
---

# --description--

數組的特殊之處在於它們被認爲是<dfn>可變的</dfn>。 這意味着你可以直接更改索引處的值。

例如，此代碼將數字 `25` 分配給數組中的第二個元素：

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

將 `rows` 數組的 **第三** 元素更新爲數字 `10`。 然後將 `rows` 數組打印到控制檯。

# --hints--

你應該再次在 `rows` 數組上使用括號表示法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

你應該訪問 `rows` 數組的第三個元素。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

你應該對 `rows` 數組的第三個元素使用賦值運算符。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

你應該將值 `10` 分配給 `rows` 數組的第三個元素。

```js
assert.equal(rows[2], 10);
```

你的代碼中應該有第二個 `console.log` 語句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你應該記錄 `rows` 數組。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
