---
id: 660f383d4c772c12ff59904b
title: Step 70
challengeType: 1
dashedName: step-70
---

# --description--

你可以將完整的表達式作爲參數傳遞。 該函數將接收該表達式的求值結果。 例如，這兩個函數調用將產生相同的結果：

```js
test(2 * 3 + 1);
test(7);
```

再次查看模式：

```md
··#··
·###·
#####
```

將 `character` 值更新爲重複 `2 * rowNumber - 1` 次。

再次打開控制檯以查看更新的結果。

# --hints--

你的函數應該將 `2 * rowNumber - 1` 傳遞給你的 `.repeat()` 調用。

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
