---
id: 660f255022991ef34ed0ee88
title: Step 44
challengeType: 1
dashedName: step-44
---

# --description--

打印數字不會產生具有視覺吸引力的金字塔。 現在你正在輸出 `rows` 數組的格式化內容，是時候更新原始循環了。

不要將 `i` 加到數組，而是添加 `character` 變量的值。

# --hints--

你不應再添加 `i` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

你應該添加你的 `character` 變量。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
