---
id: 660f3ce51f70571e1c5227c8
title: Step 76
challengeType: 1
dashedName: step-76
---

# --description--

註釋有助於解釋爲什麼你的代碼採用某種方法，或者爲你未來的自己留下待辦事項。

在 JavaScript 中，你可以使用 `//` 在代碼中留下單行註釋。

在函數上方添加單行註釋，提醒自己將代碼更改爲不同類型的循環。

# --hints--

你應該以 `//` 開始單行註釋。

```js
assert.match(code, /\/\//);
```

你的單行註釋至少應有五個字符長。

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
