---
id: 660f033cf051ebb50ea3bf48
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

在編程中，你經常需要處理大量數據。 有許多數據結構可以幫助你組織和管理數據。 最基本的數據結構之一是<dfn>數組</dfn>。

<dfn>數組</dfn>是一種非原始數據類型，可以保存一系列值。 非原始數據類型與原始數據類型的不同之處在於前者可以保存更復雜的數據。 字符串和數字等原始數據類型一次只能保存一個值。

數組使用方括號 (`[]`) 表示。 下面是一個值爲空數組的變量的示例：

```js
let array = [];
```

聲明一個 `rows` 變量併爲其分配一個空數組。

# --hints--

你應該有一個 `rows` 變量。

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

你應該使用 `let` 來聲明你的 `rows` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

你應該爲 `rows` 變量分配一個空數組。

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
