---
id: 660f359af3e32e0f1a6880b7
title: Step 68
challengeType: 1
dashedName: step-68
---

# --description--

現在你應該在控制檯中看到同一組字符。 你的 `padRow` 函數正在執行與你之前完全相同的操作，但現在它位於其自己的可重複使用部分中。

使用加法運算符將單個空格 `" "` 連接到重複 `character` 字符串的開頭和結尾。

請記住，你可以使用 `+` 運算符來連接字符串，如下所示：

```js
" " + "string"
```

# --hints--

你應該將一個空格連接到返回值的開頭。

```js
assert.match(padRow(1, 1), /^\s/);
```

你應該將一個空格連接到返回值的末尾。

```js
assert.match(padRow(1, 1), /\s$/);
```

你的 `padRow()` 函數應返回重複的 `character` 系列，並且在系列前後都有空格。

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
