---
id: 660f18f059fe0fda192ce394
title: Step 34
challengeType: 1
dashedName: step-34
---

# --description--

你的循環現在需要一個合適的迭代器。 <dfn>迭代器</dfn> 是一個變量，你可以在 `for` 循環中專門聲明它，以控制循環如何迭代或執行邏輯。

在循環中使用 `i` 作爲迭代器變量是一種常見的慣例。 `for` 循環允許你在括號 `()` 中聲明這一點。 例如，這裏有一個 `for` 循環，它聲明一個 `index` 變量併爲其賦值 `100`。

```js
for (let index = 100; "second"; "third") {

}
```

將字符串 `"iterator"` 替換爲用 `let` 聲明變量 `i`。 爲其分配值 `0` 以開始。 這將在循環**第一次**運行時賦予 `i` 變量值 `0`。

# --hints--

你應該使用 `let` 來聲明 `i` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

你應該將 `0` 分配給你的 `i` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

你的 `for` 循環應從 `0` 處啓動 `i` 迭代器。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
