---
id: 660f2a70ad6225fa503e71c3
title: Step 46
challengeType: 1
dashedName: step-46
---

# --description--

你越來越近了！ 此時，你會遇到所謂的 <dfn>差一錯誤</dfn>，這是 JavaScript 等基於零的索引語言中經常出現的問題。

`rows` 數組的第一個索引是 `0`，這就是爲什麼你用 `i = 0` 啓動 `for` 循環。 但重複字符串零次將導致不打印任何內容。

爲了解決此問題，請在 `.repeat()` 調用中給 `i` 的值添加 `1`。 不要像在循環條件中那樣將其分配回 `i`。

# --hints--

你應該在 `.repeat()` 方法中給 `i` 添加 `1`。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
