---
id: 660f4774e3e0df35a68bb5f2
title: Step 90
challengeType: 1
dashedName: step-90
---

# --description--

爲了再次生成金字塔，請將以 `done` 和 `count` 作爲參數調用 `padRow` 的結果推送到 `rows` 數組，類似於你在第一個循環中所做的操作。

# --hints--

你的循環應該在 `rows` 上調用 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

你應該在 `.push()` 方法中調用 `padRow` 函數。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

你應該將 `done` 作爲第一個參數傳遞給 `padRow` 調用。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

你應該將 `count` 作爲第二個參數傳遞給 `padRow` 調用。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
