---
id: 660f20473aef47e9b8c9afc6
title: Step 40
challengeType: 1
dashedName: step-40
---

# --description--

添加 log 語句來打印 `result` 的值。 根據你使用的控制檯，你可能看不到任何打印內容。

# --hints--

你應該添加一個 `console.log` 調用。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

你應該記錄你的 `result` 變量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
