---
id: 660f229d2dbe09ef2954a4a1
title: Step 42
challengeType: 1
dashedName: step-42
---

# --description--

回想一下，在之前的循環中，你使用了加法運算符 `+` 將 `i` 的值增加 `1`。

你可以對字符串值執行類似的操作，即將新字符串附加到現有字符串。 例如，`hello = hello + " World";` 會將字符串 `" World"` 添加到 `hello` 變量中存儲的現有字符串中。 這被稱爲<dfn>連接</dfn>。

在 `for...of` 循環中，使用加法運算符將 `row` 值連接到 `result` 值。

# --hints--

你應該在 `result` 變量上使用連接運算符。

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

你應該將 `row` 連接到 `result` 變量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

你應該將連接的結果分配回 `result` 變量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
