---
id: 660f3ba3cceef11b6ba08b59
title: Step 75
challengeType: 1
dashedName: step-75
---

# --description--

不幸的是，現在金字塔的底層已經消失了。 這是因爲你又創建了另一個 <dfn>off-by-one 錯誤</dfn>。

你的原始循環適用於 `i` 值從 `0` 到 `7`，因爲 `count` 爲 `8` 並且你的條件要求 `i` 小於 `count`。 你的循環現在正在運行 `i` 個值，從 `1` 到 `7`。

當 `i` 爲 `8` 時，你的循環也需要更新才能運行。 查看你的邏輯，這意味着當 `i` <dfn>小於或等於</dfn> `count` 時，你的循環應該運行。 你可以爲此使用小於或等於運算符 `<=`。

更新你的循環條件，使其在 `i` 小於或等於 `count` 時運行。

# --hints--

你的 `for` 循環不應檢查 `i` 是否小於 `count`。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

你的 `for` 循環應檢查 `i` 是否小於或等於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
