---
id: 660f0ee51d7460ce88cd248d
title: Step 29
challengeType: 1
dashedName: step-29
---

# --description--

你想到控制檯裏打印的會是 `4` 嗎？ 在你給數組添加一個值後，`.push()` 方法返回這個數組的新的長度。

瞭解方法返回的值很重要。 花些時間嘗試使用 `.push()` 和 `.pop()`。 準備就緒後，刪除所有 `.push()` 和 `.pop()` 調用以及 `console.log` 語句。

# --hints--

你不應該有 `.push()` 調用。

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

你不應該有 `.pop()` 調用。

```js
assert.notMatch(__helpers.removeJSComments(code), /.pull\(/);
```

你不應該有任何的 log 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

你不應該有 `popped` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

你不應該有 `pushed` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
