---
id: 660f207334fabaeac3269c38
title: Step 41
challengeType: 1
dashedName: step-41
---

# --description--

要操作 `result` 字符串，你將使用不同類型的循環。 具體來說，`for...of` 循環遍歷可迭代對象中的每個項目並將其臨時分配給變量。

`for...of` 循環的語法如下：

```js
for (const value of iterable) {

}
```

請注意，你可以使用 `const`，因爲該變量僅在一次迭代中存在，而不是在整個循環期間存在。

創建一個 `for...of` 循環來遍歷 `rows` 數組，並將每個值分配給 `row` 變量。

# --hints--

你應該使用另一個 `for` 關鍵字。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

你應該聲明一個 `row` 變量。

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

你應該使用 `const` 來聲明你的 `row` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

你的 `for...of` 循環應該聲明你的 `row` 變量。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

你應該使用 `of` 關鍵字從 `rows` 中提取 `row` 變量。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

你的 `for...of` 循環體應該是空的。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
