---
id: 660f23b53db70af0f2620e78
title: Step 43
challengeType: 1
dashedName: step-43
---

# --description--

現在，你的所有數字都出現在同一行上。 這對於創建金字塔不起作用。

你需要在每一行中添加一個新行。 但是，在 JavaScript 中按回車鍵在引號之間插入換行符將導致解析錯誤。 相反，你需要使用特殊的 <dfn>轉義序列</dfn> `\n`，它在記錄字符串時被解釋爲新行。 例如：

```js
lineOne = lineOne + "\n" + lineTwo;
```

使用第二個加法運算符在現有的 `result` 值和添加的 `row` 值之間連接一個新行。

# --hints--

你應該使用 `\n` 轉義序列。 請記住，它需要是一個字符串，因此用引號括起來。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

你應該將 `\n` 轉義序列連接到 `result` 變量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1/);
```

你應該將 `row` 變量連接到 `\n` 轉義序列。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1\s*\+\s*row/);
```

你應該將整個連接分配回你的 `result` 變量。 Don't forget your semicolon.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*result\s*\+\s*('|")\\n\1\s*\+\s*row;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
