---
id: 660f061d259bbebc37461080
title: Step 20
challengeType: 1
dashedName: step-20
---

# --description--

你可以使用值的 <dfn>index</dfn> 訪問數組內的值。 索引是一個數字，表示值在數組中的位置，從 `0` 開始表示第一個值。

你可以使用<dfn>括號表示法</dfn>訪問該值，例如 `array[0]`。

使用 `console.log` 和括號表示法打印 `rows` 數組中的第一個值。

# --hints--

你的代碼中應該有一個 `console.log()` 語句。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你應該訪問你的 `rows` 數組。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

你應該在 `rows` 數組中使用括號表示法。

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

你應該使用括號表示法來訪問 `rows` 數組的第一個元素。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

你應該記錄 `rows` 數組的第一個元素。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
