---
id: 660f4cffb1459d45e34902d1
title: Step 106
challengeType: 1
dashedName: step-106
---

# --description--

你實際上可以構建倒金字塔，而不需要像以前那樣“向後”循環。

爲此，你需要學習一些新的數組方法。 首先使用 `const` 聲明一個 `numbers` 變量。 爲其分配一個包含元素 `1`、`2` 和 `3` 的數組。 然後記錄 `numbers` 數組。

# --hints--

你應該使用 `const` 來聲明 `numbers` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

你的 `numbers` 變量應該是一個數組。

```js
assert.isArray(numbers);
```

你的 `numbers` 數組應按順序包含元素 `1`、`2` 和 `3`。

```js
assert.deepEqual(numbers, [1,2,3]);
```

你應該記錄你的 `numbers` 數組。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
