---
id: 660f39b444fd6f16d1e49c1f
title: Step 73
challengeType: 1
dashedName: step-73
---

# --description--

你不必將 `i + 1` 傳遞給 `padRow` 調用，而是可以從 `1` 開始循環。 這將允許你創建一個單索引循環。

更新你的迭代器以從 `1` 開始，而不是從 `0` 開始。

# --hints--

你的 `for` 循環應該將 `i` 初始化爲 `1`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
