---
id: 660f207334fabaeac3269c38
title: Step 41
challengeType: 1
dashedName: step-41
---

# --description--

要操作 `result` 字符串，你將使用不同類型的循環。 具體來說，`for...of` 循環遍歷可迭代對象中的每個項目並將其臨時分配給變量。

`for...of` 循環的語法如下：

```js
for (const value of iterable) {

}
```

請注意，你可以使用 `const`，因爲該變量僅在一次迭代中存在，而不是在整個循環期間存在。

創建一個 `for...of` 循環來遍歷 `rows` 數組，並將每個值分配給 `row` 變量。

# --hints--

You should use another `for` keyword.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

You should declare a `row` variable.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

You should use `const` to declare your `row` variable.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Your `for...of` loop should declare your `row` variable.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Your `row` variable should be extracted from `rows` using the `of` keyword.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Your `for...of` loop body should be empty.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
