---
id: 660f540c2176ea5dec01306d
title: Step 118
challengeType: 1
dashedName: step-118
---

# --description--

做得好！ 嘗試使用 `character`、`count` 和 `inverted` 變量的不同值。

當你準備好繼續下一個項目時，將 `character` 設置爲 `"!"`、將 `count` 設置爲 `10`，並將 `inverted` 設置爲 `false` 以繼續。

恭喜你完成你的第一個 JavaScript 項目！

# --hints--

你應該將 `character` 設置爲 `"!"`。

```js
assert.equal(character, "!");
```

你應該將 `count` 設置爲 `10`。

```js
assert.equal(count, 10);
```

你應該將 `inverted` 設置爲 `false`。

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
