---
id: 660f48e1d3682f39e81843c4
title: Step 94
challengeType: 1
dashedName: step-94
---

# --description--

你的金字塔生成器仍在運轉。 然而，它可能會再次陷入無限循環。

因爲你僅檢查 `done` 是否不等於 `count`，如果 `done` **大於** `count`，你的循環將永遠持續下去。

更新循環的條件來檢查 `done` 是否小於或等於 `count`。

# --hints--

你的 `while` 循環應檢查 `done` 是否小於或等於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
