---
id: 660f2eccfe3f820304af1b39
title: Step 64
challengeType: 1
dashedName: step-64
---

# --description--

要知道如何格式化一行，你的 `padRow` 函數需要知道你所在的行號，以及總共生成了多少行。

最好的方法是給它們創建函數參數。 爲 `padRow` 函數提供 `rowNumber` 和 `rowCount` 參數。 多個參數以逗號分隔：

```js
function name(first, second) {

}
```

# --hints--

你的 `padRow` 函數應該有一個 `rowNumber` 參數。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

你應該在 `rowNumber` 參數後添加一個逗號。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

你的 `padRow` 函數應該有一個 `rowCount` 參數。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
