---
id: 660f280dda5040f707c76b4a
title: Step 45
challengeType: 1
dashedName: step-45
---

# --description--

現在你有一系列 `#` 字符，但仍然缺少金字塔形狀。 幸運的是，`i` 變量代表循環中的當前“行”號，使你能夠使用它來製作類似金字塔的結構。

爲了實現這一點，你將使用字符串可用的 `.repeat()` 方法。 該方法接受一個數字作爲參數，指定重複目標字符串的次數。 例如使用 `.repeat()` 生成字符串 `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

在你的 `character` 上使用 `.repeat()` 方法，併爲其提供 `i` 作爲數字。

# --hints--

你應該使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

你應該在 `character` 變量上使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

你應該將 `i` 傳遞給你的 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
