---
id: 660f34626216270c682e2f7b
title: Step 66
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>函數調用</dfn>允許你實際使用函數。 你可能沒有意識到，但你一直在使用的 `.push()` 之類的方法一直是函數調用。

通過引用函數的名稱並添加 `()` 來調用函數。 調用 `test` 函數的方法如下：

```js
test();
```

將 `.push()` 調用中的 `character.repeat(i + 1)` 替換爲 `padRow` 函數的函數調用。

# --hints--

你不應該在 `push` 調用中使用 `i + 1`。

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

你不應該在 `.push()` 調用中使用 `character.repeat`。

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

你應該在 `.push()` 調用中調用 `padRow`。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*\)\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
