---
id: 660f4f79e2a82a4e92290f44
title: Step 110
challengeType: 1
dashedName: step-110
---

# --description--

有時你可能希望恢復以前註釋掉的代碼。 你可以通過刪除該代碼周圍的 `/*` 和 `*/` 來實現此目的。 這叫作<dfn>取消註釋</dfn>。

僅取消註釋第一個 `for` 循環。 保留單行註釋和另外兩個多行註釋。

# --hints--

你應該取消註釋第一個 `for` 循環。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

你不應該刪除你的單行註釋。

```js
assert.match(code, /\/\/\sTODO:/);
```

你不應取消註釋 `while` 循環。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

你不應取消註釋第二個 `for` 循環。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
