---
id: 660f46460f9c36330ebc07d8
title: Step 88
challengeType: 1
dashedName: step-88
---

# --description--

相等運算符可能會導致 JavaScript 中出現一些奇怪的行爲。 例如，`"0" == 0` 爲真，即使一個是字符串，一個是數字。

<dfn>嚴格相等</dfn>運算符 `===` 用於檢查兩個值是否相等且具有相同的類型。 一般來說，這是你應該始終使用的相等運算符。 使用嚴格相等運算符時，`"0" === 0` 變爲假，因爲雖然它們可能具有相同的零值，但它們不是同一類型。

更新你的 `done == count` 條件以使用嚴格相等運算符。

# --hints--

你的 `if` 條件應使用嚴格相等。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
