---
id: 660f455b044d3230ed971e98
title: Step 85
challengeType: 1
dashedName: step-85
---

# --description--

<dfn>`while`</dfn> 循環將不斷運行，直到指定的 `condition` 不再成立。 它的語法如下：

```js
while (condition) {
  logic;
}
```

使用該語法聲明一個以 `continueLoop` 爲條件的 `while` 循環。 主體應該是空的。

# --hints--

你應該使用 `while` 循環。

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

你的 `while` 循環應使用 `continueLoop` 作爲條件。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
