---
id: 660f1bf673487ae0bb25b900
title: Step 37
challengeType: 1
dashedName: step-37
---

# --description--

你的循環現在應該運行八次。 在循環主體內，打印 `i` 迭代器的值並觀察會發生什麼。

# --hints--

你應該使用 `console.log()`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

你應該記錄 `i` 的值。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

你應該在 `for` 循環中記錄 `i` 的值。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
