---
id: 660f0f980e98e8cf77f1ce31
title: Step 30
challengeType: 1
dashedName: step-30
---

# --description--

更改你的 `rows` 聲明，再次給它分配一個空數組。

同時，更改 `'Hello'` 字符串，再次使用雙引號。 一般來說，用單引號或者雙引號都沒問題，但是你終會發現在整個項目中，統一使用一種方式是更好的選擇。

# --hints--

你的 `rows` 數組應該是空的。

```js
assert.empty(rows);
```

你的 `"Hello"` 字符串應該用雙引號。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
