---
id: 660f45ccf4ca5c31f253005a
title: Step 86
challengeType: 1
dashedName: step-86
---

# --description--

現在，如果你將 `continueLoop` 更改爲 true，則 `while` 循環將永遠運行。 這被稱爲<dfn>無限循環</dfn>，你應該小心避免這種情況。 無限循環可能會鎖死你的系統，需要完全重新啓動才能擺脫。

爲了避免這種情況，首先使用自增運算符來增加循環內 `done` 變量的值。

# --hints--

你的 `while` 循環應該增加 `done` 變量。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
