---
id: 660f165270622fd4ec0da3f7
title: Step 31
challengeType: 1
dashedName: step-31
---

# --description--

使用 `let` 關鍵字聲明變量時，可以對其進行重新賦值。 這意味着你稍後可以將 `character` 更改爲完全不同的值。

對於這個項目，你不會想要更改這些變量值。 因此，你應該使用 `const` 來聲明它們。 `const` 變量很特殊。

首先，`const` 變量不能像 `let` 變量那樣被重新分配。 這段代碼將引發錯誤：

```js
const firstName = "Naomi";
firstName = "Jessica";
```

`const` 變量也不能未初始化。 這段代碼將引發錯誤：

```js
const firstName;
```

用 `const` 替換 `let` 關鍵字。

# --hints--

你應該使用 `const` 來聲明你的 `character` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

你應該使用 `const` 來聲明你的 `count` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

你應該使用 `const` 來聲明你的 `rows` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

你不應該在代碼中使用 `let`。

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
