---
id: 660f44f10ea40f300b896a5e
title: Step 84
challengeType: 1
dashedName: step-84
---

# --description--

現在你已經練習了使用 `if...else if...else` 語句，你可以將它們從代碼中刪除。

完成後，使用 `let` 聲明 `continueLoop` 變量併爲其分配布爾值 `false`。 然後使用 `let` 聲明一個 `done` 變量併爲其賦值 `0`。

# --hints--

你不應該有 `if` 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

你不應該有 `else if` 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

你不應該有 `else` 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

你應該使用 `let` 來聲明 `continueLoop` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

你的 `continueLoop` 變量應具有值 `false`。

```js
assert.isFalse(continueLoop);
```

你應該使用 `let` 來聲明 `done` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

你的 `done` 變量應具有值 `0`。

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
