---
id: 660f1b6e60bd9edf902c81fd
title: Step 36
challengeType: 1
dashedName: step-36
---

# --description--

你的 <dfn>iteration</dfn> 語句將告訴你的循環在每次運行後如何處理迭代器。

重新分配變量時，可以使用該變量引用重新分配之前的先前值。 這使你可以執行一些操作，比如在已有的數字上加上 3。 例如，`bees = bees + 3;` 會將 `bees` 的值增加 3。

使用該語法將 `"iteration"` 字符串替換爲給 `i` 增加 1 的新的賦值語句。

# --hints--

你應該給 `i` 變量加上 1。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

你應該給你的 `i` 變量賦值 `i + 1`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

你的 `for` 循環應在每次迭代中將 `i` 增加 `1`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
