---
id: 660f4e74f7fd3f4a99ac2e50
title: Step 108
challengeType: 1
dashedName: step-108
---

# --description--

數組還有一個 `.shift()` 方法。 這將刪除數組的 **第一個** 元素，而 `.pop()` 則會刪除最後一個元素。 以下是 `.shift()` 方法的示例：

```js
const numbers = [1, 2, 3];
numbers.shift();
```

`numbers` 數組將是 `[2, 3]`。

在你的 `numbers` 數組正下方，聲明一個 `shifted` 變量，並將在 `numbers` 數組上調用 `.shift()` 的結果分配給它。 在下一行，將 `shifted` 變量記錄到控制檯。

# --hints--

你應該使用 `const` 來聲明 `shifted` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

你應該在 `numbers` 數組上調用 `.shift()`。

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

你應該將 `.shift()` 調用的結果分配給 `shifted` 變量。

```js
assert.equal(shifted, 1);
```

你應該記錄你的 `shifted` 變量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
