---
id: 660f0da9bf1035c9097af20a
title: Step 28
challengeType: 1
dashedName: step-28
---

# --description--

你應該看到了 `"freeCodeCamp"` 被打印到了控制檯。 這是因爲 `.pop()` 返回了從數組中移除的那個值 - 這個值也正是你早先向數組末尾推入的那個值 `"freeCodeCamp"` 。

那麼 `.push()` 方法會返回什麼內容呢？ 試着爲你現有的 `rows.push()` 重新聲明一個新的 `pushed` 變量，並且打印它。

# --hints--

你應該聲明一個 `pushed` 變量。

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

你應該用 `let` 來聲明變量 `pushed`。

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

你應該用 `rows.push("freeCodeCamp")` 來聲明你的 `pushed` 變量。

```js
assert.equal(pushed, 4);
```

你應該打印你的變量 `pushed` 。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
