---
id: 660f17d4e9f227d86e834abd
title: Step 33
challengeType: 1
dashedName: step-33
---

# --description--

要生成金字塔，你需要創建多行。 當你必須重複執行某項任務直到滿足某個條件時，你將使用<dfn>循環</dfn>。 編寫循環的方法有很多種。

你將從一個基本的 `for` 循環開始。 `for` 循環使用以下語法：

```js
for (iterator; condition; iteration) {
  logic;
}
```

在接下來的步驟中，你將詳細探索循環的每個組件。 現在，構建一個 `for` 循環，其中包含三個組件的術語 `"iterator"`、`"condition"` 和 `"iteration"`。 保持循環 <dfn>主體</dfn>（花括號 `{}` 內的部分）爲空。


# --hints--

你應該有一個 `for` 循環。

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

`for` 循環的第一個組件應該是字符串 `"iterator"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

`for` 循環的第二個組件應該是字符串 `"condition"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

`for` 循環的第三個部分應該是字符串 `"iteration"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

`for` 循環的主體應該是空的。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
