---
id: 660f3b664421471aa595170f
title: Step 74
challengeType: 1
dashedName: step-74
---

# --description--

金字塔現在看起來有點滑稽。 由於你從 `1` 而不是 `0` 開始循環，因此在將 `i` 傳遞給 `padRow` 時無需爲其加一。

將 `padRow` 調用的第一個參數更新爲 `i`。

# --hints--

你不應將 `i + 1` 傳遞給 `padRow`。

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

你應該將 `i` 傳遞給 `padRow`。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
