---
id: 660f1a00ac619ddc1e259a66
title: Step 35
challengeType: 1
dashedName: step-35
---

# --description--

`for` 循環的 <dfn>條件</dfn> 告訴循環應迭代多少次。 當 `condition` 變爲假時，循環將停止。

在 JavaScript 中，布爾值可以是 `true` 或 `false`。 這些不是字符串——你稍後會了解更多有關差異的信息。

現在，你將使用 <dfn>小於</dfn> 運算符 (`<`)。 這使你可以檢查左邊的值是否小於右邊的值。 例如，如果 `count` 爲 `2`，則 `count < 3` 的計算結果爲 `true`。如果 `count` 爲 `4`，則計算結果爲 `false`。

用條件替換 `"condition"` 字符串，以檢查 `i` 是否小於 `count`。

# --hints--

你應該使用小於運算符。

```js
assert.match(__helpers.removeJSComments(code), /</);
```

你應該使用小於運算符來檢查 `i` 是否小於 `count`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

你的 `for` 循環應使用 `i < count` 作爲條件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
