---
id: 660f4cde8dd305450514a1cb
title: Step 105
challengeType: 1
dashedName: step-105
---

# --description--

使用多行註釋註釋掉此循環，爲下一步做準備。

# --hints--

你的 `for` 循環應該被註釋掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/), 1);
```

你的 `for` 循環體應該被註釋掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
