---
id: 660f374d532dc41189cc9cc2
title: Step 69
challengeType: 1
dashedName: step-69
---

# --description--

現在該進行一些數學運算了。 考慮一個三排金字塔。 如果我們想讓它居中，它看起來應該像這樣：

```md
··#··
·###·
#####
```

爲了便於閱讀，空格已被中間點取代。 如果你研究該圖案，你會發現行首和行末的空格遵循一種模式。

更新空白字符串以重複 `rowCount - rowNumber` 次。

打開控制檯查看結果。

# --hints--

你應該在 `" "` 字符串上調用 `.repeat()` 來重複 `rowCount - rowNumber` 次。

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
