---
id: 6610bbed59bc2a0194d85533
title: Step 22
challengeType: 1
dashedName: step-22
---

# --description--

注意到 `rows` 數組內的值是如何直接改變的？ 這被稱爲<dfn>突變</dfn>。 當你對數組瞭解得更多時，你就會了解何時應該改變數組，何時不應該改變數組。

在繼續之前，這是一個學習常見數組用法的好機會。 目前，你的代碼使用 `rows[2]` 訪問數組中的最後一個元素。 但是當你想要最後一個元素時，你可能不知道數組中有多少個元素。

你可以使用數組的 `.length` 屬性 - 這將返回數組中元素的數量。 要獲取任何數組的最後一個元素，可以使用以下語法：

```js
array[array.length - 1]
```

`array.length` 返回數組中元素的數量。 通過減去 `1`，你可以得到數組中最後一個元素的索引。 你可以將相同的概念應用到 `rows` 數組。

更新 `rows[2]` 以動態訪問 `rows` 數組中的最後一個元素。 請參閱上面的例子來獲得幫助。

你將不會在控制檯中看到任何變化。

# --hints--

你應該使用 `rows` 數組的 `.length` 屬性。

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

你應該在括號符號中使用 `rows.length`。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

你應該從括號符號的長度中減去 `1`。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
