---
id: 6610c48c4ea0891afa7c4696
title: Step 54
challengeType: 1
dashedName: step-54
---

# --description--

如果你再次打開控制檯，你會看到 `padRow` 函數返回 `undefined`，即使你定義了返回值！ 這是因爲當你**調用**該函數時，需要給參數賦值。

當你將一個值傳遞給函數調用時，該值被稱爲<dfn>參數</dfn>。 下面是調用 `demo` 函數並將 `"Naomi"` 作爲 `name` 參數傳遞的示例。

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

將你自己的名字作爲 `padRow` 調用中 `name` 參數傳遞。 請記住，你的名字是一個字符串，因此你需要使用引號。

# --hints--

你應該將一個字符串傳遞給你的 `padRow()` 調用。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
