---
id: 6610b741b54b90f0c0fb3d58
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

當變量名超過一個單詞時，對於如何將單詞大寫，有特定的命名約定。 在 JavaScript 中，使用的約定是 <dfn>駝峯式命名法</dfn>。

駝峯式大小寫是指名稱中的第一個單詞全部小寫，但後續單詞均採用首字母大寫。 以下是駝峯式命名法的一些示例：

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

使用駝峯式命名法聲明一個新的 `secondCharacter` 變量。

# --hints--

你應該聲明一個 `secondCharacter` 變量。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

你應該使用 `let` 來聲明你的 `secondCharacter` 變量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

你不應該爲 `secondCharacter` 變量賦值。 Do not forget your semicolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
