---
id: 6610c6541c82551f95e765ab
title: Step 58
challengeType: 1
dashedName: step-58
---

# --description--

變量也可以在函數內部聲明。 這些變量被認爲處於<dfn>局部作用域</dfn>或<dfn>塊級作用域</dfn>內。 在函數內部聲明的變量只能在該函數內部使用。 如果你嘗試在函數外部訪問它，則會收到引用錯誤。

要查看其實際效果，請使用 `const` 在 `padRow` 函數中聲明一個 `test` 變量。 使用值 `"Testing"` 對其進行初始化。

然後，在你的函數下方，嘗試將 `test` 記錄到控制檯。 你將看到一個錯誤，因爲它沒有在函數的局部作用域之外定義。 刪除 `console.log` 以通過測試並繼續。

# --hints--

你的函數應該聲明一個 `test` 變量。

```js
assert.match(padRow.toString(), /var\s+test/);
```

你應該使用值 `"Testing"` 初始化 `test`。 Don't forget the semicolon.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

你的 `test` 變量應位於 `return` 關鍵字之前。

```js
const str = padRow.toString();
assert.isBelow(str.indexOf("test"), str.indexOf("return"));
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
