---
id: 6610c2d8d67563174fcf96dc
title: Step 52
challengeType: 1
dashedName: step-52
---

# --description--

當你的代碼中有一個明確寫出的值時，例如函數中的 `"Hello!"` 字符串，它被視爲<dfn>硬編碼</dfn>。 在函數內部對值進行硬編碼可能無法使其像你希望的那樣具有可重用性。

相反，你可以爲該函數定義 <dfn>參數</dfn>。 參數是調用函數時被賦予值的特殊變量，並且可以在函數中使用來動態地改變函數代碼的結果。

要向函數添加參數，需要在括號內添加變量名。 例如，此 `demo` 函數有一個 `name` 參數：

```js
function demo(name) {

}
```

`name` 聽起來像是一個有用的參數，所以繼續將其添加到你的 `padRow` 函數中。

# --hints--

你的 `padRow` 函數應該有一個 `name` 參數。

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
