---
id: 6610b9f7619764fad5fd516d
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

你還可以將一個變量的值分配給另一個變量。 例如：

```js
let first = "One";
let second = "Two";
second = first;
```

`second` 變量現在的值爲 `"One"`。

要查看實際效果，請將 `secondCharacter` 賦值從 `"Test"` 更改爲 `character` 變量。

然後打開控制檯查看記錄的內容。

# --hints--

你不應將值 `"Test"` 分配給 `secondCharacter` 變量。

```js
assert.notEqual(secondCharacter, "Test");
```

你應該將 `character` 變量的值分配給 `secondCharacter` 變量。 Don't forget your semicolon.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

你的 `secondCharacter` 變量現在應該具有值 `"World"`。


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
