---
id: 6610c77d50636722e5b6be17
title: Step 60
challengeType: 1
dashedName: step-60
---

# --description--

在 `return` 語句下方，將字符串 `"This works!"` 記錄到控制檯。

執行此操作後，你將看到字符串 `"This works!"` 沒有顯示在控制檯中，並且 `console.log("This works!")` 行顯示爲灰色。

複製控制檯日誌並將其粘貼到 `return` 語句上方。 現在，字符串 `"This works!"` 應該出現在控制檯中。

關於 `return` 關鍵字，需要了解的一件重要事情是，它不僅定義從函數返回的值，還會停止函數或塊語句內代碼的執行。 這意味着 `return` 語句之後的任何代碼都不會運行。

# --hints--

你的 `padRow` 函數應該返回 `test` 變量。

```js
assert.equal(padRow("Naomi"), "Testing");
```

你的第一個 `console.log` 應該位於 `return` 關鍵字之後。

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

你的第二個 `console.log` 應該位於 `return` 關鍵字之前。

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";

  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
