---
id: 6610b8f6a98d25f4d485a94d
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

未初始化變量的默認值是 `undefined`。 這是一種特殊的數據類型，表示尚未定義的值。

你仍然可以爲未初始化的變量賦值。 以下是一個例子：

```js
let uninitialized;
uninitialized = "assigned";
```

將字符串 `"Test"` 分配給聲明下方的 `secondCharacter` 變量。 打開控制檯查看你的日誌是如何變化的。

# --hints--

你不應該初始化 `secondCharacter`。 請記住，初始化意味着在聲明變量時分配一個值。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

你應該在 `secondCharacter` 上使用賦值運算符。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

你應該將字符串 `"Test"` 分配給 `secondCharacter` 變量。

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
