---
id: 6610c538372aa61cc0f5b122
title: Step 57
challengeType: 1
dashedName: step-57
---

# --description--

JavaScript 中的變量在特定的<dfn>範圍</dfn>內可用。 換句話說，變量的聲明位置決定了它在代碼中的使用位置。

第一個範圍是全局範圍。 在任何“塊”（例如函數或 `for` 循環）之外聲明的變量都位於<dfn>全局範圍</dfn>內。 你的 `character`、`count` 和 `rows` 變量均位於全局範圍內。

當變量位於全局範圍內時，函數可以在其定義中訪問它。 以下是使用全局 `title` 變量的函數示例：

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

此示例將返回 `"Professor Naomi"`。 更新你的 `padRow` 函數以返回將你的 `character` 變量連接到 `name` 參數開頭的值。

# --hints--

你的 `padRow` 函數應該將 `character` 連接到 `name` 的開頭。

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

你的 `padRow` 函數應該返回 `character + name` 的結果。

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
