---
id: 6610c87eac0f0b256d7b037e
title: Step 62
challengeType: 1
dashedName: step-62
---

# --description--

因爲你的函數不再使用該參數，所以更改調用中傳遞的值不會對其產生影響。

繼續從 `padRow` 函數中刪除 `test` 聲明。 另外，刪除 `return` 語句，這樣你的函數再次變爲空。

# --hints--

你的 `padRow` 函數不應該有 `test` 變量。

```js
assert.notMatch(padRow.toString(), /test/);
```

你的 `padRow` 函數不應返回值。

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
