---
id: 6610c16c4fa0df12c0e30675
title: Step 49
challengeType: 1
dashedName: step-49
---

# --description--

你正在調用 `padRow` 函數，但沒有對該函數調用執行任何操作。 JavaScript 中的所有函數都<dfn>返回</dfn>一個值，這意味着它們提供了調用它們的定義結果供你在其他地方使用。

要查看調用 `padRow` 函數的結果，請聲明一個 `call` 變量並將現有的 `padRow` 調用分配給該變量。

# --hints--

你應該聲明一個 `call` 變量。

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+call/);
```

你應該使用 `const` 來聲明你的 `call` 變量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

你應該將 `padRow` 調用的結果分配給 `call`。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
