---
id: 6610c21b3ef82015573ffbbe
title: Step 51
challengeType: 1
dashedName: step-51
---

# --description--

你的 `call` 變量具有 `undefined` 值，即使你定義了它！ 這是因爲你的 `padRow` 函數當前沒有返回值。 默認情況下，函數返回 `undefined` 作爲其值。

爲了返回其他內容，你需要使用 `return` 關鍵字。 下面是一個返回字符串 `"Functions are cool!"` 的函數示例：

```js
function demo() {
  return "Functions are cool!";
}
```

使用 `return` 關鍵字讓你的函數返回字符串 `"Hello!"`。

# --hints--

你的 `padRow` 函數應該使用 `return` 關鍵字。

```js
assert.match(padRow.toString(), /return/);
```

你的 `padRow` 函數應該返回字符串 `"Hello!"`。

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
