---
id: 6610c424b7119919b62932f4
title: Step 53
challengeType: 1
dashedName: step-53
---

# --description--

函數不必返回硬編碼值。 它可以返回存儲在變量中的值。 參數是函數的特殊變量，因此也可以返回它們。

將你的 `padRow` 函數更改爲直接 `return` `name` 參數。

# --hints--

你的 `padRow` 函數應該返回 `name` 參數的值。

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
