---
id: 6610c71600966a2191d3a64a
title: Step 59
challengeType: 1
dashedName: step-59
---

# --description--

函數返回的值可通過調用該函數來使用。 你可以直接使用函數調用作爲其返回的值，或者在變量中捕獲返回的值。 這樣，你就可以在創建該變量的函數之外使用分配給本地作用域變量的值。

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

要使用 `"Testing"` 值，請通過更新 `return` 語句以僅返回 `test` 變量，將其從 `padRow` 函數中返回。


# --hints--

你的 `padRow` 函數應該返回 `test` 變量。

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
