---
id: 6610c83b52583e245a079217
title: Step 61
challengeType: 1
dashedName: step-61
---

# --description--

現在你的 `call` 變量具有值 `"Testing"`。 但是你的函數不再使用 `name` 參數。

從函數聲明中刪除 `name` 參數，然後從 `padRow` 調用中刪除 `"CamperChan"` 字符串。

另外，從 `padRow` 函數中刪除兩個 `console.log`。
# --hints--

你的 `padRow` 函數不應該有 `name` 參數。

```js
assert.notMatch(padRow.toString(), /name/);
```

你不應將 `"CamperChan"` 傳遞給你的 `padRow` 調用。

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

你仍然應該調用 `padRow` 函數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

在 `return` 關鍵字之前不應該有 `console.log`。

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

`return` 關鍵字之後不應該有 `console.log`。

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
