---
id: 6610c8cfe4cf4d278e35c156
title: Step 63
challengeType: 1
dashedName: step-63
---

# --description--

正如預期的那樣，你的函數現在再次返回 `undefined`。 你的 `call` 變量不再需要，因此請刪除 `call` 聲明和 `call` 變量的 `console.log`。

# --hints--

你不應該有 `call` 聲明。

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s*call/);
```

你不應該記錄你的 `call` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
