---
id: 6610b8017d1671f2814e8c77
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

當你聲明一個變量但沒有初始化它時，該變量被視爲<dfn>未初始化</dfn>。 目前，你的 `secondCharacter` 變量未初始化。

添加 `console.log()` 來查看 `secondCharacter` 變量的值。

# --hints--

你應該添加第二個 `console.log`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你應該記錄你的 `secondCharacter` 變量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
