---
id: 6610c105bbdacc114d6cdc44
title: Step 48
challengeType: 1
dashedName: step-48
---

# --description--

爲了使用函數，你需要調用它。 <dfn>函數調用</dfn> 告訴你的應用程序在你選擇調用該函數的任何地方運行該函數的代碼。 函數調用的語法是函數名稱後跟括號。 例如，此代碼定義並調用 `test` 函數。

```js
function test() {

}
test();
```

調用你的 `padRow` 函數。

# --hints--

你應該調用 `padRow` 函數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

你的 `padRow` 函數體應該是空的。

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
