---
id: 661483051820c3c1ab4595e0
title: Step 96
challengeType: 1
dashedName: step-96
---

# --description--

你還可以在 `padRow` 調用中替換 `done` 引用。

請注意，此處的 `rows.length` 會出現差一錯誤，因爲 `done` 在被調用之前 *已經* 遞增了。

因此你需要將此處的 `done` 替換爲 `rows.length + 1`。 當你執行此操作時，你可能會看到 `Range Error`，因爲我們又創建了另一個差一錯誤。

你需要更改 `while` 條件以使用小於運算符，而不是小於或等於運算符。

# --hints--

你應該將 `rows.length + 1` 作爲 `padRow` 調用的第一個參數傳遞。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(padRow\s*\(\s*rows\.length\s*\+\s*1/);
```

當 `rows.length` 小於 `count` 時，你的 `while` 循環應該運行。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
