---
id: 66458f0a05df478aa627629e
title: Step 82
challengeType: 1
dashedName: step-82
---

# --description--

文字又不見了！ 空字符串的計算結果爲 `false`，使其成爲 <dfn>假值</dfn>。 你將在未來的項目中瞭解有關真值和假值的更多信息。

除了 `if` 語句之外，JavaScript 還有 <dfn>else if</dfn> 語句。 `else if` 語句允許你在單個代碼塊中檢查多個條件。

以下是 `else if` 語句的語法：

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

如果第一個條件爲 `false`，JavaScript 將檢查鏈中的下一個條件。 如果第二個條件爲 `false`，JavaScript 將檢查第三個條件，依此類推。

在 `if` 語句下方，添加一個 `else if` 語句，檢查 `5` 是否小於 `10`。 然後在 `else if` 語句的主體內，將字符串 `"5 is less than 10"` 記錄到控制檯。

檢查控制檯來查看結果。

# --hints--

你應該有一個 `else if` 語句。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

你的 `else if` 語句應檢查 `5` 是否小於 `10`。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

你應該將字符串 `"5 is less than 10"` 記錄到控制檯。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

