---
id: 664599653fcd6e97104f9261
title: Step 83
challengeType: 1
dashedName: step-83
---

# --description--

有時，你會希望在所有 `if...else if` 條件均爲 `false` 時運行不同的代碼。 你可以通過添加 `else` 塊來實現這一點。

僅當 `if` 和 `else if` 塊中的條件不滿足時，纔會評估 `else` 塊。

這裏將 `else` 塊添加到 `else if` 塊。

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

向 `else if` 塊添加 `else` 塊。 在 `else` 塊內，將字符串 `"This is the else block"` 記錄到控制檯。

要在控制檯中查看結果，可以手動將 `else if` 語句中的 `<` 更改爲 `>`。 這將使得條件爲 `false` 並且 `else` 塊將運行。

# --hints--

你應該有一個 `else` 塊。

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

你的 `else` 塊應該將字符串 `"This is the else block"` 記錄到控制檯。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

