---
id: 664ed9344ba3495e1f67f8ea
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

現在你已經檢查了聲明和初始化，刪除你爲檢查而編寫的代碼。

不要刪除你的 `character` 變量。

# --hints--

你不應該有 `profession` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

你不應該有 `age` 變量。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

你不應該對 `age` 有 `console.log` 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

你不應該對 `profession` 有 `console.log` 語句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

你應該仍然有你的 `character` 變量。

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
