---
id: 664ed97c55a99e5ffab759dc
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

在繼續之前，你應該花點時間複習一下已經學過的概念。

使用 `let` 關鍵字聲明 `profession` 變量和 `age` 變量。 使用字符串 `"teacher"` 初始化 `profession`，但不要使用任何值初始化 `age`。

將兩個變量記錄到控制檯以查看結果。

# --hints--

你應該聲明一個 `profession` 變量。

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

你應該使用 `let` 來聲明 `profession` 變量。

```js
assert.match(code, /let\s+profession/);
```

你應該將字符串 `"teacher"` 分配給 `profession` 變量。

```js
assert.strictEqual(profession, "teacher");
```

你應該聲明一個 `age` 變量。

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

你應該使用 `let` 來聲明 `age` 變量。

```js
assert.match(code, /let\s+age/);
```

你不應該爲 `age` 變量賦值。

```js
assert.isUndefined(age);
```

你應該將 `age` 記錄到控制檯。

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

你應該將 `profession` 記錄到控制檯。

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
