---
id: 665732da4815b70bb083915e
title: Step 24
challengeType: 1
dashedName: step-24
---


# --description--

在前面幾個步驟中，你學習了有關使用數組的所有知識。 花一點時間複習一下你所學到的知識。

首先聲明一個 `cities` 變量，並將其初始化爲字符串 `"London"`、`"New York"` 和 `"Mumbai"` 的數組。 然後將該變量記錄到控制檯。

記錄完成後，將 `cities` 的最後一個元素更改爲字符串 `"Mexico City"`，然後再次記錄 `cities` 變量。

正確完成後，你應該在控制檯中看到此輸出。

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

你應該使用 `let` 來聲明 `cities` 變量。

```js
assert.match(code, /let\s+cities/);
```

你應該將字符串 `"London"`、`"New York"` 和 `"Mumbai"` 的數組分配給 `cities` 變量。

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*\]/);
```

你應該使用 `console.log()` 將整個 `cities` 數組記錄到控制檯。

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

你應該將 `cities` 數組的最後一個元素更新爲字符串 `"Mexico City"`。 請記住，你可以使用 `array[array.length - 1]` 訪問數組的最後一個元素。

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

你的代碼中應該有兩個 `console.log(cities)` 語句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
