---
id: 665732e62eb8b30c249b5c2c
title: Step 25
challengeType: 1
dashedName: step-25
---

# --description--

現在你已準備好學習下一組數組課程。

刪除上一步中的所有代碼。

# --hints--

你應該刪除對 `cities` 變量的所有引用。

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
