---
id: 66643c93e05093c728abdbe9
title: Step 55
challengeType: 1
dashedName: step-55
---

# --description--

在繼續之前，請花點時間回顧一下函數的工作原理。

聲明一個名爲 `addTwoNumbers` 的函數。 該函數應該接受兩個參數並返回這兩個參數的總和。

然後聲明一個 `sum` 變量併爲其分配使用 `5` 和 `10` 作爲參數調用 `addTwoNumbers` 函數的值。 將 `sum` 變量記錄到控制檯。


# --hints--

你應該有一個名爲 `addTwoNumbers` 的函數。

```js
assert.isFunction(addTwoNumbers);
```

你的函數應該返回兩個數字的總和。

```js
assert.strictEqual(addTwoNumbers(2,3), 5);
```

你應該聲明一個 `sum` 變量。

```js
assert.isDefined(sum);
```

你應該將以 `5` 和 `10` 作爲參數調用 `addTwoNumbers` 函數得到的值分配給 `sum`。

```js
assert.strictEqual(sum, 15);
```

你應該記錄你的 `sum` 變量。

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
