---
id: 66eaddd04a9e533fba689001
title: Design a Blog Post Card
challengeType: 14
dashedName: "lab-blog-post-card"
demoType: onClick
---

# --description--

In this lab, you'll practice how to style backgrounds and borders.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have a `div` with a class of `blog-post-card` to hold all your card elements.
2. You should have an image with a valid `alt` text and the class of `post-img`. You can use `cover-photo.jpg` for the `src` attribute of your image.
3. You should have a `div` with a class of `post-content`. Within the `.post-content` `div`:
    - You should have a `h2` with a class of `post-title`. Assign some text to act as the title of your blog post.
    - You should have a `p` with a class of `post-excerpt`. Assign some text to summarize the blog post (maximum of 20 words).
    - You should an have an `a` with a class of `read-more`. Assign the text `Read More` to it.
4. The `blog-post-card` class should add a white background, rounded corners, shadow for depth, overflow hidden to round the image borders, and control the card width and text alignment.
5. The `post-img` class should ensure that the image fills the card's width and adds a border at the bottom.
6. The `post-content` class should add padding inside the card for spacing.
7. The `post-title` and `post-excerpt` elements within `post-content` should be styled for text color and margin.
8. The `read-more` element should be styled for text color, background color, margin, `inline-block` display, `border-radius`, and `padding`. It should be styled like a button.
9. The `read-more` element should have a hover effect that changes the background color.

# --hints--

You should have a `.blog-post-card` `div`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

You should have a `.post-img` `img`. Make sure your image has an `alt` text.

```js
const img = document.querySelector('img.post-img');
assert.exists(img, 'The <img> element with class "post-img" exists');
assert.isString(img.alt, 'The alt text of the <img> element is a string');
assert.isNotEmpty(img.alt, 'The alt text of the <img> element is not empty');
```

You should have a `.post-content` `div`.

```js
assert.exists(document.querySelector("div.post-content"));
```

You should have a `.post-title` `h2`.

```js
assert.exists(document.querySelector("h2.post-title"));
```

You should have a `.post-excerpt` `p`.

```js
assert.exists(document.querySelector("p.post-excerpt"));
```

You should have a `.read-more` `a`.

```js
assert.exists(document.querySelector("a.read-more"));
```

You should have a `.read-more` `a` should have a `Read More` text.

```js
const link = document.querySelector('a.read-more');
assert.exists(link, 'The <a> element with class "read-more" exists');
assert.strictEqual(link.textContent, 'Read More', 'The <a> element has the correct text');
```

`blog-post-card` should have a `border-radius` property with a value (should not be 0 or a negative value).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card, 'The element with class "blog-post-card" exists');

const borderRadius = getComputedStyle(card).borderRadius;
assert.notStrictEqual(borderRadius, '0px', 'The card has rounded corners');
```

`blog-post-card` should have a white background.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card, 'The element with class "blog-post-card" exists');

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.include(whiteColors, backgroundColor.toLowerCase(), 'The card has a white background');
```

You should target `.blog-post-card` and set its `width` property.

```js
assert(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

You should target `.blog-post-card` and set its `text-align` property.

```js
assert(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

You should target `.blog-post-card` and set its `overflow` property.

```js
assert(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("overflow"));
```

You should target `.post-content` and set its `padding` property.

```js
assert(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

`.read-more` should have a hover effect.

```js
const link = document.querySelector('.read-more');
assert.exists(link, 'The element with class "read-more" exists');

const stylesheet = [...document.styleSheets].find(sheet => [...sheet.rules].some(rule => rule.selectorText === '.read-more:hover'));
assert.exists(stylesheet, 'The hover rule for .read-more exists');

const hoverRule = [...stylesheet.rules].find(rule => rule.selectorText === '.read-more:hover');
const hoverBackgroundColor = hoverRule.style.backgroundColor;
assert.isNotEmpty(hoverBackgroundColor, 'The hover effect changes the background color');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
