---
id: 66eaddd04a9e533fba689001
title: Design a Blog Post Card
challengeType: 14
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

In this lab, you'll practice how to style backgrounds and borders by creating a blog post card.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have a `div` with a class of `blog-post-card` to hold all your card elements.
2. You should have an image with a valid `alt` text and the class of `post-img`. You can use `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` for the `src` attribute of your image.
3. You should have a `div` with a class of `post-content`.  
   Within the `.post-content` `div`:

   - You should have a `h2` with a class of `post-title`. Assign some text to act as the title of your blog post.
   - You should have a `p` with a class of `post-excerpt`. Assign some text to summarize the blog post (maximum of 20 words).
   - You should have an `a` with a class of `read-more`. Assign the text `Read More` to it.

4. You should apply the following styles to the `blog-post-card` div:

   - A white colored background.
   - Rounded corners.
   - Shadow for depth.
   - Appropriate `overflow` property to prevent content from overflowing.
   - Width of your choice.
   - Text alignment of your choice.

5. The `post-img` class should ensure that the image fills the card's width and adds a border at the bottom.
6. The `post-content` class should add padding inside the card for spacing.
7. The `post-title` and `post-excerpt` elements within `post-content` should have a text color other than the default and a margin on all sides.
8. The `read-more` element should be styled like a button and have:

   - A text color other than the default.
   - Background color.
   - Margin on all sides.
   - A `display` property of `inline-block`.
   - Rounded corners.
   - Padding on all sides.

9. The `read-more` element should have a hover effect that changes the background color.

**Note:** Be sure to link your stylesheet in your HTML and apply your CSS.

# --hints--

You should have a `.blog-post-card` `div`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

You should have a `.post-img` `img`. Make sure your image has an `alt` text and `src` value.

```js
const img = document.querySelector('img.post-img');
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute('src')); 
assert.isNotEmpty(img.getAttribute('src'));
```

You should have a `.post-content` `div`.

```js
assert.exists(document.querySelector("div.post-content"));
```

You should have a `.post-title` `h2`. Make sure it has some text content.

```js
const postTitle = document.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

You should have a `.post-excerpt` `p`.

```js
assert.exists(document.querySelector("p.post-excerpt"));
```

You should have a `.read-more` `a`.

```js
assert.exists(document.querySelector("a.read-more"));
```

You should have a `.read-more` `a` should have a `Read More` text.

```js
const link = document.querySelector('a.read-more');
assert.exists(link);
assert.strictEqual(link.textContent, 'Read More');
```

`blog-post-card` should have a `border-radius` property with a value (should not be `0` or a negative value).

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.notStrictEqual(borderRadius, '0px');
```

`blog-post-card` should have a white background.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

You should target `.blog-post-card` and set its `width` property.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

You should target `.blog-post-card` and set its `text-align` property.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

You should target `.blog-post-card` and set its shadow property.

```js
const card = document.querySelector('.blog-post-card');

function hasBoxShadow(element) {
  const boxShadow = getComputedStyle(element).getPropertyValue('box-shadow');
  return boxShadow && boxShadow !== 'none';
}

assert.isTrue(hasBoxShadow(card));
```

You should target `.blog-post-card` and set its `overflow` property to `hidden`.

```js
const overflowVal = new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("overflow");
assert.equal(overflowVal, "hidden");
```

You should target `.post-content` and set its `padding` property.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

`.read-more` should have a hover effect.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const stylesheet = [...document.styleSheets].find(sheet => [...sheet.rules].some(rule => rule.selectorText === '.read-more:hover'));
assert.exists(stylesheet);

const hoverRule = [...stylesheet.rules].find(rule => rule.selectorText === '.read-more:hover');
const hoverBackgroundColor = hoverRule.style.backgroundColor;
assert.isNotEmpty(hoverBackgroundColor);
```

`.read-more` should be styled like a button. 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

`post-img` should fill the card's `width` and have a `border-bottom` value.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.strictEqual(imgWidth, cardWidth);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Ensure `post-title` and `post-excerpt` have text color and margin.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

const titleStyles = getComputedStyle(title);
const excerptStyles = getComputedStyle(excerpt);

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styleSheets = [...document.styleSheets];
const hasTitleColorRule = styleSheets.some(sheet => {
  try {
    return [...sheet.cssRules].some(rule => rule.selectorText === '.post-title' && rule.style.color);
  } catch (e) {
    return false;
  }
});

const hasExcerptColorRule = styleSheets.some(sheet => {
  try {
    return [...sheet.cssRules].some(rule => rule.selectorText === '.post-excerpt' && rule.style.color);
  } catch (e) {
    return false;
  }
});

assert.isTrue(hasTitleColorRule);
assert.isTrue(hasExcerptColorRule);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
