---
id: 6578b57361f2f132a02e2a18
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Você pode acessar os dados armazenados em um dicionário por meio das suas chaves:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack`
}

my_dict['name'] # 'Michael'
```

Depois do dicionário, siga o exemplo abaixo para acessar a chave `species` de `copper` e exiba o resultado.

# --hints--

Você não deve modificar seu dicionário.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Você deve usar `copper['species']` para acessar o valor da chave `species`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Você deve chamar `print()` passando `copper['species']` como argumento.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
