---
id: 6ac0ae563f6de70930f47ef7
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

When writing code, you'd often want to reuse or reference existing values in your code. Variables allows you to do that. Here's how you can assign the string `Bob` to a variable named `name`:

```py
name = 'Bob'
```

Go to a new line, create a variable named `sentence` and assign it the string `I am learning Python`.

# --hints--

You should have `sentence = 'I am learning Python'` in your code. Make sure to start a the beginning of the line.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("sentence = 'I am learning Python'")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
