---
id: 6ac0aaba914d81b32567c076
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` is a built-in function that you can use to send output to the terminal. As you can see, the text `Hello` is now visible on the terminal.

Now, print the text `freeCodeCamp` to the terminal by going on a new line and typing:

```py
print('freeCodeCamp')
```

Make sure to start at the beginning of the line.

# --hints--

Test 1

```js
({ test: () => assert(runPython(`<python expression that returns truthy>`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
