---
id: 6ac0aaba914d81b32567c076
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` is a built-in function that you can use to send output to the terminal. As you can see, the text `Hello` is now visible on the terminal.

Now, print the text `freeCodeCamp` to the terminal by going to the highlighted line and typing:

```py
print('freeCodeCamp')
```

Make sure to start at the beginning of the line.

# --hints--

You should have `print('freeCodeCamp')` in your code. Make sure to start at the beginning of the line.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
