---
id: 6ac0a4463662d9738c393e21
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

In the previous step, you created a **string**: a sequence of characters wrapped in either single or double quotes. In programming, strings are used to represent text.

The black box at the top-right on your screen is your **terminal**. This is where you can see the output of your code.

To be able to see an output, turn your existing line of code into:

```py
print('Hello')
```

Make sure to start at the beginning of the line.

# --hints--

You should have `print('Hello')` in your code. Make sure to start at the beginning of the line.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print('Hello')")`)) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
