---
id: 6ac0b35f3b29a27f7e207295
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

In the previous step, you created the variable `sentence`. Now print its value to the terminal by typing `print(sentence)` on the highlighted line.

With that, you've completed your first Python program.

# --hints--

You should have `print(sentence)` in your code. Make sure to start at the beginning of the line.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
