---
id: 6ac09f8420b4b67e075d2cd2
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will create your first Python program.

This is your code editor, where you'll write Python code. Line 1 is highlighted. Go there and type:

```py
'Hello'
```

Make sure to start at the beginning of the line and do not forget quotes.

Once you're done, click the "check your code" button to see if it's correct.

# --hints--

You should type `'Hello'`. Make sure to start at the beginning of the line and do not forget quotes.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
