---
id: 6ac09f8420b4b67e075d2cd2
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will create your first Python program.

This is your code editor, where you'll write Python code. Find line 1 and type `Hello`. Make sure to start at the begginning of the line and wrap the text in quotes.

Once you're done, click the "check your code" button to see if it's correct.

# --hints--

Test 1

```js
({ test: () => assert(runPython(`<python expression that returns truthy>`)) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
