---
id: 6ac0b35f3b29a27f7e207295
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

In the previous step, you created the variable `sentence`. Now print its value to the terminal by typing `print(sentence)` on a new line.

With that you completed your first Python program.

# --hints--

Test 1

```js
({ test: () => assert(runPython(`<python expression that returns truthy>`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
