---
id: 6ac0ae563f6de70930f47ef7
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

When writing code, you'd often want to reuse or reference existing values in your code. Variables allow you to do that. Here's how you can assign the string `Bob` to a variable named `name`:

```py
name = 'Bob'
```

Go to the highlighted line, create a variable named `sentence` and assign it the string `I am learning Python`.

# --hints--

You should have `sentence = 'I am learning Python'` in your code. Make sure to start at the beginning of the line.

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
