---
id: 6568994faf481da5d37bfa40
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

Now, you need to create the three line segments to pass to the `enumerate` function.

Use list slicing to create the three lists of equal length representing the `line` segment of each 3x3 square and pass them to the `enumerate()` call. Add `start = 1` to start the enumeration from `1` instead of `0`.

# --hints--

You should have `enumerate([line[:3], line[3:6], line[6:]], start=1)` within the inner `for` loop.

```js
assert.match(code, /\[\s*line\s*\[\s*:3\s*\]\s*,\s*line\s*\[\s*3\s*:\s*6\s*\]\s*,\s*line\s*\[\s*6\s*:\s*\]\s*\]\s*,\s*start\s*=\s*1/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

--fcc-editable-region--
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate():
--fcc-editable-region--
```
