---
id: 66f53dc2c5bd6a11d6c3282f
title: Build a Password Generator App
challengeType: 14
dashedName: lab-js-functions
demoType: onClick
---

# --description--

In this lab,

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create an arrow function named `generatePassword`. The function should have a `length` parameter. Step 2 to 5 should be created within the `generatePassword` function.
2. Using the `let` keyword, create a variable named `chars`. It should have a value of all possible characters that can be included in the password. You can use `"ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()"`.
3. You should create a variable named `password`. It should have an empty string as its value.
4. You should create a `for` loop that runs `length` times. In the body of the loop:
    - Generate a random index using `Math.floor(Math.random() * chars.length);` and store it in a variable named `randomIndex`
    - Append the random characters to `password` using `chars[randomIndex]`.
5. You should return `password`.
6. You should create an arrow function named `runPasswordGenerator`. Step 7 to 9 should be created within the `runPasswordGenerator` function.
7. You should create a prompt that asks the user to `Enter the desired length for the password:`. Store the input in a `const` variable named `length`, and convert it to an integer.
8. You should call the `generatePassword` function with `length` as a parameter and store the returned password in a `const` variable named `password`.
9. You should log the generated password to the console using the message `Generated password: ${password}`.
10. You should call the `runPasswordGenerator` function.

# --hints--

You should have a `generatePassword` function with a `length` parameter.

```js
assert.isFunction(generatePassword);
assert.isTrue(generatePassword.length === 1);
```

The `chars` variable should have a value of possible password characters.

```js
let chars = 'ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()';

assert.isNotEmpty(chars, 'chars should not be empty');

assert.include(chars, 'A');
assert.include(chars, 'a');
assert.include(chars, '1');
assert.include(chars, '@');

const expectedChars = 'ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()';
assert.strictEqual(chars, expectedChars);
```

You should have an empty string `password` variable.

```js
const initialPassword = generatePassword(0);
assert.isTrue(typeof initialPassword === 'string');
assert.strictEqual(initialPassword, '');
```

You should have a `for` loop that appends the characters

```js
const generatePasswordString = generatePassword.toString();
assert.isTrue(generatePasswordString.includes('for (') && generatePasswordString.includes('i < length'));
```

`generatePassword` should return `password`.

```js
const returnedPassword = generatePassword(5);
assert.isString(returnedPassword);
assert.isTrue(returnedPassword.length === 5);
```

You should have a `runPasswordGenerator` function.

```js
assert.isFunction(runPasswordGenerator);
```

You should prompt the user for the password length.

```js
const originalPrompt = window.prompt;
let promptCalled = false;
window.prompt = () => {
    promptCalled = true;
    return '10'; 
};
runPasswordGenerator();
assert.isTrue(promptCalled);
window.prompt = originalPrompt;
```

You should call `generatePassword`.

```js
let generatePasswordCalled = false;
const originalGeneratePassword = window.generatePassword;
window.generatePassword = (length) => {
    generatePasswordCalled = true;
    return 'mockedPassword';
};
runPasswordGenerator();
assert.isTrue(generatePasswordCalled);
window.generatePassword = originalGeneratePassword;
```

You should log the generated password to the console.

```js
let loggedMessage = '';
const originalConsoleLog = console.log;
console.log = (message) => {
    loggedMessage = message;
};
const fakePassword = 'mockedPassword';
window.generatePassword = () => fakePassword;
runPasswordGenerator();
assert.strictEqual(loggedMessage, `Generated password: ${fakePassword}`);
console.log = originalConsoleLog;
```

You should call `runPasswordGenerator`.

```js
assert.doesNotThrow(() => runPasswordGenerator());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    let chars = 'ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()';
    let password = '';
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
}

const runPasswordGenerator = () => {
    const length = parseInt(prompt("Enter the desired length for the password:"), 10);
    const password = generatePassword(length);
    console.log(`Generated password: ${password}`);
}

runPasswordGenerator();
```
