---
id: 6690b9a24f5b0f1300040c76
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

The `p` element is used to create a paragraph of text on websites. Create a `p` element below your `h2` element and give it the following text:

`See more cat photos in our gallery.`

# --hints--

Your `p` element should have an opening tag. Opening tags have the following syntax: `<elementName>`.

```js
assert(document.querySelector('p'));
```

Your `p` element should have a closing tag. Closing tags have a `/` just after the `<` character.

```js
assert(code.match(/<\/p\>/));
```

Your `p` element's text should be `See more cat photos in our gallery.` You have either omitted the text or have a typo.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

Your `p` element should be below the `h2` element. You have them in the wrong order.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
