---
id: 6690be42942f421648b119eb
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

Add an `img` element below the `p` element. At this point, no image will show up in the browser.

# --hints--

Your `img` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Your `img` element should not have a closing tag `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

You should only have one `img` element. Remove any extras.

```js
assert(document.querySelectorAll('img').length === 1);
```

Your `img` element should be below the `p` element. You have them in the wrong order.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```


