---
id: 66984fe8efbd8c53a78c9b62
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

The `title` element determines what browsers show in the title bar or tab for the page.

Add a `title` element within the `head` element using the text below: 

`CatPhotoApp`

# --hints--

You have either deleted the `head` element or it is missing an opening tag or closing tag.

```js
assert(code.match(/\<head\>/) && code.match(/\<\/head\>/));
```

Your `title` element should be nested in the `head` element. Make sure to add an opening tag and closing tag for the `title` element.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/\<head\>\<title\>.*\<\/title\>\<\/head\>/));
```

Your `title` element should have a closing tag. Closing tags have a `/` just after the `<` character.

```js
assert(code.match(/\<\/title\>/));
```

Your `title` element's text should be `CatPhotoApp`. You have either omitted the text or have a typo.

```js
assert(document.title && document.title.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


