---
id: 56533eb9ac21ba0edf2244a9
title: Inizializzare le variabili con l'operatore di assegnazione
challengeType: 1
videoUrl: 'https://scrimba.com/c/cWJ4Bfb'
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

È comune <dfn>inizializzare</dfn> una variabile ad un valore iniziale nella stessa riga in cui viene dichiarata.

```js
var myVar = 0;
```

Crea una nuova variabile chiamata `myVar` e assegna ad essa un valore iniziale di `0`.

# --instructions--

Definisci una variabile `a` con `var` e inizializzala ad un valore di `9`.

# --hints--

Dovresti inizializzare `a` ad un valore di `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(code));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
