---
id: 6982684f3a25f379e195a5fc
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now that you have stored the individual costs, you can calculate the total. In Python, you use the addition operation (`+`) to sum values together. 

Update the `running_total` variable by adding `benjis_bill`, `mortys_bill`, `marcys_bill`, and `my_bill` together. 

Finally, use `print()` to display the string `Total bill so far:` followed by a space and the value of `running_total`.

# --hints--

You should update `running_total` using the `+` operator to sum all four bills. 

```js
assert.include(code, '+');
assert.include(code, 'running_total');
assert.include(code, 'benjis_bill');
assert.include(code, 'mortys_bill');
assert.include(code, 'marcys_bill');
assert.include(code, 'my_bill');
```

Your `running_total` should equal `198.83`.

```js
assert.strictEqual(running_total, 198.83);
```

You should print the string `Total bill so far:` followed by a space and the value of `running_total`.

```js
assert.include(code, 'print(');
assert.include(code, 'Total bill so far:');
assert.include(code, 'running_total');
```

# --seed--

## --seed-contents--

```py
running_total = 0

num_of_friends = 4

benjis_bill = 37.89
mortys_bill = 57.34
marcys_bill = 39.39
my_bill = 64.21

--fcc-editable-region--

--fcc-editable-region--
```
