---
id: 65687dd4b641fdc67d7de487
title: Step 29
challengeType: 20
dashedName: step-29
---

# --description--

Move the `print` call below the `number` reassignment.

# --hints--

You should move `print(number)` below `number = (number // 10) + (number % 10)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const if_statement = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");

        const { block_body } = if_statement;
        const [number_assignment, print_pos] = block_body.split('\n').filter(line => line.trim().length > 0);
        assert.match(number_assignment, /number *=/);
        assert.match(print_pos, /print\s*\(\s*number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            number = (number // 10) + (number % 10)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
