---
id: 6551ecf597038c36637f3759
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

Variables are an essential part of Python and any programming language. A variable is a name that references or points to an object. You can declare a variable by using the assignment operator `=` to assign a value like this:

```py
variable_name = value
```

Create a variable called `number` and assign it the value `5`.

# --hints--

You should declare a variable called `number`.

```js
({ test: () => assert(__userGlobals.has("number")) })
```

You should assign the value `5` to your `number` variable.

```js
({ test: () => assert.equal(__userGlobals.get("number"), 5) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
