---
id: 6551fe3b1df7c9740f13f270
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Before we go on, here are some key aspects of variable naming in Python you should keep in mind:

- Some words like `for`, `while`, `True` and others are reserved keywords. They have a special meaning in Python and you cannot use them for naming your variables.
- Variable names cannot start with a number and can only contain alpha-numeric characters and underscores.
- Variable names are case sensitive. That means `my_var` is different from `my_Var` or `MY_VAR`.
- Finally, in Python it is a common convention to write variable names using `snake_case`, with each space replaced by an underscore character and the words written in lowercase letters.

Now remove both your `print` calls, then create another variable called `alphabet` and assign it the string value of `abcdefghijklmnopqrstuvwxyz`.

# --hints--

You should declare a variable called `alphabet`.

```js
({ test: () => assert(__userGlobals.has("alphabet")) })
```

You should assign the string `abcdefghijklmnopqrstuvwxyz` to your `alphabet` variable.

```js
({ test: () => assert.equal(__userGlobals.get("alphabet"), "abcdefghijklmnopqrstuvwxyz") })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
