---
id: 6551f5756c20146fc94f8675
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

The length of a string, that is the number of characters in the string, is stored as an *attribute* associated with the string. To access the value of this attribute you can use the built-in `len()` function. 

Modify your existing `print` call to print `len(text)`.

# --hints--


You should call the `len()` function.

```js
({ test: () => assert.match(code, /len\s*\(.*\)/) })

```

You should pass `text` to your `len()` function.

```js
({ test: () => assert.match(code, /len\s*\(\s*text\s*\)/) })

```

You should pass `len(text)` to your `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
