---
id: 6551ef504d91936d2d4e54f8
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

You can use the built-in function `print()` to print the output of a piece of code on the terminal. You will know more about functions soon. For now, go to a new line and print your `text` variable to the screen by including the variable name between the opening and closing parenthesis of the `print()` function.


# --hints--

You should call the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(.*\)/) })

```

You should pass `text` to your `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*text\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
--fcc-editable-region--
```
