---
id: 6552008c0d9d9075cbec9772
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

In this project, you are going to build a cipher. You want to take each letter in your plain message, see the actual position in the alphabet, take the letter located after 3 positions, and replace the original letter with that. This kind of cipher is called *Caesar* cipher, from Julius Caesar.

Let's start by finding the position occupied by the first letter in our string. One way is to use the built-in `find()` function. The syntax is the following:

```py
a_string.find(char)
```

In the example above, `char` is the character you want to locate and `a_string` is the string you want to parse.

At the end of your code, call `find()` on your `alphabet` string and pass `text[0]` to the function.

# --hints--

You should call the `find()` function.

```js
({ test: () => assert.match(code, /find\s*\(.*\)/) })

```

You should call the `find()` function on the `alphabet` variable.

```js
({ test: () => assert.match(code, /alphabet\.find\s*\(.*\)/) })

```

You should call the `find()` function on the `alphabet` variable and pass `text[0]` to the function.

```js
({ test: () => assert.match(code, /alphabet\.find\s*\(\s*text\[\s*\s*0\s*\]\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
