---
id: 65520f8282faf57a0db4f7fe
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

As you can see from the output, `"h"` is at index `7` inside the `alphabet` string. Now you need to find the letter at index `7` plus `shift`. For that, you can use the addition operator, `+`, in the same way you would use it for a mathematical addition.

Create a variable named `shifted` and assign it the `alphabet` letter at `index` plus `shift`.

# --hints--

You should declare a variable called `shifted`.

```js
({ test: () => assert(__userGlobals.has("shifted")) })
```

You should assign the `alphabet` letter at the index `index + shift` to your new variable.

```js
({ test: () => assert.match(code, /shifted\s*=\s*alphabet\s*\[\s*index\s*\+\s*shift\s*\]/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
