---
id: 655261b2e1f2c197093f3993
title: Step 29
challengeType: 20
dashedName: step-29
---

# --description--

There are a few remaining challenges. For instance, spaces consistently get encrypted as `c`. To maintain the original spacing in the plain message, you'll require a conditional `if` statement. This is composed of the `if` keyword, a condition, and a colon `:`.

```py
if condition:
    <code>
```

Remember, the code to be executed when the condition is met should be indented one level beneath the `if` statement.

At the top of your for loop, replace `print(char == ' ')` with an `if` statement that triggers when `char` is an empty space. Inside the loop body, print the string `space!`. Remember to indent this line.

# --hints--

You should remove `print(char == ' ')`.

```js
({ test: () => assert.isFalse(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/.test(code)) })
```

You should have an `if` statement that triggers when `char` is equal to an empty space.

```js
({ test: () => assert.match(code, /if\s+char\s*==\s*("|')\s\1\s*:/) })
```

You should print `space!` inside your `if` statement.

```js
({ test: () => assert.match(e.code.original["main.py"], /(\s*)if\s+char\s*==\s*("|')\s\2\s*:\s*?\1\1print\s*\(\s*("|')space!\3\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)    
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
