---
id: 65521203d7165c7b84b22ad4
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

Repeating this process for each character in `text` can be tedious. Thankfully, you can simplify it using a loop. A loop allows you to systematically go through a sequence of elements and execute actions on each one.

In this case, you'll employ a `for` loop. The syntax is as follows:

```py
for i in iterable:
    <code>
```

`for` is the keyword denoting the loop type. `i` sequentially takes the value of the elements in `iterable`. It can have any valid name, but it's convenient to give it a meaningful name. `iterable` is the sequence you want to iterate over. Last but not least, the code to execute at each iteration constitutes the body of the loop and must be indented — typically 4 spaces — from the loop level. This indented level is a code block.

Remove everything after the `alphabet` line. Then create a `for` loop to iterate over `text`. Inside the body print each character of your `text`.

# --hints--

You should create a `for` loop to iterate over `text` and print each its characters.

```js
({ test: () => assert.match(code, /for\s+(.*)\s+in\s+text:\s+print\s*\(\s*\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
