---
id: 65525e359ca28d938baa82c5
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

The addition assignment operator, `+=` has the following syntax:

```py
a += b
```

Its effect is equivalent to `a = a + b`. So it enables you to add something to the value of a variable and then assign the result to that variable.

In your case, you can use it to add characters to a string variable and assign the result to that variable, building the string gradually.

Instead of assigning `alphabet[new_index]` to `new_char`, use the addition assignment operator to add it and assign it to `encrypted_text`. Then, modify the `print` call to reflect this change.

# --hints--

You should use the addition assignment operator to assign `alphabet[new_index]` to `encrypted_text`.

```js
({ test: () => assert.match(code, /encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/s) })
```

You should use the addition assignment operator on the `encrypted_text` variable inside your loop body.

```js
({ test: () => assert.match(code, /:(.*?)index.*\encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/s) })
```

You should turn your `print` call into `print('char:', char, 'encrypted text:', encrypted_text)`.

```js
({
  test: () => {
    assert.isFalse(/print\s*\(\s*new_char\s*\)/.test(code));
    assert.match(code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)    
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
