---
id: 6553a7d8d05cbb1ae335a665
title: Step 33
challengeType: 20
dashedName: step-33
---

# --description--

This line is the culprit: `encrypted_text += alphabet[new_index]`. When the loop reaches the letter `Z`, the sum `index + shift` exceeds the `alphabet` length.

For that the modulo operator, `%`, can come in handy. The modulo operator returns the remainder of the division between two numbers. For example: `5 % 2` is equal to `1`, because 5 divided by 2 has a quotient of 2 and a remainder of 1.

Surround `index + shift` with parentheses and divide that expression modulo `26`, which is the `alphabet` length.

# --hints--

You should enclose `index + shift` within parentheses and divide the expression modulo (`%`) `26`.

```js
({ test: () => assert.match(code, /new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)    
        new_index = index + shift
        encrypted_text += alphabet[new_index]
        print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
