---
id: 6553f8c570f9982e013a8886
title: Step 43
challengeType: 20
dashedName: step-43
---

# --description--

Let's make the things more complex. Currently, each single letter is always encrypted with the same letter, depending on the specified offset. What if the offset is different for each letter? That would be much more difficult to decrypt.
The Vigenère cipher can be considered the combination of many Caesar ciphers, where the offset for each letter is determined by another text, called *key*.

From now on you are going to transform your Caesar cipher into a Vigenère cipher. For now, remove the two function calls.

# --hints--

You should remove your two function calls.

```js
({ test: () => {
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(e.code.original["main.py"]))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(e.code.original["main.py"]))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)    
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
