---
id: 6553f6086add4b2cbb99fd78
title: Step 51
challengeType: 20
dashedName: step-51
---

# --description--

At the bottom of your code, call again your function. This time pass in the `text` variable and the integer `13` as the arguments.

# --hints--

You should call again your function, this time passing `text` and `13` as the arguments.

```js
({ test: () => assert.match(e.code.original["main.py"], /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

You should not remove the existing function call.

```js
({ test: () => assert.match(e.code.original["main.py"], /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)    
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
