---
id: 6553995f412dd8122ed38e4a
title: Step 40
challengeType: 20
dashedName: step-40
---

# --description--

Another kind of conditional statement is the `else` statement. It is used in combination with an `if` statement to alternatively execute code when the `if` condition is not met:

```py
if condition:
    <code>
else:
    <code>
```

As you can see in your output, when the loop iterations reach the space, a space is added to the encrypted string. Then the code outside the `if` block executes and a `c` is added to the encrypted string.

To fix it, add an `else` statement after `encrypted_text += char` and indent all the subsequent lines of code.

# --hints--

You should create an `else` statement.

```js
({ test: () => assert.match(code, /else\s*:/) })
```

You should indent the lines after your `else` statement.

```js
({ test: () => assert.match(e.code.original["main.py"], /(\s*)else\s*:\s*?^\1\1\w.*?\n\s*?^\1\1\w.*?\n\s*?^\1\1\w.*?\n\s*?^\1\1\w.*?\n*/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)    
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
