---
id: 6553f3fc92741c2bf8ded140
title: Step 49
challengeType: 20
dashedName: step-49
---

# --description--

Inside your function body, replace the `text` and `shift` variables with `message` and `offset`, respectively.

# --hints--

You should replace `text` with `message`.

```js
assert.match(code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*:/)
```

You should replace `shift` with `offset`.

```js
assert.match(code, /new_index\s*=\s*\(\s*index\s*\+\s*offset\s*\)\s*%\s*len\s*\(\s*alphabet\s*\)/)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)    
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
--fcc-editable-region--
```
