---
id: 6555e12341d1c318787e28cc
title: Step 86
challengeType: 20
dashedName: step-86
---

# --description--

It's time to fix those function calls.

Delete `encryption` and pass `text` as the first argument to `decrypt`.

# --hints--

You should delete the `encryption` variable and its value.

```js
({ test: () => assert.isFalse(/encryption\s*=\s*encrypt\s*\(\s*text\s*,\s*custom_key\s*\)/.test(code)) })
```

Your `decryption` variable should have the value of `decrypt(text, custom_key)`.

```js
({ test: () => assert.match(code, /^decryption\s*=\s*decrypt\s*\(\s*text\s*,\s*custom_key\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
encryption = encrypt(text, custom_key)
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
decryption = decrypt(encryption, custom_key)
print('Decrypted text: ' + decryption)
--fcc-editable-region--
```
