---
id: 65688a50e6c998a21d8e41d3
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

Strings are immutable, which means they cannot be changed once created.

Use bracket notation to access the first letter in `text` and try to change it into a character of your choice.

# --hints--

You should assign a character of your choice to `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
