---
id: 65687d2f8c7ee27b0446cef3
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

The code to execute at each iteration — placed after the `:` — constitutes the body of the loop. This code must be indented — typically 4 spaces — from the loop level. This indented level is a code block.

```py
for i in text:
    <code>
```

Give your `for` loop a body by adding a `print(i)` call. Remember to indent the loop body.

# --hints--

You should add `print(i)` to your `for` loop body. Pay attention to the indentation. 

```js
({ test: () => assert.match(e.code.original["main.py"], /for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
