---
id: 587d7fab367417b2b2512bd9
title: Додайте позначки до кругів точкової діаграми
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

Ви можете додати текст, щоб створити позначки для точок у точковій діаграмі.

Для цього потрібно зобразити значення, розділені комами, для першого (`x`) та другого (`y`) полів кожного елемента в `dataset`.

Вузлам `text` потрібні атрибути `x` та `y`, щоб розташуватись на SVG. У цьому завданні значення `y` (яке визначає висоту) може використати те ж саме значення, що й `circle` використовує для атрибута `cy`. Значення `x` може бути дещо більшим, ніж значення `cx` елемента `circle`, тому позначку видно. Це змістить позначку праворуч від нанесеної точки.

# --instructions--

Позначте кожну точку на точковій діаграмі за допомогою елементів `text`. Текст позначки повинен складатись з двох значень, розділених комою та пробілом. Наприклад, позначкою для першої точки є `34, 78`. Встановіть атрибут `x` так, щоб його значення було на `5` більше за значення, використане для атрибута `cx` в `circle`. Встановіть атрибут `y` так само, як і значення `cy` в `circle`.

# --hints--

Код має містити 10 елементів `text`.

```js
assert($('text').length == 10);
```

Перша позначка повинна мати текст `34, 78`, `x` зі значенням `39` та `y` зі значенням `422`.

```js
assert(
  $('text').eq(0).text() == '34, 78' &&
    $('text').eq(0).attr('x') == '39' &&
    $('text').eq(0).attr('y') == '422'
);
```

Друга позначка повинна мати текст `109, 280`, `x` зі значенням `114` та `y` зі значенням `220`.

```js
assert(
  $('text').eq(1).text() == '109, 280' &&
    $('text').eq(1).attr('x') == '114' &&
    $('text').eq(1).attr('y') == '220'
);
```

Третя позначка повинна мати текст `310, 120`, `x` зі значенням `315` та `y` зі значенням `380`.

```js
assert(
  $('text').eq(2).text() == '310, 120' &&
    $('text').eq(2).attr('x') == '315' &&
    $('text').eq(2).attr('y') == '380'
);
```

Четверта позначка повинна мати текст `79, 411`, `x` зі значенням `84` та `y` зі значенням `89`.

```js
assert(
  $('text').eq(3).text() == '79, 411' &&
    $('text').eq(3).attr('x') == '84' &&
    $('text').eq(3).attr('y') == '89'
);
```

П’ята позначка повинна мати текст `420, 220`, `x` зі значенням `425` та `y` зі значенням `280`.

```js
assert(
  $('text').eq(4).text() == '420, 220' &&
    $('text').eq(4).attr('x') == '425' &&
    $('text').eq(4).attr('y') == '280'
);
```

Шоста позначка повинна мати текст `233, 145`, `x` зі значенням `238` та `y` зі значенням `355`.

```js
assert(
  $('text').eq(5).text() == '233, 145' &&
    $('text').eq(5).attr('x') == '238' &&
    $('text').eq(5).attr('y') == '355'
);
```

Сьома позначка повинна мати текст `333, 96`, `x` зі значенням `338` та `y` зі значенням `404`.

```js
assert(
  $('text').eq(6).text() == '333, 96' &&
    $('text').eq(6).attr('x') == '338' &&
    $('text').eq(6).attr('y') == '404'
);
```

Восьма позначка повинна мати текст `222, 333`, `x` зі значенням `227` та `y` зі значенням `167`.

```js
assert(
  $('text').eq(7).text() == '222, 333' &&
    $('text').eq(7).attr('x') == '227' &&
    $('text').eq(7).attr('y') == '167'
);
```

Дев’ята позначка повинна мати текст `78, 320`, `x` зі значенням `83` та `y` зі значенням `180`.

```js
assert(
  $('text').eq(8).text() == '78, 320' &&
    $('text').eq(8).attr('x') == '83' &&
    $('text').eq(8).attr('y') == '180'
);
```

Десята позначка повинна мати текст `21, 123`, `x` зі значенням `26` та `y` зі значенням `377`.

```js
assert(
  $('text').eq(9).text() == '21, 123' &&
    $('text').eq(9).attr('x') == '26' &&
    $('text').eq(9).attr('y') == '377'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
                  [ 34,    78 ],
                  [ 109,   280 ],
                  [ 310,   120 ],
                  [ 79,    411 ],
                  [ 420,   220 ],
                  [ 233,   145 ],
                  [ 333,   96 ],
                  [ 222,   333 ],
                  [ 78,    320 ],
                  [ 21,    123 ]
                ];


    const w = 500;
    const h = 500;

    const svg = d3.select("body")
                  .append("svg")
                  .attr("width", w)
                  .attr("height", h);

    svg.selectAll("circle")
       .data(dataset)
       .enter()
       .append("circle")
       .attr("cx", (d, i) => d[0])
       .attr("cy", (d, i) => h - d[1])
       .attr("r", 5);

    svg.selectAll("text")
       .data(dataset)
       .enter()
       .append("text")
       // Add your code below this line



       // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
                  [ 34,    78 ],
                  [ 109,   280 ],
                  [ 310,   120 ],
                  [ 79,    411 ],
                  [ 420,   220 ],
                  [ 233,   145 ],
                  [ 333,   96 ],
                  [ 222,   333 ],
                  [ 78,    320 ],
                  [ 21,    123 ]
                ];


    const w = 500;
    const h = 500;

    const svg = d3.select("body")
                  .append("svg")
                  .attr("width", w)
                  .attr("height", h);

    svg.selectAll("circle")
       .data(dataset)
       .enter()
       .append("circle")
       .attr("cx", (d, i) => d[0])
       .attr("cy", (d, i) => h - d[1])
       .attr("r", 5);

    svg.selectAll("text")
       .data(dataset)
       .enter()
       .append("text")
       .attr("x", (d) => d[0] + 5)
       .attr("y", (d) => h - d[1])
       .text((d) => (d[0] + ", " + d[1]))

  </script>
</body>
```
