---
id: 686d749d6f32f708fdf2b7ad
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In a prior lecture, you learned about `SVG` elements. In the real world, they are not usually
made from scratch. However for this workshop, you will be writing one by hand.

Start by creating by an `SVG` element on the page.

# --hints--

You should have an `SVG` element on the page.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
