---
id: 66faa30000000000000000b7
title: Step 17
challengeType: 0
dashedName: step-17
demoType: onLoad
---

# --description--

On desktop, the virtual buttons clutter the canvas. Detect mobile and hide the buttons unless the user is on Android or iOS:

```js
this.isMobile = Phaser.Device.os.android || Phaser.Device.os.iOS;
this.leftBtn.setVisible(this.isMobile);
this.rightBtn.setVisible(this.isMobile);
```

`Phaser.Device.os` exposes the runtime's best guess at what platform you're on. The two flags here cover the dominant touch platforms; `desktop` and others will get `false`, hiding the buttons.

# --hints--

You should detect mobile via `Phaser.Device.os.android || Phaser.Device.os.iOS` and store the result on `this.isMobile`.

```js
assert.match(code, /this\.isMobile\s*=\s*Phaser\.Device\.os\.android\s*\|\|\s*Phaser\.Device\.os\.iOS/);
```

You should call `this.leftBtn.setVisible(this.isMobile)` and `this.rightBtn.setVisible(this.isMobile)`.

```js
assert.match(code, /this\.leftBtn\.setVisible\(\s*this\.isMobile\s*\)/);
assert.match(code, /this\.rightBtn\.setVisible\(\s*this\.isMobile\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.thumbZone = this.add.rectangle(240, 600, 480, 80, 0xffffff, 0.04);
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
    this.paddle.on("drag", (pointer, dragX) => {
      this.paddle.y = 580;
      this.targetX = dragX;
    });
    this.leftBtn = this.add.rectangle(40, 320, 60, 60, 0x556677);
    this.leftBtn.setInteractive();
    this.rightBtn = this.add.rectangle(440, 320, 60, 60, 0x556677);
    this.rightBtn.setInteractive();
    this.leftBtn.on("pointerdown", () => {
      this.targetX -= 60;
    });
    this.rightBtn.on("pointerdown", () => {
      this.targetX += 60;
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
