---
id: 66faa30000000000000000ab
title: Step 5
challengeType: 0
dashedName: step-5
demoType: onLoad
---

# --description--

Without a clamp the paddle will slide off either side and disappear. After the arrow checks, force `this.paddle.x` to stay between `40` and `440`: `this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);`.

`Phaser.Math.Clamp(value, min, max)` is the workhorse you'll use any time you need a value to stay inside a range without the noise of a `Math.min(Math.max(...))` sandwich.

# --hints--

You should clamp `this.paddle.x` between `40` and `440` using `Phaser.Math.Clamp`.

```js
assert.match(code, /this\.paddle\.x\s*=\s*Phaser\.Math\.Clamp\(\s*this\.paddle\.x\s*,\s*40\s*,\s*440\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
  }

  update() {
    if (this.cursors.left.isDown) this.paddle.x -= 4;
    else if (this.cursors.right.isDown) this.paddle.x += 4;
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
