---
id: 66faa30000000000000000b9
title: Step 19
challengeType: 0
dashedName: step-19
demoType: onLoad
---

# --description--

A frantic finger can mash a button several times in one frame, and on some devices `pointerdown` fires twice for one tap. Add a 100 ms debounce shared between both buttons so back-to-back triggers within that window are dropped:

```js
this.lastButtonAt = 0;
this.leftBtn.on("pointerdown", () => {
  if (this.time.now - this.lastButtonAt < 100) return;
  this.lastButtonAt = this.time.now;
  this.targetX -= 60;
});
this.rightBtn.on("pointerdown", () => {
  if (this.time.now - this.lastButtonAt < 100) return;
  this.lastButtonAt = this.time.now;
  this.targetX += 60;
});
```

The same `lastButtonAt` is shared between both buttons — pressing left then right in quick succession is also rate-limited, which feels more deliberate.

# --hints--

You should declare `this.lastButtonAt = 0` before the button handlers.

```js
assert.match(code, /this\.lastButtonAt\s*=\s*0/);
```

The left-button handler should early-return if `this.time.now - this.lastButtonAt < 100`.

```js
assert.match(code, /this\.leftBtn\.on\(\s*["']pointerdown["'][\s\S]*?if\s*\(\s*this\.time\.now\s*-\s*this\.lastButtonAt\s*<\s*100\s*\)\s*return/);
```

The right-button handler should also early-return on the same debounce window.

```js
assert.match(code, /this\.rightBtn\.on\(\s*["']pointerdown["'][\s\S]*?if\s*\(\s*this\.time\.now\s*-\s*this\.lastButtonAt\s*<\s*100\s*\)\s*return/);
```

Both handlers should update `this.lastButtonAt` to `this.time.now` before mutating `this.targetX`.

```js
assert.match(code, /this\.lastButtonAt\s*=\s*this\.time\.now[\s\S]*?this\.targetX\s*-=\s*60/);
assert.match(code, /this\.lastButtonAt\s*=\s*this\.time\.now[\s\S]*?this\.targetX\s*\+=\s*60/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.thumbZone = this.add.rectangle(240, 600, 480, 80, 0xffffff, 0.04);
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
    this.paddle.on("drag", (pointer, dragX) => {
      this.paddle.y = 580;
      this.targetX = dragX;
    });
    this.leftBtn = this.add.rectangle(40, 320, 60, 60, 0x556677);
    this.leftBtn.setInteractive();
    this.rightBtn = this.add.rectangle(440, 320, 60, 60, 0x556677);
    this.rightBtn.setInteractive();
--fcc-editable-region--
    this.leftBtn.on("pointerdown", () => {
      this.targetX -= 60;
    });
    this.rightBtn.on("pointerdown", () => {
      this.targetX += 60;
    });
--fcc-editable-region--
    this.isMobile = Phaser.Device.os.android || Phaser.Device.os.iOS;
    this.leftBtn.setVisible(this.isMobile);
    this.rightBtn.setVisible(this.isMobile);
    this.swipeStart = null;
    this.input.on("pointerdown", pointer => {
      this.swipeStart = { x: pointer.x, t: this.time.now };
    });
    this.input.on("pointerup", pointer => {
      if (!this.swipeStart) return;
      const dx = pointer.x - this.swipeStart.x;
      const dt = this.time.now - this.swipeStart.t;
      if (Math.abs(dx) > 80 && dt < 250) {
        this.targetX = dx > 0 ? 440 : 40;
      }
      this.swipeStart = null;
    });
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
