---
id: 66faa30000000000000000b3
title: Step 13
challengeType: 0
dashedName: step-13
demoType: onLoad
---

# --description--

The paddle should ride a horizontal track, not float freely. Replace the drag handler so it pins `paddle.y` to `580` and only updates `targetX`:

```js
this.paddle.on("drag", (pointer, dragX) => {
  this.paddle.y = 580;
  this.targetX = dragX;
});
```

Two things changed: the callback drops the `dragY` parameter (we don't use it), and the body snaps `paddle.y` back to its lane. The lerp in `update` still drives the paddle's `x`, so motion stays smooth.

# --hints--

The drag handler should pin `this.paddle.y` to `580`.

```js
assert.match(code, /this\.paddle\.on\(\s*["']drag["'][\s\S]*?this\.paddle\.y\s*=\s*580/);
```

The drag handler should still set `this.targetX = dragX`.

```js
assert.match(code, /this\.paddle\.on\(\s*["']drag["'][\s\S]*?this\.targetX\s*=\s*dragX/);
```

The drag handler should no longer call `setPosition` — only `y` and `targetX` change.

```js
assert.notMatch(code, /this\.paddle\.on\(\s*["']drag["'][\s\S]*?setPosition\(/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      this.targetX = pointer.x;
    });
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
--fcc-editable-region--
    this.paddle.on("drag", (pointer, dragX, dragY) => {
      this.paddle.setPosition(dragX, dragY);
      this.targetX = dragX;
    });
--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
