---
id: 66faa30000000000000000a7
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome! In this workshop you'll build a paddle prototype that responds equally well to keyboard arrows, mouse clicks, finger drags, and on-screen virtual buttons. Then you'll layer a falling-target catch game on top so the paddle has something to do.

Start by declaring a class named `MainScene` that extends `Phaser.Scene` and gives the scene the key `"MainScene"` in its constructor. Add empty `create()` and `update()` methods.

Below the class, declare a `config` object with `type: Phaser.AUTO`, `width: 480`, `height: 640`, `parent: "game-container"`, `backgroundColor: "#000033"`, the standard arcade-physics block (`{ default: "arcade", arcade: { gravity: { y: 0 } } }`), and `scene: [MainScene]`. Then start the game with `new Phaser.Game(config)`.

# --hints--

You should declare a class named `MainScene` that extends `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /constructor\s*\(\s*\)\s*\{\s*super\(\s*["']MainScene["']\s*\)/);
```

You should declare empty `create` and `update` methods.

```js
assert.match(code, /create\s*\(\s*\)\s*\{[\s\S]*?\}/);
assert.match(code, /update\s*\(\s*\)\s*\{[\s\S]*?\}/);
```

You should configure the game at 480 by 640 with arcade physics, gravity `y: 0`, and `scene: [MainScene]`.

```js
assert.match(code, /type:\s*Phaser\.AUTO/);
assert.match(code, /width:\s*480/);
assert.match(code, /height:\s*640/);
assert.match(code, /default:\s*["']arcade["']/);
assert.match(code, /gravity:\s*\{\s*y:\s*0\s*\}/);
assert.match(code, /scene:\s*\[\s*MainScene\s*\]/);
```

You should create the game with `new Phaser.Game(config)`.

```js
assert.match(code, /new\s+Phaser\.Game\(\s*config\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
