---
id: 66faa30000000000000000b5
title: Step 15
challengeType: 0
dashedName: step-15
demoType: onLoad
---

# --description--

Mobile players who don't want to drag still need a way to steer. Below the drag handler, add a pair of on-screen buttons:

```js
this.leftBtn = this.add.rectangle(40, 320, 60, 60, 0x556677);
this.leftBtn.setInteractive();
this.rightBtn = this.add.rectangle(440, 320, 60, 60, 0x556677);
this.rightBtn.setInteractive();
```

Each button is a `60` by `60` square at the left and right edges. `setInteractive()` flips on Phaser's hit testing so the button can later receive `pointerdown` events.

# --hints--

You should add `this.leftBtn` as a `60` by `60` rectangle at `(40, 320)` with color `0x556677` and call `setInteractive` on it.

```js
assert.match(code, /this\.leftBtn\s*=\s*this\.add\.rectangle\(\s*40\s*,\s*320\s*,\s*60\s*,\s*60\s*,\s*0x556677/);
assert.match(code, /this\.leftBtn\.setInteractive\(\s*\)/);
```

You should add `this.rightBtn` as a `60` by `60` rectangle at `(440, 320)` and call `setInteractive` on it.

```js
assert.match(code, /this\.rightBtn\s*=\s*this\.add\.rectangle\(\s*440\s*,\s*320\s*,\s*60\s*,\s*60\s*,\s*0x556677/);
assert.match(code, /this\.rightBtn\.setInteractive\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.thumbZone = this.add.rectangle(240, 600, 480, 80, 0xffffff, 0.04);
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
    this.paddle.on("drag", (pointer, dragX) => {
      this.paddle.y = 580;
      this.targetX = dragX;
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
