---
id: 66faa30000000000000000a8
title: Step 2
challengeType: 0
dashedName: step-2
demoType: onLoad
---

# --description--

Inside `create`, paint a dark background and a bright paddle.

Add `this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344)` to fill the canvas. Then add `this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee)` for the paddle near the bottom edge. Give the paddle an arcade physics body with `this.physics.add.existing(this.paddle)`, then mark it immovable so the falling target won't shove it: `this.paddle.body.setImmovable(true)`.

# --hints--

You should paint a 480 by 640 background rectangle centered at `(240, 320)` with color `0x223344`.

```js
assert.match(code, /this\.bg\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*320\s*,\s*480\s*,\s*640\s*,\s*0x223344/);
```

You should create a paddle rectangle at `(240, 580)` sized `80` by `16` with color `0xeeeeee`.

```js
assert.match(code, /this\.paddle\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*580\s*,\s*80\s*,\s*16\s*,\s*0xeeeeee/);
```

You should give the paddle an arcade physics body via `this.physics.add.existing` and mark it immovable.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.paddle\s*\)/);
assert.match(code, /this\.paddle\.body\.setImmovable\(\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
