---
id: 66faa30000000000000000ac
title: Step 6
challengeType: 0
dashedName: step-6
demoType: onLoad
---

# --description--

A subtle visual cue helps the player see the paddle's lane. Above the paddle creation, add a thin horizontal track rectangle: `this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);`.

Order matters here — Phaser draws objects in the order they are added, so the track sits *under* the paddle rather than over it.

# --hints--

You should create a thin track rectangle at `(240, 580)` sized `460` by `4` with color `0x556677`.

```js
assert.match(code, /this\.track\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*580\s*,\s*460\s*,\s*4\s*,\s*0x556677/);
```

The track should be created *before* the paddle so it renders underneath.

```js
assert.match(code, /this\.track\s*=[\s\S]*?this\.paddle\s*=/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
--fcc-editable-region--

--fcc-editable-region--
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
  }

  update() {
    if (this.cursors.left.isDown) this.paddle.x -= 4;
    else if (this.cursors.right.isDown) this.paddle.x += 4;
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
