---
id: 66faa30000000000000000b2
title: Step 12
challengeType: 0
dashedName: step-12
demoType: onLoad
---

# --description--

Make the paddle draggable. Below the help text, mark the paddle interactive with the `draggable` flag and listen for the per-object `drag` event:

```js
this.paddle.setInteractive({ draggable: true });
this.paddle.on("drag", (pointer, dragX, dragY) => {
  this.paddle.setPosition(dragX, dragY);
  this.targetX = dragX;
});
```

Try grabbing the paddle and tugging it around — it follows the pointer freely on both axes. The next step constrains the vertical motion.

# --hints--

You should mark the paddle interactive with `draggable: true`.

```js
assert.match(code, /this\.paddle\.setInteractive\(\s*\{\s*draggable:\s*true\s*\}\s*\)/);
```

You should attach a `drag` handler to the paddle that calls `setPosition(dragX, dragY)` and sets `this.targetX = dragX`.

```js
assert.match(code, /this\.paddle\.on\(\s*["']drag["'][\s\S]*?\(pointer,\s*dragX,\s*dragY\)[\s\S]*?this\.paddle\.setPosition\(\s*dragX\s*,\s*dragY\s*\)[\s\S]*?this\.targetX\s*=\s*dragX/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      this.targetX = pointer.x;
    });
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
