---
id: 66faa30000000000000000ad
title: Step 7
challengeType: 0
dashedName: step-7
demoType: onLoad
---

# --description--

Time to wire pointer input. Initialize a tracked target position with `this.targetX = this.paddle.x;`. Then register a `pointermove` listener that snaps the target to the cursor or finger:

```js
this.input.on("pointermove", pointer => {
  this.targetX = pointer.x;
});
```

Nothing visible changes yet — `targetX` is a value held in scene state. The next step uses it.

# --hints--

You should initialize `this.targetX` to `this.paddle.x`.

```js
assert.match(code, /this\.targetX\s*=\s*this\.paddle\.x/);
```

You should register a `pointermove` listener via `this.input.on` that sets `this.targetX = pointer.x`.

```js
assert.match(code, /this\.input\.on\(\s*["']pointermove["'][\s\S]*?this\.targetX\s*=\s*pointer\.x/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.paddle.x -= 4;
    else if (this.cursors.right.isDown) this.paddle.x += 4;
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
