---
id: 66faa30000000000000000bd
title: Step 23
challengeType: 0
dashedName: step-23
demoType: onLoad
---

# --description--

Fill in `onCatch` so the catch is satisfying. Inside the method body, pulse the paddle with a tween and bump the score:

```js
this.tweens.add({
  targets: paddle,
  scaleX: 1.4,
  scaleY: 1.4,
  duration: 100,
  yoyo: true
});
this.score += 1;
this.scoreText.setText("Score: " + this.score);
```

`yoyo: true` runs the tween forward and then back, so the paddle expands and snaps back over 200 ms total. The HUD updates to reflect the new score.

# --hints--

The `onCatch` method should add a tween targeting `paddle` that scales it to `1.4` over `100` ms with `yoyo: true`.

```js
assert.match(code, /onCatch\s*\(\s*paddle\s*,\s*falling\s*\)[\s\S]*?this\.tweens\.add\(\s*\{[\s\S]*?targets:\s*paddle[\s\S]*?scaleX:\s*1\.4[\s\S]*?scaleY:\s*1\.4[\s\S]*?duration:\s*100[\s\S]*?yoyo:\s*true/);
```

The method should increment `this.score` and update `this.scoreText`.

```js
assert.match(code, /onCatch\s*\([\s\S]*?this\.score\s*\+=\s*1[\s\S]*?this\.scoreText\.setText\(\s*["']Score:\s*["']\s*\+\s*this\.score/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.thumbZone = this.add.rectangle(240, 600, 480, 80, 0xffffff, 0.04);
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
    this.paddle.on("drag", (pointer, dragX) => {
      this.paddle.y = 580;
      this.targetX = dragX;
    });
    this.leftBtn = this.add.rectangle(40, 320, 60, 60, 0x556677);
    this.leftBtn.setInteractive();
    this.rightBtn = this.add.rectangle(440, 320, 60, 60, 0x556677);
    this.rightBtn.setInteractive();
    this.lastButtonAt = 0;
    this.leftBtn.on("pointerdown", () => {
      if (this.time.now - this.lastButtonAt < 100) return;
      this.lastButtonAt = this.time.now;
      this.targetX -= 60;
    });
    this.rightBtn.on("pointerdown", () => {
      if (this.time.now - this.lastButtonAt < 100) return;
      this.lastButtonAt = this.time.now;
      this.targetX += 60;
    });
    this.isMobile = Phaser.Device.os.android || Phaser.Device.os.iOS;
    this.leftBtn.setVisible(this.isMobile);
    this.rightBtn.setVisible(this.isMobile);
    this.swipeStart = null;
    this.input.on("pointerdown", pointer => {
      this.swipeStart = { x: pointer.x, t: this.time.now };
    });
    this.input.on("pointerup", pointer => {
      if (!this.swipeStart) return;
      const dx = pointer.x - this.swipeStart.x;
      const dt = this.time.now - this.swipeStart.t;
      if (Math.abs(dx) > 80 && dt < 250) {
        this.targetX = dx > 0 ? 440 : 40;
      }
      this.swipeStart = null;
    });
    this.falling = this.add.rectangle(
      Phaser.Math.Between(40, 440),
      0,
      16,
      16,
      0xffaa44
    );
    this.physics.add.existing(this.falling);
    this.falling.body.setVelocityY(120);
    this.physics.add.overlap(
      this.paddle,
      this.falling,
      this.onCatch,
      null,
      this
    );
    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Score: 0", {
      fontSize: "20px",
      fill: "#ffffff"
    });
    this.scoreText.setScrollFactor(0);
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }

--fcc-editable-region--
  onCatch(paddle, falling) {
  }
--fcc-editable-region--
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
