---
id: 66faa30000000000000000ae
title: Step 8
challengeType: 0
dashedName: step-8
demoType: onLoad
---

# --description--

Make `targetX` matter. After the clamp in `update`, ease the paddle toward the target each frame: `this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);`.

`Phaser.Math.Linear(a, b, t)` returns `a + (b - a) * t`. With `t = 0.2`, the paddle moves 20% of the remaining distance per frame — smooth without lag.

You'll notice the keyboard now feels off: arrows still nudge `paddle.x`, but the lerp pulls it back toward whatever `targetX` was last set to. Step 10 fixes that. Drag the mouse for now and the paddle chases your cursor.

# --hints--

You should ease `this.paddle.x` toward `this.targetX` with `Phaser.Math.Linear` and a factor of `0.2`.

```js
assert.match(code, /this\.paddle\.x\s*=\s*Phaser\.Math\.Linear\(\s*this\.paddle\.x\s*,\s*this\.targetX\s*,\s*0\.2\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      this.targetX = pointer.x;
    });
  }

  update() {
    if (this.cursors.left.isDown) this.paddle.x -= 4;
    else if (this.cursors.right.isDown) this.paddle.x += 4;
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
