---
id: 66faa30000000000000000aa
title: Step 4
challengeType: 0
dashedName: step-4
demoType: onLoad
---

# --description--

Move the paddle in `update`. When the left arrow is held, slide the paddle left four pixels per frame: `if (this.cursors.left.isDown) this.paddle.x -= 4;`. Then add the matching `else if` branch for the right arrow that bumps `this.paddle.x` by `+4`.

Open the live preview, click into it, and tap the arrow keys — the paddle should glide left and right.

# --hints--

You should subtract `4` from `this.paddle.x` while `this.cursors.left.isDown`.

```js
assert.match(code, /if\s*\(\s*this\.cursors\.left\.isDown\s*\)[\s\S]*?this\.paddle\.x\s*-=\s*4/);
```

You should add `4` to `this.paddle.x` while `this.cursors.right.isDown` in the matching `else if`.

```js
assert.match(code, /else\s+if\s*\(\s*this\.cursors\.right\.isDown\s*\)[\s\S]*?this\.paddle\.x\s*\+=\s*4/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
  }

  update() {
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
