---
id: 66faa30000000000000000b8
title: Step 18
challengeType: 0
dashedName: step-18
demoType: onLoad
---

# --description--

Now add a swipe gesture so a quick horizontal flick dashes the paddle to that side. Below the mobile visibility block, capture the start of every press and measure the release:

```js
this.swipeStart = null;
this.input.on("pointerdown", pointer => {
  this.swipeStart = { x: pointer.x, t: this.time.now };
});
this.input.on("pointerup", pointer => {
  if (!this.swipeStart) return;
  const dx = pointer.x - this.swipeStart.x;
  const dt = this.time.now - this.swipeStart.t;
  if (Math.abs(dx) > 80 && dt < 250) {
    this.targetX = dx > 0 ? 440 : 40;
  }
  this.swipeStart = null;
});
```

Phaser dispatches multiple listeners for the same event, so this `pointerdown` listener coexists with the thumb-zone listener you added earlier — they both run.

# --hints--

You should initialize `this.swipeStart` to `null`.

```js
assert.match(code, /this\.swipeStart\s*=\s*null/);
```

You should record the swipe origin on `pointerdown` with both `x` and a timestamp from `this.time.now`.

```js
assert.match(code, /this\.input\.on\(\s*["']pointerdown["'][\s\S]*?this\.swipeStart\s*=\s*\{\s*x:\s*pointer\.x\s*,\s*t:\s*this\.time\.now\s*\}/);
```

You should detect a swipe on `pointerup` using a distance threshold of `80` and a time threshold of `250`.

```js
assert.match(code, /this\.input\.on\(\s*["']pointerup["'][\s\S]*?Math\.abs\(\s*dx\s*\)\s*>\s*80\s*&&\s*dt\s*<\s*250/);
```

A right-going swipe should set `this.targetX` to `440`; a left-going swipe to `40`.

```js
assert.match(code, /this\.targetX\s*=\s*dx\s*>\s*0\s*\?\s*440\s*:\s*40/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.thumbZone = this.add.rectangle(240, 600, 480, 80, 0xffffff, 0.04);
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
    this.paddle.on("drag", (pointer, dragX) => {
      this.paddle.y = 580;
      this.targetX = dragX;
    });
    this.leftBtn = this.add.rectangle(40, 320, 60, 60, 0x556677);
    this.leftBtn.setInteractive();
    this.rightBtn = this.add.rectangle(440, 320, 60, 60, 0x556677);
    this.rightBtn.setInteractive();
    this.leftBtn.on("pointerdown", () => {
      this.targetX -= 60;
    });
    this.rightBtn.on("pointerdown", () => {
      this.targetX += 60;
    });
    this.isMobile = Phaser.Device.os.android || Phaser.Device.os.iOS;
    this.leftBtn.setVisible(this.isMobile);
    this.rightBtn.setVisible(this.isMobile);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
