---
id: 66faa30000000000000000bc
title: Step 22
challengeType: 0
dashedName: step-22
demoType: onLoad
---

# --description--

Wire the overlap and the score HUD. Right after the falling rectangle block, add:

```js
this.physics.add.overlap(
  this.paddle,
  this.falling,
  this.onCatch,
  null,
  this
);
this.score = 0;
this.scoreText = this.add.text(10, 10, "Score: 0", {
  fontSize: "20px",
  fill: "#ffffff"
});
this.scoreText.setScrollFactor(0);
```

The overlap's fifth argument (`this`) sets the scope `onCatch` runs with, so inside the method `this.score` and `this.scoreText` resolve correctly.

`setScrollFactor(0)` locks the score text to camera coordinates — it stays at the top-left no matter how the world scrolls (relevant once you reach the camera chapter).

# --hints--

You should call `this.physics.add.overlap` with `this.paddle`, `this.falling`, `this.onCatch`, `null`, and `this`.

```js
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.paddle\s*,\s*this\.falling\s*,\s*this\.onCatch\s*,\s*null\s*,\s*this\s*\)/);
```

You should initialize `this.score = 0` and create `this.scoreText` reading `Score: 0`.

```js
assert.match(code, /this\.score\s*=\s*0/);
assert.match(code, /this\.scoreText\s*=\s*this\.add\.text\(\s*10\s*,\s*10\s*,\s*["']Score:\s*0["']/);
```

You should call `this.scoreText.setScrollFactor(0)`.

```js
assert.match(code, /this\.scoreText\.setScrollFactor\(\s*0\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.thumbZone = this.add.rectangle(240, 600, 480, 80, 0xffffff, 0.04);
    this.helpText = this.add
      .text(240, 40, "DRAG to move \u2022 ARROWS to move", {
        fontSize: "14px",
        fill: "#ffffff"
      })
      .setOrigin(0.5);
    this.paddle.setInteractive({ draggable: true });
    this.paddle.on("drag", (pointer, dragX) => {
      this.paddle.y = 580;
      this.targetX = dragX;
    });
    this.leftBtn = this.add.rectangle(40, 320, 60, 60, 0x556677);
    this.leftBtn.setInteractive();
    this.rightBtn = this.add.rectangle(440, 320, 60, 60, 0x556677);
    this.rightBtn.setInteractive();
    this.lastButtonAt = 0;
    this.leftBtn.on("pointerdown", () => {
      if (this.time.now - this.lastButtonAt < 100) return;
      this.lastButtonAt = this.time.now;
      this.targetX -= 60;
    });
    this.rightBtn.on("pointerdown", () => {
      if (this.time.now - this.lastButtonAt < 100) return;
      this.lastButtonAt = this.time.now;
      this.targetX += 60;
    });
    this.isMobile = Phaser.Device.os.android || Phaser.Device.os.iOS;
    this.leftBtn.setVisible(this.isMobile);
    this.rightBtn.setVisible(this.isMobile);
    this.swipeStart = null;
    this.input.on("pointerdown", pointer => {
      this.swipeStart = { x: pointer.x, t: this.time.now };
    });
    this.input.on("pointerup", pointer => {
      if (!this.swipeStart) return;
      const dx = pointer.x - this.swipeStart.x;
      const dt = this.time.now - this.swipeStart.t;
      if (Math.abs(dx) > 80 && dt < 250) {
        this.targetX = dx > 0 ? 440 : 40;
      }
      this.swipeStart = null;
    });
    this.falling = this.add.rectangle(
      Phaser.Math.Between(40, 440),
      0,
      16,
      16,
      0xffaa44
    );
    this.physics.add.existing(this.falling);
    this.falling.body.setVelocityY(120);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    if (this.cursors.left.isDown) this.targetX -= 4;
    else if (this.cursors.right.isDown) this.targetX += 4;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
  }

  onCatch(paddle, falling) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
