---
id: 66faa30000000000000000b0
title: Step 10
challengeType: 0
dashedName: step-10
demoType: onLoad
---

# --description--

Unify the inputs. Refactor `update` so both keyboard and pointer feed the same `targetX` channel:

- Replace `this.paddle.x -= 4` with `this.targetX -= 4` for the left arrow, and the matching `+= 4` for the right arrow.
- Clamp `this.targetX` between `40` and `440` (so arrows can't overshoot off-screen).
- Keep the `Phaser.Math.Linear` lerp line as-is.
- Add a final `this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);` after the lerp so the rendered paddle stays on-canvas even when `targetX` changes faster than the lerp can catch up.

Now the paddle responds the same whether you press arrows, click, or drag: every input writes to `targetX`, and `update` smoothly follows.

# --hints--

The keyboard handlers should mutate `this.targetX`, not `this.paddle.x`.

```js
assert.match(code, /if\s*\(\s*this\.cursors\.left\.isDown\s*\)[\s\S]*?this\.targetX\s*-=\s*4/);
assert.match(code, /else\s+if\s*\(\s*this\.cursors\.right\.isDown\s*\)[\s\S]*?this\.targetX\s*\+=\s*4/);
```

You should clamp `this.targetX` to the `[40, 440]` range *before* the lerp.

```js
assert.match(code, /this\.targetX\s*=\s*Phaser\.Math\.Clamp\(\s*this\.targetX\s*,\s*40\s*,\s*440\s*\)/);
```

You should also clamp `this.paddle.x` to `[40, 440]` *after* the lerp so the on-screen paddle never escapes the canvas.

```js
assert.match(code, /Phaser\.Math\.Linear\([\s\S]*?this\.paddle\.x\s*=\s*Phaser\.Math\.Clamp\(\s*this\.paddle\.x\s*,\s*40\s*,\s*440\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Touch Paddle Prototype with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.bg = this.add.rectangle(240, 320, 480, 640, 0x223344);
    this.track = this.add.rectangle(240, 580, 460, 4, 0x556677);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      this.targetX = pointer.x;
    });
  }

  update() {
--fcc-editable-region--
    if (this.cursors.left.isDown) this.paddle.x -= 4;
    else if (this.cursors.right.isDown) this.paddle.x += 4;
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.2);
--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
