---
id: 61329b210dac0b08047fd6ab
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بالمتابعة مع عناصر `meta`، يخبر تعريف `viewport` المتصفح كيفية عرض الصفحة. بما في ذلك تسهيل البصري (visual accessibility) على المحمول، وتحسين كفاءه محرك البحث _SEO_ أو (seach engine optimization).

قم بإضافة تعريف `viewport` مع سمة `content` مع تفاصيل `width` و `initial-scale` للصفحة.

# --hints--

يجب عليك إنشاء عنصر `meta` أخر داخل عنصر `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

يجب أن تعطي `meta` سمة `name` بقيمة `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

يجب أن تعطي `meta` سمة `content` بقيمة `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
