---
id: 61329d52e5010e08d9b9d66b
title: الخطوة ٤
challengeType: 0
dashedName: step-4
---

# --description--

عنصر `meta` آخر مهم للوصول و SEO هو `description`. يتم استخدام قيمة السمة `content` من قبل محركات البحث لتقديم وصف لصفحتك.

قم بإضافة عنصر `meta` مع تعيين السمة `name` إلى `description`، وقم بإعطائه سمة `content` مفيدة.

# --hints--

يجب عليك إنشاء عنصر `meta` جديد داخل عنصر `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

يجب أن تعطي `meta` سمة `name` بقيمة `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

يجب عليك إعطاء `meta` سمة `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

قيمة السمة `content` يجب ألا تكون أكثر من 165 حرفاً. _هذا هو الحد الأقصى لطول الوصف في جوجل_

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
