---
id: 613297a923965e0703b64796
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

قد تكون على دراية بالفعل بعنصر `meta`؛ يستخدم لتحديد معلومات حول الصفحة، مثل العنوان، الوصف، الكلمات المفتاحية، والمؤلف.

امنح صفحتك عنصر `meta` بقيمة `charset` مناسبة.

السمة `charset` تحدد ترميز الأحرف للصفحة، وفي هذه الأيام، `UTF-8` هي الترميز الوحيد المدعوم من معظم المتصفحات.

# --hints--

يجب عليك إنشاء عنصر `meta` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

يجب عليك إعطاء وسم `meta` سمة `charset` بقيمة `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
