---
id: 6133acc353338c0bba9cb553
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أخيرا في `head`، عنصر `title` مفيد لقراء الشاشة (screen readers) لفهم محتوى الصفحة. علاوة على ذلك، فهو جزء هام من _SEO_.

أعطِ صفحتك `title` وصفياً وموجزاً.

# --hints--

يجب عليك إضافة عنصر `title` إلى `head`.

```js
assert.exists(document.querySelector('head > title'));
```

يجب ألا تجعل `title` أطول من 60 حرفاً.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

حاول أن تكون أكثر وصفاً مع عنصر `title` الخاص بك. _تلميح: 15 حرف على الأقل_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
