---
id: 6133d11ef548f51f876149e3
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

التنقل (Navigation) هو جزء أساسي من تسهيل المنال (accessibility)، ويعتمد قارئ الشاشة (screen readers) عليك لتوفير بنية صفحتك. ويتم ذلك باستخدام عناصر semantic HTML.

أضف `header` و `main` إلى صفحتك.

سيتم استخدام عنصر `header` لتقديم الصفحة، بالإضافة إلى توفير قائمة التنقل (navigation menu).

عنصر `main` سوف يحتوي على المحتوى الأساسي لصفحتك.

# --hints--

يجب عليك إضافة عنصر `header` إلى `body`.

```js
assert.exists(document.querySelector('body > header'));
```

يجب عليك إضافة عنصر `main` إلى `body`.

```js
assert.exists(document.querySelector('body > main'));
```

يجب أن يأتي عنصر `header` قبل عنصر `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
