---
id: 613e275749ebd008e74bb62e
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

تعني الخاصية _SVG_ المفيدة (رُسُوم متجهة قابلة للتطوير scalable vector graphics)، هي أنها تحتوي على سمة `path` التي تسمح بتغيير حجم الصورة دون التأثير على دقة الصورة الناتجة من التغيير.

حاليًا، تأخذ `img` حجمها الافتراضي, وهو كبير جدًا. بشكل صحيح، غيّر حجم الصورة باستخدام `id` كمنتقى، وعيّن `width` ألى `max(100px, 18vw)`.

# --hints--

يجب عليك استخدام منتقي `#logo` لاستهداف عنصر `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

يجب عليك إعطاء `img` خاصية `width` بقيمة `max(100px, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
