---
id: 613e2546d0594208229ada50
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

داخل `header`، قدّم سياق حول الصفحة بواسطة أدخال عنصر `img`, و `h1`, و `nav` واحد.

ينبغي أن يشير `img` إلى `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`، ويحتوي على `id` بقيمة `logo`.

يجب أن يحتوي `h1` على النص `HTML/CSS Quiz`.

# --hints--

يجب عليك إضافة عنصر `img` إلى عنصر `header`.

```js
assert.exists(document.querySelector('header > img'));
```

يجب عليك إضافة عنصر `h1` إلى عنصر `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

يجب عليك إضافة عنصر `nav` إلى عنصر `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

يجب عليك وضع عناصر `img`, و `h1`, و `nav` بهذا الترتيب.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

يجب أن تعطي عنصر `img` سمة `src` بقيمة `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

يجب أن تعطي عنصر `img` سمة `id` بقيمة `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

يجب عليك إعطاء عنصر `h1` النص `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
