---
id: 6140827cff96e906bd38fc2b
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

كما هو موضح في <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp Style Guide</a>، يجب أن يحتفظ الشعار بنسبة عرض تبلغ `35 / 4`، وأن يحتوي على padding حول النص.

أولاً، قم بتغيير `background-color` إلى `#0a0a23` حتى تتمكن من رؤية الشعار. ثم استخدم خاصية `aspect-ratio` لتعيين نسبة العرض المطلوبة إلى `35 / 4`. أخيراً، أضف `padding` بقيمة `0.4rem` في جميع الاتجاهات.

# --hints--

يجب عليك إعطاء `#logo` الـ `background-color` بقيمة `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

يجب عليك استخدام خاصية `aspect-ratio`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

لا يجب عليك استخدام خاصية `height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

لا يجب عليك تغيير خاصية `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

يجب عليك إعطاء `img` الـ `aspect-ratio` بقيمة `35 / 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

يجب عليك إعطاء `img` الـ `padding` بقيمة `0.4rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(100px, 18vw);

}
--fcc-editable-region--

```
