---
id: 614202874ca576084fca625f
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

كل دور `region` يتطلب وصف، مما يساعد مستخدمي قارئ الشاشة على فهم غرضها. إحدى الطرق لإضافة وصف هي أن تضيف عنصر العنوان داخل المنطقة ثم الإشارة إليه مع السمة `aria-labelledby`.

أضف `aria-labelledby` التالي إلى عناصر `section`:

- `student-info`
- `html-questions`
- `css-questions`

ثم ، داخل كل عنصر `section` ادمج عنصر واحد `h2` مع `id` يطابق السمة `aria-labelledby` المقابلة. قم بإعطاء كل `h2` محتوى نصي مناسب.

# --hints--

يجب أن تعطي عنصر `section` الأول سمة `aria-labelledby` بقيمة `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

يجب أن تعطي عنصر `section` سمة `aria-labelledby` بقيمة `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

يجب عليك إعطاء ثالث عنصر `section` سمة `aria-labelledby` بقيمة `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

يجب عليك دمج عنصر `h2` واحد داخل اول عنصر `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

يجب عليك دمج عنصر `h2` واحد داخل ثاني عنصر `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

يجب عليك دمج عنصر `h2` واحد داخل ثالث عنصر `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

يجب أن تعطي عنصر `h2` سمة `id` بقيمة `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

يجب عليك إعطاء ثاني عنصر `h2`الـ `id` بقيمة `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

يجب عليك إعطاء ثالث عنصر `h2` الـ `id` بقيمة `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

يجب عليك إعطاء أول عنصر `h2` محتوى نصي مناسب. _تلميح: كنت ساختار `Student Info`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[0]?.textContent?.length, 1);
```

يجب عليك إعطاء ثاتي عنصر `h2` محتوى نصي مناسب. _تلميح: كنت ساختار `HTML`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[1]?.textContent?.length, 1);
```

يجب عليك إعطاء ثالث عنصر `h2` محتوى نصي مناسب. _تلميح: كنت ساختار `CSS`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[2]?.textContent?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
