---
id: 6605978b1b4cd18e2899c587
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

An event loop is responsible for managing the execution of asynchronous tasks. 

Remove the `pass` key word and retrieve the event loop using `asyncio.get_event_loop()` and assign it to the variable `loop`.

# --hints--

Remove the `pass` keyword.

```js
({
    test: () => assert.isFalse(runPython(`_Node(_code).find_ifs()[0].find_body().has_pass()`))
})
```

Assign `asyncio.get_event_loop()` to the variable `loop`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_variable("loop").is_equivalent("loop = asyncio.get_event_loop()")`))
})
```

# --seed--

## --seed-contents--

```py
import asyncio


async def find_primes(start, end):

    print(f"Finding prime numbers between {start} and {end}.")
    primes = []
    for val in range(start, end + 1):
        if val > 1:
            for n in range(2, val):
                if (val % n) == 0:
                    break
            else:
                primes.append(val)
                await asyncio.sleep(0.000001)

    print(f"DONE finding prime numbers between {start} and {end}. Total: {len(primes)}")
    return primes


async def main():
    primes1 = loop.create_task(find_primes(100, 10000))
    primes2 = loop.create_task(find_primes(10, 100))
    primes3 = loop.create_task(find_primes(1, 10))
    await asyncio.wait([primes1, primes2, primes3])
    return primes1, primes2, primes3

--fcc-editable-region--
if __name__ == "__main__":
    pass

--fcc-editable-region--
