---
id: 66058dac03180626609943be
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Remove the `pass` keyword and print a message indicating the start and end of the range of numbers that the function is searching for prime numbers in.

Use f-strings to print this message `'Finding prime numbers between {start} and {end}.'`

# --hints--

You should remove the `pass` keyword.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

Your print call should have `f'Finding prime numbers between {start} and {end}.'` as the argument.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

# --seed--

## --seed-contents--

```py
import asyncio

--fcc-editable-region--
async def find_primes(start, end):
    pass

--fcc-editable-region--
```
