---
id: 66054d3f3a8df3f12a749e77
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

The syntax of defining async functions is like this: `async def function_name(arg1, ...)`

Define an asynchronous function named `find_primes` that takes a `start` and an `end` value as input parameters to find prime numbers in the given range.


# --hints--

You should define a `find_primes` method.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("find_primes")`)) })
```

Use the `async` keyword before defining the function.

```js
({ test: () => assert.match(code, /async\s*def\s*find_primes/m) })

```

`find_primes` should have two parameters `start` and `end`.

```js 
({ test: () => assert(runPython(`_Node(_code).find_function("find_primes").has_args("start, end")`)) })  
```

# --seed--

## --seed-contents--

```py
import asyncio

--fcc-editable-region--

--fcc-editable-region--
```
