---
id: 66058e06f1a0ef2bcf270646
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Initialize an empty list called `primes` to store the prime numbers found in the mentioned range.

# --hints--

You should initialize the `primes` list with `primes = []`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

# --seed--

## --seed-contents--

```py
import asyncio
 
--fcc-editable-region--
async def find_primes(start, end):    
    print(f"Finding prime numbers between {start} and {end}.")
 
--fcc-editable-region--
```
