---
id: 66058eba7dca3338ad94f315
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Remove the `pass` keyword and for each value `val`, check if `val` is greater than `1`, because by definition a prime number is a natural number greater than 1 that has no positive divisors other than 1 and itself.

# --hints--

You should remove the `pass` keyword.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

The `if` condition should check              if `val > 1`.

```js
({

})
```

# --seed--

## --seed-contents--

```py
import asyncio
 
--fcc-editable-region--

async def find_primes(start, end):    
    print(f"Finding prime numbers between {start} and {end}.")
    primes = []

    for val in range(start, end + 1):
        pass
    
--fcc-editable-region--
```

