---
id: 66058fd3bcb83644da718b90
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Inside the `if` condition, remove `pass` and iterate through each value `n` using a for loop. The loop should start at `2` because the logic of finding prime numbers checks if a number is divisible by any number less than itself. The loop should end at `val` to include the end value.

# --hints--

You should remove the `pass` keyword.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

The for loop should iterate till `range(2, val)` and use `n` as the iteration variable.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

# --seed--

## --seed-contents--

```py
import asyncio
 
--fcc-editable-region--
async def find_primes(start, end):    
    print(f"Finding prime numbers between {start} and {end}.")
    primes = []

    for val in range(start, end + 1):
        if val > 1:
            pass
    
--fcc-editable-region--
```
