---
id: 6605924301176d52490e025c
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Next, you need to check if the value `val` has a divisor.

Inside the for loop, remove `pass` and check if `val` is divisible by `n`. If it is, then it means a divisor has been found, and the value is not a prime number. In this case, break out of the loop.

# --hints--

You should remove the `pass` keyword.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

The `if` condition should be placed inside the for loop and check if `val % n == 0`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

You should break out of the loop `if (val % n) == 0`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

# --seed--

## --seed-contents--

```py
import asyncio
 
--fcc-editable-region--
async def find_primes(start, end):
    print(f"Finding prime numbers between {start} and {end}.")
    primes = []

    for val in range(start, end + 1):
        if val > 1:
            for n in range(2, val):
                pass
    
--fcc-editable-region--
```
