---
id: 66058e2d68c4122efe455c7a
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Iterate through each value using a for loop. Use `val` as the iteration variable. The loop should start at `start` and end at `end + 1` to include the end value.

# --hints--

You should have a `for` loop that iterates till `range(start, end + 1)`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).`))
 })
```

# --seed--

## --seed-contents--

```py
import asyncio
 
--fcc-editable-region--
async def find_primes(start, end):    
    print(f"Finding prime numbers between {start} and {end}.")
    primes = []

--fcc-editable-region--
```
