---
id: 612e83ec2eca1e370f830511
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Adicione um elemento `link` dentro do elemento `head`. Para esse elemento `link`, defina o atributo `rel` como `stylesheet` e `href` como `./styles.css`.

# --hints--

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

Você não deve mudar as tags `head` existentes. Verifique se você não excluiu a tag de fechamento.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Você deve ter um elemento `link` de autofechamento.

```js
assert(document.querySelectorAll('link').length === 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
