---
id: 612e8279827a28352ce83a72
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Agora copie o conjunto dos sete elementos `.key` e cole mais dois conjuntos no div `.keys`.

# --hints--

Você deve ter 21 elementos `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Você deve ter 15 elementos `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Todos os elementos `.key` devem estar dentro do seu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Todos os elementos `.black--key` devem estar dentro do seu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
