---
id: 612ea4c4993aba52ab4aa32e
title: Passo 18
challengeType: 0
dashedName: step-18
---

# --description--

Agora é hora de usar os pseudosseletores que você preparou anteriormente. Para criar as teclas pretas, adicione um novo seletor `.key.black--key::after`. Ele terá como destino os elementos com a classe `key black--key`e selecionará o pseudoelemento após esses elementos no HTML.

No novo seletor, defina a `background-color` como `#1d1e22`. Defina também a propriedade `content` para `""`. Isto deixará os pseudoelementos vazios.

A propriedade `content` é usada para definir ou substituir o conteúdo do elemento. Por padrão, os pseudoelementos criados pelos pseudosseletores `::before` e `::after` estão vazios e os elementos não serão renderizados na página. Definir a propriedade `content` como uma string vazia `""` garantirá que o elemento seja renderizado na página embora ainda esteja vazio.

Se você quiser experimentar, tente remover a propriedade `background-color` e definir valores diferentes para a propriedade `content`, tais como `"♥"`. Lembre-se de desfazer essas alterações quando terminar para que os testes passem.

# --hints--

Você deve ter um seletor `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

O seletor `.key.black--key::after` deve ter a propriedade `background-color` definida como `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

O seletor `.key.black--key::after` deve ter a propriedade `content` definida como `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
