---
id: 612e89562043183c86df287c
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Navegadores podem aplicar valores padrão de margem e de preenchimento a elementos específicos. Para garantir que seu piano pareça correto, você precisa redefinir o box model.

Adicione um seletor de regras `html` ao seu arquivo CSS e defina a propriedade `box-sizing` como `border-box`.

# --hints--

Você deve ter um seletor `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

O seletor `html` deve ter a propriedade `box-sizing` com o valor de `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
