---
id: 612e804c54d5e7308d7ebe56
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro do elemento `.keys`, adicione sete elementos `div`. Dê a todas elas a classe `key`.

# --hints--

Você deve criar sete novos elementos `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

Os sete novos elementos `div` devem estar dentro do elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

Os sete novos elementos `div` devem ter `class` definida como `key`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
