---
id: 612e77aba7ca691f598feb02
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione duas tags `meta`, uma para otimizar a página para dispositivos móveis e uma para especificar um `charset` aceito para a página.

# --hints--

Você deve adicionar dois elementos `meta` à página.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Você deve ter dois elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Um elemento `meta` deve ter o atributo `name` definido como `viewport` e o atributo `content` definido como `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

O outro elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
