---
id: 656ef783bba6976de014eaa8
title: Paso 24
challengeType: 20
dashedName: step-24
---

# --description--

Put your `password` variable declaration and the following `for` loop inside a `while` loop. Use `True` as the condition for your new loop.

# --hints--

You should create a `while True` loop enclosing your existing `password` declaration and `for` loop.

```js
({ test: () => assert.match(code, /^(\s*)while\s+True\s*:\s*^\1\1password\s*=\s*("|')\2\s*#\s*Generate\spassword\s*^\1\1for\s+_\s+in\s+range\s*\(\s*length\s*\s*\)\s*:\s*^\1\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--
    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
--fcc-editable-region--        
    return password

# new_password = generate_password(8)
# print(new_password)
```
