---
id: 587d78ac367417b2b2512af5
title: 使用 flex-direction 在嵌入推文中创建一列
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cnzdVC9'
forumTopicId: 301103
---

# --description--

在上一个挑战中，把嵌入推文的`header`和`footer`的`flex-direction`属性值设为 row（行）。相似地，把`.profile-name`选择器对应的元素竖着排列会好看一点。

# --instructions--

在 header 的`.profile-name`元素添加 CSS 属性`flex-direction`，将其值设为 column。

# --hints--

`.profile-name`应有`flex-direction`属性，其值应为 column。

```js
assert($('.profile-name').css('flex-direction') == 'column');
```

# --solutions--

