---
id: 587d78ad367417b2b2512af9
title: 在推文中使用 align-items 属性
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PNfq'
forumTopicId: 301106
---

# --description--

上一个挑战介绍了`align-items`属性并给出了例子。可以对嵌入推文的一些元素使用这个属性，以调整其中 flex 子元素的位置。

# --instructions--

在 header 的`.follow-btn`添加 CSS 属性`align-items`，把值设为 center。

# --hints--

`.follow-btn`应有`align-items`属性，其值应为 `center`.

```js
assert($('.follow-btn').css('align-items') == 'center');
```

# --solutions--

