---
id: 587d78ac367417b2b2512af4
title: 使用 flex-direction 属性创建一列
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
---

# --description--

之前两个挑战使用`flex-direction`属性创建行（row）。这个属性还能创建一个列，让子元素垂直排列在 flex 容器中。

# --instructions--

给`#box-container`元素添加 CSS 属性`flex-direction`，赋值为 `column`。

# --hints--

`#box-container`应有`flex-direction`属性，其值应为 column。

```js
assert($('#box-container').css('flex-direction') == 'column');
```

# --solutions--

