---
id: 587d78ad367417b2b2512af9
title: 在推文中使用 align-items 属性
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PNfq'
forumTopicId: 301106
---

# --description--

在上一个挑战中，我们介绍了 `align-items` 属性并给出了示例。我们可以对推文的几个嵌入元素使用这个属性，以调整其中 flex 元素的对齐方式。

# --instructions--

为 header 中的 `.follow-btn` 元素添加 CSS 属性 `align-items`，并将其属性值设为 `center`。

# --hints--

`.follow-btn` 选取的元素应有 `align-items` 属性，且其属性值应为 `center`。

```js
assert($('.follow-btn').css('align-items') == 'center');
```

# --solutions--

