---
id: 6852dd8486c7ee0171322950
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

You will need a variable to hold the secret code while you decode it, so create a variable called `secret_code` inside the function and assign to it an empty string.

# --hints--

You should have a `secret_code` variable inside the `pin_extractor` function.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

The variable `secret_code` should have a value of an empty string.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    pass
--fcc-editable-region--
```
