---
id: 6852db755c9c1274d75aae9a
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop you will create a pin extractor: the pin is hidden in short poems: in the nth line the length of the nth word determines the nth digit (a missing word means 0).

Start by creating a function called `pin_extractor` that takes a parameter `poem`.

# --hints--

You should have a function named `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

The `pin_extractor` should have a parameter named `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
