---
id: 68530231c05b93324e071104
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

While you work on the function it would be useful to have a function call for debugging purposes, so create a `poem` variable, outside the function, that contains this poem:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

You can use a multi-line string for that.

And then, call the function with the `poem` variable as argument.

# --hints--

You should have a `poem` variable.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

The `poem` variable should have a value equal to `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

You should call the function `pin_extractor` with `poem` as argument.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
```
