---
id: 68557bd29f7f3c1880f98836
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

In last line of the loop, update the `print` so that it prints the word inside the `words` list at index `line_index`.

# --hints--

You should have `print(words[line_index])` at the end of the loop.

```js
runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
        print(words)


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
