---
id: 68557bd29f7f3c1880f98836
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Inside the loop, find the length of the word at index `line_index` in the `words` array, and add the digit to `secret_code`, then return `secret_code` after the loop.

Remember that to concatenate a digit to a string you need to convert the digit to string before the concatenation.

# --hints--

You should have `return secret_code` inside the function, outside the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor` should return a string.

```js
({
    test: () => runPython(`
    assert type(pin_extractor('one')) == str
    `)
})
```

`pin_extractor(poem)` should return `5262`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

`pin_extractor('The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow')` should return `3346`.


```js
({
    test: () => runPython(`
    assert pin_extractor('The grass is green\\nhere and there\\nhoping for rain\\nbefore it turns yellow') == '3346'
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split(' ')


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
