---
id: 6855918e87c33c6979d21d3f
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

There isn't a garantee that the line has enough words to have an nth word, so before accessing `words[line_index]` you should check that the length of `line` is enough.

In the case that the line is shorter than expected, add `0` to the secret_code.

# --hints--

`secret_code('There\nonce\nwas\na\ndragon')` should return `50000`. 

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

`pin_extractor(poem)` should still return `5262`.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

`pin_extractor('The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow')` should still return `3346`.


```js
({
    test: () => runPython(`
    assert pin_extractor('The grass is green\\nhere and there\\nhoping for rain\\nbefore it turns yellow') == '3346'
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split(' ')
        secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
