---
id: 685578aae65e4106ad2b4300
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

The nth digit of the secret code comes from the nth word of the nth line, so you need to know what is the number of each line.

Change the loop so that it iterates on `enumerate(lines)`, and add an other loop variable before `line` called `line_index`.

# --hints--

The loop should iterate over `enumerate(lines)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

The loop should have `line_index, line` as loop variable.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        words = line.split(' ')

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
