---
id: 68559f3561d7df287586142e
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Now the pin extractor works, but it could be more efficient! It could extract the pin from many poems at the same time!

Add two more poems in the global scope, `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` and `poem3 = 'There\nonce\nwas\na\ndragon'`, and update the function call to have `[poem, poem2, poem3]` as argument.

Update the function to accept a list of poems in a parameter called `poems`, execute the current code once for each element in the list, store each discovered code in a `secret_codes` list, and return this list.

# --hints--

The variables `poem2` and `poem3` should exist and have the correct value.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

You should call `pin_extractor` with `[poem, poem2, poem3]` as argument.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_call('pin_extractor([poem, poem2, poem3])')
    `)
})
```

The function `pin_extractor` should have a `poems` parameter.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_function('pin_extractor').has_args('poems')
    `)
})
```

Inside `pin_extractor` there should be a `secret_codes` variable, which should start as an empty list.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
        assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

You should iterate over `poems` with a for loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('poems')
    `)
})
```

The function should return `secret_codes` at the end.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_codes')
    `)
})
```

`pin_extractor([poem, poem2, poem3])` should return `['5262', '3346', '50000']`.

```js
({
    test: () => runPython(`
    assert pin_extractor([poem, poem2, poem3]) == ['5262', '3346', '50000']
    `)
})
```

The function should work correctly with other lists of poems.

```js
({
    test: () => runPython(`
plum = '''In dusky shade the plum tree grew,
Its branches bent with evening dew.
A single fruit, so ripe and round,
Falls softly to the waiting ground.'''

apple = '''Red or green on branches high,
Like lanterns hung against the sky.
Crisp with autumn's golden breath,
A bite that wakes the world from death.'''

grape = '''Clustered tight on winding vines,
Bathed in sun where silence shines.
Jewel-toned beads in twilight air,
A whispered promise hanging there.'''

assert pin_extractor([plum, apple, grape]) == ['2863', '3885', '9227']
    `)
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split(' ')
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```

# --solutions--

```py
def pin_extractor(poems):
    secret_codes = []
    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split(' ')
            if len(words) <= line_index:
                secret_code += '0'
            else:
                secret_code += str(len(words[line_index]))
        secret_codes.append(secret_code)
    return secret_codes

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

pin_extractor([poem, poem2, poem3])


```
