---
id: 685535bf527b2c77236b3a9a
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

To find the length of the nth word, the next step is to separate the line of the poem into a list of words: inside the loop, create a variable `words` and assign to it the value of `line` split into words using the `split` method.

Then, add a `print` call with `words` as its argument.

# --hints--

You should have a `words` variable inside the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

The `words` variable should have a value of `line.split(' ')`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words').is_equivalent('words = line.split(" ")')
    `)
})
```

You should have `print(words)`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
