---
id: 685926db99041c5e619ead41
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

The `len` function gives the length of a string. Update the second `print` call to print the length of that word.

# --hints--

You should have `print(len(words[line_index]))` at the end of the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line, line_index)
        words = line.split()
        print(words[line_index])


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
