---
id: 68599f0fc7479410b2269742
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Update the function argument to be `poems`, then create a `for` loop around all the current content of the function that iterates over `poems` and uses `poem` as loop variable.

# --hints--

The `pin_extractor` function should have a parameter named `poems`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poems')
    `)
})
```

You should have a `for` loop that loops over `poems`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('poems')
    `)
})
```

The loop should have `poem` as loop variable.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split(' ')
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))
--fcc-editable-region--

```
