---
id: 68599f0fc7479410b2269742
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Update the function argument to be `poems`, then create a `for` loop around all the current content of the function that iterates over `poems` and uses `poem` as loop variable.

# --hints--

The `pin_extractor` function should have a parameter named `poems`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poems')
    `)
})
```

You should have a `for` loop that loops over `poems`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('poems')
    `)
})
```

The loop should have `poem` as loop variable.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('poem')
    `)
})
```

You should move the existing body of the function into the new `for` loop.  

```js
({
    test: () => runPython(`
    loop_body = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0]
    assert loop_body.find_variable('secret_code').is_equivalent('secret_code = ""'), "secret_code definition not found"
    assert loop_body.find_variable('lines').is_equivalent("lines = poem.split('\\\\n')"), "lines definition not found"
    assert loop_body.has_return('secret_code'), "return statement not found"
    inner_loop = loop_body.find_for_loops()[0]
    assert inner_loop.find_for_vars().is_equivalent('line_index, line'), "variable of loop not found"
    assert inner_loop.find_for_iter().is_equivalent('enumerate(lines)'), "iter of loop not found"
    inner_loop_body = inner_loop.find_bodies()[0]
    assert inner_loop_body.find_variable('words').is_equivalent("words = line.split()"), "words variable not found"
    if_clause = inner_loop_body.find_ifs()[0]
    if_conditions = if_clause.find_conditions()
    assert if_conditions[0].is_equivalent('len(words) > line_index'), "if condition not found"
    assert if_conditions[1] == _Node(), "else not found"
    if_bodies = if_clause.find_bodies()
    assert if_bodies[0].find_aug_variable('secret_code').is_equivalent("secret_code += str(len(words[line_index]))"), "secret_code inside if clause not found"
    assert if_bodies[1].find_aug_variable('secret_code').is_equivalent("secret_code += '0'"), "secret_code in else clause not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))
--fcc-editable-region--

```
