---
id: 68592b42b9c7f27cbe123203
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

The third line of the poem is missing a third word, so the pin is shorter than expected.

Add an `else` statement that adds a `'0'` to `secret_code` when there are not enough words, so that all lines of the poem are used to find digits.

# --hints--

You should have an `else` statement.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

The `else` statement should contain `secret_code += '0'`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split(' ')
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
