---
id: 68592730aaf34b61ae279168
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

The length returned from `len` is an integer, you need to convert it to a string using the `str` function.

Update the `print` call to print the length as a string.

# --hints--

You should have `print(str(len(words[line_index])))` at the end of the loop.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split(' ')
        print(len(words[line_index]))


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
--fcc-editable-region--

```
