---
id: 68592b42b9c7f27cbe123203
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

The third line of the poem is missing a third word, so the pin is shorter than expected.

Add an `else` clause that appends a `'0'` to `secret_code` when there are not enough words, so that all lines of the poem are used to find digits.

# --hints--

You should have an `else` clause.


```js
runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
```

The `else` clause should contain `secret_code += '0'`.

```js
runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
```

`pin_extractor(poem)` should return `5202`.


```js
runPython(`
    assert pin_extractor(poem) == '5202'
    `)
```

`secret_code('There\nonce\nwas\na\ndragon')` should return `50000`. 

```js
runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
