---
id: 660eebd83100d37862268781
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Your `character` variable currently does not have a value. You can assign a value using the <dfn>assignment</dfn> operator `=`. For example:

```js
let hello = "Hello";
```

Assigning a value to a variable at the moment of its declaration is known as <dfn>initialization</dfn>.

Initialize your `character` variable by assigning it the value `"Hello"` during its declaration.

# --hints--

You should use the assignment operator `=`.

```js
assert.match(code, /let\s+character\s*=/);
```

You should use the string `"Hello"`.

```js
assert.match(code, /('|")Hello\1/);
```

You should use double quotes for your `"Hello"` string.

```js
assert.match(code, /"Hello"/);
```

You should assign `"Hello"` to your `character` variable.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
