---
id: 6564fd78ffff909b1531d3c3
title: Крок 55
challengeType: 20
dashedName: step-55
---

# --description--

В межах циклу `for` викличте функцію `findall()`, передавши `pattern` та `password` як аргументи.

# --hints--

Викличте `re.findall()`, передавши `pattern` та `password` як аргументи.

```js
({ test: () => assert.match(code, /^(\s*)for.+:\s*^\1\s{4}re\.findall\s*\(\s*pattern\s*,\s*password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        for constraint, pattern in constraints:
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
