---
id: 6564b8c9349bd76dc037967b
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Функція `search()` з модуля `re` аналізує рядок, переданий як аргумент, шукаючи перший випадок, де шаблон регулярного виразу відповідає рядку.

Оголосіть змінну під назвою `quote` і призначте до неї рядок `Not all those who wander are lost.`. Потім надрукуйте результат `pattern.search(quote)`.

# --hints--

Ви повинні мати змінну `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Призначте наданий рядок до нової змінної `quote`.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Надрукуйте `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
