---
id: 6564642ce403fe29cc23fee7
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Модуль Python — це файл, що містить код, розроблений для виконання певних завдань. Стандартна бібліотека Python містить багато модулів, які можна імпортувати та використовувати у своєму коді. Для цього використайте інструкцію `import`, після якої напишіть назву модуля.

Розпочніть цей проєкт, імпортувавши модуль `string`.

# --hints--

Імпортуйте модуль `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
