---
id: 6564f0279e23ce924eedd1b2
title: Крок 51
challengeType: 20
dashedName: step-51
---

# --description--

Оскільки знак підкреслення є дійсним символом для назв змінних, його включено до символьного класу `\w` (еквівалентний `[a-zA-Z0-9_]`), який можна використовувати для зіставлення назв змінних.

Таким чином, символьний клас `\W` еквівалентний `[^a-zA-Z0-9_]` зі знаком підкреслення, який не збігається. З цієї причини його не можна використовувати, щоб знайти відповідність всім спеціальним символам.

Видаліть три останні рядки коду.

# --hints--

Видаліть три останні рядки коду.

```js
({ test: () => {
  assert.isFalse(/pattern\s*=\s*r("|')\\W\1/.test(code));
  assert.isFalse(/quote\s*=\s*("|')_\1/.test(code));
  assert.isFalse(/print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/.test(code));
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
