---
id: 656479aa5f298441c190bf8f
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

 Оголосіть функцію `generate_password` та запишіть увесь код, окрім рядків `import`, в межах тіла функції.

# --hints--

Оголосіть функцію під назвою `generate_password`.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    generate_psw = __locals.get("generate_password")
    inspect.isfunction(generate_psw)
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

--fcc-editable-region--
```
