---
id: 656477845006313fbfea0ad1
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Хоча ефект може здатися рівним `random.choice()`, `secrets` гарантує найбезпечніше джерело випадковості, яке може надати операційна система.

Тепер видаліть обидва виклики `print()`.

# --hints--

Видаліть обидва виклики `print()`.

```js
({ test: () => assert.isFalse( /print/.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
