---
id: 65646ab6bced52360f44670f
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер оголосіть змінну під назвою `all_characters` та призначте до неї результат об’єднання наявних змінних.

# --hints--

Оголосіть змінну під назвою `all_characters`.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

Об’єднайте `letters`, `digits` та `symbols`, а результат призначте до змінної `all_characters`.

```js
({ test: () => assert(__pyodide.runPython(`
    import string
    from itertools import permutations
    chars = __locals.get("all_characters")
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == chars for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
