---
id: 6564f834dd717998092cfd47
title: Крок 54
challengeType: 20
dashedName: step-54
---

# --description--

Напишіть цикл `for` після нового коментаря, щоб ітерувати над списком `constraints`. Використайте `constraint` та `pattern` як змінні ітерації.

# --hints--

Напишіть цикл `for`, щоб ітерувати над списком `constraints` з `constraint` та `pattern`.

```js
assert.match(code, /for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*:/)
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints

--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
