---
id: 6564883669b5af4b69f794cc
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Змінна `all_characters` — це рядок, який складається з всіх малих та великих літер, 10 цифр та декількох спеціальних символів.

Додайте коментар `Combine all characters` одразу перед нею.

# --hints--

Додайте коментар одразу над змінною `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
