---
id: 65647c71a5d2bd431596f629
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Функція `generate_password` потребує декількох параметрів. Для початку додайте параметр `length`.

# --hints--

Функція `generate_password` має приймати параметр `length`.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    generate_psw = __locals.get("generate_password")
    sig = str(inspect.signature(generate_psw))
    sig == "(length)"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import secrets
import string


def generate_password():    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

--fcc-editable-region--
```
