---
id: 65648f4b2281ba50051ae39c
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Далі напишіть цикл `for` з `i` як змінною циклу. Використайте функцію `range()`, щоб виконати ітерацію до значення `length`.

В межах циклу використайте оператор додавання з присвоєнням, щоб додати випадковий символ з `all_characters` до поточного значення `password`. Для цього використайте функцію `choice()` з модуля `secrets`.

# --hints--

Напишіть цикл `for`, який ітерує над `range(length)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

Використайте оператор `+=`, щоб додати випадковий символ з `all_characters` до поточного значення `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password

--fcc-editable-region--
```
