---
id: 65646ffeaed2d238c562a014
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Зазвичай інструкції `import` розміщують у верхній частині коду. Крім того, якщо інструкцій `import` є декілька, то їх відсортовують за алфавітним порядком, щоб покращити читабельність.

Імпортуйте модуль `random` у верхній частині коду.

# --hints--

Імпортуйте модуль `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
