---
id: 6564f32b18480893cf7799fd
title: Крок 52
challengeType: 20
dashedName: step-52
---

# --description--

Тепер поєднайте необроблений рядок з f-рядком та інтерполюйте змінну `symbols` в межах символьного класу. Пам’ятайте, що ви можете інтерполювати змінну в f-рядку за допомогою фігурних дужок `{ }`.

# --hints--

Другим елементом у четвертому кортежі обмежень має бути рядок `fr'[{symbols}]'`.

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*fr("|')\[\{\s*symbols\s*\}\]\1/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
