---
id: 656470d517833a39bb8b5608
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Модуль `random` містить генератор псевдовипадкових чисел. Більшість його можливостей залежить від функції `random()`, яка повертає число з рухомою комою в діапазоні від `0.0` (включно) до `1.0` (виключно).

Викличте функцію `random()` та надрукуйте результат.

# --hints--

Надрукуйте `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
