---
id: 656485a8a3496d4a36b1496a
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ці три змінні являють собою можливі символи, з яких можна створити пароль.

Додайте коментар `Define the possible characters for the password` одразу перед ними.

# --hints--

Додайте коментар одразу над трьома змінними.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
