---
id: 6564683821b2ee3174e7250e
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Ви можете отримати доступ до утиліт, визначених в межах імпортованого модуля, за допомогою точкової нотації. Щоб застосувати точкову нотацію, додайте до назви модуля крапку, після якої напишіть назву утиліти. Наприклад, ось як отримати доступ до константи `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

У цьому проєкті ви будете використовувати різні константи з модуля `string`.

Оголосіть нову змінну під назвою `letters` і призначте до неї `string.ascii_letters`.

# --hints--

Оголосіть змінну під назвою `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Призначте `string.ascii_letters` до змінної `letters`.

```js
({ test: () => assert(__pyodide.runPython(`
    import string
    abc = __locals.get("letters")
    abc == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
