---
id: 6565bb128adfcd5ec362382d
title: Крок 57
challengeType: 20
dashedName: step-57
---

# --description--

У циклі `for` порівняйте `constraint` та довжину списку, який повертає `findall()`. Для цього використайте оператор `<=`.

# --hints--

Порівняйте `constraint` та довжину списку, який повертає `findall`, використовуючи оператор `<=` в межах циклу `for`.

```js
({ test: () => assert.match(code, /^(\s*)for.+:\s*^\1\s{4}constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        for constraint, pattern in constraints:
            len(re.findall(pattern, password))
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
