---
id: 656eec2f40d18056cc58b229
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Повертається значення `None`, оскільки в межах аналізованого рядка не знайдено `i`.

Тепер змініть шаблон на `l` та подивіться на результат.

# --hints--

Змініть змінну `pattern` на `re.compile('l')`.

```js
({ test: () => assert(__pyodide.runPython(`
    import re
    p = __locals.get("pattern")
    p == re.compile('l')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)    
--fcc-editable-region--
pattern = re.compile('i')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
