---
id: 656ef89b4b486b7a16077864
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Кортеж — ще одна вбудована структура даних у Python. Кортежі дуже схожі на списки, але їх визначають за допомогою круглих дужок `()`, а не квадратних. Крім того, кортежі незмінні, на відміну від списків.

```py
my_tuple = ('larch', 1, True)
```

Список `constraints` зберігатиме кортежі. Перший елемент кожного кортежу буде параметром обмеження.

Додайте кортеж до списку. Використайте `nums` як перший елемент та порожній рядок як другий елемент.

# --hints--

Додайте кортеж до списку `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Додайте `(nums, '')` до списку `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []

--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
