---
id: 656ef5bc5c0cd464be1df675
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Далі ви надасте функції більше параметрів, які діятимуть як обмеження для згенерованого пароля.

Змініть оголошення функції, додавши `nums`, `special_chars`, `uppercase` та `lowercase` (в такому ж порядку) після наявного параметра `length`.

# --hints--

Функція має приймати `length`, `nums`, `special_chars`, `uppercase` та `lowercase` як параметри. Порядок має значення.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    foo = __locals.get("generate_password")
    sig = str(inspect.signature(foo))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)

    return password

# new_password = generate_password(8)
# print(new_password)
```
