---
id: 656ef54070c72161e6feb90f
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Здається, все чудово, але було б добре мати можливість перевірити, чи згенерований пароль відповідає певним критеріям. Наприклад, мінімальній кількості спеціальних символів, цифр або великих/малих літер. Ви займетесь цим дуже скоро.

Наразі закоментуйте останні два рядки коду.

# --hints--

Перетворіть останні два рядки коду на коментарі.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)

    return password

--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
