---
id: 657732654845d8e2fb1217e6
title: Крок 62
challengeType: 20
dashedName: step-62
---

# --description--

Наразі `all()` приймає порожній список як аргумент. Заповніть цей порожній список, використавши синтаксис розуміння, щоб в списку зберігались результати обчислення виразу `constraint <= len(re.findall(pattern, password))` для кожного кортежу `constraint`-`pattern` в списку `constraints`.

Таким чином цикл `while` припинить виконуватись лише тоді, коли всі умови виконано.

# --hints--

Передайте `[constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints]` до функції `all()`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("all([constraint <= len(re.findall(pattern, password)) for constraint, pattern in constraints])")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if all([]):
--fcc-editable-region--        
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
