---
id: 657732654845d8e2fb1217e6
title: Крок 62
challengeType: 20
dashedName: step-62
---

# --description--

`all()` — це вбудована функція Python, яка повертає `True`, якщо всі елементи в межах певного ітерованого об’єкта оцінюються як `True`. В іншому випадку вона повертає `False`.

Функцію `all()` можна поєднати з синтаксисом розуміння списку, щоб зробити код більш лаконічним.

Змініть умову `if`, передавши розуміння списку, створеного в попередньому кроці, до функції `all()`.

# --hints--

Передайте розуміння списку до функції `all()`.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*\[\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\]\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if [constraint <= len(re.findall(pattern, password))
        for constraint, pattern in constraints]:
--fcc-editable-region--        
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
