---
id: 657733bc97eb83efdb7e3988
title: Крок 64
challengeType: 20
dashedName: step-64
---

# --description--

Вам більше не потрібна змінна `count`. Видаліть цю змінну та її значення.

# --hints--

Видаліть рядок `count = 0`.

```js
({ test: () => assert.isFalse( /count\s*=\s*0/.test(code)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--
        # Check constraints
        count = 0
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
