---
id: 6577319039f4f7de9251b822
title: Крок 61
challengeType: 20
dashedName: step-61
---

# --description--

Замість того, щоб використовувати цикл та лічильник, можна застосувати інший підхід, який ви виконаєте в наступних кроках.

`all()` — це вбудована функція Python, яка повертає `True`, якщо всі елементи в межах певного ітерованого об’єкта оцінюються як `True`. В іншому випадку вона повертає `False`.

Замініть наявний цикл `for` та дві інструкції `if` на одну інструкцію `if`. Для умови `if` використайте виклик до функції `all()` та передайте порожній список як аргумент до виклику функції.

# --hints--

Замініть наявний цикл `for` та дві інструкції `if` на одну інструкцію `if`.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Новою умовою `if` має бути `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])")
`)) })
```

Нова інструкція `if` повинна містити `break`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1

        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
