---
id: 6577319039f4f7de9251b822
title: Крок 61
challengeType: 20
dashedName: step-61
---

# --description--

Замість того, щоб використовувати цикл та лічильник, можна застосувати інший підхід.

Замініть наявний цикл `for` та дві інструкції `if` на одну інструкцію `if`. Для умови `if` використайте список, створений за допомогою синтаксису розуміння списків.

У списку зберігаються результати обчислення виразу `constraint <= len(re.findall(pattern, password))` для кожного кортежу `constraint`-`pattern` у списку `constraints`.

# --hints--

Замініть наявний цикл `for` та дві інструкції `if` на одну інструкцію `if`.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Новою умовою `if` має бути розуміння списку, що є результатом обчислення виразів `constraint <= len(re.findall(pattern, password))` для кожного кортежу `constraint`-`pattern` в списку `constraints`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s*\[\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s+for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\]\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1

        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
