---
id: 66faab00000000000000000c
title: Build an Endless Runner
challengeType: 26
dashedName: lab-endless-runner
---

# --description--

**Objective:** Build a complete side-scrolling endless runner using Phaser 3 that fulfills the user stories below.

The Phaser library is loaded for you on the global `Phaser` object. The hero spritesheet lives at `/curriculum-assets/phaser/spritesheets/hero-walk.png` (4×4 grid of 16×16 frames; frames 12–15 are the run cycle). Parallax layers, the cactus obstacle, the spinning coin spritesheet, and the flying-enemy spritesheet are also under `/curriculum-assets/phaser/spritesheets/`.

**User Stories:**

1. You should declare a class named `MainScene` that extends `Phaser.Scene`, with a constructor that calls `super("MainScene")` and `preload`, `create`, and `update` methods.
1. You should declare a `config` object with a `type` of `Phaser.AUTO`, a `width` of `480`, a `height` of `320`, a `parent` of `"game-container"`, and arcade physics with world gravity `y: 1500`.
1. The world should have a parallax background built from at least three `tileSprite` layers, with each layer's `tilePositionX` advanced in `update` so the world flows past at distinct rates.
1. The player should be a physics sprite anchored at a fixed X position (the world flows past, the player does not move horizontally), and the SPACE key should drive jumps via `this.input.keyboard.addKey("SPACE")`.
1. Jumps should use a coyote-time grace window (`this.lastGroundedAt = time` while grounded; allow the jump while `time - this.lastGroundedAt < 120`) and a jump buffer (`this.lastJumpPressedAt = time` on press; consume it the moment a buffered press meets a recent grounded state).
1. Obstacles should spawn off-screen on the right, be added to a `physics.add.group({ allowGravity: false })`, and scroll left at `-this.scrollSpeed` via `body.setVelocityX(-this.scrollSpeed)`.
1. `this.distance` should accumulate per frame as `(this.scrollSpeed * delta) / 1000` and a `this.distanceText` HUD label should render the current floored distance.
1. The HUD (distance + best + sessions) should live on a **second camera** added with `this.cameras.add(0, 0, 480, 320)`, with `this.cameras.main.ignore(this.hud)` so world shake/flash never rattles the HUD, and the UI camera in turn ignoring every world game object.
1. The best distance should persist across page reloads using `localStorage.getItem("runner:best")` / `localStorage.setItem("runner:best", ...)`.
1. Difficulty should ramp with distance — `this.scrollSpeed = 200 + speedLevel * 50` and `this.obstacleTimer.delay = Math.max(500, 1500 - speedLevel * 100)` based on `Math.floor(this.distance / 1000)`.
1. Hitting an obstacle should set `this.gameOver = true`, render a `Game Over` banner plus a `Press R to restart` hint, and listen for `keydown-R` to call `this.scene.restart()`.
1. The player should have at least one collectible (coin) and one power-up — both as `physics.add.group({ allowGravity: false })` instances, with the power-up granting double-jump by setting `this.canDoubleJump = true` on overlap.
1. Camera shake should fire on hit (`this.cameras.main.shake(...)`) but be suppressed when `window.matchMedia("(prefers-reduced-motion: reduce)").matches` returns true (read into `this.reduceMotion` and gate the shake behind `if (!this.reduceMotion)`).
1. The game should accept touch input on mobile via `this.input.on("pointerdown", ...)` so a tap stamps `this.lastJumpPressedAt` (or otherwise triggers a jump) just like the SPACE key.
1. The game should have at least **two distinct obstacle types** — a ground obstacle (cactus) and an airborne obstacle (flyer in its own `this.flyers` group) — and the player should be able to duck under the airborne one by pressing the DOWN key, which shrinks the player's body via `body.setSize(12, 8).setOffset(2, 8)`.

# --hints--

You should declare a class named `MainScene` that extends `Phaser.Scene` with a constructor that calls `super("MainScene")` and `preload`, `create`, and `update` methods.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{/);
assert.match(code, /super\(\s*(['"`])MainScene\1\s*\)/);
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?\bpreload\s*\(\s*\)\s*\{/);
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?\bcreate\s*\(\s*\)\s*\{/);
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?\bupdate\s*\(/);
```

You should declare a `config` object with a `type` of `Phaser.AUTO`, `width` of `480`, `height` of `320`, `parent` of `"game-container"`, and arcade physics with `gravity: { y: 1500 }`.

```js
assert.match(code, /const\s+config\s*=\s*\{/);
assert.match(code, /type\s*:\s*Phaser\.AUTO/);
assert.match(code, /width\s*:\s*480\b/);
assert.match(code, /height\s*:\s*320\b/);
assert.match(code, /parent\s*:\s*(['"`])game-container\1/);
assert.match(code, /physics\s*:\s*\{[\s\S]*?default\s*:\s*(['"`])arcade\1[\s\S]*?arcade\s*:\s*\{[\s\S]*?gravity\s*:\s*\{\s*y\s*:\s*1500\s*\}/);
assert.match(code, /scene\s*:\s*\[\s*MainScene\s*\]/);
```

You should build the parallax background from at least three `tileSprite` layers, and advance each one's `tilePositionX` in `update` so the world flows past at distinct rates.

```js
assert.match(code, /this\.add\s*\.\s*tileSprite\([\s\S]*?this\.add\s*\.\s*tileSprite\([\s\S]*?this\.add\s*\.\s*tileSprite\(/);
assert.match(code, /\.tilePositionX\s*\+=/);
assert.match(code, /\.tilePositionX\s*\+=[\s\S]{0,200}?\.tilePositionX\s*\+=/);
```

You should anchor the player at a fixed X with `this.physics.add.sprite(...)` and bind SPACE for jumps via `this.input.keyboard.addKey("SPACE")`.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\(/);
assert.match(code, /this\.input\.keyboard\.addKey\(\s*(['"`])SPACE\1\s*\)/);
```

You should implement coyote time: stamp `this.lastGroundedAt = time` while the player is grounded and allow the jump while `time - this.lastGroundedAt` is less than at least 80ms.

```js
assert.match(code, /this\.lastGroundedAt\s*=\s*time/);
assert.match(code, /time\s*-\s*this\.lastGroundedAt\s*<\s*(?:8\d|9\d|1\d{2,}|[2-9]\d{2,})/);
```

You should implement a jump buffer: stamp `this.lastJumpPressedAt = time` on a fresh press and consume it when a recent press meets a recent grounded state.

```js
assert.match(code, /this\.lastJumpPressedAt\s*=\s*time/);
assert.match(code, /time\s*-\s*this\.lastJumpPressedAt\s*<\s*(?:8\d|9\d|1\d{2,}|[2-9]\d{2,})/);
assert.match(code, /canJump\s*&&\s*buffered/);
```

You should declare `this.obstacles` as a `physics.add.group({ allowGravity: false })`, spawn obstacles on the right, and scroll them left via `body.setVelocityX(-this.scrollSpeed)`.

```js
assert.match(code, /this\.obstacles\s*=\s*this\.physics\.add\.group\(\s*\{[\s\S]*?allowGravity\s*:\s*false[\s\S]*?\}\s*\)/);
assert.match(code, /this\.obstacles\.create\(\s*\d+\s*,/);
assert.match(code, /\.body\.setVelocityX\(\s*-\s*this\.scrollSpeed\s*\)/);
```

You should accumulate `this.distance` from `(this.scrollSpeed * delta) / 1000` and render it via `this.distanceText`.

```js
assert.match(code, /this\.distance\s*\+=\s*\(\s*this\.scrollSpeed\s*\*\s*delta\s*\)\s*\/\s*1000/);
assert.match(code, /this\.distanceText\s*=\s*this\.add[\s\S]{0,80}?\.text\(/);
assert.match(code, /this\.distanceText\.setText\(/);
```

You should host the HUD on a second camera added via `this.cameras.add(0, 0, 480, 320)` and ignore the HUD on the main camera while ignoring world objects on the UI camera.

```js
assert.match(code, /this\.cameras\.add\(\s*0\s*,\s*0\s*,\s*480\s*,\s*320\s*\)/);
assert.match(code, /this\.cameras\.main\.ignore\(/);
assert.match(code, /\b\w*[Cc]am\.ignore\(\s*\[/);
```

You should persist the best distance with `localStorage.getItem("runner:best")` and `localStorage.setItem("runner:best", ...)`.

```js
assert.match(code, /localStorage\.getItem\(\s*(['"`])runner:best\1/);
assert.match(code, /localStorage\.setItem\(\s*(['"`])runner:best\1/);
```

You should ramp difficulty with distance — set `this.scrollSpeed` from `200 +` a speed-level term and adjust `this.obstacleTimer.delay` with `Math.max`.

```js
assert.match(code, /this\.scrollSpeed\s*=\s*\d+\s*\+/);
assert.match(code, /this\.obstacleTimer\.delay\s*=\s*Math\.max/);
```

You should set `this.gameOver = true` on impact, render a `Game Over` banner with a `Press R to restart` hint, and restart on `keydown-R` via `this.scene.restart()`.

```js
assert.match(code, /this\.gameOver\s*=\s*true/);
assert.match(code, /(['"`])Game Over\1/);
assert.match(code, /(['"`])Press R to restart\1/);
assert.match(code, /keydown-R/);
assert.match(code, /this\.scene\.restart\(/);
```

You should declare `this.coins` and `this.powerups` as gravity-disabled physics groups, and the power-up should grant double jump by setting `this.canDoubleJump = true`.

```js
assert.match(code, /this\.coins\s*=\s*this\.physics\.add\.group\(\s*\{[\s\S]*?allowGravity\s*:\s*false[\s\S]*?\}\s*\)/);
assert.match(code, /this\.powerups\s*=\s*this\.physics\.add\.group\(\s*\{[\s\S]*?allowGravity\s*:\s*false[\s\S]*?\}\s*\)/);
assert.match(code, /this\.canDoubleJump\s*=\s*true/);
```

You should read `prefers-reduced-motion` into `this.reduceMotion` and gate the camera shake behind `!this.reduceMotion`.

```js
assert.match(code, /window\.matchMedia\(\s*(['"`])\(prefers-reduced-motion:\s*reduce\)\1\s*\)/);
assert.match(code, /this\.reduceMotion\s*=/);
assert.match(code, /!\s*this\.reduceMotion[\s\S]{0,80}?this\.cameras\.main\.shake\(/);
```

You should accept touch input via `this.input.on("pointerdown", ...)` so a tap stamps `this.lastJumpPressedAt`.

```js
assert.match(code, /this\.input\.on\(\s*(['"`])pointerdown\1[\s\S]{0,200}?this\.lastJumpPressedAt\s*=/);
```

You should add an airborne `this.flyers` group, bind DOWN as the duck key, and shrink the player's body via `setSize(12, 8).setOffset(2, 8)` while ducking.

```js
assert.match(code, /this\.flyers\s*=\s*this\.physics\.add\.group\(/);
assert.match(code, /this\.input\.keyboard\.addKey\(\s*(['"`])DOWN\1\s*\)/);
assert.match(code, /\.body\.setSize\(\s*12\s*,\s*8\s*\)\s*\.setOffset\(\s*2\s*,\s*8\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "clouds",
      "/curriculum-assets/phaser/spritesheets/parallax-clouds.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.image(
      "cactus",
      "/curriculum-assets/phaser/spritesheets/cactus.png"
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "flyer",
      "/curriculum-assets/phaser/spritesheets/flying-enemy.png",
      { frameWidth: 24, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.clouds = this.add
      .tileSprite(240, 120, 480, 96, "clouds")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);

    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");

    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);

    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.duckKey = this.input.keyboard.addKey("DOWN");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.canDoubleJump = false;
    this.didDoubleJump = false;

    this.scrollSpeed = 200;
    this.distance = 0;
    this.gameOver = false;
    this.bestDistance = parseInt(
      localStorage.getItem("runner:best") || "0",
      10
    );
    this.sessions = parseInt(
      localStorage.getItem("runner:sessions") || "0",
      10
    );

    this.obstacles = this.physics.add.group({ allowGravity: false });
    this.coins = this.physics.add.group({ allowGravity: false });
    this.powerups = this.physics.add.group({ allowGravity: false });
    this.flyers = this.physics.add.group({ allowGravity: false });

    this.anims.create({
      key: "spin",
      frames: this.anims.generateFrameNumbers("coin", { frames: [0, 1, 2, 3, 4, 5, 6, 7] }),
      frameRate: 12,
      repeat: -1
    });
    this.anims.create({
      key: "fly",
      frames: this.anims.generateFrameNumbers("flyer", { frames: [0, 1, 2, 3] }),
      frameRate: 8,
      repeat: -1
    });

    this.lastObstacleX = -Infinity;
    this.spawnObstacle = () => {
      if (520 - this.lastObstacleX < 200) return;
      const cactus = this.obstacles.create(520, 248, "cactus");
      cactus.body.setVelocityX(-this.scrollSpeed);
      cactus.setOrigin(0.5, 1);
      cactus.body.setSize(16, 28).setOffset(4, 4);
      this.lastObstacleX = 520;
    };
    this.spawnCoin = () => {
      const c = this.coins.create(520, 200, "coin", 0);
      c.body.setVelocityX(-this.scrollSpeed);
      c.play("spin");
    };
    this.spawnDoubleJump = () => {
      const p = this.add.rectangle(520, 200, 14, 14, 0x66ddff);
      this.physics.add.existing(p);
      this.powerups.add(p);
      p.body.setVelocityX(-this.scrollSpeed);
      p.kind = "double";
    };
    this.spawnFlyer = () => {
      const f = this.flyers.create(520, 200, "flyer", 0);
      f.body.setVelocityX(-this.scrollSpeed);
      f.play("fly");
    };

    this.spawnObstacle();
    this.obstacleTimer = this.time.addEvent({
      delay: 1500,
      loop: true,
      callback: this.spawnObstacle
    });
    this.coinTimer = this.time.addEvent({
      delay: 2700,
      loop: true,
      callback: this.spawnCoin
    });
    this.powerupTimer = this.time.addEvent({
      delay: 6500,
      loop: true,
      callback: this.spawnDoubleJump
    });
    this.flyerTimer = this.time.addEvent({
      delay: 4500,
      loop: true,
      callback: this.spawnFlyer
    });

    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.clouds.tilePositionX += 12 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) {
        this.lastGroundedAt = time;
        this.didDoubleJump = false;
      }
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        this.lastJumpPressedAt = time;
      }
      const canJump = time - this.lastGroundedAt < 120;
      const buffered = time - this.lastJumpPressedAt < 120;
      if (canJump && buffered) {
        this.player.setVelocityY(-600);
        this.lastGroundedAt = -Infinity;
        this.lastJumpPressedAt = -Infinity;
      } else if (
        buffered &&
        this.canDoubleJump &&
        !this.didDoubleJump &&
        !grounded
      ) {
        this.player.setVelocityY(-520);
        this.didDoubleJump = true;
        this.lastJumpPressedAt = -Infinity;
      }
    };
    this.applyDuck = (down) => {
      if (down) {
        this.player.body.setSize(12, 8).setOffset(2, 8);
      } else {
        this.player.body.setSize(16, 16).setOffset(0, 0);
      }
    };

    const dustGfx = this.make.graphics({ x: 0, y: 0, add: false });
    dustGfx.fillStyle(0xffffff, 0.6);
    dustGfx.fillCircle(2, 2, 2);
    dustGfx.generateTexture("dust", 4, 4);
    dustGfx.destroy();
    this.dust = this.add.particles(0, 0, "dust", {
      lifespan: 400,
      speedX: { min: -120, max: -60 },
      speedY: { min: -10, max: 10 },
      scale: { start: 1, end: 0 },
      quantity: 1,
      frequency: 90,
      follow: this.player,
      followOffset: { x: -4, y: 6 }
    });

    this.reduceMotion = window.matchMedia(
      "(prefers-reduced-motion: reduce)"
    ).matches;
    if (this.reduceMotion && this.dust) this.dust.stop();

    this.gameOverDeath = () => {
      if (this.gameOver) return;
      this.gameOver = true;
      if (!this.reduceMotion) this.cameras.main.shake(220, 0.012);
      this.obstacleTimer.remove(false);
      this.coinTimer.remove(false);
      this.powerupTimer.remove(false);
      this.flyerTimer.remove(false);
      this.player.setVelocity(0, 0);
      this.player.anims.stop();
      const dist = Math.floor(this.distance);
      const isNewBest = dist > this.bestDistance;
      if (isNewBest) {
        this.bestDistance = dist;
        localStorage.setItem("runner:best", String(dist));
        this.bestText.setText("Best: " + dist);
      }
      this.sessions += 1;
      localStorage.setItem("runner:sessions", String(this.sessions));
      this.sessionsText.setText("Runs: " + this.sessions);
      const banner = this.add
        .text(240, 140, isNewBest ? "New Best!" : "Game Over", {
          fontSize: "24px",
          color: isNewBest ? "#ffd866" : "#ffffff"
        })
        .setOrigin(0.5)
        .setScrollFactor(0)
        .setDepth(200);
      const hint = this.add
        .text(240, 180, "Press R to restart", {
          fontSize: "12px",
          color: "#aaccff"
        })
        .setOrigin(0.5)
        .setScrollFactor(0)
        .setDepth(200);
      this.hud.add([banner, hint]);
    };

    this.physics.add.overlap(this.player, this.obstacles, () => {
      this.gameOverDeath();
    });
    this.physics.add.overlap(this.player, this.flyers, () => {
      this.gameOverDeath();
    });
    this.physics.add.overlap(this.player, this.coins, (_p, coin) => {
      coin.destroy();
      this.distance += 50;
    });
    this.physics.add.overlap(this.player, this.powerups, (_p, item) => {
      if (item.kind === "double") {
        this.canDoubleJump = true;
      }
      item.destroy();
      if (!this.reduceMotion) {
        this.cameras.main.flash(160, 102, 221, 255);
      }
    });

    this.distanceText = this.add
      .text(8, 8, "Dist: 0", { fontSize: "12px", color: "#ffffff" })
      .setScrollFactor(0)
      .setDepth(120);
    this.bestText = this.add
      .text(8, 24, "Best: " + this.bestDistance, {
        fontSize: "10px",
        color: "#cccccc"
      })
      .setScrollFactor(0)
      .setDepth(120);
    this.sessionsText = this.add
      .text(8, 38, "Runs: " + this.sessions, {
        fontSize: "10px",
        color: "#aaccff"
      })
      .setScrollFactor(0)
      .setDepth(120);
    this.hud = this.add.container(0, 0).setScrollFactor(0).setDepth(120);
    this.hud.add([this.distanceText, this.bestText, this.sessionsText]);

    this.uiCam = this.cameras.add(0, 0, 480, 320);
    this.cameras.main.ignore(this.hud);
    this.uiCam.ignore([
      this.sky,
      this.clouds,
      this.mountains,
      this.groundTile,
      this.player,
      this.groundBody,
      this.obstacles,
      this.coins,
      this.powerups,
      this.flyers,
      this.dust
    ]);

    this.input.on("pointerdown", () => {
      if (this.gameOver) return;
      this.lastJumpPressedAt = this.time.now;
    });
    this.input.keyboard.on("keydown-R", () => {
      if (this.gameOver) this.scene.restart();
    });
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
    this.applyDuck(this.duckKey.isDown);
    if (!this.gameOver) {
      this.distance += (this.scrollSpeed * delta) / 1000;
      const speedLevel = Math.floor(this.distance / 1000);
      this.scrollSpeed = 200 + speedLevel * 50;
      this.obstacleTimer.delay = Math.max(500, 1500 - speedLevel * 100);
    }
    this.distanceText.setText("Dist: " + Math.floor(this.distance));
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
