---
id: 66faa40000000000000000ce
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--
Track HP and surface a HUD. Add the HP and HUD text in `create`:

```js
this.hp = 3;
this.hud = this.add.text(8, 8, "", {
  fontSize: "10px",
  color: "#ffffff"
});
```

Then refresh the HUD every frame in `update`:

```js
this.hud.setText("HP " + this.hp + "  COINS " + this.score);
```

Refreshing per frame is wasteful in theory but trivial in practice for a
single text label, and it means the HUD stays correct no matter what
mutates `this.hp` or `this.score`.

# --hints--
You should initialize `this.hp` to `3`.

```js
assert.match(code, /this\.hp\s*=\s*3/);
```

You should create `this.hud` as a text at `(8, 8)` with `fontSize: "10px"`.

```js
assert.match(code, /this\.hud\s*=\s*this\.add\.text\(\s*8\s*,\s*8\s*,\s*["']{2}\s*,\s*\{[\s\S]*?fontSize\s*:\s*["']10px["']/);
```

You should call `this.hud.setText("HP " + this.hp + "  COINS " + this.score)` in `update`.

```js
assert.match(code, /this\.hud\.setText\(\s*["']HP\s*["']\s*\+\s*this\.hp\s*\+\s*["']\s*COINS\s*["']\s*\+\s*this\.score\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    this.platforms.add(this.add.rectangle(80, 160, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(240, 120, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(400, 80, 64, 16, 0x66ddff));
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.physics.add.collider(this.player, this.platforms);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.anims.create({
      key: "walk",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.coins = this.physics.add.staticGroup();
    [
      [120, 140],
      [260, 100],
      [420, 60]
    ].forEach(([x, y]) => {
      this.coins.add(this.add.rectangle(x, y, 8, 8, 0xffee44));
    });
    this.score = 0;
    this.physics.add.overlap(
      this.player,
      this.coins,
      (player, coin) => {
        coin.destroy();
        this.score += 1;
      },
      null,
      this
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-140);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(140);
    } else {
      this.player.setVelocityX(0);
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      this.player.body.touching.down
    ) {
      this.player.setVelocityY(-310);
    }
    const vx = this.player.body.velocity.x;
    if (Math.abs(vx) > 1) {
      this.player.anims.play("walk", true);
      this.player.setFlipX(vx < 0);
    } else {
      this.player.anims.stop();
      this.player.setFrame(12);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
