---
id: 66faa40000000000000000cd
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Add three coins to the level and an overlap callback that removes them on
contact. Coins do not block movement, so use `add.overlap` instead of
`add.collider`:

```js
this.coins = this.physics.add.staticGroup();
[
  [120, 140],
  [260, 100],
  [420, 60]
].forEach(([x, y]) => {
  this.coins.add(this.add.rectangle(x, y, 8, 8, 0xffee44));
});
this.score = 0;
this.physics.add.overlap(
  this.player,
  this.coins,
  (player, coin) => {
    coin.destroy();
    this.score += 1;
  },
  null,
  this
);
```

`destroy` cleans up both the rectangle and its body in one call.

# --hints--
You should create `this.coins` as a `physics.add.staticGroup()`.

```js
assert.match(code, /this\.coins\s*=\s*this\.physics\.add\.staticGroup\(\s*\)/);
```

You should initialize `this.score` to `0`.

```js
assert.match(code, /this\.score\s*=\s*0/);
```

You should register an overlap between the player and coins that calls `coin.destroy()` and increments `this.score`.

```js
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.player\s*,\s*this\.coins\s*,[\s\S]*?coin\.destroy\(\s*\)[\s\S]*?this\.score\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    this.platforms.add(this.add.rectangle(80, 160, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(240, 120, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(400, 80, 64, 16, 0x66ddff));
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.physics.add.collider(this.player, this.platforms);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.anims.create({
      key: "walk",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-140);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(140);
    } else {
      this.player.setVelocityX(0);
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      this.player.body.touching.down
    ) {
      this.player.setVelocityY(-310);
    }
    const vx = this.player.body.velocity.x;
    if (Math.abs(vx) > 1) {
      this.player.anims.play("walk", true);
      this.player.setFlipX(vx < 0);
    } else {
      this.player.anims.stop();
      this.player.setFrame(12);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
