---
id: 66faa40000000000000000cc
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--
Drive the animation from the player's velocity. While moving, play the walk
anim and flip horizontally when going left. While idle, freeze on frame 12.

```js
const vx = this.player.body.velocity.x;
if (Math.abs(vx) > 1) {
  this.player.anims.play("walk", true);
  this.player.setFlipX(vx < 0);
} else {
  this.player.anims.stop();
  this.player.setFrame(12);
}
```

The `true` second argument to `anims.play` tells Phaser "if this anim is
already playing, don't restart" — without it the walk cycle would reset to
frame 0 every frame.

# --hints--
You should bind `vx` to `this.player.body.velocity.x`.

```js
assert.match(code, /const\s+vx\s*=\s*this\.player\.body\.velocity\.x/);
```

You should play the `"walk"` animation and `setFlipX(vx < 0)` when moving.

```js
assert.match(code, /this\.player\.anims\.play\(\s*["']walk["']\s*,\s*true\s*\)/);
assert.match(code, /this\.player\.setFlipX\(\s*vx\s*<\s*0\s*\)/);
```

You should stop animations and call `setFrame(12)` when idle.

```js
assert.match(code, /this\.player\.anims\.stop\(\s*\)/);
assert.match(code, /this\.player\.setFrame\(\s*12\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    this.platforms.add(this.add.rectangle(80, 160, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(240, 120, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(400, 80, 64, 16, 0x66ddff));
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.physics.add.collider(this.player, this.platforms);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.anims.create({
      key: "walk",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
  }

  update(time, delta) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-140);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(140);
    } else {
      this.player.setVelocityX(0);
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      this.player.body.touching.down
    ) {
      this.player.setVelocityY(-310);
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
