---
id: 66faa40000000000000000f5
title: Step 50
challengeType: 0
dashedName: step-50
---

# --description--
Add a water tile that slows the player down on overlap. The overlap
multiplies horizontal velocity by `0.6` and falling Y velocity by `0.4`
each frame, giving a viscous "swimming" feel:

```js
this.water = this.physics.add.staticGroup();
this.water.add(
  this.add.rectangle(208, 168, 64, 16, 0x4488ff, 0.5)
);
this.physics.add.overlap(
  this.player,
  this.water,
  () => {
    this.player.body.velocity.x *= 0.6;
    if (this.player.body.velocity.y > 0) {
      this.player.body.velocity.y *= 0.4;
    }
  },
  null,
  this
);
```

Only damping the falling case (`velocity.y > 0`) keeps jumps out of the
water from feeling weak.

# --hints--
You should create `this.water` as a static group containing a translucent water rectangle.

```js
assert.match(code, /this\.water\s*=\s*this\.physics\.add\.staticGroup\(\s*\)/);
assert.match(code, /this\.water\.add\(\s*this\.add\.rectangle\(\s*208\s*,\s*168\s*,\s*64\s*,\s*16\s*,\s*0x4488ff\s*,\s*0\.5/);
```

You should multiply X velocity by `0.6` and falling Y velocity by `0.4` on overlap.

```js
assert.match(code, /this\.player\.body\.velocity\.x\s*\*=\s*0\.6/);
assert.match(code, /this\.player\.body\.velocity\.y\s*>\s*0[\s\S]*?this\.player\.body\.velocity\.y\s*\*=\s*0\.4/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
const SAVE_KEY = "phaser-platformer-progress";

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    const LEVEL = [
      [0, 192, 176],
      [256, 480, 176],
      [560, 768, 176],
      [832, 1280, 176],
      [96, 192, 128],
      [256, 320, 96],
      [384, 480, 128],
      [560, 624, 80],
      [720, 832, 128],
      [928, 1056, 96],
      [1136, 1216, 64]
    ];
    LEVEL.forEach(([xs, xe, y]) => {
      for (let x = xs; x < xe; x += 16) {
        this.platforms.add(this.add.rectangle(x + 8, y, 16, 16, 0x66ddff));
      }
    });
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.physics.add.collider(this.player, this.platforms);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.anims.create({
      key: "walk",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.coins = this.physics.add.staticGroup();
    [
      [120, 140],
      [260, 100],
      [420, 60]
    ].forEach(([x, y]) => {
      this.coins.add(this.add.rectangle(x, y, 8, 8, 0xffee44));
    });
    this.score = 0;
    this.physics.add.overlap(
      this.player,
      this.coins,
      (player, coin) => {
        this.tweens.add({
          targets: coin,
          scale: 2,
          alpha: 0,
          duration: 150,
          onComplete: () => coin.destroy()
        });
        this.score += 1;
      },
      null,
      this
    );
    this.hp = 3;
    this.hud = this.add.text(8, 8, "", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.physics.world.setBounds(0, 0, 1280, 192);
    this.cameras.main.setBounds(0, 0, 1280, 192);
    this.cameras.main.startFollow(this.player, true, 0.1, 0.1);
    this.movingPlatforms = this.physics.add.group({
      allowGravity: false,
      immovable: true
    });
    const moving = this.add.rectangle(240, 140, 48, 8, 0xffaa44);
    this.movingPlatforms.add(moving);
    this.tweens.add({
      targets: moving,
      x: 320,
      duration: 1500,
      yoyo: true,
      repeat: -1,
      ease: "Sine.easeInOut"
    });
    this.physics.add.collider(this.player, this.movingPlatforms);
    this.spikes = this.physics.add.staticGroup();
    [[400, 168], [800, 168], [1024, 168]].forEach(([x, y]) => {
      this.spikes.add(this.add.rectangle(x, y, 16, 8, 0xff4444));
    });
    this.physics.add.overlap(
      this.player,
      this.spikes,
      () => this.takeDamage(),
      null,
      this
    );
    this.enemies = this.physics.add.group();
    [{ x: 320, y: 160 }, { x: 700, y: 64 }].forEach(({ x, y }) => {
      const e = this.add.rectangle(x, y, 12, 12, 0xff66aa);
      this.enemies.add(e);
      e.body.setVelocityX(-40);
      e.body.setBounce(1, 0);
    });
    this.physics.add.collider(this.enemies, this.platforms);
    this.physics.add.overlap(
      this.player,
      this.enemies,
      this.handleEnemyContact,
      null,
      this
    );
    this.flyingEnemies = this.physics.add.group({
      allowGravity: false
    });
    const flyer = this.add.circle(900, 96, 8, 0xaa66ff);
    this.flyingEnemies.add(flyer);
    flyer.baseY = 96;
    this.physics.add.overlap(
      this.player,
      this.flyingEnemies,
      this.handleEnemyContact,
      null,
      this
    );
    this.oneWayPlatforms = this.physics.add.staticGroup();
    [[600, 144], [950, 144]].forEach(([x, y]) => {
      this.oneWayPlatforms.add(
        this.add.rectangle(x, y, 64, 4, 0xaaffaa)
      );
    });
    this.physics.add.collider(
      this.player,
      this.oneWayPlatforms,
      null,
      (player, platform) => {
        if (this.cursors.down.isDown) return false;
        return player.body.bottom <= platform.body.top + 2;
      },
      this
    );
    this.spawnPoint = { x: 40, y: 60 };
    this.checkpoints = this.physics.add.staticGroup();
    this.checkpoints.add(
      this.add.rectangle(640, 160, 8, 16, 0x44eeff)
    );
    this.physics.add.overlap(
      this.player,
      this.checkpoints,
      (player, cp) => {
        cp.setFillStyle(0x88ffaa);
        this.spawnPoint = { x: cp.x, y: cp.y - 16 };
      },
      null,
      this
    );
    this.hasKey = false;
    this.key = this.add.rectangle(1120, 48, 10, 10, 0xffdd00);
    this.physics.add.existing(this.key, true);
    this.physics.add.overlap(
      this.player,
      this.key,
      () => {
        this.key.destroy();
        this.hasKey = true;
      },
      null,
      this
    );
    this.goal = this.add.rectangle(1240, 152, 12, 32, 0x44ee88);
    this.physics.add.existing(this.goal, true);
    this.physics.add.overlap(
      this.player,
      this.goal,
      () => {
        if (this.hasKey) this.winLevel();
      },
      null,
      this
    );
    this.startTime = this.time.now;
    this.bestTime = parseFloat(
      localStorage.getItem("phaser-platformer-progress:best") || "Infinity"
    );
    this.dust = this.add.particles(0, 0, "hero", {
      frame: 12,
      speed: { min: 40, max: 80 },
      scale: { start: 0.4, end: 0 },
      lifespan: 200,
      quantity: 4,
      emitting: false
    });
    this.hasWallJumpPower = false;
    this.wallPower = this.add.rectangle(880, 160, 10, 10, 0x88ffff);
    this.physics.add.existing(this.wallPower, true);
    this.physics.add.overlap(
      this.player,
      this.wallPower,
      () => {
        this.wallPower.destroy();
        this.hasWallJumpPower = true;
      },
      null,
      this
    );
    this.playSfx = (name) => {};
    this.input.keyboard.on("keydown-P", () => {
      this.scene.pause();
      this.scene.launch("PauseScene");
    });
    this.input.keyboard.on("keydown-F1", () => {
      const w = this.physics.world;
      w.drawDebug = !w.drawDebug;
      if (!w.drawDebug) w.debugGraphic.clear();
    });
    this.ladders = this.physics.add.staticGroup();
    this.ladders.add(
      this.add.rectangle(496, 144, 8, 64, 0xddaa66, 0.6)
    );
    this.onLadder = false;
    this.physics.add.overlap(
      this.player,
      this.ladders,
      () => {
        this.onLadder = true;
      },
      null,
      this
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-140);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(140);
    } else {
      this.player.setVelocityX(0);
    }
    const onGround =
      this.player.body.touching.down || this.player.body.blocked.down;
    if (onGround) this.lastGroundedAt = time;
    if (onGround) this.canDoubleJump = true;
    if (Phaser.Input.Keyboard.JustDown(this.cursors.up)) {
      this.lastJumpPressedAt = time;
    }
    const canJump = time - this.lastGroundedAt < 120;
    const buffered = time - this.lastJumpPressedAt < 120;
    if (canJump && buffered) {
      this.player.setVelocityY(-310);
      this.lastGroundedAt = -Infinity;
      this.lastJumpPressedAt = -Infinity;
    }
    const vx = this.player.body.velocity.x;
    if (Math.abs(vx) > 1) {
      this.player.anims.play("walk", true);
      this.player.setFlipX(vx < 0);
    } else {
      this.player.anims.stop();
      this.player.setFrame(12);
    }
    this.hud.setText("HP " + this.hp + "  COINS " + this.score);
    if (
      Phaser.Input.Keyboard.JustUp(this.cursors.up) &&
      this.player.body.velocity.y < -120
    ) {
      this.player.setVelocityY(-120);
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      !canJump &&
      this.canDoubleJump
    ) {
      this.player.setVelocityY(-260);
      this.canDoubleJump = false;
    }
    this.movingPlatforms.getChildren().forEach((plat) => {
      if (
        this.player.body.touching.down &&
        this.player.y < plat.y &&
        Math.abs(this.player.x - plat.x) < plat.width / 2 + 8
      ) {
        this.player.x += plat.body.deltaX();
      }
    });
    this.enemies.getChildren().forEach((e) => {
      if (e.body.blocked.left) e.body.setVelocityX(40);
      else if (e.body.blocked.right) e.body.setVelocityX(-40);
    });
    this.flyingEnemies.getChildren().forEach((f) => {
      f.y = f.baseY + Math.sin(time / 300) * 24;
    });
    if (onGround && this.prevVelocityY > 280) {
      this.cameras.main.shake(120, 0.005);
      this.dust.emitParticleAt(this.player.x, this.player.y + 8);
    }
    this.prevVelocityY = this.player.body.velocity.y;
    const pressingWall =
      (this.cursors.left.isDown && this.player.body.touching.left) ||
      (this.cursors.right.isDown && this.player.body.touching.right);
    this.wallSliding =
      this.hasWallJumpPower &&
      pressingWall &&
      !onGround &&
      this.player.body.velocity.y > 0;
    if (this.wallSliding) {
      this.player.body.velocity.y = Math.min(
        this.player.body.velocity.y,
        60
      );
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      this.wallSliding
    ) {
      const dir = this.player.body.touching.left ? 1 : -1;
      this.player.setVelocityX(200 * dir);
      this.player.setVelocityY(-260);
    }
    if (this.onLadder) {
      this.player.body.allowGravity = false;
      if (this.cursors.up.isDown) this.player.setVelocityY(-80);
      else if (this.cursors.down.isDown) this.player.setVelocityY(80);
      else this.player.setVelocityY(0);
    } else {
      this.player.body.allowGravity = true;
    }
    this.onLadder = false;
  }

  takeDamage() {
    if (this.time.now < this.invincibleUntil) return;
    this.invincibleUntil = this.time.now + 800;
    this.hp -= 1;
    this.player.setTint(0xff4444);
    this.time.delayedCall(150, () => this.player.clearTint());
    this.cameras.main.flash(150, 200, 50, 50);
    if (this.hp <= 0) this.gameOver();
  }

  handleEnemyContact(player, enemy) {
    if (player.body.bottom <= enemy.body.top + 4 && player.body.velocity.y > 0) {
      enemy.destroy();
      player.setVelocityY(-220);
    } else {
      this.takeDamage();
    }
  }

  gameOver() {
    this.scene.start("GameOverScene");
  }

  winLevel() {
    if (this.won) return;
    this.won = true;
    this.physics.pause();
    localStorage.setItem("phaser-platformer-progress:complete", "1");
    const elapsed = (this.time.now - this.startTime) / 1000;
    if (elapsed < this.bestTime) {
      this.bestTime = elapsed;
      localStorage.setItem(
        "phaser-platformer-progress:best",
        elapsed.toFixed(2)
      );
    }
    this.add
      .text(
        this.player.x,
        80,
        "LEVEL COMPLETE\n" + elapsed.toFixed(2) + "s\nbest " +
          (this.bestTime === Infinity ? "--" : this.bestTime.toFixed(2)),
        { fontSize: "12px", color: "#44ee88", align: "center" }
      )
      .setOrigin(0.5);
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create() {
    this.add.rectangle(240, 96, 480, 192, 0x000000, 0.85);
    this.add
      .text(240, 70, "GAME OVER", { fontSize: "20px", color: "#ff6644" })
      .setOrigin(0.5);
    this.add
      .text(240, 110, "Press R to retry", { fontSize: "10px", color: "#ffee44" })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-R", () => {
      this.scene.start("MainScene");
    });
  }
}

class PauseScene extends Phaser.Scene {
  constructor() {
    super("PauseScene");
  }

  create() {
    this.add.rectangle(240, 96, 480, 192, 0x000000, 0.6);
    this.add
      .text(240, 80, "PAUSED", { fontSize: "20px", color: "#ffffff" })
      .setOrigin(0.5);
    this.add
      .text(240, 110, "Press P to resume", { fontSize: "10px", color: "#ffee44" })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-P", () => {
      this.scene.stop("PauseScene");
      this.scene.resume("MainScene");
    });
  }
}

class LevelSelectScene extends Phaser.Scene {
  constructor() {
    super("LevelSelectScene");
  }

  create() {
    this.add.rectangle(240, 96, 480, 192, 0x111122);
    this.add
      .text(240, 60, "LEVEL SELECT", { fontSize: "16px", color: "#ffffff" })
      .setOrigin(0.5);
    const completed = localStorage.getItem(SAVE_KEY + ":complete") === "1";
    const label = completed ? "Level 1 [CLEAR]" : "Level 1";
    this.add
      .text(240, 100, label, { fontSize: "12px", color: "#44ee88" })
      .setOrigin(0.5);
    this.add
      .text(240, 130, "Press ENTER to play", { fontSize: "10px", color: "#ffee44" })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-ENTER", () => {
      this.scene.start("MainScene");
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene, GameOverScene, PauseScene, LevelSelectScene]
};

new Phaser.Game(config);
```
