---
id: 66faa40000000000000000c4
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the platformer workshop. Over the next 55 steps you will build a
complete side-scrolling platformer with a polished game-feel jump (variable
height, coyote time, jump buffering, double jump), enemies, hazards,
collectibles, a goal flag, a pause scene, and more.

Start by declaring the scene class. Inside the editable region declare a
`MainScene` that extends `Phaser.Scene`, with a constructor that calls
`super("MainScene")` and empty `preload`, `create`, and `update(time, delta)`
methods.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {}

  create() {}

  update(time, delta) {}
}
```

The `config` object below already references `MainScene` and sets gravity to
`800` — once the class is defined the canvas mounts.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["']MainScene["']\s*\)/);
```

You should declare a `preload` method.

```js
assert.match(code, /preload\s*\(\s*\)\s*\{/);
```

You should declare a `create` method.

```js
assert.match(code, /create\s*\(\s*\)\s*\{/);
```

You should declare an `update` method that takes `time` and `delta`.

```js
assert.match(code, /update\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
--fcc-editable-region--

--fcc-editable-region--


const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
