---
id: 66faa40000000000000000d2
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--
Add variable jump height. If the player releases UP early while still moving
upward, clamp the velocity to a small upward number so the rest of the jump
arc is shorter:

```js
if (
  Phaser.Input.Keyboard.JustUp(this.cursors.up) &&
  this.player.body.velocity.y < -120
) {
  this.player.setVelocityY(-120);
}
```

Tap UP for a short hop, hold UP for a full jump. This is the cheapest piece
of polish that makes a platformer feel responsive.

# --hints--
You should listen for `JustUp(this.cursors.up)` and clamp `velocity.y` to `-120` when it is below `-120`.

```js
assert.match(code, /Phaser\.Input\.Keyboard\.JustUp\(\s*this\.cursors\.up\s*\)[\s\S]*?this\.player\.body\.velocity\.y\s*<\s*-120[\s\S]*?this\.player\.setVelocityY\(\s*-120\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    const LEVEL = [
      [0, 192, 176],
      [256, 480, 176],
      [560, 768, 176],
      [832, 1280, 176],
      [96, 192, 128],
      [256, 320, 96],
      [384, 480, 128],
      [560, 624, 80],
      [720, 832, 128],
      [928, 1056, 96],
      [1136, 1216, 64]
    ];
    LEVEL.forEach(([xs, xe, y]) => {
      for (let x = xs; x < xe; x += 16) {
        this.platforms.add(this.add.rectangle(x + 8, y, 16, 16, 0x66ddff));
      }
    });
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.physics.add.collider(this.player, this.platforms);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.anims.create({
      key: "walk",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.coins = this.physics.add.staticGroup();
    [
      [120, 140],
      [260, 100],
      [420, 60]
    ].forEach(([x, y]) => {
      this.coins.add(this.add.rectangle(x, y, 8, 8, 0xffee44));
    });
    this.score = 0;
    this.physics.add.overlap(
      this.player,
      this.coins,
      (player, coin) => {
        coin.destroy();
        this.score += 1;
      },
      null,
      this
    );
    this.hp = 3;
    this.hud = this.add.text(8, 8, "", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.physics.world.setBounds(0, 0, 1280, 192);
    this.cameras.main.setBounds(0, 0, 1280, 192);
    this.cameras.main.startFollow(this.player, true, 0.1, 0.1);
  }

  update(time, delta) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-140);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(140);
    } else {
      this.player.setVelocityX(0);
    }
    if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      this.player.body.touching.down
    ) {
      this.player.setVelocityY(-310);
    }
    const vx = this.player.body.velocity.x;
    if (Math.abs(vx) > 1) {
      this.player.anims.play("walk", true);
      this.player.setFlipX(vx < 0);
    } else {
      this.player.anims.stop();
      this.player.setFrame(12);
    }
--fcc-editable-region--

--fcc-editable-region--
    this.hud.setText("HP " + this.hp + "  COINS " + this.score);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
