---
id: 66faa40000000000000000c5
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
Inside `preload`, load the hero spritesheet. The asset is a 4×4 grid of 16×16
frames; you will use frames 12–15 (right-facing walk cycle) for the player's
walk animation later.

```js
this.load.spritesheet(
  "hero",
  "/curriculum-assets/phaser/spritesheets/hero-walk.png",
  { frameWidth: 16, frameHeight: 16 }
);
```

# --hints--
You should load the hero spritesheet under the key `"hero"` with 16×16 frames.

```js
assert.match(code, /this\.load\.spritesheet\(\s*["']hero["']\s*,\s*["']\/curriculum-assets\/phaser\/spritesheets\/hero-walk\.png["']\s*,\s*\{\s*frameWidth\s*:\s*16\s*,\s*frameHeight\s*:\s*16\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
