---
id: 66faa40000000000000000c6
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Open `create` and lay down some level geometry. Static groups are the right
container for tiles that never move — gravity ignores them and they never
budge when collided with.

```js
this.platforms = this.physics.add.staticGroup();
this.platforms.add(this.add.rectangle(80, 160, 64, 16, 0x66ddff));
this.platforms.add(this.add.rectangle(240, 120, 64, 16, 0x66ddff));
this.platforms.add(this.add.rectangle(400, 80, 64, 16, 0x66ddff));
```

Three placeholder platforms is enough for now — you'll replace them with a
proper 30-tile level in step 12.

# --hints--
You should create `this.platforms` as a `physics.add.staticGroup()`.

```js
assert.match(code, /this\.platforms\s*=\s*this\.physics\.add\.staticGroup\(\s*\)/);
```

You should add three rectangles to `this.platforms` at the specified positions.

```js
assert.match(code, /this\.platforms\.add\(\s*this\.add\.rectangle\(\s*80\s*,\s*160\s*,\s*64\s*,\s*16/);
assert.match(code, /this\.platforms\.add\(\s*this\.add\.rectangle\(\s*240\s*,\s*120\s*,\s*64\s*,\s*16/);
assert.match(code, /this\.platforms\.add\(\s*this\.add\.rectangle\(\s*400\s*,\s*80\s*,\s*64\s*,\s*16/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
