---
id: 66faa40000000000000000c7
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
Drop the hero into the world as a physics sprite. `physics.add.sprite` does
two things at once: it adds the sprite to the display list and gives it an
Arcade body. The body inherits world gravity, so the hero will fall toward
the platforms.

```js
this.player = this.physics.add.sprite(40, 60, "hero", 12);
this.player.setCollideWorldBounds(true);
```

`setCollideWorldBounds(true)` keeps the player inside the canvas rectangle
while you wire up controls. Frame 12 is the first walk-right pose in the
sheet — the hero faces right by default.

# --hints--
You should add `this.player` as a physics sprite at `(40, 60)` using frame `12` of `"hero"`.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\(\s*40\s*,\s*60\s*,\s*["']hero["']\s*,\s*12\s*\)/);
```

You should call `setCollideWorldBounds(true)` on the player.

```js
assert.match(code, /this\.player\.setCollideWorldBounds\(\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    this.platforms.add(this.add.rectangle(80, 160, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(240, 120, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(400, 80, 64, 16, 0x66ddff));
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
