---
id: 66faa40000000000000000c9
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Wire up the arrow keys for left/right movement. First grab the cursor keys
in `create`:

```js
this.cursors = this.input.keyboard.createCursorKeys();
```

Then in `update`, drive the player's horizontal velocity:

```js
if (this.cursors.left.isDown) {
  this.player.setVelocityX(-140);
} else if (this.cursors.right.isDown) {
  this.player.setVelocityX(140);
} else {
  this.player.setVelocityX(0);
}
```

Setting the velocity to zero in the else branch stops the player crisply on
key release — without it the player would drift indefinitely.

# --hints--
You should call `createCursorKeys()` and assign it to `this.cursors`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)/);
```

You should set the player's velocity to `-140` when left is pressed, `140` when right is pressed, and `0` otherwise.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*?this\.player\.setVelocityX\(\s*-140\s*\)/);
assert.match(code, /this\.cursors\.right\.isDown[\s\S]*?this\.player\.setVelocityX\(\s*140\s*\)/);
assert.match(code, /this\.player\.setVelocityX\(\s*0\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Side-Scrolling Platformer with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    this.platforms.add(this.add.rectangle(80, 160, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(240, 120, 64, 16, 0x66ddff));
    this.platforms.add(this.add.rectangle(400, 80, 64, 16, 0x66ddff));
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.physics.add.collider(this.player, this.platforms);
  }

  update(time, delta) {
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
