---
id: 60a3e3396c7b40068ad6996d
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Вміст оточений простором, який називається відступом, що подібно до того, як бульбашкова плівка відокремлює товар від коробки.

Думайте про кордон як про картонну коробку, в якій доставили товар.

Змініть атрибут `src` на `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

Елемент `img` повинен мати `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
