---
id: 60a3e3396c7b40068ad6997f
title: Крок 22
challengeType: 0
dashedName: step-22
---

# --description--

Додавання 1 пікселя відступу вгорі, внизу, ліворуч та праворуч від полотна змінило його розміри на 502 пікселі x 602 пікселі.

Замініть властивість `padding` на `overflow` зі значенням `hidden`, повертаючи полотно до початкових розмірів.

# --hints--

Ви повинні видалити властивість `padding` із селектора `.canvas`.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Ви повинні встановити властивість `overflow` на `hidden`.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

Ваш елемент `.canvas` повинен мати `overflow` зі значенням `hidden`.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
