---
id: 60a3e3396c7b40068ad69974
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Кожній картині потрібна рамка.

Обгорніть елемент `.canvas` в інший `div`. Надайте цьому `div` клас `frame`.
# --hints--

Ви повинні додати новий елемент `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

Ваш елемент `.canvas` повинен бути вкладений в новий елемент `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Ваш новий `div` повинен мати `class` зі значенням `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Ваш новий `div` повинен бути в межах елемента `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
