---
id: 60a3e3396c7b40068ad69977
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Використайте відступи, щоб налаштувати простір всередині елемента.

В `.frame` використайте скорочену властивість `padding`, щоб збільшити простір між елементами `.frame` та `.canvas` на `50px`. Скорочення збільшить простір вгорі, внизу, ліворуч та праворуч від кордону елемента та полотна всередині.

# --hints--

Ви повинні встановити властивість `padding` на `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Ваш елемент `.frame` повинен мати `padding` зі значенням `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
