---
id: 60a3e3396c7b40068ad6998b
title: Крок 33
challengeType: 0
dashedName: step-33
---

# --description--

Було б добре, якби ваші поля рухались в одному напрямку.

В цьому випадку нижнє поле елемента `.one` зміщує `.two` на 20 пікселів вниз.

В селекторі `.two` використайте скорочену властивість `margin`, щоб встановити верхнє поле на `0`, горизонтальне поле на `auto` та нижнє поле на `20px`. Це видалить верхнє поле, горизонтально зцентрує його та встановить нижнє поле на 20 пікселів.

# --hints--

Ви повинні встановити властивість `margin` на `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

Ваш елемент `.two` повинен мати `margin` зі значенням `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
