---
id: 60a3e3396c7b40068ad69971
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Час для CSS.

Хоча ваш `<div>` не має тексту, він однаково розглядається як блок з вмістом. Напишіть CSS-правило, яке використовує селектор класу `.canvas` та встановіть його `width` на 500 пікселів. Ось CSS-правило, яке встановлює ширину класу `card` на 300 пікселів:

```css
.card {
  width: 300px;
}
```

# --hints--

Ваш код повинен мати селектор `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Ви повинні встановити властивість `width` на `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Ваш елемент `.canvas` повинен мати `width` зі значенням `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
