---
id: 60a3e3396c7b40068ad6996a
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Ви вже повинні бути знайомими з основними елементами, які повинна мати сторінка HTML.

Налаштуйте свій код з допомогою декларації `DOCTYPE`, елемента `html` з мовою встановлену на англійську, елемента `head` та елемента `body`.

# --hints--

Ваш код повинен мати декларацію `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ви повинні мати початковий теґ `<html>` з атрибутом `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш код повинен мати елемент `head` в межах елемента `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Ваш код повинен мати елемент `body` в межах елемента `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

Ваш елемент `head` повинен йти перед елементом `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
