---
id: 60a3e3396c7b40068ad69975
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Напишіть нове правило з допомогою селектора класу `.frame`.

Використайте скорочену декларацію `border`, щоб надати елементу `.frame` суцільний чорний кордон шириною `50px`.

# --hints--

Ваш код повинен мати селектор `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Ви повинні встановити властивість `border` на `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Ваш елемент `.frame` повинен мати `border` зі значенням `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
