---
id: 60b69a66b6ddb80858c51582
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

在該 `.characters` 元素中，創建另一個 `div`，其 `id` 爲 `offwhite-character`。

# --hints--

應該只創建一個額外的 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 3);
```

新的 `div` 元素應該嵌套在 `.characters` 元素中。

```js
assert(document.querySelector('.characters div'));
```

新的 `div` 元素應具有 `offwhite-character` 的 `id`。

```js
assert(document.querySelector('.characters div')?.getAttribute('id') === 'offwhite-character');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
