---
id: 60b69a66b6ddb80858c5157f
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

通常，HTML 以自上而下的方式呈現。 代碼頂部的元素位於頁面頂部。 但是，很多時候你可能希望將元素移動到不同的位置。 你可以使用 `position` 屬性來執行此操作。

將 `#back-wall` 元素的 `position` 屬性設置爲 `absolute`。 `absolute` 位置將元素從自上而下的文檔流中取出，並允許你相對於其容器調整它。

手動定位元素時，你可以使用 `top`、`left`、`right` 和 `bottom` 移動其佈局。 設置 `#back-wall` 的 `top` 值爲 `0`，`left` 值爲 `0`。


# --hints--

`#back-wall` 選擇器應該將 `position` 屬性設置爲 `absolute`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.position === 'absolute');
```

`#back-wall` 選擇器應該將 `top` 屬性設置爲 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.top === '0px');
```

`#back-wall` 選擇器應該將 `left` 屬性設置爲 `0`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.left === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
