---
id: 60b69a66b6ddb80858c51579
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，添加一個 `meta` 標籤，並將 `charset` 屬性設置爲 `utf-8`。 同時添加一個包含文本 `Picasso Painting` 的 `title` 元素。

# --hints--

應該只添加一個 `meta` 元素。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 元素應該有一個 `charset` 屬性。

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

`charset` 屬性應設置爲 `utf-8`。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

應該只添加一個 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 元素應該包含文本 `Picasso Painting`。 請注意拼寫和大小寫的問題。

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
