---
id: 60b69a66b6ddb80858c51580
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

`z-index` 屬性用於爲你的 HTML 元素創建“層”。 如果你熟悉圖像編輯工具，你可能以前使用過圖層。 這是一個類似的概念。

`z-index` 值較高的元素會顯示在 `z-index` 值較低的元素之上。 這可以與上一課中的定位相結合，創造出獨特的效果。

由於 `back-wall` 元素需要出現在你將要創建的其他元素的“後面”，因此爲 `back-wall` 元素提供一個 `z-index` 爲 `-1` 的屬性。

# --hints--

`#back-wall` 選擇器應該將 `z-index` 屬性設置爲 `-1`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.zIndex === '-1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
