---
id: 60b69a66b6ddb80858c5157c
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在你的 body 標籤中，添加一個 `div` 元素。 給它一個 `back-wall` 的 `id`。

# --hints--

應該準確添加 1 個 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 1);
```

`div` 元素應該具有 `back-wall` 的 `id` 值。

```js
assert(document.querySelector('div')?.getAttribute('id') === 'back-wall');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}
```
