---
id: 60b69a66b6ddb80858c51587
title: 步驟 17
challengeType: 0
dashedName: step-17
---

# --description--

將 `#offwhite-character` 移至適當位置，爲此，給它一個 `absolute` 的 `position`，一個 `top` 值爲 `20%`，一個 `left` 值爲 `17.5%`。

# --hints--

`#offwhite-character` 選擇器應該有一個 `position` 屬性設置爲 `absolute`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.position === 'absolute');
```

`#offwhite-character` 選擇器應該有一個 `top` 屬性設置爲 `20%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.top === '20%');
```

`#offwhite-character` 選擇器應該有一個 `left` 屬性設置爲 `17.5%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#offwhite-character')?.left === '17.5%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
