---
id: 60b69a66b6ddb80858c5157b
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

要開始繪畫，請爲你的 `body` 元素提供顏色爲 `rgb(184, 132, 46)` 的 `background-color`。

# --hints--

應該使用 `body` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` 元素應將 `background-color` 屬性設置爲 `rgb (184, 132, 46)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(184, 132, 46)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  </body>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

```
