---
id: 60b69a66b6ddb80858c5157a
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

FontAwesome 是一個 SVG 圖標庫，其中許多都可以免費使用。 你將在此項目中使用其中一些圖標，因此你需要將外部樣式錶鏈接到你的 HTML。

添加一個 `link` 元素，其中 `rel` 爲 `stylesheet`，`href` 爲 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

# --hints--

你應該有兩個 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 2);
```

`link` 元素應該具有 `stylesheet` 的 `rel`。

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('rel') === 'stylesheet');
```

`link` 元素應該有一個 `href` 爲 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
