---
id: 60b69a66b6ddb80858c5158c
title: 步驟 22
challengeType: 0
dashedName: step-22
---

# --description--

使用 id 選擇器爲 id 爲 `black-mask` 的元素創建規則。 給它一個 `100%` 的 `width`、一個 `50px` 的 `height`，和一個 `rgb(45, 31, 19)` 的 `background-color`。

# --hints--

應該有一個 `#black-mask` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask'));
```

`#black-mask` 選擇器應該將 `width` 屬性設置爲 `100%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask')?.width === '100%');
```

`#black-mask` 選擇器應該將 `height` 屬性設置爲 `50px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask')?.height === '50px');
```

`#black-mask` 選擇器應該有一個 `background-color` 屬性設置爲 `rgb(45, 31, 19)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('#black-mask')?.backgroundColor === 'rgb(45, 31, 19)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

#white-hat {
  width: 0;
  height: 0;
  border-style: solid;
  border-width: 0 120px 140px 180px;
  border-top-color: transparent;
  border-right-color: transparent;
  border-bottom-color: GhostWhite;
  border-left-color: transparent;
  position: absolute;
  top: -140px;
  left: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
