---
id: 60b80da8676fb3227967a731
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

鏈接你的 CSS 文件，即使你還沒有編寫任何 CSS。

添加一個 `link` 元素，其中 `rel` 值爲 `stylesheet`，`href` 值爲 `styles.css`。

# --hints--


你的代碼應該有一個 `link` 元素。

```js
assert.match(code, /<link/)
```

你應該有一個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
