---
id: 60b69a66b6ddb80858c51579
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge in dein `head`-Element ein `meta`-Tag mit dem `charset`-Attribut von `utf-8` hinzu. Füge auch ein `title`-Element mit dem Text `Picasso Painting` hinzu.

# --hints--

Du solltest genau ein `meta`-Element hinzufügen.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Your `meta` element should have a `charset` attribute.

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

Dein `charset`-Attribut sollte auf `utf-8` gesetzt sein.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Du solltest genau ein `title`-Element hinzufügen.

```js
assert(document.querySelectorAll('title').length === 1);
```

Dein `title`-Element sollte den Text `Picasso Painting` enthalten. Beachte, dass Recht- und Großschreibung von Bedeutung sind.

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
