---
id: 60b69a66b6ddb80858c51584
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

This character needs eyes. Create two `div` elements in the `#black-mask` element. Give them the classes `eyes left` and `eyes right`, in that order.

# --hints--

You should create 2 `div` elements within your `#black-mask` element.

```js
assert(document.querySelectorAll('#black-mask div').length === 2);
```

Your first new `div` should have the classes `eyes` and `left`.

```js
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('left'));
```

Your second new `div` should have the classes `eyes` and `right`.

```js
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
        <div id="gray-instrument"></div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
