---
id: 60b69a66b6ddb80858c5157c
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Füge innerhalb deines Body-Tags ein `div`-Element hinzu. Give it an `id` of `back-wall`.

# --hints--

Du solltest genau 1 `div`-Element hinzufügen.

```js
assert(document.querySelectorAll('div').length === 1);
```

Dein `div`-Element sollte den `id`-Wert von `back-wall` haben.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'back-wall');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}
```
