---
id: 60b80da8676fb3227967a731
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Go ahead and link your CSS file now, even though you have not written any CSS yet.

Add a `link` element with a `rel` of `stylesheet` and an `href` of `styles.css`.

# --hints--


Your code should have a `link` element.

```js
assert.match(code, /<link/)
```

You should have one self-closing `link` element.

```js
assert(document.querySelectorAll('link').length === 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
