---
id: 60b80da8676fb3227967a731
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fahre fort und verlinke jetzt deine CSS-Datei, obwohl du noch kein CSS geschrieben hast.

Füge ein `link`-Element mit einem `rel` von `stylesheet` und einen `href` von `styles.css` hinzu.

# --hints--


Dein Code sollte ein `link`-Element haben.

```js
assert.match(code, /<link/)
```

Du solltest ein selbstabschließendes `link`-Element haben.

```js
assert(document.querySelectorAll('link').length === 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
