---
id: 68f1fac259a4afabb478b960
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

You might have noticed you didn't link to the actual video.
You will do that in the next phase. When it comes to video file types,
there are differences in browser support. To accommodate this, you can
use `source` elements inside the `video` element and the browser will select
the first compatible `source`.

The source element is a void element, like the `img` element, so it contains no content. 
Do not use a closing tag.

Now, create a `source` element. Don't put anything in it yet.
You'll recall that its a void element.


# --hints--

You have a `source` element.

```js
const source = document.querySelector('source');
assert.isNotNull(source);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--

--fcc-editable-region--
  </video>
  </body>
</html>
```
