---
id: 68f18724d71db1dda8d6fec7
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Now that you finished the previous lesson, you'll build a HTML Video Player. 
You have been provided the HTML5 starter code. First, create a `h1` tag and
insert the text `Working with the HTML Video Element`.

# --hints--

Your `h1` element should contain the text `Working with the HTML Video Element`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent, 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
