---
id: 68f27139bddf9f6bf2821acf
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

You will add one final `source` element. This one will use the `video/quicktime` MIME type.

QuickTime is an extensible multimedia architecture created by Apple,
which supports playing, streaming, encoding, and transcoding a variety
of digital media formats. Not as popular as the MP4 format, you may need it
for legacy application support.

Now, add the `source` element with the self-closing tag. Next, add the `src` attribute
and the value `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`.
Finally, add the `type` attribute and the value `video/quicktime`.

Congratulations! You completed the HTML Video Player Workshop.
You learned to create a HTML5 video player for users to consume
video content. In addition, you learned to work with different video
formats such as MP4, WebM, and Ogg. Furthermore, you learned to offer
different video format fallbacks, so everyone can access your content.
Use your new found knowledge to complete the HTML Audio and Video Player lab.

Good Luck and Keep Learning!

# --hints--

You have a `source` element with the `src` attribute and
the value `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` and
the `type` attribute and the value `video/quicktime`.

```js
const source = document.querySelector('source:nth-of-type(4)');

assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
    type="video/mp4"
  />
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
    type="video/webm"
  />
  <source
    src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
    type="video/ogg"
  />
--fcc-editable-region--

--fcc-editable-region--
  </video>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    />
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    />
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    />
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    />
  </video>
</body>
</html>
```
