---
id: 680b3ef395479b0e449ecb6e
title: Build a Selection Sort
challengeType: 0
dashedName: build-selection-sort
---

# --description--

## Updated Description

Selection sort is another popular sorting algorithm taught in most computer science courses.

This algorithm works by repeatedly finding the smallest element from the unsorted portion of the list and swapping it with the first unsorted element. It begins by selecting the minimum value in the entire list and swapping it with the first element. Then it moves to the second position, finds the smallest value in the remaining unsorted elements, and swaps it with the second element. This process continues, moving through the list one element at a time, until the entire list is sorted.

Selection sort results in a quadratic time complexity in the best, average, and worst case scenarios. The space complexity will be constant `O(1)` because the sorting is done in place and a constant amount of memory is being used regardless of the size of the list.

## User Stories

1. You should define a function named `selection_sort`.
1. Your `selection_sort` function should have one parameter that represents the list of items.
1. Your `selection_sort` function should take a list and sort the items in place from smallest to largest.
1. Your `selection_sort` function should not use the built-in `sort()` method.

# --hints--

You should have a function named `selection_sort`

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

Your `selection_sort` function should have one parameter.

```js

```

Placeholder

```js
# Test cases shown to campers


print(selection_sort([33, 1, 89, 2, 67, 245]))
print(selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]))
print(selection_sort([1, 4, 2, 8, 345, 123, 43,
      32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]))

# Test cases not shown to campers to prevent attempts at hardcoded solutions

print(selection_sort([42, 17, 93, 8, 61, 29]))
print(selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]))
print(selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]))
print(selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]))      
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
