---
id: 614387cbefeeba5f3654a291
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Erstelle nach deinem `header`-Element einen `div` mit der `class` auf `author` gesetzt.

Erstelle in diesem `div` ein `p`-Element mit der `class` auf `author-name` gesetzt und gib ihm den Text `By freeCodeCamp`. Schließe den `freeCodeCamp` Teil in einem `a`-Element ein, indem du `href` auf `https://freecodecamp.org` und das `target` auf `_blank` setzt.


Füge unter dieses ein zweites `p`-Element mit der Klasse `publish-date` und dem Text `March 7, 2019` hinzu.

# --hints--

Du solltest ein neues `div`-Element erstellen.

```js
assert.exists(document.querySelector('div'));
```

Dein `div`-Element sollte nach deinem `header`-Element stehen.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

Dein `div`-Element sollte die `class` auf `author` gesetzt haben.

```js
assert(document.querySelector('div')?.className === 'author');
```

Du solltest zwei neue `p`-Elemente erstellen.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

Deine zwei neuen `p`-Elemente sollten sich innerhalb deines `div`-Elements befinden.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Dein erstes neues `p`-Element sollte eine `class` auf `author-name` gesetzt haben.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

Dein erstes neues `p`-Element sollte den Text `By freeCodeCamp` haben.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

Dein zweites neues `p`-Element sollte eine `class` auf `publish-date` gesetzt haben.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

Dein zweites neues `p`-Element sollte den Text `March 7, 2019` enthalten.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

Du solltest ein neues `a`-Element erstellen.

```js
assert.exists(document.querySelector('a'));
```

Dein `a`-Element sollte sich innerhalb deines ersten neuen `p`-Elements befinden.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

Dein `a`-Element sollte den `href` auf `https://freecodecamp.org` gesetzt haben.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

Dein `a`-Element sollte den `target` auf `_blank` gesetzt haben.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

Dein `a`-Element sollte den Text `freeCodeCamp` umschließen.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
