---
id: 614385513d91ae5c251c2052
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge ein `title`-Element mit dem Text `Magazine` hinzu, ein `link`-Element für das `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`-Font-Stylesheet, einen `link` für das `https://use.fontawesome.com/releases/v5.8.2/css/all.css`-FontAwesome-Stylesheet und einen `link` für dein `./styles.css`-Stylesheet.

Dein Font-Stylesheet lädt drei verschiedene Fonts: `Anton`, `Baskervville` und `Raleway`.

# --hints--

Dein Code sollte drei selbstschließende `link`-Elemente haben.

```js
assert(document.querySelectorAll('link').length === 3);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert(document.querySelectorAll('head > link').length === 3);
```

Deine drei `link`-Elemente sollten ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines deiner Link-Elemente sollte `href` auf `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` eingestellt haben.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Eines deiner Link-Elemente sollte `href` auf `https://use.fontawesome.com/releases/v5.8.2/css/all.css` eingestellt haben.

```js
const links = [...document.querySelectorAll('link')];
assert(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Eines deiner `link`-Elemente sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
assert.match(code, /<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/)
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Magazine`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'magazine')
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
