---
id: 6143862a5e54455d262c212e
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Within your `body`, create a `main` element. Then in that element, create a `section` with a `class` set to `heading`.

# --hints--

Du solltest ein `main`-Element haben.

```js
assert.exists(document.querySelector('main'));
```

Dein `main`-Element sollte sich innerhalb deines `body`-Elements befinden.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

Du solltest ein `section`-Element haben.

```js
assert.exists(document.querySelector('section'));
```

Dein `section`-Element sollte sich innerhalb deines `main`-Elements befinden.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

Dein `section`-Element sollte die `class` auf `heading` eingestellt haben.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
