---
id: 6143b9e1f5035c6e5f2a8231
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Erstelle innerhalb deines `ul`-Elements sechs `li`-Elemente. Füge ein `h4`-Element mit einer `class` von `list-subtitle` und ein `p`-Element zu jedem deiner `li`-Elemente hinzu.

Gib dann den `h4`- und `p`-Elementen in Reihenfolge den folgenden Textinhalt, wobei jedes `h4`-Element das aufnimmt, was auf der linken Seite des Doppelpunkts steht und jedes `p`-Element das, was auf der rechten Seite steht:

- `V1 - 2014`: `We launched freeCodeCamp with a simple list of 15 resources, including Harvard's CS50 and Stanford's Database Class.`
- `V2 - 2015`: `We added interactive algorithm challenges.`
- `V3 - 2015`: `We added our own HTML+CSS challenges (before we'd been relying on General Assembly's Dash course for these).`
- `V4 - 2016`: `We expanded the curriculum to 3 certifications, including Front End, Back End, and Data Visualization. They each had 10 required projects, but only the Front End section had its own challenges. For the other certs, we were still using external resources like Node School.`
- `V5 - 2017`: `We added the back end and data visualization challenges.`
- `V6 - 2018`: `We launched 6 new certifications to replace our old ones. This was the biggest curriculum improvement to date.`

# --hints--

Dein `ul`-Element sollte sechs `li`-Elemente haben.

```js
assert(document.querySelectorAll('.lists li')?.length === 6);
```

Jedes deiner neuen `li`-Elemente sollte ein `h4`- und ein `p`-Element haben.

```js
const lis = [...document.querySelectorAll('.lists li')];
assert(lis?.every(li => li?.children?.[0]?.localName === 'h4' && li?.children?.[1]?.localName === 'p'));
```

Dein erstes `h4` sollte den Text `V1 - 2014` haben.

```js
assert(document.querySelectorAll('.lists li h4')?.[0]?.innerText === 'V1 - 2014');
```

Dein erstes `p` sollte den Text `We launched freeCodeCamp with a simple list of 15 resources, including Harvard's CS50 and Stanford's Database Class.` haben

```js
assert(document.querySelectorAll('.lists li p')?.[0]?.innerText === 'We launched freeCodeCamp with a simple list of 15 resources, including Harvard\'s CS50 and Stanford\'s Database Class.');
```

Dein zweites `h4` sollte den Text `V2 - 2015` haben.

```js
assert(document.querySelectorAll('.lists li h4')?.[1]?.innerText === 'V2 - 2015');
```

Dein zweites `p` sollte den Text `We added interactive algorithm challenges.` haben.

```js
assert(document.querySelectorAll('.lists li p')?.[1]?.innerText === 'We added interactive algorithm challenges.');
```

Dein drittes `h4` sollte den Text `V3 - 2015` haben.

```js
assert(document.querySelectorAll('.lists li h4')?.[2]?.innerText === 'V3 - 2015');
```

Dein drittes `p` sollte den Text `We added our own HTML+CSS challenges (before we'd been relying on General Assembly's Dash course for these).` haben.

```js
assert(document.querySelectorAll('.lists li p')?.[2]?.innerText === 'We added our own HTML+CSS challenges (before we\'d been relying on General Assembly\'s Dash course for these).');
```

Dein viertes `h4` sollte den Text `V4 - 2016` haben.

```js
assert(document.querySelectorAll('.lists li h4')?.[3]?.innerText === 'V4 - 2016');
```

Dein viertes `p` sollte den Text `We expanded the curriculum to 3 certifications, including Front End, Back End, and Data Visualization. They each had 10 required projects, but only the Front End section had its own challenges. For the other certs, we were still using external resources like Node School` haben.

```js
assert(document.querySelectorAll('.lists li p')?.[3]?.innerText === 'We expanded the curriculum to 3 certifications, including Front End, Back End, and Data Visualization. They each had 10 required projects, but only the Front End section had its own challenges. For the other certs, we were still using external resources like Node School.');
```

Dein fünftes `h4` sollte den Text `V5 - 2017` haben.

```js
assert(document.querySelectorAll('.lists li h4')?.[4]?.innerText === 'V5 - 2017');
```

Dein fünftes `p` sollte den Text `We added the back end and data visualization challenges.` haben.

```js
assert(document.querySelectorAll('.lists li p')?.[4]?.innerText === 'We added the back end and data visualization challenges.');
```

Dein sechstes `h4` sollte den Text `V6 - 2018` haben.

```js
assert(document.querySelectorAll('.lists li h4')?.[5]?.innerText === 'V6 - 2018');
```

Dein sechstes `p` sollte den Text `We launched 6 new certifications to replace our old ones. This was the biggest curriculum improvement to date.` haben.

```js
assert(document.querySelectorAll('.lists li p')?.[5]?.innerText === 'We launched 6 new certifications to replace our old ones. This was the biggest curriculum improvement to date.');
```

Deine sechs `h4`-Elemente sollten die Klasse `list-subtitle` haben.

```js
const h4s = [...document.querySelectorAll('.lists li h4')];
assert(h4s?.every(h4 => h4?.classList?.contains('list-subtitle')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/freecodecamp/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/freecodecamp">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/freecodecamp">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
      <section class="text">
        <p class="first-paragraph">
          Soon the freeCodeCamp curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you'll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.
        </p>
        <p>
          After years - years - of pondering these two problems and how to solve them, I slipped, hit my head on the sink, and when I came to I had a revelation! A vision! A picture in my head! A picture of this! This is what makes time travel possible: the flux capacitor!
        </p>
        <p>
          It wasn't as dramatic as Doc's revelation in Back to the Future. It
          just occurred to me while I was going for a run. The revelation: the entire curriculum should be a series of projects. Instead of individual coding challenges, we'll just have projects, each with their own seamless series of tests. Each test gives you just enough information to figure out how to get it to pass. (And you can view hints if that isn't enough.)
        </p>
        <blockquote>
          <hr />
          <p class="quote">
            The entire curriculum should be a series of projects
          </p>
          <hr />
        </blockquote>
        <p>
          No more walls of explanatory text. No more walls of tests. Just one
          test at a time, as you build up a working project. Over the course of passing thousands of tests, you build up projects and your own understanding of coding fundamentals. There is no transition between lessons and projects, because the lessons themselves are baked into projects. And there's plenty of repetition to help you retain everything because - hey - building projects in real life has plenty of repetition.
        </p>
        <p>
          The main design challenge is taking what is currently paragraphs of explanation and instructions and packing them into a single test description text. Each project will involve dozens of tests like this. People will be coding the entire time, rather than switching back and forth from "reading mode" to "coding mode".
        </p>
        <p>
          Instead of a series of coding challenges, people will be in their code editor passing one test after another, quickly building up a project. People will get into a real flow state, similar to what they experience when they build the required projects at the end of each certification. They'll get that sense of forward progress right from the beginning. And freeCodeCamp will be a much smoother experience.
        </p>
      </section>
      <section class="text text-with-images">
--fcc-editable-region--
        <article class="brief-history">
          <h3 class="list-title">A Brief History</h3>
          <p>Of the Curriculum</p>
          <ul class="lists">

          </ul>
        </article>
--fcc-editable-region--
        <aside class="image-wrapper"></aside>
      </section>
    </main>
  </body>
</html>
```

```css

```
