---
id: 61437d575fb98f57fa8f7f36
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Beginne mit deinem standardmäßigen HTML-Boilerplate. Füge eine `DOCTYPE`-Deklaration ein, ein `html`-Element, das angibt, dass diese Seite auf Eglisch ist, ein `head`-Element und ein `body`-Element.

Füge innerhalb des `head`-Elements ein `<meta>`-Tag mit dem entsprechenden `charset` und ein `<meta>`-Tag für mobile Reaktionsfähigkeit hinzu.

# --hints--

Dein Code sollte die `DOCTYPE`-Referenz beinhalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Du solltest ein öffnendes `<html>`-Tag haben und es sollte eine `lang` von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein abschließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest ein öffnendes `<head>`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein abschließendes `</head>`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `<body>`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein abschließendes `</body>`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Deine `head` und `body`-Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Du solltest zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ein `meta`-Element sollte einen `name` mit `viewport` haben und `content` mit dem Wert `width=device-width, initial-scale=1.0` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset`-Attribut auf `UTF-8` gesetzt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
