---
id: 61438b5b66d76a6264430f2a
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Füge in jedes deiner neuen `a`-Elemente ein `i`-Element hinzu und gib ihnen die folgenden Klassen:

- Dein erstes `i`-Element sollte die Klasse `fab fa-facebook-f` haben
- Dein zweites `i`-Element sollte die Klasse `fab fa-twitter` haben
- Dein drittes `i`-Element sollte die Klasse `fab fa-instagram` haben
- Dein viertes `i`-Element sollte die Klasse `fab fa-linkedin-in` haben
- Dein fünftes `i`-Element sollte die Klasse `fab fa-youtube` haben

# --hints--

Du solltest fünf `i`-Elemente haben.

```js
assert(document.querySelectorAll('i')?.length === 5);
```

Jedes `a`-Element sollte nur ein `i`-Element haben.

```js
const aaaaa = [...document.querySelectorAll('.social-icons a')];
assert(aaaaa?.every(a => a?.children?.length === 1 && a?.children?.[0]?.localName === 'i'));
```

Jedes `i`-Element sollte ein `class`-Attribut haben, welches `fab` enthält.

```js
const iiiii = [...document.querySelectorAll('i')];
assert(iiiii?.every(i => i?.classList?.contains('fab')));
```

Das erste `i`-Element sollte ein `class`-Attribut haben, welches `fa-facebook-f` enthält.

```js
assert(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

Das zweite `i`-Element sollte ein `class`-Attribut haben, welches `fa-twitter` enthält.

```js
assert(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

Das dritte `i`-Element sollte ein `class`-Attribut haben, welches `fa-instagram` enthält.

```js
assert(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

Das vierte `i`-Element sollte ein `class`-Attribut haben, welches `fa-linkedin-in` enthält.

```js
assert(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

Das fünfte `i`-Element sollte ein `class`-Attribut haben, welches `fa-youtube` enthält.

```js
assert(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
          </a>
          <a href="https://twitter.com/freecodecamp/">
          </a>
          <a href="https://instagram.com/freecodecamp">
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
          </a>
          <a href="https://www.youtube.com/freecodecamp">
          </a>
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
