---
id: 6551ef504d91936d2d4e54f8
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

An *argument* is an object or an expression passed to a function — added between the opening and closing parentheses — when it is called:

```py
greet = 'Hello!'
print(greet)
```

The code in the example above would print the string `'Hello!'`, which is the value of the variable `greet` passed to `print()` as the argument.

Print your `text` variable to the screen by passing the `text` variable as the argument to the `print()` function.

# --hints--

You should pass `text` to the `print()` function by adding the name of this variable within the opening and closing parentheses. Pay attention to place the function call at the beginning of the line.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
