---
id: 68f64e5f51efb45e693f4be2
title: Step 1
challengeType: 25
dashedName: step-1
demoType: onClick
---

# --description--

In this project you will use the data from `https://cdn.freecodecamp.org/project-data/bar-chart/GDP-data.json` to build a bar chart.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**



# --hints--

The chart should have an element with the `id` of `title` to contain the title.

```js
assert.exists(document.getElementById('title'));
```

The chart should have a `g` element x-axis with a corresponging id `x-axis`.

```js
__helpers.retryingTest(() => document.querySelector('g#x-axis'), "g element with id x-axis not found");
```

The chart should have a `g` element y-axis with a corresponding id `y-axis`.

```js
__helpers.retryingTest(() => document.querySelector('g#y-axis'), 'element g#y-axis not found');
```

The x axe should contain multiple tick labels, eaxh with the corresponding `class="tick"`.

```js
__helpers.retryingTest(() => document.querySelectorAll('#x-axis .tick').length > 1, 'There are not enough tick labels on the x-axis');
```

The y axe should contain multiple tick labels, eaxh with the corresponding `class="tick"`.

```js
__helpers.retryingTest(() => document.querySelectorAll('#y-axis .tick').length > 1, 'There are not enough tick labels on the y-axis');
```

My Chart should have a <rect> element for each data point with a corresponding class="bar" displaying the data

```js
__helpers.retryingTest(() => document.querySelectorAll('rect.bar').length === 275, "The number of bars is not equal to the number of data points");
```

Each bar should have the properties "data-date" and "data-gdp" containing date and GDP values

```js
__helpers.retryingTest(() => {
  const bars = document.querySelectorAll('rect.bar');
  if (!bars) return;

  return [...bars].every(bar => bar.hasAttribute('data-date') && has.getAttribute('data-gdp'));
}, "data-date or data-gdp attribute not found in all .bar elements");
```

The bar elements' "data-date" properties should match the order of the provided data

```js
const res = await fetch(
          'https://raw.githubusercontent.com/FreeCodeCamp/ProjectReferenceData/master/GDP-data.json'
        ); // TODO probably should mock this. or need updating to link on cdn. anyway move to an hook

if (res.ok) {
  const json = await res.json();
  __helpers.retryingTest(() => {
    const bars = document.querySelectorAll('rect.bar');
    if (!bars) return;

    return [...bars].every((bar, i) => bar.getAttribute('data-date') === json.data[i][0]);

  },  'Bars should have date data in the same order as the ' + 
                'provided data ');
} else {
  assert.fail('JSON data not available')
}
```

The bar elements' "data-gdp" properties should match the order of the provided data

```js
const res = await fetch(
          'https://raw.githubusercontent.com/FreeCodeCamp/ProjectReferenceData/master/GDP-data.json'
        ); // TODO probably should mock this. or need updating to link on cdn. anyway move to an hook

if (res.ok) {
  const json = await res.json();
  __helpers.retryingTest(() => {
    const bars = document.querySelectorAll('rect.bar');
    if (!bars) return;

    return [...bars].every((bar, i) => bar.getAttribute('data-gdp') === json.data[i][1]);

  },  'Bars should have gdp data in the same order as the ' + 
                'provided data ');
} else {
  assert.fail('JSON data not available')
}
```

Each bar element's height should accurately represent the data's corresponding GDP

```js
__helpers.retryingTest(() => {
  const bars = document.querySelectorAll('rect.bar');
  // get the ratio of the first data point to the height of the first bar
  const firstRatio =
    parseFloat(bars[0]?.getAttribute('data-gdp')) /
    parseFloat(bars[0]?.getAttribute('height'));

  return [...bars].every(function (bar) {
    const dataValue = bar.getAttribute('data-gdp');
    const barHeight = bar.getAttribute('height');
    const ratio = parseFloat(dataValue) / parseFloat(barHeight);
    return firstRatio.toFixed(3) === ratio.toFixed(3);
  })
}, 'The heights of the bars should correspond to the data values');
```

The data-date attribute and its corresponding bar element should align with the corresponding value on the x-axis. https://github.com/freeCodeCamp/testable-projects-fcc/blob/main/src/project-tests/bar-chart-tests.js

```js
assert.fail()
```

The data-gdp attribute and its corresponding bar element should align with the corresponding value on the y-axis.

```js
assert.fail()
```

# --seed--

## --seed-contents--

```html

```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
  </head>
  <body>
    <div class="main">
      <div class="container">
        <div id="title">United States GDP</div>
        <div class="visHolder"></div>
      </div>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
}

.main {
  height: 100vh;
  width: 100vw;
  display: flex;
  align-items: center;
  justify-content: center;
  font-family: 'Roboto';
  background-color: #708090;
  background-size: 64px 128px;
}
.main .container {
  height: 560px;
  width: 900px;
  background-color: #fff;
  display: flex;
  flex-direction: column;
  padding: 20px 20px 20px 20px;
  align-self: center;
  position: relative;
}
@media (min-width: 1000px) {
  .main .container {
    box-shadow: 2px 2px 20px;
  }
}
.main .container #title {
  text-align: center;
  font-size: 2.5em;
}
.main .container .visHolder {
  position: absolute;
  top: 6em;
}

#tooltip {
  display: flex;
  align-items: center;
  justify-content: center;
  position: absolute;
  text-align: center;
  width: 150px;
  height: 50px;
  padding: 2px;
  font: 12px;
  background: lightsteelblue;
  box-shadow: 1px 1px 10px;
  border-radius: 2px;
  pointer-events: none;
}

.overlay {
  position: absolute;
  background: #fff;
  pointer-events: none;
}

#y-axis path {
  stroke: black;
  stroke-width: 1;
  fill: none;
}

#x-axis path {
  stroke: black;
  stroke-width: 1;
  fill: none;
}

.info {
  font-size: 0.8em;
}
```

```js
const projectName = 'bar-chart';

// coded by @Christian-Paul

var width = 800,
  height = 400,
  barWidth = width / 275;

var tooltip = d3
  .select('.visHolder')
  .append('div')
  .attr('id', 'tooltip')
  .style('opacity', 0);

var overlay = d3
  .select('.visHolder')
  .append('div')
  .attr('class', 'overlay')
  .style('opacity', 0);

var svgContainer = d3
  .select('.visHolder')
  .append('svg')
  .attr('width', width + 100)
  .attr('height', height + 60);

d3.json(
  'https://cdn.freecodecamp.org/project-data/bar-chart/GDP-data.json'
)
  .then(data => {
    svgContainer
      .append('text')
      .attr('transform', 'rotate(-90)')
      .attr('x', -200)
      .attr('y', 80)
      .text('Gross Domestic Product');

    svgContainer
      .append('text')
      .attr('x', width / 2 + 120)
      .attr('y', height + 50)
      .text('More Information: http://www.bea.gov/national/pdf/nipaguid.pdf')
      .attr('class', 'info');

    var years = data.data.map(function (item) {
      var quarter;
      var temp = item[0].substring(5, 7);

      if (temp === '01') {
        quarter = 'Q1';
      } else if (temp === '04') {
        quarter = 'Q2';
      } else if (temp === '07') {
        quarter = 'Q3';
      } else if (temp === '10') {
        quarter = 'Q4';
      }

      return item[0].substring(0, 4) + ' ' + quarter;
    });

    var yearsDate = data.data.map(function (item) {
      return new Date(item[0]);
    });

    var xMax = new Date(d3.max(yearsDate));
    xMax.setMonth(xMax.getMonth() + 3);
    var xScale = d3
      .scaleTime()
      .domain([d3.min(yearsDate), xMax])
      .range([0, width]);

    var xAxis = d3.axisBottom().scale(xScale);

    svgContainer
      .append('g')
      .call(xAxis)
      .attr('id', 'x-axis')
      .attr('transform', 'translate(60, 400)');

    var GDP = data.data.map(function (item) {
      return item[1];
    });

    var scaledGDP = [];

    var gdpMax = d3.max(GDP);

    var linearScale = d3.scaleLinear().domain([0, gdpMax]).range([0, height]);

    scaledGDP = GDP.map(function (item) {
      return linearScale(item);
    });

    var yAxisScale = d3.scaleLinear().domain([0, gdpMax]).range([height, 0]);

    var yAxis = d3.axisLeft(yAxisScale);

    svgContainer
      .append('g')
      .call(yAxis)
      .attr('id', 'y-axis')
      .attr('transform', 'translate(60, 0)');

    d3.select('svg')
      .selectAll('rect')
      .data(scaledGDP)
      .enter()
      .append('rect')
      .attr('data-date', function (d, i) {
        return data.data[i][0];
      })
      .attr('data-gdp', function (d, i) {
        return data.data[i][1];
      })
      .attr('class', 'bar')
      .attr('x', function (d, i) {
        return xScale(yearsDate[i]);
      })
      .attr('y', function (d) {
        return height - d;
      })
      .attr('width', barWidth)
      .attr('height', function (d) {
        return d;
      })
      .attr('index', (d, i) => i)
      .style('fill', '#33adff')
      .attr('transform', 'translate(60, 0)')
      .on('mouseover', function (event, d) {
        // d or datum is the height of the
        // current rect
        var i = this.getAttribute('index');

        overlay
          .transition()
          .duration(0)
          .style('height', d + 'px')
          .style('width', barWidth + 'px')
          .style('opacity', 0.9)
          .style('left', i * barWidth + 0 + 'px')
          .style('top', height - d + 'px')
          .style('transform', 'translateX(60px)');
        tooltip.transition().duration(200).style('opacity', 0.9);
        tooltip
          .html(
            years[i] +
              '<br>' +
              '$' +
              GDP[i].toFixed(1).replace(/(\d)(?=(\d{3})+\.)/g, '$1,') +
              ' Billion'
          )
          .attr('data-date', data.data[i][0])
          .style('left', i * barWidth + 30 + 'px')
          .style('top', height - 100 + 'px')
          .style('transform', 'translateX(60px)');
      })
      .on('mouseout', function () {
        tooltip.transition().duration(200).style('opacity', 0);
        overlay.transition().duration(200).style('opacity', 0);
      });
  })
  .catch(e => console.log(e));
  ```
