---
id: 615f5c1cb7575c7551ed8a40
title: الخطوة 35
challengeType: 0
dashedName: step-35
---

# --description--

حجم الخط الكبير للرقم `230` يتسبب في overflow. اعطي `.calories-info h1` خاصية `overflow` بقيمة `hidden` لتجنب هذا.

# --hints--

يجب أن يكون لمنتقي `.calories-info h1` الخاص بك الخاصية `overflow` بقيمة `hidden`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.calories-info h1')?.overflow === 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
      <div class="calories-info">
        <p class="bold sm-text">Amount per serving</p>
        <h1>Calories <span class="right">230</span></h1>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}

.sm-text {
  font-size: 0.85rem;
}

--fcc-editable-region--
.calories-info h1 {
  margin: -5px -2px;
}
--fcc-editable-region--

.calories-info span {
  font-size: 1.2em;
}
```
