---
id: 615f2d4150fe0d4cbd0f2628
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

تحت عنصر `h1` الخاص بك، أضف عنصر `p` مع النص `8 servings per container`.

# --hints--

يجب عليك إضافة عنصر `p` جديد.

```js
assert.exists(document.querySelector('p'));
```

عنصر `p` الخاص بك يجب أن يكون داخل عنصر `body` الخاص بك.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

عنصر `p` الخاص بك يجب أن يأتي بعد عنصر `h1` الخاص بك.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

يجب أن يحتوي عنصر `p` الخاص بك على النص `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
