---
id: 615f38279e5c3d53692ea441
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

قم بتغليف عنصري `h1` و `p` في عنصر `div`. اعطي هذا الـ `div` السمة `class` بقيمة `label`.

# --hints--

يجب عليك إضافة عنصر `div` جديد.

```js
assert(document.querySelector('div'));
```

يجب أن يحتوي عنصر `div` الجديد على سمة `class` بقيمة `label`.

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

يجب أن تكون عناصر `h1` و `p` ضمن عنصر `.label` الجديد الخاص بك.

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

--fcc-editable-region--
<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
--fcc-editable-region--
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
