---
id: 615f7ad94380408d971d14f6
title: الخطوة ٤٨
challengeType: 0
dashedName: step-48
---

# --description--

Pseudo-selector الـ `:not` يمكن ان يستخدم لتحديد جميع العناصر التي لا تتطابق مع قاعدة الـ CSS المحددة.

```css
div:not(#example) {
  color: red;
}
```

ما ورد أعلاه يحدد جميع عناصر `div` بدون `id` بقيمة `example`.

قم بتعديل منتقي `.daily-value p` لاستبعاد عناصر `.no-divider`.

# --hints--

يجب أن يكون لديك منتقي `.daily-value p:not(.no-divider)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.daily-value p:not(.no-divider)'));
```

يجب ألا يكون لديك منتقي `.daily-value p`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.daily-value p'));
```

لا يجب عليك تغيير الخصائص في منتقي `.daily-value p:not(.no-divider)`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.daily-value p:not(.no-divider)');
assert(style?.borderBottom === '1px solid rgb(136, 137, 137)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
      <div class="calories-info">
        <p class="bold sm-text">Amount per serving</p>
        <h1>Calories <span class="right">230</span></h1>
      </div>
      <div class="divider md"></div>
      <div class="daily-value sm-text">
        <p class="right bold no-divider">% Daily Value *</p>
        <div class="divider"></div>
        <p><span class="bold">Total Fat</span> 8g<span class="bold right">10%</span></p>
        <p class="indent no-divider">Saturated Fat 1g <span class="bold right">5%</span></p>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
  clear: right;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}

.md {
  height: 5px;
}

.sm-text {
  font-size: 0.85rem;
}

.calories-info h1 {
  margin: -5px -2px;
  overflow: hidden;
}

.calories-info span {
  font-size: 1.2em;
  margin-top: -7px;
}

.indent {
  margin-left: 1em;
}

--fcc-editable-region--
.daily-value p {
  border-bottom: 1px solid #888989;
}
--fcc-editable-region--
```
