---
id: 615f808d85793195b0f53be9
title: الخطوة ٥٦
challengeType: 0
dashedName: step-56
---

# --description--

قم بإنشاء عنصر آخر `p` بعد آخر `.divider`، وقم بإعطائه النص `Total Sugars 12g`. قم بتعيين عنصر `p` هذا، قيم الـ `class` اللازمة لاضافة indent وإزالة الحدود السفلي (bottom border). ثم قم بإنشاء عنصر `.divider` آخر تحت عنصر `p` الجديد.

# --hints--

يجب عليك إنشاء عنصر `p` و `div` جديدان في نهاية عنصر `.daily-value.sm-text` الخاص بك.

```js
assert(document.querySelector('.daily-value.sm-text')?.lastElementChild?.previousElementSibling?.localName === 'p');
assert(document.querySelector('.daily-value.sm-text')?.lastElementChild?.localName === 'div');
```

عنصر `p` الجديد الخاص بك يجب أن يحتوي على النص `Total Sugars 12g`.

```js
assert(document.querySelector('.daily-value.sm-text p:last-of-type')?.innerText.match(/Total Sugars[\s|\n]+12g/));
```

يجب أن يحتوي عنصر `p` الجديد على السمة `class` بقيمة `indent no-divider`.

```js
assert(document.querySelector('.daily-value.sm-text p:last-of-type')?.classList?.contains('indent'));
assert(document.querySelector('.daily-value.sm-text p:last-of-type')?.classList?.contains('no-divider'));
```

يجب أن يحتوي عنصر `div` الجديد على السمة `class` بقيمة `divider`.

```js
assert(document.querySelector('.daily-value.sm-text')?.lastElementChild?.classList?.contains('divider'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
      <div class="calories-info">
        <p class="bold sm-text">Amount per serving</p>
        <h1>Calories <span class="right">230</span></h1>
      </div>
      <div class="divider md"></div>
      <div class="daily-value sm-text">
--fcc-editable-region--
        <p class="right bold no-divider">% Daily Value *</p>
        <div class="divider"></div>
        <p><span class="bold">Total Fat</span> 8g<span class="bold right">10%</span></p>
        <p class="indent no-divider">Saturated Fat 1g <span class="bold right">5%</span></p>
        <div class="divider"></div>
        <p class="indent no-divider"><i>Trans</i> Fat 0g</p>
        <div class="divider"></div>
        <p><span class="bold">Cholesterol</span> 0mg <span class="right bold">0%</span></p>
        <p><span class="bold">Sodium</span> 160mg <span class="right bold">7%</span></p>
        <p><span class="bold">Total Carbohydrate</span> 37g <span class="right bold">13%</span></p>
        <p class="indent no-divider">Dietary Fiber 4g</p>
        <div class="divider"></div>
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
  clear: right;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}

.md {
  height: 5px;
}

.sm-text {
  font-size: 0.85rem;
}

.calories-info h1 {
  margin: -5px -2px;
  overflow: hidden;
}

.calories-info span {
  font-size: 1.2em;
  margin-top: -7px;
}

.indent {
  margin-left: 1em;
}

.daily-value p:not(.no-divider) {
  border-bottom: 1px solid #888989;
}
```
