---
id: 615f2abbe7d18d49a1e0e1c8
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

لقد قدمنا لك مستند HTML الأساسي.

قم بإنشاء عنصر `h1` داخل عنصر `body` وأعطيه النص `Nutrition Facts`.

# --hints--

يجب عليك إضافة عنصر `h1` جديد.

```js
assert.exists(document.querySelector('h1'));
```

عنصر `h1` الخاص بك يجب أن يكون داخل عنصر `body` الخاص بك.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

يجب أن يحتوي العنصر `h1` على النص `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
