---
id: 615f357957e370510f21ea16
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

قم بإنشاء منتقي `body` وإعطائه `font-family` بقيمة `Open Sans` مع fallback بقيمة `sans-serif`.

تذكر أن الخطوط (fonts) التي تحتوي على مسافات في اسمها، يجب تغليفها في علامات اقتباس (quotes) للـ CSS.

# --hints--

يجب أن يكون لديك منتقي `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب أن يكون لمنتقي `body` الخاص بك خاصية `font-family` بقيمة `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
