---
id: 615f3cafd794015aa9547a6d
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

تذكر أن استخدام `h1` و `h2` و الـ tags مشابهة تحدد البنية الدلالية (semantic structure) لـ HTML. ومع ذلك، يمكنك ضبط الـ CSS لهذه العناصر للتحكم في التدفق البصري (visual flow) والتسلسل الهرمي (hierarchy).

قم بإنشاء قاعدة `h1` وتعيين خاصية `font-weight` إلى `800`. هذا سيجعل نص `h1` أكثر سمكاً (bolder).

# --hints--

يجب عليك إنشاء منتقي `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

يجب أن يكون لمنتقي `h1` الخاص بك الخاصية `font-weight` بقيمة `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
