---
id: 66c07238b01053abaf812065
title: Build a Factorial Calculator
challengeType: 14
dashedName: build-a-factorial-calculator
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

<!-- To-Do:
- Fix last 2 tests.
- Make loop variable flexible in tests -->

A factorial is the product of an integer and all the integers below it. For example, the factorial of `5` is `5 * 4 * 3 * 2 * 1 = 120`. In this lab, you will create a factorial calculator that takes a number from the user and calculates the factorial of that number.

1. You should ask the user for a number using a prompt window and store the number in a variable called `inputNumber`.
   
2. Declare a variable `result` and assign it the value of `1`.
   
3. Using a for loop, you should loop through all numbers from `1` to `inputNumber` (inclusive) and for each number, multiply the `result` variable by the current number and assign the result to the `result` variable.
   
4. You should log the result of the factorial of `inputNumber` to the console. The output should be in the format `Factorial of [inputNumber] is [result]`.

# --hints--

You should prompt the user to enter a number using the `prompt` method.

```js
assert.match(code, /prompt\(["']Enter a number:\s*["']\)/);
```

You should create a variable `inputNumber` using `const` to store the number entered by the user through the `prompt` method.

```js
assert.match(code, /const\s+inputNumber\s*=\s*prompt\(["']Enter a number:\s*["']\);/);
```

You should create a variable `result` and assign it the value of `1`.

```js
assert.match(code, /let\s+result\s*=\s*1;/);
```

You should use a `for` loop to loop through all numbers from `1` to `inputNumber` (inclusive).

```js
assert.match(code, /for\s*\(\s*let\s+([a-zA-Z_$][\w$]*)\s*=\s*1\s*;\s*\1\s*<=\s*inputNumber\s*;\s*\1\s*\+\+\s*\)\s*{/);
```

For each number in the loop, multiply the `result` variable by the current number and assign the result to the `result` variable.

```js
// Check the first possible pattern: `result *= i;`
let condition1 = assert.match(code, /for\s*\(\s*let\s+i\s*=\s*1\s*;\s*i\s*<=\s*inputNumber\s*;\s*i\s*\+\+\s*\)\s*{\s*result\s*\*=\s*i;/);

// Check the first possible pattern: `result = result * i;`
let condition2 = assert.match(code, /for\s*\(\s*let\s+i\s*=\s*1\s*;\s*i\s*<=\s*inputNumber\s*;\s*i\s*\+\+\s*\)\s*{\s*result\s*=\s*result\s*\*\s*i;/);

assert.isTrue(condition1 || condition2);
```

You should output the result of the factorial of `inputNumber` to the console using the message `"Factorial of [inputNumber] is [result]"`.

```js

// Check the first possible pattern: console.log(`Factorial of ${inputNumber} is: ${result}`); 
let condition1 = assert.match(code, /console\.log\(["']Factorial\s+of\s*["']\s*\+\s*inputNumber\s*\+\s*["']\s*is:\s*["']\s*\+\s*result\s*\);/);

// Check the second possible pattern: console.log("Factorial of " + inputNumber + " is: "+ result);
let condition2 = assert.match(code, /console\.log\(\`Factorial\s*of\s*\$\{inputNumber\}\s*is:\s*\$\{result\}\`\);\s*/);

// Assert that at least one condition is true
assert(condition1 || condition2);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const num = 5

let result = 1;

for (let i = 1; i <= num; i++) {
  result *= i;
}

const resultMsg = `Factorial of ${num} is: ${result}`;
console.log(resultMsg); 
```
