---
id: 687d8c6f066877768f7ee441
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside your `LinkedList` class, remove the pass keyword and create an inner class called `Node`. This will represent individual nodes that store data in the list.

After creating the `Node` class, you should add the `pass` keyword inside it.

# --hints--

You should remove the `pass` keyword inside the `LinkedList` class and create a `Node` class. Remember to add the `pass` keyword inside the `Node` class.

```js
({
    test: () => runPython(`assert not _Node(_code).find_class("LinkedList").has_pass() and hasattr(LinkedList, "Node") and _Node(_code).find_class("LinkedList").find_class("Node").has_pass()`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
