---
id: 687d8e05a529a4838f1c961c
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Give the `Node` class an `__init__` method that takes an `element` parameter. Inside the constructor, assign `element` to `self.element` and set `self.next` to `None`.

# --hints--

You should set the `element` attribute correctly when a `Node` is created.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

You should initialize the `next` attribute to `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    class Node:
        pass
--fcc-editable-region--
```
