---
id: 687d8c6f066877768f7ee441
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside your `LinkedList` class, remove the pass keyword and create an inner class called `Node`. This will represent individual nodes that store data in the list.

After creating the `Node` class, you should add the `pass` keyword inside it.

# --hints--

You should create a `Node` class inside the `LinkedList` class.

```js
({
    test: () => runPython(`assert hasattr(LinkedList, "Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
