---
id: 687d7deea6a2a22fa8816a73
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In the previous lecture, you learned about the common data structures. In this workshop, you'll implement a linked list.

First, create a class named `LinkedList` to represent the structure that will hold your list nodes. Add the `pass` keyword inside the class so that the test can pass.

# --hints--

You should create a class named `LinkedList`. Remember to add the `pass` keyword.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList") and _Node(_code).find_class("LinkedList").has_pass()`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
