---
id: 687d8e05a529a4838f1c961c
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Next, remove the `pass` keyword inside the `Node` class, then give the `Node` class an `__init__` method that takes an `element` parameter. 

Inside the constructor, assign `element` to `self.element`. This stores the value passed when creating a node.

# --hints--

You should assign `element` to `self.element` inside the `__init__` method. Make sure you remove the `pass` keyword.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A" and not _Node(_code).find_class("LinkedList").find_class("Node").has_pass()
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
