---
id: 687e3be72b1efea97b704c1a
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Define a `remove` method that removes the first occurrence of an element from the list. It should handle three cases: removing the head, a middle node, or a non-existent node. Decrease the length if a node is removed.

# --hints--

You should be able to remove the head node and update the head reference.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    ll.add("A")
    ll.add("B")
    ll.remove("A")
    assert ll.head.element == "B"
    assert ll.length == 1`) 
})
```

You should be able to remove the second node and update the first node's next reference.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    ll.add("A")
    ll.add("B")
    ll.remove("B")
    assert ll.head.element == "A"
    assert ll.head.next is None`) 
})
```

You should handle trying to remove an element that does not exist in the list.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    ll.add("A")
    ll.remove("X")
    assert ll.length == 1`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1
--fcc-editable-region--
```
