---
id: 687e366b3004a9831e4cbd41
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Add a method called `is_empty` that returns `True` if the list is empty (length is `0`) and `False` otherwise.

# --hints--

You should create an `is_empty` method that returns `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None
--fcc-editable-region--
```
