---
id: 687e366b3004a9831e4cbd41
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Add a method called `is_empty` that returns `self.length == 0`.

We'll use this to keep track of how many nodes are in the list. So, with `self.length == 0`, it returns `True`, meaning that the list has no nodes. 

If `self.length` is greater than `0`, it returns `False`, meaning that the list has at least one node.

# --hints--

You should create an `is_empty` method that returns `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
--fcc-editable-region--
    def __init__(self):
        self.length = 0
        self.head = None
--fcc-editable-region--
```
