---
id: 687e3970088c9097de03ef3b
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Create a method named `add` that takes an `element`, creates a `Node`, and adds it to the end of the linked list. If the list is empty, this node becomes the head. Increase the length by `1`.

# --hints--

You should be able to add one element to an empty list.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    ll.add("X")
    assert ll.length == 1
    assert ll.head.element == "X"`) 
})
```

You should be able to add multiple elements and confirm that the new element is placed at the end.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    ll.add("X")
    ll.add("Y")
    assert ll.head.element == "X"
    assert ll.head.next.element == "Y"`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    def is_empty(self):
        return self.length == 0
--fcc-editable-region--
```
