---
id: 687e344cc1993f7316af030c
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Inside the `LinkedList` class, add an `__init__` method that sets `self.length` to `0` and `self.head` to `None`.

# --hints--

You should set `self.length` to `0` and `self.head` to `None` using the `init` method `LinkedList` class.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0 and ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
--fcc-editable-region--
```
