---
id: 68800a64b06e80d9a31c131b
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Still in the `__init__` method of the `LinkedList` class, set `self.head` to `None`. This acts as a reference to the first node in the list. The list is empty at the moment so we have to set it to `None`.


# --hints--

You should set `self.head` to `None` inside the `__init__` method of the `LinkedList` class.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
