---
id: 68857e0a49baf4751f737a2a
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

You should create a method called `add` inside the `LinkedList` class. This method will used to add new elements to the end of the list. 

You should add the `pass` keyword inside the method.

# --hints--

You should create an `add` method and add the `pass` keyword inside it.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").has_pass()`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--
```
