---
id: 6885851e3a743ca7f4ef6c38
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Inside the `add` method, you should create a new `Node` using the given `element`. This creates a new node to be added to the linked list.

Start by creating a `node` variable and assigning `self.Node(element)` to it. You should remove the `pass` keyword first.

# --hints--

You should assign `self.Node(element)` to `node` in the `add` method. Remember to remove the `pass` keyword first.

```js
({
  test: () =>
    assert(
      runPython(
        `not _Node(_code).find_class("LinkedList").find_function("add").has_pass() and _Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        pass
--fcc-editable-region--
```
