---
id: 68857e0a49baf4751f737a2a
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Now, create a method called `add` inside the `LinkedList` class.

Inside your new method, create a variable named `node` and assign it a new instance of `Node` by using `self.Node(element)`. This will create a new node to be added to the linked list.

# --hints--

You should create an `add` method and add the `pass` keyword inside it.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").has_pass()`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--
```
