---
id: 68865003446e2f5d97ff2813
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Still inside the `else` block, create a `while` loop that checks that `current_node.next` is not `None`, then add the `pass` keyword inside it.

# --hints--

You should create a `while` loop that check whether `current_node.next` is not `None`. Remeber to add the `pass` keyword inside the loop.

```js

({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0]
assert _cond.is_equivalent("while current_node.next is not None:\\n pass")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
--fcc-editable-region--
```
