---
id: 688675c4cab82f6eafc8fe4a
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Inside the `while` loop, remove the pass keyword and assign `current_node.next` to `current_node`.

# --hints--

You should remove the `pass` keyword and assign `current_node.next` to `current_node` inside the `while` loop.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0]
assert _cond.is_equivalent("while current_node.next is not None:\\n current_node = current_node.next") and not _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].find_whiles()[0].has_pass()
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                pass
--fcc-editable-region--
```
