---
id: 6888f9827637322d4f309192
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

Next, create a variable called `current_node` inside the `remove` method. It should have a value of `self.head`.

# --hints--

`current_node` should have a value of `self.head`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--
```
