---
id: 6888f53fe737ae0bac9abdc7
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Inside the `remove` method, remove the `pass` keywored, then create a variable called `previous_node` and assign a value of `None` to it.

# --hints--

You should assign `None` to `previous_node` in the `remove` method. Remember to remove the `pass` keyword first.

```js
({
  test: () =>
    assert(
      runPython(
        `not _Node(_code).find_class("LinkedList").find_function("remove").has_pass() and _Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        pass
--fcc-editable-region--
```
