---
id: 6888d05ba074461b64f8eaca
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Next, you're going to create another method that will be used to remove a given `element` from the linked list.

You should create a method called `remove`, and add the `pass` keyword inside it.

# --hints--

You should create a method called `remove`. Add the `pass` keyword inside it.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

--fcc-editable-region--

--fcc-editable-region--
```
