---
id: 688920d7b71de9ac61d59448
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

Inside the `while` loop, remove the pass keyword and assign `current_node` to `previous_node`.

# --hints--

You should remove the `pass` keyword and assign `current_node` to `previous_node` inside the `while` loop.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node") and not _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].has_pass()
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
        while current_node is not None and current_node.element != element:
            pass
--fcc-editable-region--
```
