---
id: 688acb2053ac7978a06345f1
title: Step 23
challengeType: 20
dashedName: step-23
---

# --description--

Next, create an `elif` statement that checks that `previous_node` is not `None`. Add the `pass` keyword inside it.

# --hints--

You should create an `elif` statement that checks that `previous_node` is not `None`, and add the `pass` keyword inside it.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n pass")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
        if current_node is None:
            return        
--fcc-editable-region--
        
--fcc-editable-region--
```
