---
id: 69708d59833ceded7d00541f
title: Implement Linked List Operations
challengeType: 26
dashedName: lab-linked-list-operations
demoType: onClick
---

# --description--

In this lab, you will implement additional operations for a linked list data structure, building on the basic functionality of adding and removing nodes.

**Objective**: Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories**

1. Your `contains` function should return `true` if a specified element exists in the linked list, and `false` otherwise.
1. Your `getAt` function should return the element at a given index in the linked list. If the index is out of bounds, it should return `undefined`.
1. Your `insertAt` function should insert an element at a specific position in the linked list. If the index is out of bounds, it should not modify the list.
1. Your `removeAt` function should remove the node at a given index in the linked list. If the index is out of bounds, it should not modify the list.
1. Your `clear` function should remove all elements from the linked list, effectively resetting it to an empty state.

# --hints--

You should have a `contains` function.

```js
assert.isFunction(contains);
```

`contains` should return `true` if the element exists in the list.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` should return `false` if the element does not exist in the list.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` should return `false` for an empty list.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

You should have a `getAt` function.

```js
assert.isFunction(getAt);
```

`getAt` should return the element at a given index.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` should return `undefined` for invalid indices.

```js
const list = initList();
add(list, 'apple');
assert.isUndefined(getAt(list, -1));
assert.isUndefined(getAt(list, 5));
```

You should have an `insertAt` function.

```js
assert.isFunction(insertAt);
```

`insertAt` should insert an element at the beginning of the list.

```js
const list = initList();
add(list, 'second');
add(list, 'third');
insertAt(list, 0, 'first');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(list.length, 3);
```

`insertAt` should insert an element at the middle of the list.

```js
const list = initList();
add(list, 'first');
add(list, 'third');
insertAt(list, 1, 'second');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(list.length, 3);
```

`insertAt` should insert an element at the end of the list.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` should not modify the list for invalid indices.

```js
const list = initList();
add(list, 'apple');
insertAt(list, -1, 'invalid');
insertAt(list, 10, 'invalid');
assert.strictEqual(list.length, 1);
```

You should have a `removeAt` function.

```js
assert.isFunction(removeAt);
```

`removeAt` should remove the element at a specified index.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
removeAt(list, 1);
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'third');
assert.strictEqual(list.length, 2);
```

`removeAt` should remove the first element correctly.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
```

`removeAt` should not modify invalid indices.

```js
const list = initList();
add(list, 'apple');
removeAt(list, -1);
removeAt(list, 10);
assert.strictEqual(list.length, 1);
```

You should have a `clear` function.

```js
assert.isFunction(clear);
```

`clear` should remove all elements from the list.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  // Return true if the element exists, otherwise false.
}

function getAt(list, index) {
 // Return the element at a given index 
}

function insertAt(list, index, element) {
  // Insert an element at a specific position.
}

function removeAt(list, index) {
 // Remove the node at a given index.
}


function clear(list) {
  // Remove all elements from the list.
}
```

# --solutions--

```js
function contains(list, element) {
  let current = list.head;

  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }

  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;

  let current = list.head;
  let i = 0;

  while (i < index) {
    current = current.next;
    i++;
  }

  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;

  const node = { element, next: null };

  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }

  let current = list.head;
  let previous = null;
  let i = 0;

  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }

  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;

  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }

  let current = list.head;
  let previous = null;
  let i = 0;

  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }

  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
