---
id: 6a89c69c6af682856558d78b
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

The recursive call returns a list containing the smaller numbers. Use `append()` to add `number` to `count_list`, then return `count_list`.

# --hints--

You should append `number` to `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

You should return `count_list` from your `countup` function.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

Your `countup` function should return `[1, 2, 3]` when called with `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
