---
id: 6a89c69c111694edc6d1c628
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Test your `countup` function by passing `5` to it inside a `print()` call. The output should be `[1, 2, 3, 4, 5]`.

With that, the countup function workshop is complete.

# --hints--

You should call `countup(5)` inside the `print()` function.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
