---
id: 6a89c69c99298438d6fbe60b
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Now add the recursive case. After the `if` statement, create a variable named `count_list`. Assign it the result of calling `countup` with `number - 1`.

Each call uses a smaller number, so the recursion moves toward the base case.

# --hints--

You should assign `countup(number - 1)` to a variable named `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
