---
id: 6a89c69c8dbfda74151fa0b6
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will create a `countup` function that uses recursion to create a list of numbers from `1` to a given number.

Start by defining a function named `countup` with one parameter named `number`. Add the `pass` keyword to the function body so that the function is valid while it is empty.

# --hints--

You should have a function named `countup` with a parameter named `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

Your `countup` function should contain the `pass` keyword.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
