---
id: 6a89c69c91bc7365fbd4b733
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Add the base case to your `countup` function. Replace `pass` with an `if` statement that checks if `number` is less than `1`. Inside the `if` statement, return an empty list.

# --hints--

Your `countup` function should return an empty list when `number` is less than `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

Your `countup` function should not return an empty list when `number` is greater than or equal to `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
