---
id: 6565a536ba1f9f25bd30e88b
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви реалізуєте алгоритм Луна. Цей алгоритм є формулою для перевірки різних ідентифікаційних номерів.

<!-- TODO: Maybe cover `if __name__ == '__main__'` -->
<!-- Might not make sense, because import is not covered yet -->
<!-- https://docs.python.org/3/reference/import.html?highlight=__name__#import-related-module-attributes -->

Почніть з оголошення функції під назвою `main`. Використайте ключове слово `pass`, щоб уникнути помилки.

# --hints--

Код повинен містити `def main():`.

```js
({
    test: () => {
        assert(__pyodide.runPython(`
        import inspect
        a = __locals.get("main")
        inspect.isfunction(a)
        `))
    }
})
```

Код повинен містити `pass`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
