---
id: 656877f71bba2b97acedf9af
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Python постачається з вбудованими класами, які можуть допомогти маніпулювати рядками. Одним із них є клас `str`. Він має метод під назвою `maketrans`, який може допомогти створити таблицю перекладу. Цю таблицю можна використовувати для заміни символів у рядку:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

Вищезазначене під час виклику на рядку замінить всі символи `t` на `c` та всі символи `l` на `b`.

Створіть змінну під назвою `card_translation` і призначте до неї таблицю перекладу, щоб замінити всі символи `-` та `` на порожній рядок.

# --hints--

Створіть змінну `card_translation` в межах `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

Призначте значення `str.maketrans({'-': '', ' ': ''})` до `card_translation`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'

--fcc-editable-region--
```
