---
id: 65687b48f2201ab32e06c37c
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Наразі сценарій видає помилку `TypeError`, оскільки ви намагаєтесь додати рядок до цілого числа. Це можна виправити, перетворивши змінну `digit` на ціле число за допомогою вбудованої функції `int`, перш ніж додавати її до `sum_of_odd_digits`:

```python
my_string = '123'
my_int = int(my_string)
```

Перетворіть змінну `digit` на ціле число, перш ніж додати її до `sum_of_odd_digits`. Потім перемістіть виклик `print` у кінець функції `verify_card_number`, щоб надрукувати значення `sum_of_odd_digits`.

# --hints--

Цикл `for` повинен містити `sum_of_odd_digits += int(digit)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

Функція `verify_card_number` повинна містити `print(sum_of_odd_digits)` вкінці.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

Вам більше не знадобиться `print(digit)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
