---
id: 656879c6f35a85a59c06b3a7
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Раніше ви отримали елементи (символи) рядка за допомогою оператора індексу `[]`. Оператор індексу також можна використати, щоб отримати символи рядка в діапазоні. Для цього використайте `string[x:y:h]`:

```python
my_string = 'camperbot'
camper = my_string[0:6]
cp = my_string[0:6:3]
```

Де `x` — початковий індекс, `y` — кінцевий індекс, а `h` — крок (кількість символів, які потрібно пропустити).

Створіть змінну під назвою `card_number_reversed` та призначте до неї значення перших 4 символів `card_number`.

# --hints--

Функція `verify_card_number` повинна містити `card_number_reversed = card_number[0:4]`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
