---
id: 656880227dab4bd8fbc02d41
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

Вітаємо із завершенням проєкту.

Як останній крок, видаліть виклик `print` з функції `verify_card_number` та змініть `card_number` на щось дійсне.

# --hints--

Функція `verify_card_number` не повинна містити виклик `print`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;
        const no_comments = __helpers.python.removeComments(function_body);
        assert.notMatch(no_comments, /print\(/);
    }
})
```

Змініть `card_number` назад на щось дійсне.

```js
({
    test: () => {
        const spyCode = "a = ''\ndef printSpy(*x):\n    global a\n    a = str(x)\n" + code.replaceAll("print(", "printSpy(") + "\na";
        const out = __pyodide.runPython(spyCode);
        assert.notInclude(out, "INVALID!");
        assert.include(out, "VALID!");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number
    total = sum_of_odd_digits + sum_of_even_digits
    print(total)
    return total % 10 == 0

def main():
    card_number = '4111-1111-4555-1141'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    if verify_card_number(translated_card_number):
        print('VALID!')
    else:
        print('INVALID!')

main()
--fcc-editable-region--
```

# --solutions--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    return total % 10 == 0

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    if verify_card_number(translated_card_number):
        print('VALID!')
    else:
        print('INVALID!')

main()
```
