---
id: 656878585631369a6b2d2191
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Визначення перекладу саме по собі не перекладає рядок. Метод `translate` потрібно викликати на рядку, який потрібно перекласти, де аргументом є таблиця перекладу:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

Створіть змінну під назвою `translated_card_number` та призначте до неї результат виклику методу `translate` на `card_number`, де аргументом є `card_translation`.

# --hints--

Створіть змінну `translated_card_number` в межах `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

Призначте значення `card_number.translate(card_translation)` до `translated_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})

--fcc-editable-region--
```
