---
id: 66faa40000000000000000b6
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--
Add a "wind zone" — a translucent column on the left edge that nudges anything inside it to the right:

```js
this.windZone = this.add.rectangle(80, 140, 30, 120, 0x88ccff, 0.2);
this.physics.add.existing(this.windZone, true);
this.physics.add.overlap(this.ball, this.windZone, (ball) => {
  ball.body.velocity.x += 4;
});
```

The fifth argument to `add.rectangle` is alpha — `0.2` makes the column translucent so you can see the ball pass through it. The overlap callback bumps `velocity.x` by `4` every frame the ball is inside, accumulating into a soft horizontal push that the X-axis drag from step 7 eventually wears down.

# --hints--
You should create `this.windZone` as a rectangle at `(80, 140)` sized `30 × 120` with color `0x88ccff` and alpha `0.2`.

```js
assert.match(code, /this\.windZone\s*=\s*this\.add\.rectangle\(\s*80\s*,\s*140\s*,\s*30\s*,\s*120\s*,\s*0x88ccff\s*,\s*0\.2\s*\)/);
```

You should give the wind zone a static body.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.windZone\s*,\s*true\s*\)/);
```

You should register an overlap that adds `4` to `ball.body.velocity.x` per frame of contact.

```js
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.ball\s*,\s*this\.windZone\s*,[\s\S]*?ball\.body\.velocity\.x\s*\+=\s*4/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
    this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
    this.physics.add.existing(this.ceiling, true);
    this.physics.add.collider(this.ball, this.ceiling);
    this.trampoline = this.add.rectangle(420, 320, 60, 6, 0xff44aa);
    this.physics.add.existing(this.trampoline, true);
    this.physics.add.overlap(
      this.ball,
      this.trampoline,
      (ball) => ball.body.setVelocityY(-500),
      (ball) => ball.body.velocity.y > 0
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
