---
id: 66faa40000000000000000b7
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--
Wire up an on-screen HUD that shows the ball's velocity and Y position so you can see the physics engine's numbers in real time. Add a Phaser text object in `create`:

```js
this.hud = this.add.text(8, 8, "", {
  fontSize: "12px",
  color: "#ffffff",
  backgroundColor: "#00000088",
  padding: { x: 4, y: 2 }
});
```

Then, in the `update` method (which runs every frame), refresh the text from the body's current values:

```js
const v = this.ball.body.velocity;
this.hud.setText(
  "v.x " + v.x.toFixed(0) +
  "\nv.y " + v.y.toFixed(0) +
  "\npos.y " + this.ball.y.toFixed(0) +
  "\ngravity " + this.physics.world.gravity.y
);
```

`toFixed(0)` keeps the readout integer-clean. The HUD updates every frame, so you can watch the numbers spike on impact and ease as drag does its work.

# --hints--
You should create `this.hud` as an empty text at `(8, 8)` with the specified style options.

```js
assert.match(code, /this\.hud\s*=\s*this\.add\.text\(\s*8\s*,\s*8\s*,\s*["']{2}/);
assert.match(code, /fontSize\s*:\s*["']12px["']/);
assert.match(code, /backgroundColor\s*:\s*["']#00000088["']/);
```

You should bind `const v = this.ball.body.velocity` in `update`.

```js
assert.match(code, /const\s+v\s*=\s*this\.ball\.body\.velocity/);
```

You should call `this.hud.setText(...)` with `v.x.toFixed(0)`, `v.y.toFixed(0)`, `this.ball.y.toFixed(0)`, and `this.physics.world.gravity.y`.

```js
assert.match(code, /this\.hud\.setText\([\s\S]*?v\.x\.toFixed\(\s*0\s*\)[\s\S]*?v\.y\.toFixed\(\s*0\s*\)[\s\S]*?this\.ball\.y\.toFixed\(\s*0\s*\)[\s\S]*?this\.physics\.world\.gravity\.y/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
    this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
    this.physics.add.existing(this.ceiling, true);
    this.physics.add.collider(this.ball, this.ceiling);
    this.trampoline = this.add.rectangle(420, 320, 60, 6, 0xff44aa);
    this.physics.add.existing(this.trampoline, true);
    this.physics.add.overlap(
      this.ball,
      this.trampoline,
      (ball) => ball.body.setVelocityY(-500),
      (ball) => ball.body.velocity.y > 0
    );
    this.windZone = this.add.rectangle(80, 140, 30, 120, 0x88ccff, 0.2);
    this.physics.add.existing(this.windZone, true);
    this.physics.add.overlap(this.ball, this.windZone, (ball) => {
      ball.body.velocity.x += 4;
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
