---
id: 66faa40000000000000000ac
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Crank the bounce up to `1` to see what perfect elasticity looks like. With `setBounce(1)` the ball returns at the full speed it arrived at — and because gravity adds the same delta on the way down each cycle, it bounces in place forever. This is intentional: it's the easiest way to *feel* the difference a single coefficient makes.

Update the editable line to:

```js
this.ball.body.setBounce(1);
```

# --hints--
You should set the bounce coefficient to `1`.

```js
assert.match(code, /this\.ball\.body\.setBounce\(\s*1\s*\)/);
```

You should no longer call `setBounce(0.7)`.

```js
assert.notMatch(code, /setBounce\(\s*0\.7\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
--fcc-editable-region--
    this.ball.body.setBounce(0.7);
--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
