---
id: 66faa40000000000000000ae
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--
Wire up a click-to-launch handler. On `pointerdown`, set the ball's velocity to the vector pointing from the ball to the click:

```js
this.input.on("pointerdown", (pointer) => {
  this.ball.body.setVelocity(
    pointer.x - this.ball.x,
    pointer.y - this.ball.y
  );
});
```

This gives you a quick aim-and-fire feel. The launch speed is whatever the raw `dx`/`dy` distance is in pixels — clicks far from the ball send it harder than nearby clicks. The next step replaces this with a calibrated, distance-scaled launch.

# --hints--
You should register a `pointerdown` listener on `this.input`.

```js
assert.match(code, /this\.input\.on\(\s*["']pointerdown["']/);
```

You should set the ball's velocity using the raw `pointer.x - this.ball.x` and `pointer.y - this.ball.y` deltas.

```js
assert.match(code, /this\.ball\.body\.setVelocity\([\s\S]*?pointer\.x\s*-\s*this\.ball\.x[\s\S]*?pointer\.y\s*-\s*this\.ball\.y/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
