---
id: 66faa40000000000000000b5
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--
Add a process callback as the fourth argument to `overlap`. The process callback runs first, every frame the bodies touch, and decides whether the contact counts. Return `true` to allow it, `false` to skip — and we only want it to fire when the ball is moving downward:

```js
this.physics.add.overlap(
  this.ball,
  this.trampoline,
  (ball) => ball.body.setVelocityY(-500),
  (ball) => ball.body.velocity.y > 0
);
```

After this change, an upward launch can pass through the pad without triggering the boost. Only a falling ball gets thrown back up.

# --hints--
You should pass a process callback as the fourth argument to the trampoline overlap that returns `ball.body.velocity.y > 0`.

```js
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.ball\s*,\s*this\.trampoline\s*,\s*\(\s*ball\s*\)\s*=>\s*ball\.body\.setVelocityY\(\s*-500\s*\)\s*,\s*\(\s*ball\s*\)\s*=>\s*ball\.body\.velocity\.y\s*>\s*0\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
    this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
    this.physics.add.existing(this.ceiling, true);
    this.physics.add.collider(this.ball, this.ceiling);
--fcc-editable-region--
    this.trampoline = this.add.rectangle(420, 320, 60, 6, 0xff44aa);
    this.physics.add.existing(this.trampoline, true);
    this.physics.add.overlap(
      this.ball,
      this.trampoline,
      (ball) => ball.body.setVelocityY(-500)
    );
--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
