---
id: 66faa40000000000000000b8
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--
Add a clickable RESET button that drops the ball back at its starting position. First, the button itself in `create`:

```js
this.resetBtn = this.add
  .rectangle(440, 24, 60, 22, 0x3344aa)
  .setInteractive();
this.add
  .text(440, 24, "RESET", { fontSize: "11px", color: "#ffffff" })
  .setOrigin(0.5);
this.resetBtn.on("pointerdown", () => this.resetBall());
```

Then, *outside* both `create` and `update` but still inside the `MainScene` class, add a `resetBall` method (a sibling of `create`/`update`):

```js
resetBall() {
  this.ball.setPosition(240, 60);
  this.ball.body.setVelocity(0, 0);
}
```

`setPosition` teleports the visual; `setVelocity(0, 0)` zeroes any momentum so the ball doesn't immediately fly back into wherever it was headed.

# --hints--
You should create `this.resetBtn` as a rectangle at `(440, 24)` sized `60 × 22` with color `0x3344aa` and call `setInteractive()` on it.

```js
assert.match(code, /this\.resetBtn\s*=\s*this\.add[\s\S]*?\.rectangle\(\s*440\s*,\s*24\s*,\s*60\s*,\s*22\s*,\s*0x3344aa\s*\)[\s\S]*?\.setInteractive\(\s*\)/);
```

You should add a centered `RESET` text label.

```js
assert.match(code, /\.text\(\s*440\s*,\s*24\s*,\s*["']RESET["']/);
```

You should bind `pointerdown` on `this.resetBtn` to call `this.resetBall()`.

```js
assert.match(code, /this\.resetBtn\.on\(\s*["']pointerdown["']\s*,\s*\(\s*\)\s*=>\s*this\.resetBall\(\s*\)\s*\)/);
```

You should declare a `resetBall` method on the scene that resets position and velocity.

```js
assert.match(code, /resetBall\s*\(\s*\)\s*\{[\s\S]*?this\.ball\.setPosition\(\s*240\s*,\s*60\s*\)[\s\S]*?this\.ball\.body\.setVelocity\(\s*0\s*,\s*0\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
    this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
    this.physics.add.existing(this.ceiling, true);
    this.physics.add.collider(this.ball, this.ceiling);
    this.trampoline = this.add.rectangle(420, 320, 60, 6, 0xff44aa);
    this.physics.add.existing(this.trampoline, true);
    this.physics.add.overlap(
      this.ball,
      this.trampoline,
      (ball) => ball.body.setVelocityY(-500),
      (ball) => ball.body.velocity.y > 0
    );
    this.windZone = this.add.rectangle(80, 140, 30, 120, 0x88ccff, 0.2);
    this.physics.add.existing(this.windZone, true);
    this.physics.add.overlap(this.ball, this.windZone, (ball) => {
      ball.body.velocity.x += 4;
    });
    this.hud = this.add.text(8, 8, "", {
      fontSize: "12px",
      color: "#ffffff",
      backgroundColor: "#00000088",
      padding: { x: 4, y: 2 }
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
    const v = this.ball.body.velocity;
    this.hud.setText(
      "v.x " + v.x.toFixed(0) +
      "\nv.y " + v.y.toFixed(0) +
      "\npos.y " + this.ball.y.toFixed(0) +
      "\ngravity " + this.physics.world.gravity.y
    );
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
