---
id: 66faa40000000000000000af
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--
The raw `dx`/`dy` launch is too sensitive at long range and too soft up close. Replace it with a clamped, distance-scaled version that uses `Phaser.Math.Distance.Between` to get the magnitude and trig to split it across X and Y:

```js
this.input.on("pointerdown", (pointer) => {
  const dx = pointer.x - this.ball.x;
  const dy = pointer.y - this.ball.y;
  const dist = Phaser.Math.Distance.Between(
    this.ball.x,
    this.ball.y,
    pointer.x,
    pointer.y
  );
  const angle = Math.atan2(dy, dx);
  const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
  this.ball.body.setVelocity(
    Math.cos(angle) * speed,
    Math.sin(angle) * speed
  );
});
```

`Phaser.Math.Clamp(dist * 3, 100, 600)` guarantees a minimum oomph for tap-near launches and a hard cap for tap-far ones, so you never overflow the simulation.

# --hints--
You should compute the click distance with `Phaser.Math.Distance.Between(this.ball.x, this.ball.y, pointer.x, pointer.y)`.

```js
assert.match(code, /Phaser\.Math\.Distance\.Between\(\s*this\.ball\.x\s*,\s*this\.ball\.y\s*,\s*pointer\.x\s*,\s*pointer\.y\s*\)/);
```

You should clamp the launch speed with `Phaser.Math.Clamp(dist * 3, 100, 600)`.

```js
assert.match(code, /Phaser\.Math\.Clamp\(\s*dist\s*\*\s*3\s*,\s*100\s*,\s*600\s*\)/);
```

You should split the speed across X and Y with `Math.cos(angle) * speed` and `Math.sin(angle) * speed`.

```js
assert.match(code, /Math\.cos\(\s*angle\s*\)\s*\*\s*speed/);
assert.match(code, /Math\.sin\(\s*angle\s*\)\s*\*\s*speed/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
--fcc-editable-region--
    this.input.on("pointerdown", (pointer) => {
      this.ball.body.setVelocity(
        pointer.x - this.ball.x,
        pointer.y - this.ball.y
      );
    });
--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
