---
id: 66faa40000000000000000b1
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--
Register a collider so the ball physically bounces off the floor segments:

```js
this.physics.add.collider(this.ball, this.floors);
```

A collider blocks movement and resolves the collision (using the bounce coefficient you set earlier). No callback is needed yet — the engine handles the rebound for you.

# --hints--
You should call `this.physics.add.collider(this.ball, this.floors)`.

```js
assert.match(code, /this\.physics\.add\.collider\(\s*this\.ball\s*,\s*this\.floors\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
