---
id: 66faa40000000000000000b0
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Add five floor segments so the ball has things to bounce off besides the world bounds. Use a static physics group — every member gets an immovable physics body for free:

```js
this.floors = this.physics.add.staticGroup();
for (let i = 0; i < 5; i++) {
  this.floors.add(
    this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
  );
}
```

`physics.add.staticGroup` produces a group whose children are all static bodies. Adding a plain rectangle to it via `.add(rect)` automatically promotes the rectangle to a static body — no explicit `physics.add.existing` call needed.

# --hints--
You should create `this.floors` with `this.physics.add.staticGroup()`.

```js
assert.match(code, /this\.floors\s*=\s*this\.physics\.add\.staticGroup\(\s*\)/);
```

You should loop five times and add a rectangle to `this.floors` each iteration.

```js
assert.match(code, /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*5\s*;\s*i\+\+\s*\)/);
assert.match(code, /this\.floors\.add\(\s*[\s\S]*?this\.add\.rectangle\([\s\S]*?70\s*,\s*8\s*,\s*0x66ddff/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
