---
id: 66faa40000000000000000b2
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--
Add a *visually* tilted floor segment to the same group:

```js
const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
tilted.setAngle(25);
this.physics.add.existing(tilted, true);
this.floors.add(tilted);
```

Now bounce the ball off it — and watch the rebound. The ball treats it like a flat horizontal block, not a 25° ramp. That's because arcade physics bodies are axis-aligned bounding boxes (AABB) and ignore the visual `angle`. Real slopes need a different physics engine — the Matter.js chapter later in the certification handles that. For now, this is your reminder that arcade is fast and rectangular by design.

# --hints--
You should declare a `tilted` rectangle at `(240, 200)` sized `80 × 8` with color `0xff8866`.

```js
assert.match(code, /const\s+tilted\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*200\s*,\s*80\s*,\s*8\s*,\s*0xff8866\s*\)/);
```

You should rotate the visual with `tilted.setAngle(25)`.

```js
assert.match(code, /tilted\.setAngle\(\s*25\s*\)/);
```

You should give the tilted rectangle a static arcade body with `this.physics.add.existing(tilted, true)`.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*tilted\s*,\s*true\s*\)/);
```

You should add the tilted rectangle to `this.floors`.

```js
assert.match(code, /this\.floors\.add\(\s*tilted\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
