---
id: 66faa40000000000000000ba
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--
You're one step from done. With gravity adjustable on the fly and a wind zone that adds velocity each frame, the ball can build up speeds the simulation isn't tuned for — visible as twitching, tunneling, or runaway bounces. Cap both axes right after you read the velocity and before the HUD prints it:

```js
v.x = Phaser.Math.Clamp(v.x, -800, 800);
v.y = Phaser.Math.Clamp(v.y, -800, 800);
```

`Phaser.Math.Clamp` enforces a hard ceiling and floor on each axis, keeping the body within speeds the engine handles cleanly. The HUD now always shows realistic numbers, and the trampoline + wind combo can't crash the physics step.

Reload one final time, click around, ride the trampoline, drift through the wind zone, mash Q/E to flip gravity. You've shipped a complete arcade-physics sandbox: a ball with bounce and drag, click-and-drag launch, five floors plus a tilted teaching-moment ramp, a ceiling, a directional trampoline, a wind zone, a live HUD, a reset button, runtime-tunable gravity, and velocity clamping. Every dial in arcade physics is now visible, audible, and yours.

# --hints--
You should clamp `v.x` between `-800` and `800` with `Phaser.Math.Clamp` after reading the velocity.

```js
assert.match(code, /v\.x\s*=\s*Phaser\.Math\.Clamp\(\s*v\.x\s*,\s*-800\s*,\s*800\s*\)/);
```

You should clamp `v.y` between `-800` and `800` the same way.

```js
assert.match(code, /v\.y\s*=\s*Phaser\.Math\.Clamp\(\s*v\.y\s*,\s*-800\s*,\s*800\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
    this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
    this.physics.add.existing(this.ceiling, true);
    this.physics.add.collider(this.ball, this.ceiling);
    this.trampoline = this.add.rectangle(420, 320, 60, 6, 0xff44aa);
    this.physics.add.existing(this.trampoline, true);
    this.physics.add.overlap(
      this.ball,
      this.trampoline,
      (ball) => ball.body.setVelocityY(-500),
      (ball) => ball.body.velocity.y > 0
    );
    this.windZone = this.add.rectangle(80, 140, 30, 120, 0x88ccff, 0.2);
    this.physics.add.existing(this.windZone, true);
    this.physics.add.overlap(this.ball, this.windZone, (ball) => {
      ball.body.velocity.x += 4;
    });
    this.hud = this.add.text(8, 8, "", {
      fontSize: "12px",
      color: "#ffffff",
      backgroundColor: "#00000088",
      padding: { x: 4, y: 2 }
    });
    this.resetBtn = this.add
      .rectangle(440, 24, 60, 22, 0x3344aa)
      .setInteractive();
    this.add
      .text(440, 24, "RESET", { fontSize: "11px", color: "#ffffff" })
      .setOrigin(0.5);
    this.resetBtn.on("pointerdown", () => this.resetBall());
    this.gravityKeys = this.input.keyboard.addKeys("Q,E");
  }

  update() {
    if (Phaser.Input.Keyboard.JustDown(this.gravityKeys.Q)) {
      this.physics.world.gravity.y = Math.max(
        0,
        this.physics.world.gravity.y - 100
      );
    }
    if (Phaser.Input.Keyboard.JustDown(this.gravityKeys.E)) {
      this.physics.world.gravity.y = Math.min(
        1500,
        this.physics.world.gravity.y + 100
      );
    }
    const v = this.ball.body.velocity;
--fcc-editable-region--

--fcc-editable-region--
    this.hud.setText(
      "v.x " + v.x.toFixed(0) +
      "\nv.y " + v.y.toFixed(0) +
      "\npos.y " + this.ball.y.toFixed(0) +
      "\ngravity " + this.physics.world.gravity.y
    );
  }

  resetBall() {
    this.ball.setPosition(240, 60);
    this.ball.body.setVelocity(0, 0);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```

# --solutions--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
    this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
    this.physics.add.existing(this.ceiling, true);
    this.physics.add.collider(this.ball, this.ceiling);
    this.trampoline = this.add.rectangle(420, 320, 60, 6, 0xff44aa);
    this.physics.add.existing(this.trampoline, true);
    this.physics.add.overlap(
      this.ball,
      this.trampoline,
      (ball) => ball.body.setVelocityY(-500),
      (ball) => ball.body.velocity.y > 0
    );
    this.windZone = this.add.rectangle(80, 140, 30, 120, 0x88ccff, 0.2);
    this.physics.add.existing(this.windZone, true);
    this.physics.add.overlap(this.ball, this.windZone, (ball) => {
      ball.body.velocity.x += 4;
    });
    this.hud = this.add.text(8, 8, "", {
      fontSize: "12px",
      color: "#ffffff",
      backgroundColor: "#00000088",
      padding: { x: 4, y: 2 }
    });
    this.resetBtn = this.add
      .rectangle(440, 24, 60, 22, 0x3344aa)
      .setInteractive();
    this.add
      .text(440, 24, "RESET", { fontSize: "11px", color: "#ffffff" })
      .setOrigin(0.5);
    this.resetBtn.on("pointerdown", () => this.resetBall());
    this.gravityKeys = this.input.keyboard.addKeys("Q,E");
  }

  update() {
    if (Phaser.Input.Keyboard.JustDown(this.gravityKeys.Q)) {
      this.physics.world.gravity.y = Math.max(
        0,
        this.physics.world.gravity.y - 100
      );
    }
    if (Phaser.Input.Keyboard.JustDown(this.gravityKeys.E)) {
      this.physics.world.gravity.y = Math.min(
        1500,
        this.physics.world.gravity.y + 100
      );
    }
    const v = this.ball.body.velocity;
    v.x = Phaser.Math.Clamp(v.x, -800, 800);
    v.y = Phaser.Math.Clamp(v.y, -800, 800);
    this.hud.setText(
      "v.x " + v.x.toFixed(0) +
      "\nv.y " + v.y.toFixed(0) +
      "\npos.y " + this.ball.y.toFixed(0) +
      "\ngravity " + this.physics.world.gravity.y
    );
  }

  resetBall() {
    this.ball.setPosition(240, 60);
    this.ball.body.setVelocity(0, 0);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
