---
id: 66faa40000000000000000b3
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--
Cap the world with a ceiling so a strong upward launch eventually rebounds:

```js
this.ceiling = this.add.rectangle(240, 8, 480, 8, 0x444466);
this.physics.add.existing(this.ceiling, true);
this.physics.add.collider(this.ball, this.ceiling);
```

This ceiling is a separate static body, not part of `this.floors`, because the next steps add directional triggers (a trampoline, a wind zone) that should *only* react to the ball — adding the ceiling to the floor group would conflate them.

# --hints--
You should create `this.ceiling` as a rectangle at `(240, 8)` sized `480 × 8` with color `0x444466`.

```js
assert.match(code, /this\.ceiling\s*=\s*this\.add\.rectangle\(\s*240\s*,\s*8\s*,\s*480\s*,\s*8\s*,\s*0x444466\s*\)/);
```

You should give the ceiling a static body with `this.physics.add.existing(this.ceiling, true)`.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.ceiling\s*,\s*true\s*\)/);
```

You should register a collider between the ball and the ceiling.

```js
assert.match(code, /this\.physics\.add\.collider\(\s*this\.ball\s*,\s*this\.ceiling\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setCollideWorldBounds(true);
    this.ball.body.setBounce(1);
    this.ball.body.setDragX(50);
    this.input.on("pointerdown", (pointer) => {
      const dx = pointer.x - this.ball.x;
      const dy = pointer.y - this.ball.y;
      const dist = Phaser.Math.Distance.Between(
        this.ball.x,
        this.ball.y,
        pointer.x,
        pointer.y
      );
      const angle = Math.atan2(dy, dx);
      const speed = Phaser.Math.Clamp(dist * 3, 100, 600);
      this.ball.body.setVelocity(
        Math.cos(angle) * speed,
        Math.sin(angle) * speed
      );
    });
    this.floors = this.physics.add.staticGroup();
    for (let i = 0; i < 5; i++) {
      this.floors.add(
        this.add.rectangle(50 + i * 90, 280, 70, 8, 0x66ddff)
      );
    }
    this.physics.add.collider(this.ball, this.floors);
    const tilted = this.add.rectangle(240, 200, 80, 8, 0xff8866);
    tilted.setAngle(25);
    this.physics.add.existing(tilted, true);
    this.floors.add(tilted);
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 300 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
