---
id: 66faa40000000000000000a8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
Inside `create`, drop a ball at the top of the canvas and give it an arcade physics body. The visual is a circle and the body is added to it via `physics.add.existing`.

```js
this.ball = this.add.circle(240, 60, 12, 0xffaa44);
this.physics.add.existing(this.ball);
```

Gravity is still `0`, so the ball will hang in mid-air for now. You'll fix that in the next step.

# --hints--
You should add `this.ball` as a circle at `(240, 60)` with radius `12` and color `0xffaa44`.

```js
assert.match(code, /this\.ball\s*=\s*this\.add\.circle\(\s*240\s*,\s*60\s*,\s*12\s*,\s*0xffaa44\s*\)/);
```

You should give the ball an arcade physics body with `this.physics.add.existing`.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.ball\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
