---
id: 66faa40000000000000000a9
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
The ball needs to fall. Open up the editable region in the config and set `gravity.y` to `300`. Arcade physics measures gravity as a constant downward acceleration in pixels per second squared, so `300` gives a calm, observable fall — fast enough to see, slow enough to study.

# --hints--
You should set the world gravity Y to `300` in the config's `physics.arcade` block.

```js
assert.match(code, /arcade\s*:\s*\{\s*gravity\s*:\s*\{\s*y\s*:\s*300\s*\}\s*\}/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Bouncing Ball Lab with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    this.ball = this.add.circle(240, 60, 12, 0xffaa44);
    this.physics.add.existing(this.ball);
  }

  update() {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  physics: {
    default: "arcade",
--fcc-editable-region--
    arcade: { gravity: { y: 0 } }
--fcc-editable-region--
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
