---
id: 657e928716b77b2277980276
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you are going to learn about list comprehensions in Python by building a program that can take a `camelCase` or `PascalCase` formatted string and convert that to a `snake_case` formatted string.

List comprehensions in Python are a concise way to construct a list without using loops or the `.append()` method. Apart from being briefer, list comprehensions often run faster.

Start defining a new function named `convert_to_snake_case()` that accepts a string named `pascal_or_camel_cased_string` as input.

# --hints--

You should define a new function named `convert_to_snake_case()`. Don't forget the colon at the end and use `pass` to fill the function body.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get("convert_to_snake_case")
      inspect.isfunction(a)
    `));
  }
})
```

Your function should take a parameter named `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get('convert_to_snake_case')
      'pascal_or_camel_cased_string' in inspect.signature(a).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
