---
id: 646ae477486c4d02bf4e78f3
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

A set of theme colors will be provided for you to use in the `:root` CSS selector. You will target the `body` element in your CSS and: 

* Set the `background-color` property to `var(--dark-grey)`. 
* Then set the `color` property to `var(--white)`.

# --hints--

Your `body` selector should have a `background-color` property set to `var(--dark-grey)`. Make sure to add a semicolon `;` at the end of the code.

```js
const background = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert.equal(background, 'var(--dark-grey)');
```

Your `body` selector should have a `color` property set to `var(--white)`. Make sure to add a semicolon `;` at the end of the code.

```js
const color = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('color');
assert.equal(color, 'var(--white)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}
--fcc-editable-region--

--fcc-editable-region--
```
