---
id: 646ae02fc520da01e52e0d1e
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Create a `div` HTML element with the class name `title-container`

# --hints--

You Should have a `div` HTML Element 

```js
const div = document.querySelector('div');
assert.exists(div);
```

Your code should have a `div` element with the class `title-container`

```js
const titleContainer = document.querySelector('.title-container');
assert.exists(titleContainer);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
  <script src="script.js"></script>
</html>
```
