---
id: 646ae0eb559a40021a07bae7
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Inside the `div` with the class named `title-container`, You should create a `h1` HTML element with the class name `title` and inside write `Today's Date`

# --hints--

You should have a `h1` HTML Element 

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

You should have a `h1` HTML Element with the class `title`

```js
const title = document.querySelector('.title');
assert.exists(title);
```

Your `h1` HTML element should have the text `Today's Date`

```js
const title = document.querySelector('.title');
assert.exists(title);
assert.equal(title.textContent, "Today's Date");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```
