---
id: 646ae2205fe2e9024cc414ed
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

You will create an HTML `img` tag which will have the `src` attribute set to `an image link here`

# --hints--

You should have an `img` HTML Element 

```js
const img = document.querySelector('img');
assert.exists(img);
```

You should have an `img` HTML Element with the `src` attribute set to `an image link here`

```js
const img = document.querySelector('img');
assert.exists(img);
assert.equal(img.getAttribute('src'), 'an image link here');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```
