---
id: 646afc523275b302f16f948e
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Now, you can see some gaps all around due to the default margin and padding of the project. you can fix that by targeting the universal selector `*` and setting the `margin` and `padding` to `0`. Then set the `box-sizing` property to `border-box` to make sure the padding and border are included in the total width and height of the element.

# --hints--

Your universal selector `*` should have a `margin` property set to `0`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('*')?.getPropertyValue('margin');
assert.ok(margin === '0' || margin === '0px');
```

Your universal selector `*` should have a `padding` property set to `0`.

```js
const padding = new __helpers.CSSHelp(document).getStyle('*')?.getPropertyValue('padding');
assert.ok(padding === '0' || padding === '0px');
```

Your universal selector `*` should have a `box-sizing` property set to `border-box`.

```js
const bs = new __helpers.CSSHelp(document).getStyle('*')?.getPropertyValue('box-sizing');
assert.ok(bs === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}
```
