---
id: 646add141ce9f901b37f6d2b
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

You will create a `main` element inside the `body` element that will contain all the elements of the project.

# --hints--

Your code should have a `main` element

```js
const main = document.querySelector('main');
assert.exists(main);
```

You should have an opening `main` tag.

```js
assert(code.match(/<main\s*>/i));
```

You should have a closing `main` tag.

```js
assert(code.match(/<\/main\s*>/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
  <script src="script.js"></script>
</html>
```
