---
id: 646afe22e06a520335f6fffa
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Inside the CSS file, target the class selector named `title-container` and:

* set the `margin` to `10px 0`
* set the `display` to `flex`
* set the `align-items` to `center`
* set the `justify-content` to `center`
* and finally, set the `gap` to `8px`

# --hints--

Your `title-container` class selector should have a `margin` property set to `10px 0`.

```js
const cssHelper = new __helpers.CSSHelp(document);
const marginTop = cssHelper.getStyle('.title-container')?.getPropertyValue('margin-top');
const marginBottom = cssHelper.getStyle('.title-container')?.getPropertyValue('margin-bottom');
const marginleft = cssHelper.getStyle('.title-container')?.getPropertyValue('margin-left');
const marginright = cssHelper.getStyle('.title-container')?.getPropertyValue('margin-right');
assert.equal(marginTop, '10px');
assert.equal(marginBottom, '10px');
assert.equal(marginleft, '0px');
assert.equal(marginright, '0px');
```

Your `title-container` class selector should have a `display` property set to `flex`.

```js
const display = new __helpers.CSSHelp(document).getStyle('.title-container')?.getPropertyValue('display');
assert.ok(display === 'flex');
```

Your `title-container` class selector should have a `align-items` property set to `center`.

```js
const align = new __helpers.CSSHelp(document).getStyle('.title-container')?.getPropertyValue('align-items');
assert.ok(align === 'center');
```

Your `title-container` class selector should have a `justify-content` property set to `center`.

```js
const justify = new __helpers.CSSHelp(document).getStyle('.title-container')?.getPropertyValue('justify-content');
assert.ok(justify === 'center');
```

Your `title-container` class selector should have a `gap` property set to `8px`.

```js
const gap = new __helpers.CSSHelp(document).getStyle('.title-container')?.getPropertyValue('gap');
assert.ok(gap === '8px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}
--fcc-editable-region--

--fcc-editable-region--
```
