---
id: 646ae2c40c98e8028dc7072f
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Create a HTML `div` element with the class name `divider` that will be later used as an underline for the title.

# --hints--

You should have a `div` HTML Element 

```js
const div = document.querySelector('div');
assert.exists(div);
```

Your code should have a `div` element with the class `divider`

```js
const divider = document.querySelector('.divider');
assert.exists(divider);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
  <script src="script.js"></script>
</html>
```
