---
id: 646b050091d91e036d4135a4
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

target the class named `title` and set the `font-size` property to `2.5rem`

# --hints--

Your `.title` class should have a `font-size` property set to `2.5rem`.

```js
const fontSize = new __helpers.CSSHelp(document).getStyle('.title')?.getPropertyValue('font-size');
assert.ok(fontSize === '2.5rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

--fcc-editable-region--

--fcc-editable-region--
```
