---
id: 646baf7968b91006bc1856c9
title: Step 26
challengeType: 0
dashedName: step-26
---

# --description--

In the same class selector named `.select-icon` in the CSS file, add the following properties:

* set the `width` to `30px`
* set the `height` to `36px`
* set the `padding-left` to `4px`

# --hints--

Your `select-icon` class selector should have a `width` property set to `30px`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('width');
assert.equal(dropdowndisplay, '30px');
```

Your `select-icon` class selector should have a `height` property set to `36px`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('height');
assert.equal(dropdowndisplay, '36px');
```

Your `select-icon` class selector should have a `padding-left` property set to `4px`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('padding-left');
assert.equal(dropdowndisplay, '4px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

.dropdown-container {
  width: 100%;
  margin: 40px auto;
  position: relative;
}

select {
  display: block;
  margin: 20px auto;
  height: 50px;
  font-size: 100%;
  font-weight: bold;

  cursor: pointer;
  background-color: var(--white);
  border: none;
  color: var(--dark-grey);
  appearance: none;
  padding: 10px 38px 10px 10px;
}
--fcc-editable-region--
.select-icon{
  position: absolute;
  top: 4px;
  right: 4px;
  pointer-events: none;

}
--fcc-editable-region--
```
