---
id: 646b6217f2fa8a06818fd8aa
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

You will target the class named `.select-icon` in the CSS file and add the following properties:

* set the `position` to `absolute`
* set the `top` to `4px`
* set the `right` to `4px`
* set the `pointer-events` to `none`

# --hints--

Your `select-icon` class selector should have a `position` property set to `absolute`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('position');
assert.equal(dropdowndisplay, 'absolute');
```

Your `select-icon` class selector should have a `top` property set to `4px`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('top');
assert.equal(dropdowndisplay, '4px');
```

Your `select-icon` class selector should have a `right` property set to `4px`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('right');
assert.equal(dropdowndisplay, '4px');
```

Your `select-icon` class selector should have a `pointer-events` property set to `none`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('.select-icon')?.getPropertyValue('pointer-events');
assert.equal(dropdowndisplay, 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

.dropdown-container {
  width: 100%;
  margin: 40px auto;
  position: relative;
}

select {
  display: block;
  margin: 20px auto;
  height: 50px;
  font-size: 100%;
  font-weight: bold;

  cursor: pointer;
  background-color: var(--white);
  border: none;
  color: var(--dark-grey);
  appearance: none;
  padding: 10px 38px 10px 10px;
}
--fcc-editable-region--

--fcc-editable-region--
```
