---
id: 646b0582494fbe03ab07971c
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

You should target the class named `date-svg` and set the `width` to `30px` and `height` to `30px`.

# --hints--

Your `.date-svg` class should have a `width` property set to `30px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.date-svg')?.getPropertyValue('width');
assert.ok(width === '30px');
```

Your `.date-svg` class should have a `height` property set to `30px`.

```js
const height = new __helpers.CSSHelp(document).getStyle('.date-svg')?.getPropertyValue('height');
assert.ok(height === '30px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

--fcc-editable-region--

--fcc-editable-region--
```
