---
id: 646b56e96902be0551c38b7a
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

You will create an `img` HTML tag with the `src` attribute set to `an image link here`. Then, set the class name of the `img` to `select-icon`. This will be the dropdown icon.

# --hints--

You should have an `img` HTML Element 

```js
const img = document.querySelectorAll('img');
assert.exists(img[1]);
```

You should have an `img` HTML Element with the `src` attribute set to `an image link here`

```js
const img = document.querySelectorAll('img');
assert.exists(img);
assert.equal(img[1].getAttribute('src'), 'an image link here');
```

Your `img` HTML Element should have the class name `select-icon`

```js
const img = document.querySelectorAll('img');
assert.equal(img[1].className, 'select-icon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

```
