---
id: 646b07faa74eb103dd52d206
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Target the class named `divider` and do the following:

* set the `margin` to `auto`
* set the width to `150px`
* set the `height` to `10px`
* set the `background-color` to `var(--yellow)`

# --hints--

Your `.divider` class should have a `margin` property set to `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('.divider')?.getPropertyValue('margin');
assert.ok(margin === 'auto');
```

Your `.divider` class should have a `width` property set to `150px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.divider')?.getPropertyValue('width');
assert.ok(width === '150px');
```

Your `.divider` class should have a `height` property set to `10px`.

```js
const height = new __helpers.CSSHelp(document).getStyle('.divider')?.getPropertyValue('height');
assert.ok(height === '10px');
```

Your `.divider` class should have a `background-color` property set to `var(--yellow)`.

```js
const backgroundColor = new __helpers.CSSHelp(document).getStyle('.divider')?.getPropertyValue('background-color');
assert.ok(backgroundColor === 'var(--yellow)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

--fcc-editable-region--

--fcc-editable-region--
```
