---
id: 646b601805ebd90649999f96
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

In the same selector named `select`, add these styles:

* set the `cursor` to `pointer`
* set the `background-color` to `var(--white)`
* set the `border` to `none`
* set the `color` to `var(--dark-grey)`
* set the `appearance` to `none`
* set the `padding` to `10px 38px 10px 10px`

# --hints--

Your `select` selector should have a `cursor` property set to `pointer`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('cursor');
assert.equal(dropdowndisplay, 'pointer');
```

Your `select` selector should have a `background-color` property set to `var(--white)`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('background-color');
assert.equal(dropdowndisplay, 'var(--white)');
```

Your `select` selector should have a `border` property set to `none`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('border');
assert.equal(dropdowndisplay, 'none');
```

Your `select` selector should have a `color` property set to `var(--dark-grey)`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('color');
assert.equal(dropdowndisplay, 'var(--dark-grey)');
```

Your `select` selector should have a `appearance` property set to `none`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('appearance');
assert.equal(dropdowndisplay, 'none');
```

Your `select` selector should have a `padding` property set to `10px 38px 10px 10px`.

```js
const cssHelper = new __helpers.CSSHelp(document);
const paddingTop = cssHelper.getStyle('select')?.getPropertyValue('padding-top');
const paddingBottom = cssHelper.getStyle('select')?.getPropertyValue('padding-bottom');
const paddingleft = cssHelper.getStyle('select')?.getPropertyValue('padding-left');
const paddingright = cssHelper.getStyle('select')?.getPropertyValue('padding-right');
assert.equal(paddingTop, '10px');
assert.equal(paddingBottom, '10px');
assert.equal(paddingleft, '10px');
assert.equal(paddingright, '38px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

.dropdown-container {
  width: 100%;
  margin: 40px auto;
  position: relative;
}

--fcc-editable-region--
select {
  display: block;
  margin: 20px auto;
  height: 50px;
  font-size: 100%;
  font-weight: bold;

}
--fcc-editable-region--
```
