---
id: 646b4ae6f9d0dc04a3705dcd
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

You will create a `option` HTML element with the `value` attribute set to `dd-mm-yyyy` and set the Text inside to `Day, Month, Year`

# --hints--

You should create a `option` HTML element

```js
const options = document.querySelectorAll('option');
assert.exists(options[0]);
```

Your code should have a `option` element with the `value` attribute set to `dd-mm-yyyy`

```js
const options = document.querySelectorAll('option');
assert.equal(options[0].value, 'dd-mm-yyyy');
```

Your code should have a `option` element with the Text set to `Day, Month, Year`

```js
const options = document.querySelectorAll('option');
assert.equal(options[0].textContent, 'Day, Month, Year');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
--fcc-editable-region--

--fcc-editable-region--
        </select>
      </div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

```
