---
id: 646b5a85776edb05d02d23c4
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

You will target the class selector named `dropdown-container` inside the CSS file and add these:

* set the `width` property to `100%`
* set the `margin` to `40px auto`
* set the `position` to `relative`

# --hints--

Your `dropdown-container` selector should have a `width` property set to `100%`.

```js
const dropdownWidth = new __helpers.CSSHelp(document).getStyle('.dropdown-container')?.getPropertyValue('width');
assert.equal(dropdownWidth, '100%');
```

Your `dropdown-container` selector should have a `margin` property set to `40px auto`.

```js
const cssHelper = new __helpers.CSSHelp(document);
const marginTop = cssHelper.getStyle('.dropdown-container')?.getPropertyValue('margin-top');
const marginBottom = cssHelper.getStyle('.dropdown-container')?.getPropertyValue('margin-bottom');
const marginleft = cssHelper.getStyle('.dropdown-container')?.getPropertyValue('margin-left');
const marginright = cssHelper.getStyle('.dropdown-container')?.getPropertyValue('margin-right');
assert.equal(marginTop, '40px');
assert.equal(marginBottom, '40px');
assert.equal(marginleft, 'auto');
assert.equal(marginright, 'auto');
```

Your `dropdown-container` selector should have a `position` property set to `relative`.

```js
const position = new __helpers.CSSHelp(document).getStyle('.dropdown-container')?.getPropertyValue('position');
assert.equal(position, 'relative');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

--fcc-editable-region--

--fcc-editable-region--
```
