---
id: 646b0bac4577060412809a6c
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Create a `div` with a class name `dropdown-container` that will be used to hold the dropdown menu.

# --hints--

You Should have a `div` HTML Element 

```js
const div = document.querySelector('div');
assert.exists(div);
```

Your code should have a `div` element with the class `dropdown-container`

```js
const titleContainer = document.querySelector('.dropdown-container');
assert.exists(titleContainer);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

```
