---
id: 646b5c9f34975b060e8e5ae5
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

You will target the `select` element and add these styles: 

* set the `display` property to `block`
* set the `margin` to `20px auto`
* set the `height` to `50px`
* set the `font-size` to `100%`
* set the `font-weight` to `bold`

# --hints--

Your `select` selector should have a `display` property set to `block`.

```js
const dropdowndisplay = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('display');
assert.equal(dropdowndisplay, 'block');
```

Your `select` selector should have a `margin` property set to `20px auto`.

```js
const cssHelper = new __helpers.CSSHelp(document);
const marginTop = cssHelper.getStyle('select')?.getPropertyValue('margin-top');
const marginBottom = cssHelper.getStyle('select')?.getPropertyValue('margin-bottom');
const marginleft = cssHelper.getStyle('select')?.getPropertyValue('margin-left');
const marginright = cssHelper.getStyle('select')?.getPropertyValue('margin-right');
assert.equal(marginTop, '20px');
assert.equal(marginBottom, '20px');
assert.equal(marginleft, 'auto');
assert.equal(marginright, 'auto');
```

Your `select` selector should have a `height` property set to `50px`.

```js
const dropdownheight = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('height');
assert.equal(dropdownheight, '50px');
```

Your `select` selector should have a `font-size` property set to `50px`.

```js
const fs = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('font-size');
assert.equal(fs, '100%');
```

Your `select` selector should have a `font-weight` property set to `bold`.

```js
const dropdownwidth = new __helpers.CSSHelp(document).getStyle('select')?.getPropertyValue('font-weight');
assert.equal(dropdownwidth, 'bold');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

.dropdown-container {
  width: 100%;
  margin: 40px auto;
  position: relative;
}

--fcc-editable-region--

--fcc-editable-region--
```
