---
id: 646bb182a47c3f0759416db7
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Javascript starts from this step. Use the `.getElementById()` method to get the `#current-date` element. Assign it to a `currentDateParagraph` variable.

# --hints--

You should use `const` to declare a `currentDateParagraph` variable.

```js
assert.match(code, /const\s+currentDateParagraph\s+=/);
```

You should use `document.getElementById()` to get the `#current-date` element.

```js
assert.match(code, /document\.getElementById\(('|"|`)current-date\1\)/);
```

You should assign the `#current-date` element to your `currentDateParagraph` variable.

```js
assert(code.match(/const\s+currentDateParagraph\s+=\s+document\.getElementById\(('|"|`)current-date\1\)/g));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

.dropdown-container {
  width: 100%;
  margin: 40px auto;
  position: relative;
}

select {
  display: block;
  margin: 20px auto;
  height: 50px;
  font-size: 100%;
  font-weight: bold;

  cursor: pointer;
  background-color: var(--white);
  border: none;
  color: var(--dark-grey);
  appearance: none;
  padding: 10px 38px 10px 10px;
}

.select-icon{
  position: absolute;
  top: 4px;
  right: 4px;
  pointer-events: none;
  width: 30px;
  height: 36px;
  padding-left: 4px;
}

#current-date{
  font-size: 2rem;
  text-align:center
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
