---
id: 646b4a1adb5a4304654d0347
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Create a `select` HTML element with the `name` attribute set to `date-options` and the `id` attribute set to `date-options`.

# --hints--

You should create a `select` HTML element

```js
const select = document.querySelector('select');
assert.exists(select);
```

Your code should have a `select` element with the `name` attribute set to `date-options`

```js
const select = document.querySelector('select');
assert.equal(select.name, 'date-options');
```

Your code should have a `select` element with the `id` attribute set to `date-options`

```js
const select = document.querySelector('select');
assert.equal(select.id, 'date-options');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

```
