---
id: 646e482c620eae02878e9d4c
title: Step 48
challengeType: 0
dashedName: step-48
---

# --description--

In the second case, you will the condition `mm-dd-yyyy-h-mm`. Then you will assign empty template literal to the constant named `currentDateParagraph` using the `textContent` property.

# --hints--

You need to use the `case` where the condition is `mm-dd-yyyy-h-mm` to assign empty template literal to the constant named `currentDateParagraph` using the `textContent` property.

```js
assert(code.match(/case\s*['"]mm-dd-yyyy-h-mm['"]\s*:\s*currentDateParagraph.textContent\s*=\s*``/g));
```

You need to use the `break` after the end of your code 

```js
assert(code.match(/break/g));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Date Formatter</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <main>
      <div class="title-container">
        <h1 class="title">Today's Date</h1>
        <img src="an image link here"/>
      </div>
      <div class="divider"></div>
      <div class="dropdown-container">
        <select name="date-options" id="date-options">
          <option value="dd-mm-yyyy">Day, Month, Year</option>
          <option value="yyyy-mm-dd">Year, Month, Day</option>
          <option value="mm-dd-yyyy-h-mm">
            Month, Day, Year, Hours, Minutes
          </option>
        </select>
        <img src="an image link here" class="select-icon"/>
      </div>
      <p id="current-date"></p>
    </main>
  </body>
  <script src="script.js"></script>
</html>
```

```css
*{
  margin: 0;
  padding: 0;
  box-sizing: border-box
}
:root {
  --dark-grey: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
}

body{
  background-color: var(--dark-grey);
  color: var(--white);
}

.title-container{
  margin: 10px 0;
  display: flex;
  align-items: center;
  justify-content: center;
  gap: 8px;
}

.title {
  font-size: 2.5rem;
}

.date-svg{
  width: 30px;
  height: 30px;
}

.divider {
  margin: auto;
  width: 150px;
  height: 10px;
  background-color: var(--yellow);
}

.dropdown-container {
  width: 100%;
  margin: 40px auto;
  position: relative;
}

select {
  display: block;
  margin: 20px auto;
  height: 50px;
  font-size: 100%;
  font-weight: bold;

  cursor: pointer;
  background-color: var(--white);
  border: none;
  color: var(--dark-grey);
  appearance: none;
  padding: 10px 38px 10px 10px;
}

.select-icon{
  position: absolute;
  top: 4px;
  right: 4px;
  pointer-events: none;
  width: 30px;
  height: 36px;
  padding-left: 4px;
}

#current-date{
  font-size: 2rem;
  text-align:center
}
```

```js
const currentDateParagraph = document.getElementById("current-date")
const dateOptionsSelectElement = document.getElementById("date-options")

const date = new Date();
const day = date.getDate();
const month = date.getMonth() + 1;
const year = date.getFullYear();
const hours = date.getHours();
const minutes = date.getMinutes();

const formattedDate = `${day}-${month}-${year}`
currentDateParagraph.textContent = formattedDate

dateOptionsSelectElement.addEventListener("change", ()=>{
  switch(dateOptionsSelectElement.value){
    case 'yyyy-mm-dd':
      currentDateParagraph.textContent = formattedDate
      .split("-")
      .reverse()
      .join("-")
    break;
--fcc-editable-region--

--fcc-editable-region--
    default: 
      currentDateParagraph.textContent = formattedDate
    break
  }
})
```
