---
id: 614387cbefeeba5f3654a291
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`header` 要素の後に、`div` を作成し、`class` として `author` を設定してください。

その `div` の中に `p` 要素を作成し、`class` を `author-name` に、テキストを `By freeCodeCamp` に設定してください。 この `freeCodeCamp` の部分を `a` 要素で囲み、`href` を `https://freecodecamp.org` に、`target` を `_blank` に設定してください。


その下に 2 つ目の `p` 要素を追加して、クラスを `publish-date` に、テキストを `March 7, 2019` に設定してください。

# --hints--

新しい `div` 要素を作成してください。

```js
assert.exists(document.querySelector('div'));
```

新しい `div` 要素は `header` 要素の後に置く必要があります。

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

`div` 要素の `class` を `author` に設定する必要があります。

```js
assert(document.querySelector('div')?.className === 'author');
```

新しい `p` 要素を 2 つ作成する必要があります。

```js
assert(document.querySelectorAll('p')?.length === 3);
```

新しい 2 つの `p` 要素は `div` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

1 つ目の新しい `p` 要素の `class` を `author-name` に設定する必要があります。

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

1 つ目の新しい `p` 要素には `By freeCodeCamp` というテキストが必要です。

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

2 つ目の新しい `p` 要素の `class` を `publish-date` に設定する必要があります。

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

2 つ目の新しい `p` 要素には `March 7, 2019` というテキストが必要です。

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

新しい `a` 要素を 1 つ作成する必要があります。

```js
assert.exists(document.querySelector('a'));
```

`a` 要素は 1 つ目の新しい `p` 要素の中にある必要があります。

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

`a` 要素の `href` を `https://freecodecamp.org` に設定する必要があります。

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

`a` 要素の `target` を `_blank` に設定する必要があります。

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

`a` 要素でテキスト `freeCodeCamp` を囲む必要があります。

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.textContent === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
