---
id: 6143862a5e54455d262c212e
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`body` 内に `main` 要素を作成してください。 次に、その要素の中に `section` を作成し、`class` を `heading` に設定してください。

# --hints--

`main` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('main'));
```

`main` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

`section` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('section'));
```

`section` 要素は `main` 要素の中にある必要があります。

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

`section` 要素の `class` を `heading` に設定する必要があります。

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
