---
id: 614389f601bb4f611db98563
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`.author` 要素の下に新しい `div` 要素を作成し、クラスを `social-icons` としてください。

その新しい `div` の中に `a`要素を 5 つ追加し、次の `href` 属性を設定してください。

- 1 つ目の `a` 要素の `href` を `https://www.facebook.com/freecodecamp` に設定
- 2 つ目の `a` 要素の `href` を `https://twitter.com/freecodecamp` に設定
- 3 つ目の `a` 要素の `href` を `https://instagram.com/freecodecamp` に設定
- 4 つ目の `a` 要素の `href` を `https://www.linkedin.com/school/free-code-camp` に設定
- 5 つ目の `a` 要素の `href` を `https://www.youtube.com/freecodecamp` に設定

# --hints--

新しい `div` 要素を作成してください。

```js
assert(document.querySelectorAll('div')?.length === 2);
```

新しい `div` 要素を `.author` 要素の後に置く必要があります。

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

新しい `div` 要素にはクラス `social-icons` が必要です。

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

`.social-icons` 要素には `a` 要素が 5 つ必要です。

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

1 つ目の `a` 要素の `href` を `https://www.facebook.com/freecodecamp` に設定する必要があります。

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

2 つ目の `a` 要素の `href` を `https://twitter.com/freecodecamp` に設定する必要があります。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

3 つ目の `a` 要素の `href` を `https://instagram.com/freecodecamp` に設定する必要があります。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

4 つ目の `a` 要素の `href` を `https://www.linkedin.com/school/free-code-camp` に設定する必要があります。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

5 つ目の `a` 要素の `href` を `https://www.youtube.com/freecodecamp` に設定する必要があります。

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
