---
id: 6143869bb45bd85e3b1926aa
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`.heading` 要素の中に、`header` 要素を作成し、`class` 属性として `hero` を設定してください。

この要素の中に、`img` 要素を作成します。`src` 属性には `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` を指定してください。`alt` 属性には `freecodecamp logo` を指定してください。さらに、`class` 属性として `hero-img` を指定してください。

`img` 要素の `loading` 属性として `lazy` を指定すると、ブラウザはその画像が必要になるまで画像のリソースの取得を行いません (ユーザーがスクロールして画像の表示が必要になったら、画像のリソースを取得します)。 遅延なし要素の読み込みが終わるまで、遅延読み込みが指定された要素は読み込みが開始されません。このため、追加のメリットとして、インターネットの接続性が遅いユーザーであっても、画像の読み込みの完了を待たずにページのコンテンツを表示できるようになります。

`img` 要素に `loading` 属性として `lazy` を指定してください。

# --hints--

`header` 要素を 1 つ作成する必要があります。

```js
assert.exists(document.querySelector('header'));
```

`header` 要素は `.heading` 要素の中にある必要があります。

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

`header` 要素の `class` を `hero` に設定する必要があります。

```js
assert(document.querySelector('header')?.className === 'hero');
```

`img` 要素を 1 つ作成する必要があります。

```js
assert.exists(document.querySelector('img'));
```

`img` 要素は `header` 要素の中にある必要があります。

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

`img` 要素の `src` を `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` に設定する必要があります。

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

`img` 要素の `alt` を `freecodecamp logo` に設定する必要があります。

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

`img` 要素の `loading` 属性を `lazy` に設定する必要があります。

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

`img` 要素の `class` を `hero-img` に設定する必要があります。

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
