---
id: 61437d575fb98f57fa8f7f36
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

標準的な HTML ボイラープレートから始めましょう。 `DOCTYPE` 宣言と、このページが英語であることを指定する `html` 要素、`head` 要素 および `body` 要素を追加してください。

`head` 要素内には、適切な `charset` を持つ `<meta>` タグと、モバイル対応にするための `<meta>` タグを追加してください。

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`lang` を `en` に設定した `<html>` の開始タグが必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`<head>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`</head>` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`<body>` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

終了タグ `</body>` が 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

`meta` 要素のひとつは `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

もうひとつの `meta` 要素は `charset` 属性を `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
