---
id: 6140c7e645d8e905819f1dd4
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

从标准样板开始。 添加 `DOCTYPE` 声明、语言设置为英语的 `html` 元素、 `head` 和 `body` 元素。

在 `meta` 元素中添加正确的 `charset`、`title` 元素和链接 `./styles.css` 文件的 `link` 元素。

将 `title` 设置为 `Ferris Wheel`。

# --hints--

你的代码应包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你应该在 `DOCTYPE` 引用后添加一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你应该将文档类型定义为 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你应该在类型之后使用 `>` 结束 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素应该有一个带有值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你应该有一个 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

你应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

你应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

你应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的代码应该有一个 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你的 `meta` 元素应该有一个 `charset` 属性，其值为 `UTF-8`。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的项目应该有一个 `Ferris Wheel` 标题。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'ferris wheel')
```

记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Ferris Wheel');
```

你的代码应该有一个 `link` 元素。

```js
assert.match(code, /<link/)
```

你应该有一个自闭合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
