---
id: 6140c7e645d8e905819f1dd4
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

从标准样板开始。 添加 `DOCTYPE` 声明、语言设置为英语的 `html` 元素、 `head` 和 `body` 元素。

在 `meta` 元素中添加正确的 `charset`、`title` 元素和链接 `./styles.css` 文件的 `link` 元素。

将 `title` 设置为 `Ferris Wheel`。

# --hints--

Your code should contain the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` declaration.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

你应该将文档类型定义为 `html`。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

你应该在类型之后使用 `>` 结束 `DOCTYPE` 声明。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

你的 `html` 元素应该有一个带有值为 `en` 的 `lang` 属性的开始标签。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你的 `html` 元素应该有一个结束标签。

```js
assert.match(code, /<\/html\s*>/);
```

你的 `DOCTYPE` 声明应位于 HTML 的开头。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

你应该有一个 `head` 开始标签。

```js
assert.match(code, /<head\s*>/i);
```

你应该有一个 `head` 结束标签。

```js
assert.match(code, /<\/head\s*>/i);
```

你应该有一个 `body` 开始标签。

```js
assert.match(code, /<body\s*>/i);
```

你应该有一个 `body` 结束标签。

```js
assert.match(code, /<\/body\s*>/i);
```

Your `body` element should be the first element after the `head` element.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Your `head` and `body` elements should be siblings.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Your `head` element should be the first element inside your `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Your `head` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Your `body` element should be within the `html` element.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的代码应该有一个 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你的 `meta` 元素应该有一个 `charset` 属性，其值为 `UTF-8`。

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Your `meta` element should be inside the `head` element.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的项目应该有一个 `Ferris Wheel` 标题。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

记住，标题的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

你的代码应该有一个 `link` 元素。

```js
assert.match(code, /<link/);
```

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
