---
id: 6140cd32d018ed0f600eefce
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

将 `.wheel` 选择器的 `position` 属性设置为 `absolute`。 将 `height` 和 `width` 都设置为 `55vw`。

# --hints--

你的 `.wheel` 选择器应该有一个 `position` 属性设置为 `absolute`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

你的 `.wheel` 选择器应该有一个 `height` 属性设置为 `55vw`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

你的 `.wheel` 选择器应该有一个 `width` 属性设置为 `55vw`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
