---
id: 6140d1a351e88f159ed54fca
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

使用相同的模式，为第三个 `.line`、第四个 `.line`、第五个 `.line` 和第六个 `.line` 创建一个单独的选择器。

将第三个 `.line` 的 `transform` 属性设置为 `rotate(120deg)`，将第四个的属性设置为 `rotate(180deg)` ，第五个的属性设置为 `rotate(240deg)`，第六个的属性设置为 `rotate(300deg)`。

# --hints--

你应该创建一个 `.line:nth-of-type(3)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

你的 `.line:nth-of-type(3)` 选择器应该有一个 `transform` 属性设置为 `rotate(120deg)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform === 'rotate(120deg)');
```

你应该创建一个 `.line:nth-of-type(4)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

你的 `.line:nth-of-type(4)` 选择器应该有一个 `transform` 属性设置为 `rotate(180deg)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform === 'rotate(180deg)');
```

你应该创建一个 `.line:nth-of-type(5)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

你的 `.line:nth-of-type(5)` 选择器应该有一个 `transform` 属性设置为 `rotate(240deg)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform === 'rotate(240deg)');
```

你应该创建一个 `.line:nth-of-type(6)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

你的 `.line:nth-of-type(6)` 选择器应该有一个 `transform` 属性设置为 `rotate(300deg)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform === 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
