---
id: 6140d10d50636e14695013b2
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

为你的 `.line` 元素创建一个选择器。 设置 `transform` 属性为 `rotate(60deg)`。

请记住，`transform`属性允许你改变一个元素的形状。 在这种情况下，使用 `rotate(60deg)` 值会将元素围绕其 `transform-origin` 点顺时针旋转 60 度。

# --hints--

你应该创建一个`.line:nth-of-type(2)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

你的 `.line:nth-of-type(2)` 选择器应该有一个 `transform` 属性值为 `rotate(60deg)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
