---
id: 6140d4bc9db3c81c51a09ab7
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

继续该模式，选择以下 `.cabin` 元素并将特定规则应用于它们：

- 第二个 `.cabin` 应该将 `right` 属性设置为 `17%`，并将 `top` 属性设置为 `93.5%`。
- 第三个 `.cabin` 应将 `right` 属性设置为 `67%`，并将 `top` 属性设置为 `93.5%`。
- 第四个 `.cabin` 应将 `left` 属性设置为 `-8.5%`，并将 `top` 属性设置为 `50%`。
- 第五个 `.cabin` 应将 `left` 属性设置为 `17%`，并将 `top` 属性设置为 `7%`。
- 第六个 `.cabin` 应将 `right` 属性设置为 `17%`，并将 `top` 属性设置为 `7%`。

# --hints--

你应该有一个 `.cabin:nth-of-type(2)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

你的 `.cabin:nth-of-type(2)` 选择器应该有一个 `right` 属性值为 `17%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right === '17%');
```

你的 `.cabin:nth-of-type(2)` 选择器应该有一个 `top` 属性值为 `93.5%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top === '93.5%');
```

你应该有一个 `.cabin:nth-of-type(3)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

你的 `.cabin:nth-of-type(3)` 选择器应该有一个 `right` 属性值为 `67%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right === '67%');
```

你的 `.cabin:nth-of-type(3)` 选择器应该有一个 `top` 属性值为 `93.5%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top === '93.5%');
```

你应该有一个 `.cabin:nth-of-type(4)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

你的 `.cabin:nth-of-type(4)` 选择器应该有一个 `left` 属性值为 `-8.5%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left === '-8.5%');
```

你的 `.cabin:nth-of-type(4)` 选择器应该有一个 `top` 属性设置为 `50%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top === '50%');
```

你应该有一个 `.cabin:nth-of-type(5)` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

你的 `.cabin:nth-of-type(5)` 选择器应该有一个 `left` 属性值为 `17%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left === '17%');
```

你的 `.cabin:nth-of-type(5)` 选择器应该有一个 `top` 属性值为 `7%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top === '7%');
```

你应该有一个 `.cabin:nth-of-type(6)` 选择器。

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

您的 `.cabin:nth-of-type(6)` 选择器应该有一个 `right` 属性值为 `17%`。

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

你的 `.cabin:nth-of-type(6)` 选择器应该有一个 `top` 属性值为 `7%`。

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
