---
id: 6140d0069049f5139d78da40
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

`transform-origin`属性用于设置应用CSS变换点。 例如，当执行 `rotate`时（你将在本项目的后面进行），`transform-origin`决定了元素围绕哪一点旋转。

给 `.line` 选择器一个 `0% 0%` 的 `transform-origin` 属性。 这将使原点从左边偏移 `0%`，从上面偏移 `0%`，把它设置在元素的左上角。

# --hints--

你的 `.line` 选择器应该有一个 `transform-origin` 属性，设置为 `0% 0%`。

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
