---
id: 6140d3dc359b371b1a21d783
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

是时候将机舱定位在车轮周围了。 选择第一个 `.cabin` 元素。 将 `right` 属性设置为 `-8.5%`，将 `top` 属性设置为 `50%`。

# --hints--

你应该有一个 `.cabin:nth-of-type(1)` 选择器。

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s) 
assert.exists(def('.cabin:first-of-type') || def('.cabin:nth-of-type(1)'));
```

你的 `.cabin:nth-of-type(1)` 选择器应该有一个 `right` 属性值为 `-8.5%`。

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(1)') || right('.cabin:first-of-type'), '-8.5%');
```

你的 `.cabin:nth-of-type(1)` 选择器应该有一个 `top` 属性值为 `50%`。

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(1)') || top('.cabin:first-of-type'),'50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
