---
id: 6140cdebd39d6a101e747529
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

为你的 `.line` 元素创建一个选择器。 首先将 `background-color` 设置为 `black`，将 `width` 设置为 `50%`，并将 `height` 设置为 `2px`。

# --hints--

你应该有一个 `.line` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line'));
```

你的 `.line` 选择器应该将 `background-color` 属性设置为 `black`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor === "black");
```

你的 `.line` 选择器应该有一个 `width` 属性值为 `50%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.width === "50%");
```

你的 `.line` 选择器应该有一个 `height` 属性值为 `2px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.height === "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
