---
id: 6140df547f09402144e40b92
title: 步骤 20
challengeType: 0
dashedName: step-20
---

# --description--

`animation-name` 属性用于将 `@keyframes` 规则链接到 CSS 选择器。 此属性的值应与 `@keyframes` 规则的名称匹配。 给你的 `.wheel` 选择器一个 `animation-name` 属性值为 `wheel`。

`animation-duration` 属性用于设置动画应该在多长时间内完成。 时间应以秒（`s`）或毫秒（`ms`）为单位指定。 给你的 `.wheel` 选择器添加值为 `10s` 的 `animation-duration` 属性。

# --hints--

你的 `.wheel` 选择器应该将 `animation-name` 属性值设置为 `wheel`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationName === 'wheel');
```

你的 `.wheel` 选择器应该将 `animation-duration` 属性值设置为 `10s`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.animationDuration === '10s');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}
--fcc-editable-region--

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
   0% {
     transform: rotate(0deg);
   }
   100% {
     transform: rotate(360deg);
   }
}
```
