---
id: 6140cbeec34e970dfe75e710
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

为你的 `.wheel` 元素创建一个选择器。 首先将 `border` 设置为 `2px solid black`，将 `border-radius` 设置为 `50%`，并将 `margin-left` 设置为 `50px`。

# --hints--

你应该有一个 `.wheel` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

你的 `.wheel` 选择器应该有一个 `border` 属性设置为 `2px solid black`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

你的 `.wheel` 选择器应该将 `border-radius` 属性设置为 `50%`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

你的 `.wheel` 选择器应该有一个 `margin-left` 属性设置为 `50px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
