---
id: 61fd7b3fcaa5406257abc5d1
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

È ora di passare alla seconda tabella. Inizia dandole un elemento `caption` con il testo `Liabilities`. Poi aggiungi gli elementi `thead` e `tbody`.

# --hints--

Il secondo elemento `table` dovrebbe avere un elemento `caption`.

```js
assert(document.querySelectorAll('table')?.[1]?.children?.[0]?.localName === 'caption');
```

Il secondo elemento `caption` dovrebbe avere il testo `Liabilities`.

```js
assert(document.querySelectorAll('table')?.[1]?.querySelector('caption')?.textContent === 'Liabilities');
```

Il secondo elemento `table` dovrebbe avere un elemento `thead`.

```js
assert(document.querySelectorAll('table')?.[1]?.querySelector('thead'));
```

Il secondo elemento `table` dovrebbe avere un elemento `tbody`.

```js
assert(document.querySelectorAll('table')?.[1]?.querySelector('tbody'));
```

L'elemento `thead` dovrebbe essere immediatamente al di sotto dell'elemento `caption`.

```js
assert(document.querySelectorAll('table')?.[1]?.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

L'elemento `tbody` dovrebbe essere immediatamente al di sotto dell'elemento `thead`.

```js
assert(document.querySelectorAll('table')?.[1]?.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
              <tr class="total">
                <th>Total <span class="sr-only">Assets</span></th>
                <td>$579</td>
                <td>$736</td>
                <td class="current">$809</td>
              </tr>
            </tbody>
          </table>
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
