---
id: 61fd70336ebb3e4f62ee81ba
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

L'elemento `caption` viene usato per descrivere il contenuto delle tabelle HTML. L'elemento `caption` dovrebbe essere sempre il primo discendente di un elemento `table`, ma può essere posizionato con la proprietà CSS `caption-side`.

Aggiungi un elemento `caption` al primo elemento `table` e dagli il testo `Assets`.

# --hints--

Il primo elemento `table` dovrebbe avere un elemento `caption`.

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

L'elemento `caption` dovrebbe avere il testo `Assets`.

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
