---
id: 61fd719788899952e67692b9
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Gli elementi `thead` e `tbody` sono utilizzati per indicare quale parte della tabella è l'intestazione, e quale porzione contiene i dati o i contenuti primari.

Aggiungi un `thead` e un `tbody` al primo elemento `table`, sotto l'elemento `caption`.

# --hints--

Il primo elemento `table` dovrebbe avere un elemento `thead`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Il primo elemento `table` dovrebbe avere un elemento `tbody`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

L'elemento `thead` dovrebbe essere immediatamente sotto l'elemento `caption`.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

L'elemento `tbody` dovrebbe essere immediatamente sotto l'elemento `thead`.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
