---
id: 61fd71d596e8f253b9408b39
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

L'elemento `tr` è usato per indicare una riga di una tabella. Aggiungi un elemento `tr` all'interno dell'elemento `thead`. Nel nuovo elemento `tr`, aggiungi un elemento`td`, seguito da tre elementi `th`.

L'elemento `td` indica una cella con dati, mentre l'elemento `th` indica una cella di intestazione.

# --hints--

L'elemento `thead` dovrebbe avere un elemento `tr`.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

L'elemento `tr` dovrebbe avere un elemento `td` come primo figlio.

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

L'elemento `tr` dovrebbe avere tre elementi `th` dopo l'elemento `td`.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
