---
id: 61fd6343fb42013d99bcd7f3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno dell'elemento `body`, annida un elemento `section` all'interno di un `main`.

# --hints--

Il primo elemento all'interno dell'elemento `body` dovrebbe essere l'elemento `main`.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

Dovresti avere un elemento `section` all'interno dell'elemento `main`.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
