---
id: 6551ef504d91936d2d4e54f8
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Functions can take *arguments*. You can pass an object as an argument to a function by adding it between the opening and closing parentheses of the function call.

```py
greet = 'Hello!'
print(greet)
```

The code in the example above would print the string `Hello!`, which is the value of the variable `greet` passed to `print()` as the argument.

Print your `text` variable to the screen by passing the `text` variable as an argument to the `print()` function.

# --hints--

You should pass `text` to the `print()` function by including the name of this variable within the opening and closing parentheses. Pay attention to place the function call at the beginning of the line.

```js
({test: () => assert.match(code, /^print\s*\(\s*text\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
