---
id: 6551ef504d91936d2d4e54f8
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Functions can take *arguments*. You can pass an object as an argument to a function by including it between the opening and closing parentheses of the function call.

Print your `text` variable to the screen by passing the `text` variable as an argument to the `print()` function.

# --hints--

You should pass `text` to the `print()` function by including the name of this variable within opening and closing parentheses.

```js
assert.match(code, /print\s*\(\s*text\s*\)/)
```

Your `print()` function call should be placed at the beginning of the line. Remove any white space before it.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
