---
id: 6551f5756c20146fc94f8675
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

You can access the number of characters in a string with the built-in `len()` function.

Modify your existing `print()` call by passing `len(text)` instead of `text[-1]`.

# --hints--


You should call the `len()` function.

```js
assert.match(code, /len\s*\(.*\)/)
```

You should pass the variable `text` to the `len()` function by including it between the parentheses.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

You should pass `len(text)` to the `print()` function by including it between the parentheses. Pay attention to place the function call at the beginning of the line.

```js
({test: () => assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
