---
id: 655220953ba90d80514d7ee2
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

An *argument* is an object or an expression passed to a function — between the parentheses — when it is called. The `print` function can take multiple arguments, separated by a comma.

Add a second argument to `print(char)` so that it prints the character and its index inside the alphabet.

# --hints--

You should add a second argument to the existing `print()` call.

```js
({ test: () => assert.match(code, /print\s*\(\s*char\s*,\s*.*\)/) })
```

You should print `char` and `index` in the same `print()` call.

```js
({ test: () => assert.match(code, /print\s*\(\s*char\s*,\s*index\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
