---
id: 65521203d7165c7b84b22ad4
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

A loop allows you to systematically go through a sequence of elements and execute actions on each one.

In this case, you'll employ a `for` loop. Here's how you can iterate over `text`:

```py
for i in text:
```

`for` is the keyword denoting the loop type. `i` is a variable that sequentially takes the value of the elements in `text`. The statement ends with a colon, `:`.

Remove everything after the `alphabet` line. Then write a `for` loop to iterate over `text`.

# --hints--

You should use the `for` keyword to create a loop. Make sure to place the `for` keyword at the beginning of the line and leave a white space after the keyword.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

You should create a `for` loop to iterate over `text`. Don't forget the final `:`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+\w+\s+in\s+text\s*:\s*$/m)
```

Your `for` loop should be placed below the line of code `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+\w+\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
