---
id: 655261b2e1f2c197093f3993
title: Step 38
challengeType: 20
dashedName: step-38
---

# --description--

Currently, spaces get encrypted as `c`. To maintain the original spacing in the plain message, you'll require a conditional `if` statement. This is composed of the `if` keyword, a condition, and a colon `:`.

```py
if <condition>:
    <code>
```
The `<code>` block is indented to show that it is included in the `if` statement. Only if the `<condition>` evaluates to `true` then the `<code>` is executed. In this example, if `x` is a value greater than 5 then the sentence will be printed. 

```py
if x > 5:
    print("x is greater than five!")
```
Take note of the output of the `print` statement at the top of your `for` loop. You can use the same comparison to check if `char` is equal to an empty space in an `if` statement.

Remove `print(char == ' ')` and replace it with an `if` statement. If `char` is equal to an empty space then print the string `'space!'`. Remember to indent this line.

# --hints--

You should remove `print(char == ' ')`.

```js
({ test: () => assert.isFalse(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/.test(code)) })
```

You should have an `if` statement that triggers when `char` is equal to an empty space.

```js
({ test: () => assert.match(code, /if\s+char\s*==\s*("|')\s\1\s*:/) })
```

You should print `space!` inside your `if` statement.

```js
({ test: () => assert.match(code, /(\s*)if\s+char\s*==\s*("|')\s\2\s*:\s*?\1\1print\s*\(\s*("|')space!\3\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
