---
id: 6553a44b1801991847d8cc69
title: Step 44
challengeType: 20
dashedName: step-44
---

# --description--

Next, modify your `print()` call to print the encrypted message only and put it outside the `for` loop, so that the encrypted string is printed one time.

# --hints--

You should print `'encrypted text:', encrypted_text` after your for loop.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
        print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
