---
id: 6555df0c81300b175308557d
title: Step 88
challengeType: 20
dashedName: step-88
---

# --description--

Below the `print()` call you just added, add another `print()` call to print `Key: python` by concatenating the string `'Key: '` and the value of the `custom_key` variable.

# --hints--

You should have a `print()` call that prints `Key: python` by concatenating the string `'Key: '` and the value of the `custom_key` variable.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Key:\s\1\s*\+\s*custom_key\s*\)/m);
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
