---
id: 655522883e66f618e03a9411
title: Step 80
challengeType: 20
dashedName: step-80
---

# --description--

Modify your comment into `Append any non-letter character to the message`.

# --hints--

You should turn your first comment into `Append any non-letter character to the message`.

```js
({ test: () => assert.match(code, /#\s*Append\sany\snon-letter\scharacter\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if not char.isalpha():
            final_message += char
--fcc-editable-region--
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)

```
