---
id: 656886ec4e4d9b96ded3c565
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--

The iteration variable can have any valid name, but it's convenient to give it a meaningful name.

Rename your `i` variable to `char`.

# --hints--

You should replace `i` with `char`. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
