---
id: 66f55eac933ff64ce654ca74
title: Build a Markdown to HTML Converter
challengeType: 14
dashedName: build-a-markdown-to-html-converter
demoType: onClick
---

# --description--

For this lab, all the HTML and CSS has been provided to you. You will use JavaScript to complete the Markdown to HTML Converter app so that it can handle the conversion of basic Markdown constructs into HTML elements.

**Note:** The final result won't be a comprehensive Markdown to HTML converter, but you can add extra functionalities to it if you would like.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a function named `convertMarkdown` that takes no parameters.
1. The `convertMarkdown` function should convert the markdown input from `#markdown-input` into HTML and should return a string containing the HTML code.
1. When you input text inside `#markdown-input`, the raw HTML code returned by `convertMarkdown` should be displayed inside `#html-output`.

Here's a table containing all the markdown that `convertMarkdown` should be able to handle and the expected HTML after conversion:

<table>
  <thead>
    <tr>
      <th>Markdown</th>
      <th>HTML</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td><code># heading 1</code></th>
      <td><code>&lth1&gtheading 1&lt/h1&gt</code></th>
    </tr>
    <tr>
      <td><code>## heading 2</code></th>
      <td><code>&lth2&gtheading 2&lt/h2&gt</code></th>
    </tr>
    <tr>
      <td><code>### heading 3</code></th>
      <td><code>&lth3&gtheading 3&lt/h3&gt</code></th>
    </tr>
    <tr>
      <td><code>**bold text**</code> or <code>__bold text__</code></th>
      <td><code>&ltstrong&gtbold text&lt/strong&gt</code></th>
    </tr>
    <tr>
      <td><code>*italic text*</code> or <code>_italic text_</code></th>
      <td><code>&ltem&gtitalic text&lt/em&gt</code></th>
    </tr>
    <tr>
      <td><code>![alt-text](image-source)</code></th>
      <td><code>&ltimg alt="alt-text" src="image-source"&gt</code></th>
    </tr>
    <tr>
      <td><code>[link text](URL)</code></th>
      <td><code>&lta href="URL"&gtlink text&lt/a&gt</code></th>
    </tr>
    <tr>
      <td><code>> quote</code></th>
      <td><code>&ltblockquote&gtquote&lt/blockquote&gt</code></th>
    </tr>
  </tbody>
</table>

**Note:** Be sure to link your JavaScript file in your HTML.

# --hints--

You should have a function named `convertMarkdown`.

```js
assert.isFunction(convertMarkdown);
```

When the value of `#markdown-input` is `# title 1`, `convertMarkdown()` should return `<h1>title 1</h1>`.

```js
input.value = "# title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH1 = testDiv.querySelector("h1")
assert.lengthOf(testDiv.children, 1)
assert.exists(testH1);
assert.equal(testH1.innerText, "title 1");
```

When the value of `#markdown-input` is `# title 1`, `<h1>title 1</h1>` should be displayed inside `#html-output`.

```js
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testH1 = testDiv.querySelector("h1")
assert.lengthOf(testDiv.children, 1)
assert.exists(testH1);
assert.equal(testH1.innerText, "title 1");
```

When the value of `#markdown-input` is `some text # title 1`, `convertMarkdown()` should not convert `# title 1` into an `h1` element.

```js
input.value = "some text # title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH1 = testDiv.querySelector("h1")
assert.notExists(testH1);
```

When the value of `#markdown-input` is `## title 2`, `convertMarkdown()` should return `<h2>title 2</h2>`.

```js
input.value = "## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH2 = testDiv.querySelector("h2")
assert.lengthOf(testDiv.children, 1)
assert.exists(testH2);
assert.equal(testH2.innerText, "title 2");
```

When the value of `#markdown-input` is `## title 2`, `<h2>title 2</h2>` should be displayed inside `#html-output`.

```js
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testH2 = testDiv.querySelector("h2")
assert.lengthOf(testDiv.children, 1)
assert.exists(testH2);
assert.equal(testH2.innerText, "title 2");
```

When the value of `#markdown-input` is `some text ## title 2`, `convertMarkdown()` should not convert `## title 2` into an `h2` element.

```js
input.value = "some text ## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH2 = testDiv.querySelector("h2")
assert.notExists(testH2);
```

When the value of `#markdown-input` is `### title 3`, `convertMarkdown()` should return `<h3>title 3</h3>`.

```js
input.value = "### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH3 = testDiv.querySelector("h3")
assert.lengthOf(testDiv.children, 1)
assert.exists(testH3);
assert.equal(testH3.innerText, "title 3");
```

When the value of `#markdown-input` is `### title 3`, `<h3>title 3</h3>` should be displayed inside `#html-output`.

```js
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testH3 = testDiv.querySelector("h3")
assert.lengthOf(testDiv.children, 1)
assert.exists(testH3);
assert.equal(testH3.innerText, "title 3");
```

When the value of `#markdown-input` is `some text ### title 3`, `convertMarkdown()` should not convert `### title 3` into an `h3` element.

```js
input.value = "some text ### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH3 = testDiv.querySelector("h3")
assert.notExists(testH3);
```

When the value of `#markdown-input` is `**this is bold**`, `convertMarkdown()` should return `<strong>this is bold</strong>`.

```js
input.value = "**this is bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testStrong = testDiv.querySelector("strong")
assert.lengthOf(testDiv.children, 1)
assert.exists(testStrong);
assert.equal(testStrong.innerText, "this is bold");
```

When the value of `#markdown-input` is `**this is bold**`, `<strong>this is bold</strong>` should be displayed inside `#html-output`.

```js
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testStrong = testDiv.querySelector("strong")
assert.lengthOf(testDiv.children, 1)
assert.exists(testStrong);
assert.equal(testStrong.innerText, "this is bold");
```

When the value of `#markdown-input` is `__this is bold__`, `convertMarkdown()` should return `<strong>this is bold</strong>`.

```js
input.value = "__this is bold__";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testStrong = testDiv.querySelector("strong")
assert.lengthOf(testDiv.children, 1)
assert.exists(testStrong);
assert.equal(testStrong.innerText, "this is bold");
```

When the value of `#markdown-input` is `__this is bold__`, `<strong>this is bold</strong>` should be displayed inside `#html-output`.

```js
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testStrong = testDiv.querySelector("strong")
assert.lengthOf(testDiv.children, 1)
assert.exists(testStrong);
assert.equal(testStrong.innerText, "this is bold");
```

When the value of `#markdown-input` is `*this is italic*`, `convertMarkdown()` should return `<em>this is italic</em>`.

```js
input.value = "*this is italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testItalic = testDiv.querySelector("em")
assert.lengthOf(testDiv.children, 1)
assert.exists(testItalic);
assert.equal(testItalic.innerText, "this is italic");
```

When the value of `#markdown-input` is `*this is italic*`, `<em>this is italic</em>` should be displayed inside `#html-output`.

```js
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testItalic = testDiv.querySelector("em")
assert.lengthOf(testDiv.children, 1)
assert.exists(testItalic);
assert.equal(testItalic.innerText, "this is italic");
```

When the value of `#markdown-input` is `_this is italic_`, `convertMarkdown()` should return `<em>this is italic</em>`.

```js
input.value = "_this is italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testItalic = testDiv.querySelector("em")
assert.lengthOf(testDiv.children, 1)
assert.exists(testItalic);
assert.equal(testItalic.innerText, "this is italic");
```

When the value of `#markdown-input` is `_this is italic_`, `<em>this is italic</em>` should be displayed inside `#html-output`.

```js
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testItalic = testDiv.querySelector("em")
assert.lengthOf(testDiv.children, 1)
assert.exists(testItalic);
assert.equal(testItalic.innerText, "this is italic");
```

When the value of `#markdown-input` is either `# **title 1**` or `# __title 1__`, `convertMarkdown()` should return `<h1><strong>title 1</strong></h1>`.

```js
input.value = "# **title 1**";
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = convertMarkdown();
let testH1 = testDiv1.querySelector("h1");
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.exists(testH1);
assert.lengthOf(testH1.children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

input.value = "# __title 1__";
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = convertMarkdown();
testH1 = testDiv2.querySelector("h1");
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.exists(testH1);
assert.lengthOf(testH1.children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

When the value of `#markdown-input` is either `# **title 1**` or `# __title 1__`, `<h1><strong>title 1</strong></h1>` should be displayed inside `#html-output`.

```js
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = output.innerText;
let testH1 = testDiv1.querySelector("h1")
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.exists(testH1);
assert.lengthOf(testH1.children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = output.innerText;
testH1 = testDiv2.querySelector("h1")
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.exists(testH1);
assert.lengthOf(testH1.children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

When the value of `#markdown-input` is `![alt-text](image-source)`, `convertMarkdown()` should return `<img alt="alt-text" src="image-source">`.

```js
input.value = "![alt-text](image-source)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testImg = testDiv.querySelector("img")
assert.lengthOf(testDiv.children, 1)
assert.exists(testImg);
assert.equal(testImg.alt, "alt-text");
assert.isTrue(testImg.src.endsWith("image-source"));
```

When the value of `#markdown-input` is `![alt-text](image-source)`, `<img alt="alt-text" src="image-source">` should be displayed inside `#html-output`.

```js
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testImg = testDiv.querySelector("img")
assert.lengthOf(testDiv.children, 1)
assert.exists(testImg);
assert.equal(testImg.alt, "alt-text");
assert.isTrue(testImg.src.endsWith("image-source"));
```

When the value of `#markdown-input` is `[link text](URL)`, `convertMarkdown()` should return `<a href="URL">link text</a>`.

```js
input.value = "[link text](URL)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testAnchor = testDiv.querySelector("a")
assert.lengthOf(testDiv.children, 1)
assert.exists(testAnchor);
assert.isTrue(testAnchor.href.endsWith("URL"));
assert.equal(testAnchor.innerText, "link text");
```

When the value of `#markdown-input` is `[link text](URL)`, `<a href="URL">link text</a>` should be displayed inside `#html-output`.

```js
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testAnchor = testDiv.querySelector("a")
assert.lengthOf(testDiv.children, 1)
assert.exists(testAnchor);
assert.isTrue(testAnchor.href.endsWith("URL"));
assert.equal(testAnchor.innerText, "link text");
```

When the value of `#markdown-input` is `> this is a quote`, `convertMarkdown()` should return `<blockquote>this is a quote</blockquote>`.

```js
input.value = "> this is a quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testQuote = testDiv.querySelector("blockquote")
assert.lengthOf(testDiv.children, 1)
assert.exists(testQuote);
assert.equal(testQuote.innerText, "this is a quote");
```

When the value of `#markdown-input` is `> this is a quote`, `<blockquote>this is a quote</blockquote>` should be displayed inside `#html-output`.

```js
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testQuote = testDiv.querySelector("blockquote")
assert.lengthOf(testDiv.children, 1)
assert.exists(testQuote);
assert.equal(testQuote.innerText, "this is a quote");
```

When the value of `#markdown-input` is `some text > not a quote anymore`, `convertMarkdown()` should not convert `> not a quote anymore` into a `blockquote` element.

```js
input.value = "some text > not a quote anymore";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testQuote = testDiv.querySelector("blockquote")
assert.notExists(testQuote);
```

When the value of `#markdown-input` is `> **this is a *quote***`, `convertMarkdown()` should return `<blockquote><strong>this is a <em>quote</em></strong></blockquote>`.

```js
input.value = "> **this is a *quote***";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testQuote = testDiv.querySelector("blockquote");
const testStrong = testDiv.querySelector("blockquote>strong");
const testItalic = testDiv.querySelector("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.exists(testQuote);
assert.lengthOf(testQuote.children, 1);
assert.exists(testStrong);
assert.lengthOf(testStrong.children, 1);
assert.equal(testStrong.innerText, "this is a quote");
assert.exists(testItalic);
assert.equal(testItalic.innerText, "quote");
```

When the value of `#markdown-input` is `> **this is a *quote***`, `<blockquote><strong>this is a <em>quote</em></strong></blockquote>` should be displayed inside `#html-output`.

```js
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const testQuote = testDiv.querySelector("blockquote")
const testStrong = testDiv.querySelector("blockquote>strong");
const testItalic = testDiv.querySelector("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.exists(testQuote);
assert.lengthOf(testQuote.children, 1);
assert.exists(testStrong);
assert.lengthOf(testStrong.children, 1);
assert.equal(testStrong.innerText, "this is a quote");
assert.exists(testItalic);
assert.equal(testItalic.innerText, "quote");
```

# --seed--

## --after-user-code--

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
```

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Markdown to HTML Converter</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>Markdown to HTML Converter</h1>
    <textarea id="markdown-input" placeholder="Enter your markdown here..."></textarea>
    <h2>HTML:</h2>
    <div id="html-output"></div>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

body {
    font-family: Arial, sans-serif;
    padding: 20px;
}
#markdown-input {
    width: 100%;
    height: 200px;
}
#html-output {
    display: inline-block;
    width: 100%;
    border: 1px solid #ccc;
    padding: 10px;
    margin: 20px auto 20px;
    white-space: pre-wrap;
    background-color: #f9f9f9;
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Markdown to HTML Converter</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>Markdown to HTML Converter</h1>
    <textarea id="markdown-input" placeholder="Enter your markdown here..."></textarea>
    <h2>HTML:</h2>
    <div id="html-output"></div>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

body {
    font-family: Arial, sans-serif;
    padding: 20px;
}
#markdown-input {
    width: 100%;
    height: 200px;
}
#html-output {
    display: inline-block;
    width: 100%;
    border: 1px solid #ccc;
    padding: 10px;
    margin: 20px auto 20px;
    white-space: pre-wrap;
    background-color: #f9f9f9;
}
```

```js
const patternReplacementArray = [
    { h6: [/^[ \t]*######[ \t](.+$)/gm, "<h6>$1</h6>\n"] },
    { h5: [/^[ \t]*#####[ \t](.+$)/gm, "<h5>$1</h5>\n"] },
    { h4: [/^[ \t]*####[ \t](.+$)/gm, "<h4>$1</h4>\n"] },
    { h3: [/^[ \t]*###[ \t](.+$)/gm, "<h3>$1</h3>\n"] },
    { h2: [/^[ \t]*##[ \t](.+$)/gm, "<h2>$1</h2>\n"] },
    { h1: [/^[ \t]*#[ \t](.+$)/gm, "<h1>$1</h1>\n"] },
    { hr: [/(?<=\n|^)---(?=\n|$)/, "<hr>\n"] },
    { strong: [/((?:\*|_){2})(.*?)\1/g, "<strong>$2</strong>"] },
    { em: [/(\*|_)(.*?)\1/g, "<em>$2</em>"] },
    { img: [/!\[\s*(.*?)\s*\]\(\s*(.*?)\s*\)/g, '<img alt="$1" src="$2">'] },
    { a: [/\[(.*?)\]\(\s*(.*?)\s*\)/g, '<a href="$2">$1</a>'] },
    { blockquote1: [/^[ \t]*> ([^\n]+?)$/gm, "<blockquote>$1</blockquote>\n"] },
    { blockquote2: [/(?<=<blockquote>)(.+?)<\/blockquote>\n\n^([^<\n]+?)$/gm, "$1 $2</blockquote>\n"] },
    { ul: [/(?:^[-*+]\s+.*?$\n)+\s*?/gm, "<ul>\n$&</ul>\n\n"] },
    { ol: [/(?:^\d*\.\s+?\w.*?$\n)+\s*?/gm, "<ol>\n$&</ol>\n\n"] },
    { uli: [/^[-*+]\s+(.*?$)\n/gm, "<li>$1</li>\n"] },
    { oli: [/^\d*\.\s+?(.*?$)\n/gm, "<li>$1</li>\n"] },
    { p1: [/^([^<>\s].*?)(?=\n\s*$)/gsm, "<p>$1</p>\n"] },
    { p2: [/^\w*?(?:<strong>|<em>).+(?:<\/strong>|<\/em>)\s*\w*?[ \t]*$/gm, "<p>$&</p>\n"] }
]

const markdownInput = document.getElementById('markdown-input');
const convertMarkdown = () => {
    let markdown = markdownInput.value;
    let matchUList = false;
    let matchOList = false;
    patternReplacementArray.forEach((item) => {
        const [key] = Object.keys(item);
        const [pattern, replacement] = Object.values(item)[0];
        if (key === "blockquote2") {
            while (markdown.match(pattern)) {
                markdown = markdown.replace(pattern, replacement);
            }
        } else if (key === "ul") {
            if (markdown.match(pattern)) {
                matchUList = true;
                markdown = markdown.replace(pattern, replacement);
            }

        } else if (key === "ol") {
            if (markdown.match(pattern)) {
                matchOList = true;
                markdown = markdown.replace(pattern, replacement);
            }

        } else if (key === "uli" && !matchUList) {
            return;
        } else if (key === "oli" && !matchOList) {
            return;
        } else {
            markdown = markdown.replace(pattern, replacement);
        }
    })
    return markdown;
}

markdownInput.addEventListener('input', () => {
    const html = convertMarkdown();
    const out = document.getElementById('html-output');
    out.textContent = html;
});
```
