---
id: 6448d62ce222044458b75931
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Your `Convert` button should be working now. But it could get tiring for users to enter in a number, then click that button each time they want to convert from decimal to binary. It would be much more convenient to perform the conversion when the `Enter` or `Return` key is pressed.

The `keydown` event fires every time a user presses a key on their keyboard, and is a good way to add more interactivity to `input` elements.

Chain `.addEventListener()` to `numberInput`. The event listener should listen for `keydown` events and take an empty arrow function as a callback.

# --hints--

You should call the `.addEventListener()` method on your `numberInput` element.

```js
assert.match(code, /numberInput\s*\.\s*addEventListener\s*\(/);
```

Your `.addEventListener()` method should listen for the `keydown` event.

```js
assert.match(code, /numberInput\s*\.\s*addEventListener\s*\(\s*('|"|`)keydown\1\s*,?/);
```

Your `.addEventListener()` method should take a callback function using the arrow syntax.

```js
assert.match(code, /('|"|`)keydown\1\s*,\s*\(\s*\)\s*=>/);
```

Your callback function should be empty.

```js
assert.match(code, /('|"|`)keydown\1\s*,\s*\(\s*\)\s*=>\s*{\s*}\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Decimal to Binary Converter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Decimal to Binary Converter</h1>
    <div class="input-container">
      <label for="number-input">Enter a decimal number:</label>
      <input
        value=""
        type="number"
        name="decimal number input"
        id="number-input"
        class="number-input"
      />
      <button class="convert-btn" id="convert-btn">Convert</button>
    </div>
    <output id="result" for="number-input"></output>
    <h2>Call stack</h2>
    <div id="animation-container"></div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #f5f6f7;
  --dark-blue: #1b1b32;
  --orange: #f1be32;
}

body {
  display: flex;
  flex-direction: column;
  align-items: center;
  font-family: Menlo, Consolas, Monaco, Liberation Mono, Lucida Console,
    monospace;
  font-size: 1.125rem;
  color: var(--light-grey);
  background-color: var(--dark-blue);
  padding: 0 15px;
}

h1 {
  font-size: 2.3rem;
  text-align: center;
  margin: 20px 0;
}

h2 {
  font-size: 1.5rem;
  text-align: center;
  margin: 10px 0 20px 0;
}

.input-container {
  display: flex;
  flex-direction: column;
  gap: 10px;
  justify-content: center;
  align-items: center;
  margin: 10px 0;
}

.convert-btn {
  font-size: inherit;
  font-family: inherit;
  background-color: var(--orange);
  width: 5rem;
  height: 2rem;
  padding: 4px;
  border: none;
  cursor: pointer;
}

.number-input {
  font-size: inherit;
  padding: 0.3rem;
}

#result {
  display: inline-flex;
  align-items: center;
  justify-content: center;
  font-size: 2rem;
  text-align: center;
  min-width: 200px;
  width: fit-content;
  min-height: 80px;
  margin: 10px 0;
  padding: 15px;
  border: 5px solid var(--orange);
}

#animation-container {
  display: flex;
  flex-direction: column-reverse;
  justify-content: end;
  gap: 1rem;
  margin-inline: auto;
  margin-block-end: 1rem;
  width: clamp(320px, 30vw, 420px);
  min-height: 40vh;
  border: 2px solid var(--orange);
  padding: 1rem;
}

.animation-frame {
  font-family: -apple-system, BlinkMacSystemFont, avenir next, avenir, segoe ui,
    helvetica neue, helvetica, Cantarell, Ubuntu, roboto, noto, arial,
    sans-serif;
  padding: 15px 10px;
  border: 5px solid var(--orange);
  font-size: 1.2rem;
  text-align: center;
}

@media screen and (min-width: 540px) {
  .input-container {
    flex-direction: row;
  }

  #result {
    max-width: 460px;
  }
}
```

```js
const numberInput = document.getElementById("number-input");
const convertBtn = document.getElementById("convert-btn");
const result = document.getElementById("result");

const checkUserInput = () => {
  console.log(numberInput.value);
};

convertBtn.addEventListener("click", checkUserInput);

--fcc-editable-region--

--fcc-editable-region--
```
