---
id: 644a37d773b7610ac5ede66e
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Now that you're familiar with binary numbers, it's time to finish building the function to do the conversion for you. You'll start off with a simpler solution first, then refactor that into a recursive solution.

First, you need to create some arrays to store the inputs and results of the division you'll do in the following steps. These will make it easier to see how the decimal to binary conversion works.

Remove the `return` statement from your `decimalToBinary` function. Then, declare variables named `inputs`, `quotients`, and `remainders`, and assign an empty array to each of them.

# --hints--

You should remove the `return` statement from your `decimalToBinary` function.

```js
assert.notMatch(String(decimalToBinary), /return\s*('|"|`)(\s|\d)*\1\s*;?\s*\}$/);
```

You should declare a variable named `inputs`.

```js
assert.match(String(decimalToBinary), /(var|let|const)\s+inputs/);
```

You should assign an empty array to `inputs`.

```js
assert.match(String(decimalToBinary), /inputs\s*=\s*\[\s*\]/);
```

You should declare a variable named `quotients`.

```js
assert.match(String(decimalToBinary), /(var|let|const)\s+quotients/);
```

You should assign an empty array to `quotients`.

```js
assert.match(String(decimalToBinary), /quotients\s*=\s*\[\s*\]/);
```

You should declare a variable named `remainders`.

```js
assert.match(String(decimalToBinary), /(var|let|const)\s+remainders/);
```

You should assign an empty array to `remainders`.

```js
assert.match(String(decimalToBinary), /remainders\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Decimal to Binary Converter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Decimal to Binary Converter</h1>
    <div class="input-container">
      <label for="number-input">Enter a decimal number:</label>
      <input
        value=""
        type="number"
        name="decimal number input"
        id="number-input"
        class="number-input"
      />
      <button class="convert-btn" id="convert-btn">Convert</button>
    </div>
    <output id="result" for="number-input"></output>
    <h2>Call stack</h2>
    <div id="animation-container"></div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  box-sizing: border-box;
  margin: 0;
  padding: 0;
}

:root {
  --light-grey: #f5f6f7;
  --dark-blue: #1b1b32;
  --orange: #f1be32;
}

body {
  display: flex;
  flex-direction: column;
  align-items: center;
  font-family: Menlo, Consolas, Monaco, Liberation Mono, Lucida Console,
    monospace;
  font-size: 1.125rem;
  color: var(--light-grey);
  background-color: var(--dark-blue);
  padding: 0 15px;
}

h1 {
  font-size: 2.3rem;
  text-align: center;
  margin: 20px 0;
}

h2 {
  font-size: 1.5rem;
  text-align: center;
  margin: 10px 0 20px 0;
}

.input-container {
  display: flex;
  flex-direction: column;
  gap: 10px;
  justify-content: center;
  align-items: center;
  margin: 10px 0;
}

.convert-btn {
  font-size: inherit;
  font-family: inherit;
  background-color: var(--orange);
  width: 5rem;
  height: 2rem;
  padding: 4px;
  border: none;
  cursor: pointer;
}

.number-input {
  font-size: inherit;
  padding: 0.3rem;
}

#result {
  display: inline-flex;
  align-items: center;
  justify-content: center;
  font-size: 2rem;
  text-align: center;
  min-width: 200px;
  width: fit-content;
  min-height: 80px;
  margin: 10px 0;
  padding: 15px;
  border: 5px solid var(--orange);
}

#animation-container {
  display: flex;
  flex-direction: column-reverse;
  justify-content: end;
  gap: 1rem;
  margin-inline: auto;
  margin-block-end: 1rem;
  width: clamp(320px, 30vw, 420px);
  min-height: 40vh;
  border: 2px solid var(--orange);
  padding: 1rem;
}

.animation-frame {
  font-family: -apple-system, BlinkMacSystemFont, avenir next, avenir, segoe ui,
    helvetica neue, helvetica, Cantarell, Ubuntu, roboto, noto, arial,
    sans-serif;
  padding: 15px 10px;
  border: 5px solid var(--orange);
  font-size: 1.2rem;
  text-align: center;
}

@media screen and (min-width: 540px) {
  .input-container {
    flex-direction: row;
  }

  #result {
    max-width: 460px;
  }
}
```

```js
const numberInput = document.getElementById("number-input");
const convertBtn = document.getElementById("convert-btn");
const result = document.getElementById("result");

const decimalToBinary = (input) => {
--fcc-editable-region--
  return "01110110";
--fcc-editable-region--
};

const checkUserInput = () => {
  if (
    !numberInput.value ||
    isNaN(parseInt(numberInput.value)) ||
    parseInt(numberInput.value) < 0
  ) {
    alert("Please provide a decimal number greater than or equal to 0");
    return;
  }

  decimalToBinary(parseInt(numberInput.value));
  numberInput.value = "";
};

convertBtn.addEventListener("click", checkUserInput);

numberInput.addEventListener("keydown", (e) => {
  if (e.key === "Enter") {
    checkUserInput();
  }
});
```
