---
id: 65649122c7f77f519aaf0975
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

Finally, call the `generate_password` function with `8` as the argument and assign the function call to a `new_password` variable.

# --hints--

You should call `generate_password` passing `8` as the argument.

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

You should assign `generate_password(8)` to a `new_password` variable.

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)

    return password

--fcc-editable-region--

--fcc-editable-region--
```
