---
id: 61537485c4f2a624f18d7794
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

標準的な HTML ボイラープレートから始めましょう。 `DOCTYPE` 宣言、`html` 要素、`head` 要素、および `body` 要素を追加してください。

`<html>` の開始タグに `lang` 属性を追加して、値を `en` に設定してください。

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプに `html` を定義する必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素には開始タグが必要です。

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`<html>` の開始タグに `lang` 属性と、値として `en` が必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
