---
id: 61537d86bdc3dd343688fceb
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

次に、各 `img` にドキュメント内の順番に従って `src` 属性を追加します。 1 番目の `img` の `src` は `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` としてください。 残りも同様に、`1` の部分をドキュメント内の `img` の順番に置き換えたものを設定してください。

# --hints--

9 個ある `img` 要素のすべてに `src` 属性が必要です。

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

1 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

2 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

3 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

4 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

5 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

6 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

7 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

8 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

9 番目の `img` 要素には `src` 属性値として `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` を設定する必要があります。

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
