---
id: 61537c5f81f0cf325b4a854c
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`body` 要素内に `header` 要素を追加して、`header` というクラスを割り当ててください。

そして、`header` 内に `h1` 要素を作成して、`css flexbox photo gallery` というテキストを持たせてください。

# --hints--

`body` 要素の中に `header` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

`header` 要素の `class` の値として `header` を設定する必要があります。

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

`header` 要素内に `h1` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

`h1` 要素内に `css flexbox photo gallery` というテキストが必要です。

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
