---
id: 61537bb9b1a29430ac15ad38
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`head` 要素内に、テキストが `Photo Gallery` に設定された `title` 要素と、`styles.css` ファイルをページにリンクするための `link` 要素を追加してください。

# --hints--

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトには `Photo Gallery` というタイトルが必要です。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
