---
id: 6153908a366afb4d57185c8d
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

フレックスボックスには主軸と交差軸があります。 主軸は `flex-direction` プロパティで定義され、以下 4 種類の値を受け付けます。

- `row` (デフォルト): フレックスアイテムを左から右へ配置する、水平方向の軸
- `row-reverse`: フレックスアイテムを右から左へ配置する、水平方向の軸
- `column`: フレックスアイテムを上から下へ配置する、垂直方向の軸
- `column-reverse`: フレックスアイテムを下から上へ配置する、垂直方向の軸

**注**: 軸と方向は、テキストの方向によって異なります。 上記の値はテキストの方向が左から右の場合です。

異なる値を設定して、レイアウトにどう影響するか試してみてください。

確認し終わったら、`.gallery` 要素の `flex-direction` を明示的に `row` に設定しましょう。

# --hints--

`.gallery` セレクターの `flex-direction` プロパティの値を `row` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection === 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
