---
id: 61537a8054753e2f1f2a1574
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素内に、`name` が `viewport` 、`content` が `width=device-width, initial-scale=1` に設定された `meta` タグを追加してください。

更に、`charset` を `UTF-8` に設定した `meta` タグも追加してください。

# --hints--

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

`meta` 要素のひとつは `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

もうひとつの `meta` 要素は `charset` 属性が `UTF-8` に設定されている必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
