---
id: 6153897c27f6334716ee5abe
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

`.header` のテキストを中央揃えにしてください。 また、`text-transform` プロパティと `uppercase` の値を使用して、テキストを大文字にしてください。

各辺のパディングは `32px` にしてください。 背景色を `#0a0a23` に、文字色を `#fff` に設定してください。

そして `border-bottom` を追加して、値を `4px solid #fdb347` に設定してください。


# --hints--

`.header` セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

`.header` セレクターの `text-align` プロパティの値を `center` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign === 'center');
```

`.header` セレクターの `text-transform` プロパティの値を `uppercase` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform === 'uppercase');
```

`.header` セレクターの `padding` プロパティの値を `32px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.padding === '32px');
```

`.header` セレクターの `background-color` プロパティの値を `#0a0a23` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor === 'rgb(10, 10, 35)');
```

`.header` セレクターの `color` プロパティの値を `#fff` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.color === 'rgb(255, 255, 255)');
```

`.header` セレクターの `border-bottom` プロパティの値を `4px solid #fdb347` に設定する必要があります。

```js
console.log(new __helpers.CSSHelp(document).getStyle('.header'))
assert(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom === '4px solid rgb(253, 179, 71)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
