---
id: 615f171d05def3218035dc85
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

画像が大きすぎますね。 画像を選択する `.gallery img` セレクターを作成してください。 すべての画像の `width` を `100%` に、`max-width` を `350px` に設定することで、必要に応じて縮小されながらも大きくなりすぎないようにしてください。

また、`height` プロパティを `300px` に設定して画像のサイズを均一にしましょう。

# --hints--

`.gallery img` セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

`.gallery img` セレクターの `width` プロパティの値を `100%` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

`.gallery img` セレクターの `max-width` プロパティの値を `350px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

`.gallery img` セレクターの `height` プロパティの値を `300px` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
