---
id: 6494da0daf5df5197963671d
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

画像に追加した青い境界線が、ギャラリーに追加した赤い境界線を越えて広がっていることに注目してください。 これは、ブラウザーがコンテナー要素のサイズを計算する方法に起因します。

`box-sizing` プロパティを使って、この挙動を設定できます。 デフォルトでは `content-box` モデルが使われます。 このモデルでは、要素に特定の幅が指定されているとき、幅は要素のコンテンツ領域だけを使って計算されます。 その幅にパディングや境界線の値が加わり、要素はそれを含めた大きさになります。

全称セレクター `*` を使って、`box-sizing` を明示的に `content-box` に設定してみましょう。 デフォルト値を使用しているので、この時点ではまだ目に見える変化はありません。

# --hints--

`*` セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` セレクターには `box-sizing` プロパティが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

`*` セレクターの `box-sizing` プロパティを `content-box` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
