---
id: 64dd076a1b33c86d84b4232f
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

要素のサイズを確認しやすくするためには、境界線を追加すると便利です。

`.gallery` 要素の `width` を `50%` に、`border` を `5px solid red` に設定してください。

次に、`img` 要素の `width` を `100%` に、`padding` を `5px` に、`border` を `5px solid blue` に設定してください。

# --hints--

`.gallery` セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

`.gallery` セレクターには `width` プロパティが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

`.gallery` セレクターの `width` プロパティを `50%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

`.gallery` セレクターには `border` プロパティが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

`.gallery` セレクターの `border` プロパティを `5px solid red` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

`img` 要素のセレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

`img` セレクターには `width` プロパティが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

`img` セレクターの `width` プロパティを `100%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

`img` セレクターが `padding` プロパティを持たないようにしてください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

`img` セレクターの `padding` プロパティを `5px` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

`img` セレクターには `border` プロパティが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

`img` セレクターの `border` プロパティを `5px solid blue` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
