---
id: 587d781b367417b2b2512abd
title: 调整标题与段落的大小
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3bRPTz'
forumTopicId: 301037
---

# --description--

标题（`h1` 到 `h6`）的文字应该比段落的文字大，这样可以让用户更直观的看到页面布局，同时能根据不同元素的层级梳理出页面的结构，更方便用户获取关键的信息。你可以使用 `font-size` 属性来设置元素内文字的大小。

# --instructions--

把 `h4` 标签的 `font-size` 的属性值改成 27 像素，让标题更醒目。

# --hints--

应给 `h4` 元素添加一个 `font-size` 属性并将属性值设置为 `27px`。

```js
assert($('h4').css('font-size') == '27px');
```

# --solutions--

