---
id: 587d781c367417b2b2512ac2
title: 设置多个标题元素的 font-size
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPpQNT3'
forumTopicId: 301067
---

# --description--

`font-size` 属性用来指定元素内文字的大小。这个规则可以应用到多个元素上，合理的使用，能让页面的文字显示的错落有致。在本挑战里，你将要设置 `h1` 到 `h6` 的每个标题文字的大小。

# --instructions--

<ul><li>设置 <code>h1</code> 标签的 <code>font-size</code> 为 68px。</li><li>设置 <code>h2</code> 标签的 <code>font-size</code> 为 52px。</li><li>设置 <code>h3</code> 标签的 <code>font-size</code> 为 40px。</li><li>设置 <code>h4</code> 标签的 <code>font-size</code> 为 32px。</li><li>设置 <code>h5</code> 标签的 <code>font-size</code> 为 21px。</li><li>设置 <code>h6</code> 标签的 <code>font-size</code> 为 14px。</li></ul>

# --hints--

你应该设置 `h1` 标签的 `font-size` 为 `68px`。

```js
assert($('h1').css('font-size') == '68px');
```

你应该设置 `h2` 标签的 `font-size` 为 `52px`。

```js
assert($('h2').css('font-size') == '52px');
```

你应该设置 `h3` 标签的 `font-size` 为 `40px`。

```js
assert($('h3').css('font-size') == '40px');
```

你应该设置 `h4` 标签的 `font-size` 为 `32px`。

```js
assert($('h4').css('font-size') == '32px');
```

你应该设置 `h5` 标签的 `font-size` 为 `21px`。

```js
assert($('h5').css('font-size') == '21px');
```

你应该设置 `h6` 标签的 `font-size` 为 `14px`。

```js
assert($('h6').css('font-size') == '14px');
```

# --solutions--

