---
id: 587d781d367417b2b2512ac5
title: 设置段落的 line-height
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVWdcv'
forumTopicId: 301070
---

# --description--

CSS 提供 `line-height` 属性来设置行间的距离。行高，顾名思义，可以用来设置每行文字所占据的垂直空间。

# --instructions--

给 `p` 标签添加 `line-height` 属性并赋值 25px。

# --hints--

`p` 标签的 `line-height` 属性值应为 `25px`。

```js
assert($('p').css('line-height') == '25px');
```

# --solutions--

