---
id: 587d781c367417b2b2512ac3
title: 设置多个标题元素的 font-weight
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVWRHq'
forumTopicId: 301069
---

# --description--

在上一个挑战里你已经设置了每个标题的 `font-size`，接下来你将设置 `font-weight`。

`font-weight` 属性用于设置文本中所用的字体的粗细。

# --instructions--

<ul><li>设置 <code>h1</code> 标签的 <code>font-weight</code> 为 800。</li><li>设置 <code>h2</code> 标签的 <code>font-weight</code> 为 600。</li><li>设置 <code>h3</code> 标签的 <code>font-weight</code> 为 500。</li><li>设置 <code>h4</code> 标签的 <code>font-weight</code> 为 400。</li><li>设置 <code>h5</code> 标签的 <code>font-weight</code> 为 300。</li><li>设置 <code>h6</code> 标签的 <code>font-weight</code> 为 200。</li></ul>

# --hints--

你应该设置 `h1` 标签的 `font-weight` 为 800。

```js
assert($('h1').css('font-weight') == '800');
```

你应该设置 `h2` 标签的 `font-weight` 为 600。

```js
assert($('h2').css('font-weight') == '600');
```

你应该设置 `h3` 标签的 `font-weight` 为 500。

```js
assert($('h3').css('font-weight') == '500');
```

你应该设置 `h4` 标签的 `font-weight` 为 400。

```js
assert($('h4').css('font-weight') == '400');
```

你应该设置 `h5` 标签的 `font-weight` 为 300。

```js
assert($('h5').css('font-weight') == '300');
```

你应该设置 `h6` 标签的 `font-weight` 为 200。

```js
assert($('h6').css('font-weight') == '200');
```

# --solutions--

