---
id: 587d781c367417b2b2512ac4
title: 设置段落文本的 font-size
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJ36Cr'
forumTopicId: 301068
---

# --description--

CSS 里面的 `font-size` 属性不只限于标题，还可以应用于任何包含文字的元素内。

# --instructions--

把段落的 `font-size` 属性值设置为 16px，让它看起来更清晰。

# --hints--

`p` 标签的 `font-size` 属性值应为 `16px`。

```js
assert($('p').css('font-size') == '16px');
```

# --solutions--

