---
id: 68e65ec5348e40a739903e6d
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now, create a paragraph element below the `h1` element. Inside this paragraph add the text `Learning to code is hard, but as Quincy Larson says, You can become a developer.`.

# --hints--

You should have an `p` element.

```js
assert.exists(document.querySelector('p'));
```

Your paragraph should have the text `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. Double check your spelling.

```js
assert.equal(document.querySelector('p')?.textContent, "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

Your paragraph should be below the `h1` element.

```js
assert.match(code, /<\/h1>\s*<p>/g);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
