---
id: 68e90335bd53a901c41dfc13
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In the previous step you added a paragraph that contains a quote from Quincy. As you learned in the lesson on block and inline quotes, in HTML quoted elements are used to distinguish quoted text from the surrounding content.

The quote in this case is: `You can become a developer.`. It is part of the text of the larger paragraph. Thus, in order to distinguish the quote from the rest of the paragraph we need to use the inline quotation element `q`. Most modern browsers will add quotation marks around an inline quote automatically when you use this element.

This is an example: 

```html
<p>
  Quincy said, <q>Happy coding!</q>
</p>
```

Inside the paragraph wrap `You can become a developer.` in an inline quotation element, keeping the rest of the paragraph unchanged.

# --hints--

You should have a `q` element.

```js
assert.exists(document.querySelector('q'));
```

Your `q` element should be nested inside the paragraph, that is below the `h1` element.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Your `q` element should have the text `You can become a developer.`.

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Ignoring the `q` element you just added, the text of the `p` element should still be `Learning to code is hard, but as Quincy Larson says, You can become a developer.`.

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replace(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
