---
id: 68e90335bd53a901c41dfc13
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In a previous lesson you learned that the `q` element is used to distinguish quoted text from the surrounding content.

Here is an example:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

Most modern browsers will add quotation marks around an inline quote automatically when you use the `q` element.

Inside the paragraph wrap `You can become a developer.` in an inline quotation element, keeping the rest of the paragraph unchanged.

# --hints--

You should have a `q` element.

```js
assert.exists(document.querySelector('q'));
```

Your `q` element should be nested inside the paragraph that is below the `h1` element.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

Your `q` element should have the text `You can become a developer.`.

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

Inside your `q` element, the text must not have leading or trailing whitespace (no spaces, tabs, or line breaks before the first character or after the last).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

You should only add `q` tags around the appropriate part of the sentence without changing the original text. The `p` element should still read: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
