---
id: 68eab6a2b37d2f13602d0c3f
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

You can write quoted text directly within the block quotation element, like you did in the first `section` element. Alternatively, you can wrap it within one or more paragraph elements. This is helpful when the text has multiple paragraphs, but you want to keep them within the same block quotation element.

Here's an example with two paragraphs:

```html
<blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
  <p>So many people ask me each week: is college still worth it? In this 1-hour video I answer this question and other commonly asked questions about university.</p>
  <p>I've been in adult education for two decades at this point, and even though I'm not a labor market economist, I do feel confident enough to answer these questions.</p>
</blockquote>
```

In the second section, inside the existing block quotation element, add four `p` elements with the following texts, in order:

- `So much of getting a job is who you know.`
- `It's OK to be an introvert, but you do need to push your boundaries.`
- `Create GitHub, Twitter, LinkedIn, and Discord accounts.`
- `Go to tech meetups and conferences. Travel if you have to.`

# --hints--

You should have four `p` elements nested inside the `blockquote` element inside the second section.

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(2) > blockquote > p');
assert.lengthOf(pElements, 4);
```

Your first `p` element should have the text `So much of getting a job is who you know.`.

```js
const firstPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'So much of getting a job is who you know.');
```

Your second `p` element should have the text `It's OK to be an introvert, but you do need to push your boundaries.`.

```js
const secondPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "It's OK to be an introvert, but you do need to push your boundaries.");
```

Your third `p` element should have the text `Create GitHub, Twitter, LinkedIn, and Discord accounts.`.

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Create GitHub, Twitter, LinkedIn, and Discord accounts.');
```

Your forth `p` element should have the text `Go to tech meetups and conferences. Travel if you have to.`.

```js
const forthPEl = document.querySelector('main > section:nth-of-type(2) > blockquote > p:nth-of-type(4)');
assert.equal(forthPEl?.innerText.trim(), 'Go to tech meetups and conferences. Travel if you have to.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          —Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
        --fcc-editable-region--
          
        --fcc-editable-region--
        </blockquote>
      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
