---
id: 612e83ec2eca1e370f830511
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Agrega un elemento `link` dentro de tu elemento `head`. Para ese elemento `link`, establece el atributo `rel` en `stylesheet` y el atributo `href` en `./styles.css`.

# --hints--

Tu código debe tener un elemento `link`.

```js
assert.match(code, /<link/)
```

No debes cambiar tus etiquetas `head` existentes. Asegúrate de no haber eliminado la etiqueta de cierre.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Debes tener un elemento `link` de autocierre.

```js
assert(document.querySelectorAll('link').length === 1);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
