---
id: 612e8279827a28352ce83a72
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Ahora copia el conjunto de siete elementos `.key`, y pega otros dos conjuntos en el div `.keys`.

# --hints--

Debes tener 21 elementos `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Debes tener 15 elementos `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Todos los elementos `.key` deben estar dentro de tu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Todos los elementos `.black--key` deben estar dentro de tu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
