---
id: 612e89d254fe5d3df7d6693d
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Ahora que has restablecido el modelo de caja del `html`, tienes que pasarlo también a los elementos de su interior. Para ello, puedes establecer la propiedad `box-sizing` en `inherit`, lo que indicará a los elementos apuntados que utilicen el mismo valor que el elemento padre.

También tendrás que apuntar a los pseudo-elementos, que son palabras clave especiales que siguen a un selector. Los dos pseudo-elementos que usaras son `::before` y `::after`.

El selector `::before` crea un pseudo-elemento que es el primer hijo del elemento seleccionado, mientras que el selector `::after` crea un pseudo-elemento que es el último hijo del elemento seleccionado. Estos pseudo-elementos a menudo se usan para crear contenido cosmético, que verás más adelante en este proyecto.

Por ahora, crea un selector CSS para apuntar a todos los elementos con `*`, e incluye los pseudo-elementos con `::before` y `::after`. Establece la propiedad `box-sizing` en `inherit`.

# --hints--

Debes tener un selector `*, ::before, ::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Tu selector `*, ::before, ::after` debe tener la propiedad `box-sizing` establecida en `inherit`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
