---
id: 612e78af05201622d4bab8aa
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Es hora de empezar a trabajar en el piano. Crea un elemento `div` dentro de tu elemento `body` con el `id` establecido en `piano`.

# --hints--

Debes crear un nuevo elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Tu `div` debe estar dentro de tu elemento `body`.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

Tu `div` debe tener el `id` establecido en `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
