---
id: 612e804c54d5e7308d7ebe56
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro de tu elemento `.keys`, agrega siete elementos `div`. Dales a todos la clase `key`.

# --hints--

Debes crear siete nuevos elementos `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

Tus siete nuevos elementos `div` deben estar dentro de tu elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

Tus siete nuevos elementos `div` deben tener el `class` establecido en `key`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
