---
id: 612ea4c4993aba52ab4aa32e
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

Ahora es el momento de utilizar los pseudo-selectores que has preparado antes. Para crear las teclas negras, agrega un nuevo selector `.key.black--key::after`. Esto apuntará a los elementos con la clase `key black--key`, y seleccionará el pseudo-elemento después de estos elementos en el HTML.

En el nuevo selector, establece el `background-color` en `#1d1e22`. También establece la propiedad `content` en `""`. Esto hará que los pseudo-elementos estén vacíos.

La propiedad `content` se utiliza para establecer o anular el contenido de un elemento. Por defecto, los pseudo-elementos creados por los pseudo-selectores `::before` y `::after` están vacíos, y los elementos no serán renderizados en la página. Establecer la propiedad `content` a una cadena vacía `""` asegurará que el elemento se renderice en la página aunque siga estando vacío.

Si quieres experimentar, prueba eliminando la propiedad `background-color` y a establecer valores diferentes para la propiedad `content`, como `"♥"`. Recuerde deshacer estos cambios cuando haya terminado para que las pruebas pasen.

# --hints--

Debes tener un selector `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Tu selector `.key.black--key::after` debe tener una propiedad `background-color` establecida en `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

Tu selector `.key.black--key::after` debe tener una propiedad `content` establecida en `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
