---
id: 612e89562043183c86df287c
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Los navegadores pueden aplicar valores de margen y relleno por defecto a elementos específicos. Para asegurarte de que tu piano tiene un aspecto correcto, debe reajustar el modelo de caja.

Agrega una regla que tenga el selector `html` a tu archivo CSS y establece la propiedad `box-sizing` en `border-box`.

# --hints--

Debes tener un selector `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

Tu selector `html` debe tener una propiedad `box-sizing` establecida en `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
