---
id: 612e9d142affc44a453655db
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Es hora de estilizar las teclas en sí. Crea un selector `class` para los elementos `.key`. Establece el `background-color` al valor `#ffffff`, la propiedad `position` a `relative`, la propiedad `width` a `41px` y la propiedad `height` a `175px`.

# --hints--

Debes tener un selector `.key`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

Tu selector `.key` debe tener una propiedad `background-color` establecida en `#ffffff`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

Tu selector `.key` debe tener la propiedad `position` establecida en `relative`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

Tu selector `.key` debe tener una propiedad `width` establecida en `41px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

Tu selector `.key` debe tener una propiedad `height` establecida en `175px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
