---
id: 612e7d1c29fb872d6384379c
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Anida un segundo `div` dentro de tu `div` existente y establece que el `class` sea `keys`.

# --hints--

Debes crear un segundo elemento `div`.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

Tu nuevo elemento `div` debe estar dentro de tu elemento `div` existente.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

Tu nuevo elemento `div` debe tener el `class` establecido `keys`.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
