---
id: 6a57c331ae94fa4631b60862
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Each artifact also tracks where it has been, how it is described, and whether it is on display.

Add three more properties to the first artifact:

- `locations`: an array holding two objects, `{ gallery: "Hall A", year: 2020 }` and `{ gallery: "Hall C", year: 2024 }`
- `tags`: an array holding the strings `"gold"` and `"egypt"`
- `onDisplay`: the boolean `true`

# --hints--

The first artifact's `locations` should be an array of the two given location objects.

```js
assert.deepEqual(collection[0].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

The first artifact's `tags` should be `["gold", "egypt"]`.

```js
assert.deepEqual(collection[0].tags, ["gold", "egypt"]);
```

The first artifact's `onDisplay` should be `true`.

```js
assert.strictEqual(collection[0].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--

--fcc-editable-region--
  },
];
```
