---
id: 6a57c331ae94fa4631b60861
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Artifacts have a person responsible for them. The value of a property can be another object.

Add a `curator` property to the first artifact. Set its value to an object with an `id` of `201` and a `name` of `"Earl Sinclair"`.

# --hints--

The first artifact should have a `curator` property that is an object.

```js
assert.isObject(collection[0].curator);
```

The `curator` should have an `id` of `201`.

```js
assert.strictEqual(collection[0].curator.id, 201);
```

The `curator` should have a `name` of `"Earl Sinclair"`.

```js
assert.strictEqual(collection[0].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--

--fcc-editable-region--
  },
];
```
