---
id: 6a57c331ae94fa4631b60866
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

If no artifact is stored under that `id`, `collection[id]` is `undefined`, and reading `.title` from it would throw an error.

Update the return statement to use a conditional. Return `artifact.title` when an artifact is found, and the string `"Artifact not found"` when it is not.

# --hints--

`getArtifactTitle(101)` should still return `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` should return `"Artifact not found"` when no artifact matches the `id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
