---
id: 6a57c265abbc95dfebc05098
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will build an artifact provenance auditor. You will store museum artifacts as objects, then write functions to look them up, tag them, move them between galleries, and summarize their history.

Each artifact is stored as an object. Start by creating an array named `collection`. Add one object to it with three properties: `id` set to `101`, `title` set to `"Golden Mask"`, and `category` set to `"Ceremonial"`.

# --hints--

You should declare a variable named `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` should be an array.

```js
assert.isArray(collection);
```

`collection` should contain one object.

```js
assert.isObject(collection[0]);
assert.lengthOf(collection, 1);
```

The first artifact should have an `id` of `101`.

```js
assert.strictEqual(collection[0].id, 101);
```

The first artifact should have a `title` of `"Golden Mask"`.

```js
assert.strictEqual(collection[0].title, "Golden Mask");
```

The first artifact should have a `category` of `"Ceremonial"`.

```js
assert.strictEqual(collection[0].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
