---
id: 6a57c331ae94fa4631b6086f
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

Update the curator of the artifact with an `id` of `101` to `"Fran Sinclair"`, then log that artifact's curator `name`.

# --hints--

You should call `updateCurator` with `101` and `"Fran Sinclair"`.

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)/);
```

You should log the first artifact's curator `name` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*0\s*\]\s*\.\s*curator\s*\.\s*name\s*\)/);
```

The first artifact's curator `name` should be `"Fran Sinclair"`.

```js
assert.strictEqual(collection[0].curator.name, "Fran Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
console.log(collection[0].title);
console.log(collection[0].curator.name);

function getArtifactTitle(id) {
  const artifact = collection.find((item) => item.id === id);
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[0].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
moveArtifact(102, "Hall B", 2026);
console.log(collection[1].locations);

function toggleDisplayStatus(id) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log("DISPLAY STATUS");
console.log(collection[1].onDisplay);
toggleDisplayStatus(102);
console.log(collection[1].onDisplay);

function updateCurator(id, name) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.curator.name = name;
  }
}

console.log("UPDATED CURATOR");
--fcc-editable-region--

--fcc-editable-region--
```
