---
id: 6a57c331ae94fa4631b60872
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

Now return the summary as a template literal. Put the artifact's title on the first line, then one line each for the category, curator name, current gallery, and display status. Use `currentLocation.gallery` for the current gallery.

For example, `buildSummary(101)` should return:

```md
Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true
```

# --hints--

`buildSummary(101)` should return the multi-line summary string for the first artifact.

```js
assert.strictEqual(buildSummary(101), `Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true`);
```

`buildSummary(102)` should return the summary using the most recent location.

```js
assert.strictEqual(buildSummary(102), `Bronze Tablet
Category: Inscription
Curator: Robert Sinclair
Current Gallery: Hall B
On Display: true`);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
console.log(collection[0].title);
console.log(collection[0].curator.name);

function getArtifactTitle(id) {
  const artifact = collection.find((item) => item.id === id);
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[0].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
moveArtifact(102, "Hall B", 2026);
console.log(collection[1].locations);

function toggleDisplayStatus(id) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log("DISPLAY STATUS");
console.log(collection[1].onDisplay);
toggleDisplayStatus(102);
console.log(collection[1].onDisplay);

function updateCurator(id, name) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.curator.name = name;
  }
}

console.log("UPDATED CURATOR");
updateCurator(101, "Fran Sinclair");
console.log(collection[0].curator.name);

function buildSummary(id) {
  const artifact = collection.find((item) => item.id === id);
  if (!artifact) {
    return "Artifact not found";
  }
  const currentLocation = artifact.locations[artifact.locations.length - 1];
--fcc-editable-region--

--fcc-editable-region--
}
```
