---
id: 6a57c331ae94fa4631b60868
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

A tag should never be added twice. The `includes` method returns `true` when an array already contains a value.

Create a function named `addTag` with parameters `id` and `tag`. Look up the artifact with `collection[id]`. If it exists and its `tags` array does not already include `tag`, push `tag` onto `artifact.tags`.

# --hints--

You should have a function named `addTag` with parameters named `id` and `tag`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

Calling `addTag` with a new tag should add that tag to the artifact's `tags` array.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

Calling `addTag` with a tag the artifact already has should not add a duplicate.

```js
const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
