---
id: 6a57c331ae94fa4631b60867
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Try your function. Call `getArtifactTitle` with an `id` of `102` and log the result.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map((call) => call?.[0]);
```

# --hints--

You should call `getArtifactTitle` with `102`.

```js
assert.match(__helpers.removeJSComments(code), /getArtifactTitle\s*\(\s*102\s*\)/);
```

You should log the result of `getArtifactTitle(102)`.

```js
assert.include(getLogs(), getArtifactTitle(102));
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
--fcc-editable-region--

--fcc-editable-region--
```
