---
id: 6a57c265abbc95dfebc05098
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will build an artifact provenance auditor. You will store museum artifacts in an object, then write functions to look them up, tag them, move them between galleries, and summarize their history.

Each artifact is stored under its id. Create an object named `collection`. Add one entry whose key is `101` and whose value is an object with a `title` of `"Golden Mask"` and a `category` of `"Ceremonial"`.

# --hints--

You should declare a variable named `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` should be an object.

```js
assert.isObject(collection);
```

`collection` should have one entry.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

The artifact stored under `101` should have a `title` of `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

The artifact stored under `101` should have a `category` of `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
