---
id: 6a57c331ae94fa4631b6086e
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

As your museum grows, the person responsible for an artifact can change.

Create an `updateCurator` function that takes `id` and `name` as parameters. Look up the artifact using `collection[id]`. If it exists, update the `name` property of its `curator` object to the provided `name`.

# --hints--

You should have a function named `updateCurator` with parameters named `id` and `name`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.updateCurator);
const params = explorer.allFunctions.updateCurator.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "name");
```

Calling `updateCurator` should change the `name` of the artifact's `curator` object.

```js
updateCurator(102, "Test Curator");
assert.strictEqual(collection[102].curator.name, "Test Curator");
```

`updateCurator` should change the curator of whichever artifact the `id` refers to.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
updateCurator(457, "Fresh Curator");
assert.strictEqual(collection[457].curator.name, "Fresh Curator");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

--fcc-editable-region--

--fcc-editable-region--
```
