---
id: 6a57c331ae94fa4631b60864
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Auditing starts with reading what is already on record. Since each artifact is stored under its ID, you can access one by putting that ID in square brackets, like `collection[101]`.

Log the `title` of the artifact with the `id` `101`, then log the `name` of that artifact's `curator`.

# --hints--

You should log the first artifact's `title` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*101\s*\]\s*\.\s*title\s*\)/);
```

You should log the first artifact's curator `name` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*101\s*\]\s*\.\s*curator\s*\.\s*name\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
