---
id: 6a57c331ae94fa4631b60869
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Now use your function. Add the tag `"royal"` to the artifact with an `id` of `101`, then log that artifact's `tags`.

# --hints--

You should call `addTag` with `101` and `"royal"`.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

You should log the first artifact's `tags` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*0\s*\]\s*\.\s*tags\s*\)/);
```

The first artifact's `tags` should include `"royal"`.

```js
assert.include(collection[0].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
console.log(collection[0].title);
console.log(collection[0].curator.name);

function getArtifactTitle(id) {
  const artifact = collection.find((item) => item.id === id);
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
--fcc-editable-region--

--fcc-editable-region--
```
