---
id: 6a57c331ae94fa4631b60867
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

The front desk needs to know which artifact is filed under `102`. Call `getArtifactTitle(102)` inside `console.log()` to log the returned title.

# --hints--

You should call `getArtifactTitle(102)` inside `console.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*getArtifactTitle\s*\(\s*102\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

--fcc-editable-region--

--fcc-editable-region--
```
