---
id: 6a57c331ae94fa4631b60865
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Because `collection` is keyed by id, you can look up any artifact with a variable in the square brackets: `collection[id]`.

Create a function named `getArtifactTitle` with a parameter `id`. Use `collection[id]` to get the artifact, store it in a variable named `artifact`, and return `artifact.title`.

# --hints--

You should have a function named `getArtifactTitle` with a parameter named `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

`getArtifactTitle(101)` should return `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` should return `"Bronze Tablet"`.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
