---
id: 6a57c331ae94fa4631b60864
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

With the data in place, you can read values from it. Use dot notation to reach a property, and chain it to reach a property inside a nested object.

Log the first artifact's `title` to the console. Then, on the next line, log the first artifact's curator `name`.

# --hints--

You should log the first artifact's `title` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*0\s*\]\s*\.\s*title\s*\)/);
```

You should log the first artifact's curator `name` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*0\s*\]\s*\.\s*curator\s*\.\s*name\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
--fcc-editable-region--

--fcc-editable-region--
```
