---
id: 6a57c331ae94fa4631b6086a
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Artifacts move between galleries over time, and each move is a new entry in the `locations` array.

Create a function named `moveArtifact` with parameters `id`, `gallery`, and `year`. Look up the artifact with `collection[id]`. If the artifact exists, push a new object with the given `gallery` and `year` onto its `locations` array.

# --hints--

You should have a function named `moveArtifact` with parameters named `id`, `gallery`, and `year`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.moveArtifact);
const params = explorer.allFunctions.moveArtifact.parameters;
assert.equal(params.length, 3);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "gallery");
assert.equal(params[2].toString(), "year");
```

Calling `moveArtifact` should push a new `{ gallery, year }` object onto the artifact's `locations` array.

```js
const start = collection[101].locations.length;
moveArtifact(101, "Test Hall", 2030);
assert.lengthOf(collection[101].locations, start + 1);
assert.deepEqual(collection[101].locations[collection[101].locations.length - 1], { gallery: "Test Hall", year: 2030 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[101].tags);

--fcc-editable-region--

--fcc-editable-region--
```
