---
id: 6a57c331ae94fa4631b60870
title: Step 17
challengeType: 1
dashedName: step-17
---

# --description--

Finally, you will build a readable summary for any artifact. Start by handling the missing case.

Create a function named `buildSummary` with a parameter `id`. Look up the artifact with `collection[id]`. If no artifact is found, return the string `"Artifact not found"`.

# --hints--

You should have a function named `buildSummary` with a parameter named `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary);
const params = explorer.allFunctions.buildSummary.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

`buildSummary` should return `"Artifact not found"` when no artifact matches the `id`.

```js
assert.strictEqual(buildSummary(999), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log("DISPLAY STATUS");
console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

console.log("UPDATED CURATOR");
updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
