---
id: 6a57c331ae94fa4631b60872
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

Now return a template literal summarizing the artifact. Put the title on the first line, then one line each for the category, curator name, current gallery, and display status. Use `currentLocation.gallery` for the gallery.

For example, `buildSummary(101)` returns:

```md
Golden Mask
Category: Ceremonial
Curator: Fran Sinclair
Current Gallery: Hall C
On Display: true
```

# --hints--

`buildSummary(101)` should include the artifact's title, `Golden Mask`.

```js
assert.include(buildSummary(101), "Golden Mask");
```

`buildSummary(101)` should include a `Category: Ceremonial` line.

```js
assert.include(buildSummary(101), "Category: Ceremonial");
```

`buildSummary(101)` should include a `Curator: Fran Sinclair` line.

```js
assert.include(buildSummary(101), "Curator: Fran Sinclair");
```

`buildSummary(101)` should include a `Current Gallery: Hall C` line, using `currentLocation.gallery`.

```js
assert.include(buildSummary(101), "Current Gallery: Hall C");
```

`buildSummary(101)` should include an `On Display: true` line.

```js
assert.include(buildSummary(101), "On Display: true");
```

`buildSummary(101)` should put each field on its own line.

```js
assert.isAtLeast(buildSummary(101).split("\n").length, 5);
```

`buildSummary(102)` should use the most recent location, showing `Current Gallery: Hall B`.

```js
assert.include(buildSummary(102), "Current Gallery: Hall B");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log("DISPLAY STATUS");
console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

console.log("UPDATED CURATOR");
updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
  const currentLocation = artifact.locations[artifact.locations.length - 1];
--fcc-editable-region--

--fcc-editable-region--
}
```
