---
id: 6a57c331ae94fa4631b60863
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

A catalog needs more than one artifact. Add a second object to `collection` with the same shape as the first.

Give it an `id` of `102`, a `title` of `"Bronze Tablet"`, and a `category` of `"Inscription"`. Add a `curator` object with an `id` of `202` and a `name` of `"Robert Sinclair"`. Add a `locations` array holding one object, `{ gallery: "Archive Wing", year: 2019 }`. Add a `tags` array holding `"bronze"` and `"writing"`, and set `onDisplay` to `false`.

# --hints--

`collection` should contain two objects.

```js
assert.lengthOf(collection, 2);
```

The second artifact should have an `id` of `102`.

```js
assert.strictEqual(collection[1].id, 102);
```

The second artifact should have a `title` of `"Bronze Tablet"`.

```js
assert.strictEqual(collection[1].title, "Bronze Tablet");
```

The second artifact should have a `category` of `"Inscription"`.

```js
assert.strictEqual(collection[1].category, "Inscription");
```

The second artifact's `curator` should be `{ id: 202, name: "Robert Sinclair" }`.

```js
assert.deepEqual(collection[1].curator, { id: 202, name: "Robert Sinclair" });
```

The second artifact's `locations` should be `[{ gallery: "Archive Wing", year: 2019 }]`.

```js
assert.deepEqual(collection[1].locations, [
  { gallery: "Archive Wing", year: 2019 },
]);
```

The second artifact's `tags` should be `["bronze", "writing"]`.

```js
assert.deepEqual(collection[1].tags, ["bronze", "writing"]);
```

The second artifact's `onDisplay` should be `false`.

```js
assert.strictEqual(collection[1].onDisplay, false);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--

--fcc-editable-region--
];
```
