---
id: 6a57c331ae94fa4631b6086d
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

Check that the toggle works. Log the `onDisplay` value of the artifact with an `id` of `102`, then call `toggleDisplayStatus` on it, then log its `onDisplay` value again.

# --hints--

You should call `toggleDisplayStatus` with `102`.

```js
assert.match(__helpers.removeJSComments(code), /toggleDisplayStatus\s*\(\s*102\s*\)/);
```

You should log the second artifact's `onDisplay` value to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*1\s*\]\s*\.\s*onDisplay\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
console.log(collection[0].title);
console.log(collection[0].curator.name);

function getArtifactTitle(id) {
  const artifact = collection.find((item) => item.id === id);
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[0].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
moveArtifact(102, "Hall B", 2026);
console.log(collection[1].locations);

function toggleDisplayStatus(id) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log("DISPLAY STATUS");
--fcc-editable-region--

--fcc-editable-region--
```
