---
id: 6a57c331ae94fa4631b6086c
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Artifacts go on and off display.

Create a function named `toggleDisplayStatus` with a parameter `id`. Look up the artifact by `id` with `collection.find()`. If the artifact exists, set its `onDisplay` property to the opposite of its current value.

# --hints--

You should have a function named `toggleDisplayStatus` with a parameter named `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.toggleDisplayStatus);
const params = explorer.allFunctions.toggleDisplayStatus.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Calling `toggleDisplayStatus` should flip the artifact's `onDisplay` value.

```js
const before = collection[0].onDisplay;
toggleDisplayStatus(101);
assert.strictEqual(collection[0].onDisplay, !before);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
console.log(collection[0].title);
console.log(collection[0].curator.name);

function getArtifactTitle(id) {
  const artifact = collection.find((item) => item.id === id);
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[0].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection.find((item) => item.id === id);
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
moveArtifact(102, "Hall B", 2026);
console.log(collection[1].locations);

--fcc-editable-region--

--fcc-editable-region--
```
