---
id: 6a57c331ae94fa4631b60867
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Try your function. Call `getArtifactTitle` with an `id` of `102` and log the result.

# --hints--

You should call `getArtifactTitle` with `102` and log the result.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*getArtifactTitle\s*\(\s*102\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = [
  {
    id: 101,
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  {
    id: 102,
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
];

console.log("FIRST ARTIFACT");
console.log(collection[0].title);
console.log(collection[0].curator.name);

function getArtifactTitle(id) {
  const artifact = collection.find((item) => item.id === id);
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
--fcc-editable-region--

--fcc-editable-region--
```
