---
id: 6a57c331ae94fa4631b6086b
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Move the artifact with an `id` of `102` to `"Hall B"` in the year `2026`, then log that artifact's `locations`.

# --hints--

You should call `moveArtifact` with `102`, `"Hall B"`, and `2026`.

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)/);
```

You should log the second artifact's `locations` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*102\s*\]\s*\.\s*locations\s*\)/);
```

The second artifact's most recent location should be `{ gallery: "Hall B", year: 2026 }`.

```js
assert.deepEqual(collection[102].locations[collection[102].locations.length - 1], { gallery: "Hall B", year: 2026 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log("FIRST ARTIFACT");
console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log("LOOKUP BY ID");
console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

console.log("UPDATED TAGS");
addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

console.log("UPDATED LOCATIONS");
--fcc-editable-region--

--fcc-editable-region--
```
