---
id: 6a57c331ae94fa4631b60863
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

A lookup table is more useful with more than one entry. Add a second artifact to `collection` under the key `102`, following the same shape as the first.

Give it a `title` of `"Bronze Tablet"` and a `category` of `"Inscription"`. Add a `curator` object with an `id` of `202` and a `name` of `"Robert Sinclair"`. Add a `locations` array holding one object, `{ gallery: "Archive Wing", year: 2019 }`. Add a `tags` array holding `"bronze"` and `"writing"`, and set `onDisplay` to `false`.

# --hints--

`collection` should have two entries.

```js
assert.strictEqual(Object.keys(collection).length, 2);
```

The artifact stored under `102` should match the described object.

```js
assert.deepEqual(collection[102], {
  title: "Bronze Tablet",
  category: "Inscription",
  curator: {
    id: 202,
    name: "Robert Sinclair",
  },
  locations: [{ gallery: "Archive Wing", year: 2019 }],
  tags: ["bronze", "writing"],
  onDisplay: false,
});
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
--fcc-editable-region--

--fcc-editable-region--
};
```
