---
id: 64b171849f925b0773aa434c
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Створіть список чисел від 1 до 3 та призначте його до змінної `xs`.

# --hints--

`xs` має бути списком `[1,2,3]`.

```js
({ test: () => assert.deepEqual(Array.from(__userGlobals.get("xs")), [1,2,3])})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
xs = [1,2,3]
```
