---
id: 65ef181dad3d96b14a21a03e
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you will find the approximate square root of a given number using the bisection method.

The bisection method is a technique for finding the roots of a real-valued function. It works by narrowing down an interval where the square root lies until it converges to a value within a specified tolerance.

Begin by creating a function named `square_root_bisection`. Leave the parameters empty for now.

# --hints--

You should define the function `square_root_bisection` with no parameters.

```js
({
    test: () => {
        assert(runPython(`
        import inspect        
        inspect.isfunction(square_root_bisection)
        `))
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
