---
id: 66b1ec7e2bb25947ca579bf3
title: Build a Rainbow Painting
challengeType: 14
dashedName: build-a-rainbow-painting
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `div` element with the `id` of `container`.
1. Your `#container` should have nine `div` elements nested inside it.
1. Each nested `div` element should have an `id` with the value of `red`, `orange`, `yellow`, `green`, `blue`, `indigo`, `violet`, `white`, and `cover`, respectively.
1. Each `div` element should have `position` set to `absolute`.
1. Each `div` element should have `border-radius` set to `50%` except for `#cover`.
1. You should target `#container` and set its `width` and `height` properties to `350px` and center the element in the viewport.
1. `#red`, `#orange`, `#yellow`, `#green`, `#blue`, `#indigo`, `#violet`, and `#white` should have the color of their `id` and the `width` and `height` properties of each of them should be `15px` less than the previous one, starting from `250px`.
1. You should use absolute positioning to center `#red` respect to its parent, and each of the remaining `div` elements (except `#cover`) respect to the previous element.
1. You should target `#cover` and set its background color to `white` and its size and position to cover the bottom half of the colored `div` elements.

# --hints--

You should have a `div` element with the `id` of `container`.

```js
const container = document.getElementById('container');
assert(!!container && container?.tagName === 'DIV');
```

Your `#container` should have nine `div` elements nested inside it.

```js
const containerChildren = document.querySelectorAll('#container>div');
assert(!!containerChildren && containerChildren.length === 9);
```

You should have a `div` element with the `id` of `red` as the first child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(1)');
assert(!!containerChild && containerChild.id === "red");
```

You should have a `div` element with the `id` of `orange` as the second child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(2)');
assert(!!containerChild && containerChild.id === "orange");
```

You should have a `div` element with the `id` of `yellow` as the third child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(3)');
assert(!!containerChild && containerChild.id === "yellow");
```

You should have a `div` element with the `id` of `green` as the fourth child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(4)');
assert(!!containerChild && containerChild.id === "green");
```

You should have a `div` element with the `id` of `blue` as the fifth child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(5)');
assert(!!containerChild && containerChild.id === "blue");
```

You should have a `div` element with the `id` of `indigo` as the sixth child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(6)');
assert(!!containerChild && containerChild.id === "indigo");
```

You should have a `div` element with the `id` of `violet` as the seventh child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(7)');
assert(!!containerChild && containerChild.id === "violet");
```

You should have a `div` element with the `id` of `white` as the eighth child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(8)');
assert(!!containerChild && containerChild.id === "white");
```

You should have a `div` element with the `id` of `cover` as the ninth child of `#container`.

```js
const containerChild = document.querySelector('#container div:nth-child(9)');
assert(!!containerChild && containerChild.id === "cover");
```

Each `div` element should have the `position` property set to `absolute`.

```js
const div = new __helpers.CSSHelp(document).getStyle('div');
assert(!!div && div.position === 'absolute')
```

Each `div` element should have `border-radius` set to `50%` except for `#cover`.

```js
const divsNotCover = new __helpers.CSSHelp(document).getStyle('div:not([id="cover"])');
assert(!!divsNotCover && divsNotCover.borderRadius === '50%');
```

You should target `#container` and set its `width` property to `350px`.

```js
const container = new __helpers.CSSHelp(document).getStyle('#container');
assert(!!container && container.width === '350px');
```

You should target `#container` and set its `height` property to `350px`.

```js
const container = new __helpers.CSSHelp(document).getStyle('#container');
assert(!!container && container.height === '350px');
```

You should set `top`, `bottom`, `left`, and `right` to `0` for your `#container` element.

```js
const container = new __helpers.CSSHelp(document).getStyle('#container');
assert(!!container && container.top === '0px' && container.bottom === '0px' && container.right === '0px' && container.left === '0px');
```

Your `#red` element should have a background color set to `red`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#red');
assert(!!containerChild && containerChild.backgroundColor === "red");
```

Your `#red` element should have a width of `250px`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#red');
assert(!!containerChild && containerChild.width === "250px");
```

Your `#red` element should have a height of `250px`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#red');
assert(!!containerChild && containerChild.height === "250px");
```

You should use absolute positioning to center `#red` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#red');
assert(!!containerChild);
const val = '50px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#orange` element should have a background color set to `orange`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#orange');
assert(!!containerChild && containerChild.backgroundColor === "orange");
```

The size of your `#orange` element should be `15px` less than the size of `#red`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#orange');
assert(!!containerChild && containerChild.width === "235px" && containerChild.height === "235px");
```

You should use absolute positioning to center `#orange` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#orange');
assert(!!containerChild);
const val = '57.5px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#yellow` element should have a background color set to `yellow`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#yellow');
assert(!!containerChild && containerChild.backgroundColor === "yellow");
```

The size of your `#yellow` element should be `15px` less than the size of `#orange`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#yellow');
assert(!!containerChild && containerChild.width === "220px" && containerChild.height === "220px");
```

You should use absolute positioning to center `#yellow` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#yellow');
assert(!!containerChild);
const val = '65px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#green` element should have a background color set to `green`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#green');
assert(!!containerChild && containerChild.backgroundColor === "green");
```

The size of your `#green` element should be `15px` less than the size of `#yellow`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#green');
assert(!!containerChild && containerChild.width === "205px" && containerChild.height === "205px");
```

You should use absolute positioning to center `#green` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#green');
assert(!!containerChild);
const val = '72.5px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#blue` element should have a background color set to `blue`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#blue');
assert(!!containerChild && containerChild.backgroundColor === "blue");
```

The size of your `#blue` element should be `15px` less than the size of `#green`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#blue');
assert(!!containerChild && containerChild.width === "190px" && containerChild.height === "190px");
```

You should use absolute positioning to center `#blue` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#blue');
assert(!!containerChild);
const val = '80px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#indigo` element should have a background color set to `indigo`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#indigo');
assert(!!containerChild && containerChild.backgroundColor === "indigo");
```

The size of your `#indigo` element should be `15px` less than the size of `#blue`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#indigo');
assert(!!containerChild && containerChild.width === "175px" && containerChild.height === "175px");
```

You should use absolute positioning to center `#indigo` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#indigo');
assert(!!containerChild);
const val = '87.5px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#violet` element should have a background color set to `violet`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#violet');
assert(!!containerChild && containerChild.backgroundColor === "violet");
```

The size of your `#violet` element should be `15px` less than the size of `#indigo`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#violet');
assert(!!containerChild && containerChild.width === "160px" && containerChild.height === "160px");
```

You should use absolute positioning to center `#violet` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#violet');
assert(!!containerChild);
const val = '95px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

Your `#white` element should have a background color set to `white`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#white');
assert(!!containerChild && containerChild.backgroundColor === "white");
```

The size of your `#white` element should be `15px` less than the size of `#violet`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#white');
assert(!!containerChild && containerChild.width === "145px" && containerChild.height === "145px");
```

You should use absolute positioning to center `#white` respect to its parent.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#white');
assert(!!containerChild);
const val = '102.5px';
const topOrBottom = containerChild.top === val || containerChild.bottom === val;
const leftOrRight = containerChild.left === val || containerChild.right === val;
assert(topOrBottom && leftOrRight);
```

You should target `#cover` and set its background color to `white`.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#cover');
assert(!!containerChild && containerChild.backgroundColor === 'white');
```

You should target `#cover` and set its size and position to cover the bottom half of the colored `div` elements.

```js
const containerChild = new __helpers.CSSHelp(document).getStyle('#cover');
const leftOrRight = containerChild.left === '50px' || containerChild.right === '50px';
const topOrBottom = containerChild.top === '175px' || containerChild.bottom === '50px';
const size = containerChild.width === '250px' && containerChild.height === '125px';
assert(!!containerChild && size && leftOrRight && topOrBottom);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>Rainbow</title>
</head>

<body>
    
</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>Rainbow</title>
</head>

<body>
    <main>
        <div id="container">
            <div id="red"></div>
            <div id="orange"></div>
            <div id="yellow"></div>
            <div id="green"></div>
            <div id="blue"></div>
            <div id="indigo"></div>
            <div id="violet"></div>
            <div id="white"></div>
            <div id="cover"></div>
        </div>
    </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

div {
    position: absolute;
}

div:not([id="cover"]) {
    border-radius: 50%;
}

#container {
    width: 350px;
    height: 350px;
    right: 0;
    left: 0;
    top: 0;
    bottom: 0;
    margin: auto;
}

#red {
    background-color: red;
    width: 250px;
    height: 250px;
    top: 50px;
    left: 50px;
}

#orange {
    background-color: orange;
    width: 235px;
    height: 235px;
    left: 57.5px;
    top: 57.5px;
}

#yellow {
    background-color: yellow;
    width: 220px;
    height: 220px;
    left: 65px;
    top: 65px;
}

#green {
    background-color: green;
    width: 205px;
    height: 205px;
    left: 72.5px;
    top: 72.5px;
}

#blue {
    background-color: blue;
    width: 190px;
    height: 190px;
    left: 80px;
    top: 80px;
}

#indigo {
    background-color: indigo;
    width: 175px;
    height: 175px;
    left: 87.5px;
    top: 87.5px;
}

#violet {
    background-color: violet;
    width: 160px;
    height: 160px;
    left: 95px;
    top: 95px;
}

#white {
    background-color: white;
    width: 145px;
    height: 145px;
    left: 102.5px;
    top: 102.5px;
}

#cover {
    background-color: white;
    width: 250px;
    height: 125px;
    left: 50px;
    top: 175px;
}
```
