---
id: 66bf6bacf178eac7b96d4f5e
title: Build a Favourite Icon Toggler
challengeType: 14
dashedName: build-a-favourite-icon-toggler
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**



# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title></title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

```js

```

# --solutions--

```html

```

```css

```

```js

```
