---
id: 69a92f0b9cc04bb0d5327bb5
title: Implement a Stack
challengeType: 26
dashedName: implement-a-stack
---

# --description--

In this lab, you will implement a stack data structure using functions. A stack is a Last-In-First-Out (LIFO) data structure where elements are added and removed from the top.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have a `push` function that adds an element to the top of the stack.
1. You should have a `pop` function that removes and returns the top element of the stack.
1. You should have a `peek` function that returns the top element of the stack without removing it.
1. You should have a `isEmpty` function that returns `true` if the stack contains no elements, and `false` otherwise.
1. You should have a `clear` function that removes all elements from the stack.

# --hints--

You should have a `push` function.

```js
assert.isFunction(push);
```

`push` function should add an element to the top of the stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
```

You should have a `pop` function.

```js
assert.isFunction(pop);
```

`pop` function should remove and return the top element of the stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

`pop` function should return `undefined` if the stack is empty.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

You should have a `peek` function.

```js
assert.isFunction(peek);
```

`peek` function should return the top element of the stack without removing it.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

`peek` function should return `undefined` if the stack is empty.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

You should have an `isEmpty` function.

```js
assert.isFunction(isEmpty);
```

`isEmpty` function should return `true` for an empty stack.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

`isEmpty` function should return `false` for a non-empty stack.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

You should have a `clear` function.

```js
assert.isFunction(clear);
```

`clear` function should remove all elements from the stack.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.isTrue(isEmpty(stack));
assert.isUndefined(peek(stack));
```

# --seed--

## --seed-contents--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  
}

function pop(stack) {
  
}

function peek(stack) {
  
}

function isEmpty(stack) {
  
}

function clear(stack) {
  
}
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
