---
id: 61969aa6acef5b12200f672e
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Para criar um cenário no fundo, você adicionará duas montanhas.

Acima do elemento `.penguin`, adicione uma `div` com a `class` `left-mountain`.

# --hints--

Você deve adicionar uma nova `div` dentro do `body`. O esperado eram os elementos `--fcc-expected--` `div`, mas o que foi encontrado foi `--fcc-actual--`.

```js
assert.equal(document.querySelectorAll('body > div')?.length, 3);
```

Você deve dar à `div` uma `class` de `left-mountain`.

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

Você deve colocar o `.left-mountain` antes de `.penguin`.

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
