---
id: 6196d1ac33c68d27dcda5796
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

Gire o elemento `.back-mountain` `45deg` no sentido horário. Então, dê a propriedade `left` de `110px` e uma propriedade `top` de `225px`.

# --hints--

Você deve usar a propriedade `transform` para girar o elemento.

```js
const backMountain = new __helpers.CSSHelp(document).getStyle('.back-mountain');

if (backMountain?.transform) {
  assert.notEmpty(backMountain?.transform);
} else {
  assert.notEmpty(backMountain?.rotate);
}
```

Você deve dar aos elementos `.back-mountain` um `transform` de `--fcc-expected--`, mas foi encontrado `--fcc-actual--`.

```js
const backMountain = new __helpers.CSSHelp(document).getStyle('.back-mountain');

if (backMountain?.transform) {
  assert.equal(backMountain?.transform, 'rotate(45deg)');
} else {
  assert.equal(backMountain?.rotate, '45deg');
}
```

Você deve dar a `.back-mountain` uma propriedade `left`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.left);
```

Você deve dar à classe`.back-mountain` uma propriedade `left` de `--fcc-expected--`, mas foi encontrado `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.left, '110px');
```

Você deve dar a `.back-mountain` uma propriedade `top`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.top);
```

Você deve dar à classe`.back-mountain` uma propriedade `top` de `--fcc-expected--`, mas foi encontrado `--fcc-actual--`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.back-mountain')?.top, '225px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

--fcc-editable-region--
.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;

}
--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
