---
id: 61967e74a8e3690ab6292daa
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Mire o elemento `body` para definir `background` como um gradiente linear inclinado a 45 graus no sentido horário, começando em `rgb(118, 201, 255)` e terminando em `rgb(247, 255, 222)`.

# --hints--

Você deve usar o seletor de elemento `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Você deve usar a propriedade `background` no seletor `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Você deve definir `background` como `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
