---
id: 657efce98e958b75df86b305
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Use the `.append()` list method to add the converted character to the list you created earlier.

```py
snake_cased_char_list.append(converted_character)
```

The `.append()` method adds a given object to the end of the list you invoke it on.

# --hints--

You should use the `.append()` method to add the converted character to the `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\(converted_character\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
