---
id: 657ee28cefc4945568287673
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Now that you have an empty list in place, you can start iterating through the input string and start converting each character to snake case. Use a `for` loop to iterate through the `pascal_or_camel_cased_string`. Make sure to name the target variable `char` which is short for character.

# --hints--

You should write a new `for` loop with the target variable named `char`. Don't forget the colon at the end.

```js
({
    test: () => {
        const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s{1}char\s{1}in\s{1}pascal_or_camel_cased_string:/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
