---
id: 657ed53c19461d4b95c4757a
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

I hope you remember from the list of steps in the previous step, the first task of this function is to create a list to store all the characters of the output snake cased string.

To create a new list, start by writing out the name of the list, `snake_cased_char_list` in this case. Then write an equal sign, known as the assignment operator in Python, followed by an empty set of square braces.

Make sure you put a space before and after the assignment operator. You can also use the built-in `list()` function to create a list instead of the empty set of square braces.

# --hints--

You should name the list `snake_cased_char_list`.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get('convert_to_snake_case')
      src = inspect.getsource(a)
      'snake_cased_char_list' in src
    `));
  }
})
```

You should use a set of empty square braces or the built-in `list()` function to create the list.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get('convert_to_snake_case')
      src = inspect.getsource(a)
      'snake_cased_char_list = []' in src
      or
      'snake_cased_char_list = list()' in src
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
--fcc-editable-region--
    
--fcc-editable-region--
    pass
```
