---
id: 657efa642593c5746acc5c81
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Inside the if statement body, you need to convert any uppercase character to lowercase and prepend an underscore to this lowercase character.

Use the `.lower()` string method to convert uppercase characters to lowercase characters. You can then concatenate an underscore to the character using the plus sign.

```python
'_' + char.lower()
```

Assign the modified character to a variable named `converted_character` inside the if statement body.

# --hints--

Test 1

```js
({
    test: () => {
        const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +converted_character\s{1}=\s{1}char\.lower\(\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
--fcc-editable-region--


--fcc-editable-region--
            pass
```
