---
id: 657ed53c19461d4b95c4757a
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Now create a new list named `snake_cased_char_list` inside the function. You can use a set of empty square braces to create the new list. This list will hold the characters of the string after you have converted them to snake case.

# --hints--

You should declare a variable named `snake_cased_char_list`.

```js
assert.match(code, /snake_cased_char_list = \[\]/i)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
