---
id: 657f0353c9523d7d896873ea
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Inside the `main()` function, call the `convert_to_snake_case()` function. Pass `aLongAndComplexString` as input to the function and print out the output using the `print()` function.

# --hints--

You should call `convert_to_snake_case()` inside the `main()` function and pass `aLongAndComplexString` as input to the function.

```js
({
    test: () => {
        const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(convert_to_snake_case\('aLongAndComplexString'\)\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            snake_cased_char_list.append(f'_{char.lower()}')
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    pass
--fcc-editable-region--
```
