---
id: 5dc174fcf86c76b9248c6eb2
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML 要素は、`<h1>` のような「開始タグ」と、`</h1>` のような「終了タグ」を持ちます。

要素のテキストは開始タグと終了タグの間に入ります。

`h1` 要素を見つけて、要素のテキストを下記に変更してください:

`CatPhotoApp`

# --hints--

コード内には `CatPhotoApp` というテキストが必要です。 スペルが正しいか確認してください。

```js
assert(code.match(/catphotoapp/i));
```

`h1` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('h1'));
```

`h1` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/h1\>/));
```

`h1` 要素が複数あります。 余分な `h1` 要素を削除してください。

```js
assert(document.querySelectorAll('h1').length === 1);
```

`h1` 要素のテキストは `CatPhotoApp` でなければなりません。 テキストに誤字脱字があるか、テキストが `h1` 要素の開始タグと終了タグの間に入っていない可能性があります。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

ウェブページを変更するブラウザ拡張機能が使用されているようです。 拡張機能をオフにしてください。

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

