---
id: 5dc1798ff86c76b9248c6eb3
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` から `h6` の見出し要素は、その下にあるコンテンツの重要度を表すために使われます。 数字が小さいほど重要度が高くなるので、`h2` 要素は `h1` 要素より重要度が低いです。 `h1` 要素は各ページに 1 つだけ使用し、重要度の高い見出しの下に、重要度の低い見出しを置くようにします。

`h1` 要素の下に、次のテキストを持つ `h2` 要素を追加してください:

`Cat Photos`

# --hints--

`h1` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('h1'));
```

`h1` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/h1\>/));
```

`h1` 要素が 1 つだけ必要です。 余分なものは削除してください。

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

`h1` 要素のテキストは 'CatPhotoApp' でなければなりません。 テキストに誤字脱字があります。

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

`h2` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('h2'));
```

`h2` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/h2\>/));
```

`h2` 要素のテキストは 'Cat Photos' でなければなりません。 `Cat Photos` というテキストだけを、`h2` の開始タグと終了タグの間に置いてください。

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

`h2` 要素は `h1` 要素の下に置く必要があります。 `h1` 要素の方が重要度が高いので、`h2` 要素より上になければなりません。

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

