---
id: 5dc17dc8f86c76b9248c6eb5
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

コメントを使うと、ブラウザーの表示に影響を及ぼさずにメッセージを残すことができます。 また、コメントを使ってコードを非アクティブにすることも可能です。 HTML のコメントは `<!--` で始まり、テキストの行を囲み、`-->` で終了します。 例えば `<!-- TODO: Remove h1 -->` というコメントは、テキスト `TODO: Remove h1` を含んでいます。

`p` 要素の上に、次のテキストのコメントを追加してください:

`TODO: Add link to cat photos`

# --hints--

コメントは `<!--` で始める必要があります。 コメントの開始を定義する文字が 1 つ以上欠けています。

```js
assert(code.match(/<!--/));
```

コメントは `-->` で終わる必要があります。  コメントの終了を定義する文字が 1 つ以上欠けています。

```js
assert(code.match(/-->/));
```

コード内に余分なコメントを開始 / 終了する文字があってはいけません。 余分な `<!--` または `-->` がブラウザーに表示されています。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

コメントには `TODO: Add link to cat photos` というテキストを含めてください。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

コメントが `p` 要素の上にあるようにしてください。 順番が誤っているようです。

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

