---
id: 5dc17d3bf86c76b9248c6eb4
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`p` 要素は、ウェブサイト上に文章の段落 (paragraph) を作成するために使われます。 `p` 要素を `h2` 要素の下に作成し、それに次のテキストを設定してください:

`See more cat photos in our gallery.`

# --hints--

`p` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('p'));
```

`p` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/p\>/));
```

`p` 要素のテキストは `See more cat photos in our gallery.` でなければなりません。 テキストに誤字脱字があります。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

`p` 要素は `h2` 要素の下に置く必要があります。 順番が誤っているようです。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

