---
id: 5dc24165f86c76b9248c6ebc
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`img` 要素には常に `alt` 属性を設定するようにしましょう。 `alt` 属性のテキストは、アクセシビリティを向上させるためスクリーンリーダーに使用されたり、画像の読み込みに失敗した際に表示されたりします。 例えば `<img src="cat.jpg" alt="A cat">` と書いた場合、`alt` 属性のテキストは `A cat` です。

`img` 要素に、`alt` 属性を追加してこのテキストを設定してください:

`A cute orange cat lying on its back`

# --hints--

コードには `img` 要素が 1 つ必要です。 以前のステップで追加した `img` 要素が削除されているようです。

```js
assert(document.querySelector('img'));
```

`img` 要素に `alt` 属性がありません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

`img` 要素の `alt` 属性の値が 'A cute orange cat lying on its back' 以外に設定されています。 `alt` 属性の値が引用符 (") で囲まれていることを確認してください。

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

