---
id: 5dc24073f86c76b9248c6ebb
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML の<dfn>属性 (attributes)</dfn> とは、要素の開始タグ内で使用され、要素の動作を制御する特別な単語です。 `img` 要素の `src` 属性は、画像の URL (画像がある場所) を指定します。 `src` 属性を使用している `img` 要素の例は `<img src="https://www.example.com/the-image.jpg">` のようになります。

既存の `img` 要素に、`src` 属性を追加してこの URL を指定してください:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

コードには `img` 要素が 1 つ必要です。 `img` 要素を削除してしまったか、`src` 属性の値を引用符 (") で囲んでいない可能性があります。

```js
assert(document.querySelector('img'));
```

`img` 要素には `src` 属性が必要です。 属性が設定されていないか、誤字脱字があります。 要素名と属性名の間には必ず半角スペースを入れてください。

```js
assert(document.querySelector('img').src);
```

`img` 要素の `src` 属性は '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`' に設定される必要があります。 URL が設定されていないか、誤字脱字があります。 URL の大文字小文字も関係があります。

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

`img` 要素の `src` は正しい URL に設定されていますが、属性の値は常に引用符 (") で囲むことが推奨されています。

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

