---
id: 5dc23f9bf86c76b9248c6eba
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`img` 要素を使うと、ウェブサイトに画像を追加することができます。 `img` 要素は開始タグを持ちますが、終了タグを持ちません。 終了タグを持たない要素のタグは<dfn>自己終了タグ (self-cloding tag)</dfn> として知られています。

`p` 要素の下に、`img` 要素を追加してください。 この時点では、ブラウザーに画像は表示されません。

# --hints--

`img` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('img'));
```

`img` 要素は終了タグを持たないようにしてください。 終了タグは `<` の直後に `/` があるタグです。

```js
assert(!code.match(/<\/img\>/));
```

`img` 要素は 1 つだけにしてください。 余分なものは削除してください。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 要素は `p` 要素の下に置く必要があります。 順番が誤っているようです。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

