---
id: 5dc24614f86c76b9248c6ebd
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

アンカー (`a`) 要素を使うと、他のページにリンクすることができます。 例えば、`<a href='https://freecodecamp.org'></a>` と書くと `freecodecamp.org` にリンクします。

段落要素の後に、`https://freecatphotoapp.com` へとリンクするアンカー要素を追加してください。 この時点では、プレビューにリンクは表示されません。

# --hints--

アンカー (`a`) 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('a'));
```

アンカー (`a`) 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/a\>/));
```

アンカー (`a`) 要素は `p` 要素の下に置く必要があります。 順番が誤っているようです。

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

アンカー (`a`) 要素に `href` 属性がありません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
assert(document.querySelector('a').hasAttribute('href'));
```

アンカー (`a`) 要素は `https://freecatphotoapp.com` にリンクされる必要があります。 URL が設定されていないか、誤字脱字があります。

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

アンカー ('a') 要素の `href` は正しいリンクに設定されていますが、属性の値は常に引用符 (") で囲むことが推奨されています。

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

