---
id: 5dc23991f86c76b9248c6eb8
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

In the previous step, you put the `h1`, `h2`, comment, and `p` elements inside the `main` element. This is called *nesting*. Nested elements should be placed two spaces further to the right of the element they are nested in. This spacing is called indentation and it is used to make HTML easier to read.

The `h1` element, `h2` element and the comment are indented two spaces more than the `main` element in the code below. Use the space bar on your keyboard to add two more spaces in front of the `p` element so that it is indented properly as well.

# --hints--

コード内には、テキストが `Cat Photos` の `h2` 要素が必要です。 誤って削除されたか、開始タグや終了タグが欠けているか、テキストが変更された可能性があります。

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

例示されたコード内の `ul` 要素や `li` 要素は追加しないでください。

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

`h2` 要素の行のインデントは変更しないでください。 Its opening tag should start 6 spaces over from the start of the line. ステップをやり直すと元のインデントを復元することができます。

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

コード内にはコメントが必要です。 前のステップで追加されたコメントが削除されているようです。

```js
assert(code.match(/<!--.*-->/));
```

コメントのテキストは `TODO: Add link to cat photos` でなければなりません。 コメントのテキストやスペースを変更しないでください。

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

コメント要素の行のインデントは変更しないでください。 Its opening tag should start 6 spaces over from the start of the line. ステップをやり直すと元のインデントを復元することができます。

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

コードには `p` 要素が 1 つ必要です。 前のステップで追加された `p` 要素が削除されているようです。

```js
assert(document.querySelector('p'));
```

`p` 要素のテキストは `Click here to view more cat photos.` でなければなりません。 `p` 要素のテキストやスペース、句読点を変更しないでください。

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

The opening `p` tag should have indentation that matches your `h2` and comment elements. Its opening tag should start 6 spaces over from the start of the line.

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

