---
id: 5ddbd81294d8ddc1510a8e56
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

リンクのテキストは、アンカー (`a`) 要素の開始タグと終了タグの間に配置します。 例えば `<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` と書くと、`click here to go to freeCodeCamp.org` というテキストのリンクになります。

アンカー要素に、`cat photos` というアンカーテキストを追加してください。 これがリンクのテキストになります。

# --hints--

アンカー (`a`) 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('a'));
```

アンカー (`a`) 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/a\>/));
```

アンカー (`a`) 要素のテキストは `cat photos` でなければなりません。 リンクテキストがアンカー (`a`) 要素の開始タグと終了タグの間に置かれていることを確認してください。

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Click here to view more cat photos.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

