---
id: 5dfa3589eacea3f48c6300ae
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

2 つ目の `section` 要素内に、新しい `h2` 要素を追加して、テキストを `Cat Lists` としてください。

# --hints--

`section` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(
  document.querySelectorAll('section').length === 2 &&
    code.match(/<\/section>/g).length === 2
);
```

`h2` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelectorAll('h2').length === 2);
```

`h2` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/h2\>/g).length === 2);
```

2 つ目の `h2` 要素は、2 つ目の `section` 要素の終了タグのすぐ上にある必要があります。 位置が正しくないようです。

```js
const secondSection = document.querySelectorAll('section')[1];
assert(secondSection.lastElementChild.nodeName === 'H2');
```

2 つ目の `h2` 要素には `Cat Lists` というテキストが必要です。 テキストに誤字脱字があります。

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase() === 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

