---
id: 5dfa22d1b521be39a3de7be0
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

前のステップでは、`link to cat pictures` という言葉をアンカー (`a`) の開始タグと終了タグの間に配置することで、リンクに変えることができました。 ある要素、例えば `p` 要素の中の言葉に対しても、同じようにリンクを設定することができます。

`p` 要素内のテキストについて、`cat photos` という言葉の前後にアンカータグ (`a`) の開始タグと終了タグを追加して、この部分を `https://freecatphotoapp.com` へのリンクにしてください。


# --hints--

新しいアンカー (`a`) 要素を `p` 要素の中にネストする必要があります。

```js
assert($('p > a').length);
```

このリンクの `href` 属性の値は `https://freecatphotoapp.com` としてください。 `href` の値が設定されていないか、タイプミスがあります。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

リンクのテキストは `cat photos` でなければなりません。 テキストに誤字脱字があります。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

アンカー (`a`) 要素をネストした後は、ブラウザーに表示される `p` 要素の中身は `See more cat photos in our gallery.` だけになるはずです。`p` 要素およびネストされたアンカー要素の、テキスト、スペース、句読点を再確認してください。

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

