---
id: 5dfa22d1b521be39a3de7be0
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

プレビュー内の画像の隣に、`cat photos` という語が下線付きで表示されています。 これがあなたが作成したリンクです。試しにクリックしても構いません。

`p` 要素内のテキストについて、`cat photos` という部分を新しいアンカー要素の中にネストして、`https://freecatphotoapp.com` へのリンクにしてください。

# --hints--

新しいアンカー (`a`) 要素を `p` 要素の中にネストする必要があります。

```js
assert($('p > a').length);
```

リンクの href の値は `https://freecatphotoapp.com` である必要があります。 href の値が設定されていないか、誤字脱字があります。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

リンクのテキストは `cat photos` でなければなりません。 テキストに誤字脱字があります。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

アンカー (`a`) 要素をネストした後は、ブラウザーに表示される `p` 要素の中身は `Click here to view more cat photos.` だけになるはずです。`p` 要素およびネストされたアンカー要素の、テキスト、スペース、句読点を再確認してください。

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

