---
id: 5dfa37b9eacea3f48c6300b0
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

テキスト `Things cats love:` を持つ `h3` 要素の後に、順序なしリスト (unordered list、`ul`) 要素を追加してください。 この時点では何も表示されません。

# --hints--

`ul` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('ul'));
```

`ul` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/ul>/));
```

`ul` 要素は、2 つ目の `section` 要素の終了タグの上にある必要があります。

```js
const secondSectionLastElemNode = $('main > section')[1].lastElementChild;
assert(secondSectionLastElemNode.nodeName === 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

