---
id: 5dfa2407b521be39a3de7be1
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

リンクを新しいタブで開くには、アンカー (`a`) 要素に `target` 属性を設定します。

`target` 属性は、リンクされたドキュメントをどこで開くかを指定します。 `target="_blank"` を指定すると、リンクされたドキュメントは新しいタブまたはウィンドウで開かれます。

`target` 属性を持つ `a` 要素の基本的な構文は次の通りです:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

リンクが新しいタブで開くようにするために、アンカー (`a`) 要素の開始タグに `target` 属性を追加し、値 `_blank` を設定してください。

# --hints--

`p` 要素内にはアンカー (`a`) 要素とテキスト `cat photos` がネストされているはずです。 削除されたか、誤字脱字がある可能性があります。

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

アンカー (`a`) 要素に `target` 属性がありません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
assert(document.querySelector('a').hasAttribute('target'));
```

`target` 属性の値は `_blank` である必要があります。 値が設定されていないか、誤字脱字があります。 属性値は引用符で囲むことを忘れないでください。

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

