---
id: 5dfb6a35eacea3f48c6300b4
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

図キャプション (`figcaption`) 要素は、`figure` 要素内の画像を説明するキャプションを追加するために利用されます。 例えば `<figcaption>A cute cat</figcaption>` と書くことで、`A cute cat` というキャプションを追加できます。

`figure` 要素内にネストされた画像の後に、`figcaption` 要素を追加して下記のテキストを設定してください:

`Cats love lasagna.`

# --hints--

`figcaption` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('figcaption'));
```

`figcaption` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/figcaption\>/));
```

`figcaption` 要素は `figure` 要素内にネストされている必要があります。

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

ラザニアの `img` 要素は `figure` 要素内にネストされている必要があります。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

`figure` 要素内にネストされた `figcaption` 要素は、`img` 要素の下に置く必要があります。 順番が誤っているようです。

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

`figcaption` 要素のテキストは `Cats love lasagna.` でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

