---
id: 5dfb655eeacea3f48c6300b3
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

`figure` 要素はそれ自体で完結する 1 つのコンテンツを表します。これを使って画像をキャプションと関連付けることができます。

先ほど追加した画像を `figure` 要素内にネストしてください。

# --hints--

`figure` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('figure'));
```

`figure` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/figure\>/));
```

2 つ目の `section` 要素の終了タグのすぐ上に `figure` 要素が必要です。

```js
assert($('section')[1].lastElementChild.nodeName === 'FIGURE');
```

ラザニアの `img` 要素は `figure` 要素内にネストされている必要があります。

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

