---
id: 5dfb6250eacea3f48c6300b2
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

順序なしリストの後に、`src` 属性値を下記の値に設定した新しい画像を追加してください:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

`alt` 属性値は下記の値に設定してください:

`A slice of lasagna on a plate.`

# --hints--

終了タグ `</ul>` の直後に、`img` 要素が必要です。

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

新しい画像に `alt` 属性がありません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

新しい画像の `alt` の値は `A slice of lasagna on a plate.` に設定してください。`alt` 属性の値が引用符で囲まれていることを確認してください。

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

新しい画像に `src` 属性がありません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

新しい画像の `src` の値は `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg` に設定してください。 `src` 属性の値が引用符で囲まれていることを確認してください。

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

新しい画像の `src` は正しい URL に設定されていますが、属性の値は常に引用符で囲むことが推奨されています。

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

