---
id: 5ef9b03c81a63668521804d6
title: ステップ 35
challengeType: 0
dashedName: step-35
---

# --description--

では、ユーザーから情報を集めるためにウェブフォームを追加しましょう。

`Cat Form` の見出しの後に、`form` 要素を追加してください。

# --hints--

`form` 要素には開始タグおよび終了タグが必要です。 必要なタグの片方、または両方が欠けているか、順番が誤っている可能性があります。

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

`form` 要素のタグの順番が正しくありません。

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

最後の `section` 要素内にネストされた `form` 要素は、`h2` 要素の下に置く必要があります。 `h2` 要素と `form` の順番が誤っているようです。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

`form` 要素には中身がないようにしてください。 `form` 要素のタグの間にある HTML 要素やテキストは削除してください。

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

