---
id: 5ef9b03c81a63668521804da
title: ステップ 41
challengeType: 0
dashedName: step-41
---

# --description--

`button` 要素を使ってクリック可能なボタンを作りましょう。 例えば `<button>Click Here</button>` と書くと、`Click Here` というテキストが表示されたボタンを作成できます。

`input` 要素の下に、`button` 要素を追加して `Submit` というテキストを設定してください。 属性が何も設定されていないボタンは、クリックされるとデフォルトの動作として、フォームの `action` 属性で指定された場所にフォームを送信します。

# --hints--

`button` 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('button'));
```

`button` 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/button\>/));
```

`button` 要素のテキストは `Submit` でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

`button` 要素は `input` 要素の下に置く必要があります。 順番が誤っているようです。

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

