---
id: 5ef9b03c81a63668521804d0
title: ステップ 25
challengeType: 0
dashedName: step-25
---

# --description--

`figcaption` 要素内の単語 `love` を、強調 (emphasis) 要素 `em` で囲んで強調しましょう。

# --hints--

`em` (強調) 要素には開始タグが必要です。 開始タグは `<elementName>` のような形式の構文です。

```js
assert(document.querySelector('em'));
```

`em` (強調) 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
assert(code.match(/<\/em\>/));
```

`figcaption` 要素が削除されたか、開始タグまたは終了タグがありません。

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

`em` (強調) 要素がテキスト `love` を囲む必要があります。 テキストが設定されていないか、誤字脱字があります。

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

`figcaption` のテキストは `Cats love lasagna` でなければなりません。 誤字がないか、`em` 要素の開始タグと終了タグの前後に適切なスペースがあるか確認してください。

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

