---
id: 5ef9b03c81a63668521804d7
title: ステップ 36
challengeType: 0
dashedName: step-36
---

# --description--

`action` 属性は、フォームのデータがどこへ送られるべきかを示します。 例えば `<form action="/submit-url"></form>` は、フォームのデータを `/submit-url` というパスに送信するようブラウザーに指示します。

`form` 要素に、`action` 属性と値 `https://freecatphotoapp.com/submit-cat-photo` を追加してください。

# --hints--

`form` 要素には開始タグおよび終了タグが正しい順序で必要です。 必要なタグの片方、または両方が欠けているか、順番が誤っている可能性があります。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

最後の `section` 要素内にネストされた `form` 要素は、`h2` 要素の下に置く必要があります。 `h2` 要素と `form` の順番が誤っているようです。

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

`form` 要素には中身がないようにしてください。 `form` 要素のタグの間にある HTML 要素やテキストは削除してください。

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

`form` 要素に `action` 属性がありません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

`form` 要素には値が `https://freecatphotoapp.com/submit-cat-photo` に設定されている `action` 属性が必要です。

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

`action` 属性は正しい URL に設定されていますが、属性の値は常に引用符で囲むことが推奨されています。

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

