---
id: 5ef9b03c81a63668521804e2
title: ステップ 54
challengeType: 0
dashedName: step-54
---

# --description--

フォームでは一般的に、複数の回答を得る可能性のある質問にはチェックボックスを使用します。 例えば `tacos` という選択肢のチェックボックスは次のように書けます: `<input type="checkbox"> tacos`

先ほど追加した `legend` 要素の下に `input` 要素を追加してください。`type` 属性は `checkbox` に設定し、選択肢は下記に設定してください:

`Loving`

# --hints--

チェックボックス用の `input` 要素には開始タグが必要ですが、終了タグは必要ありません。

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

チェックボックス用の input 要素は 1 つだけ追加してください。 余分なものは削除してください。

```js
assert($('fieldset > input').length === 1);
```

新しい `input` 要素は、`What's your cat's personality?` というテキストを持つ `legend` 要素の下にあるべきです。 順番に誤りがあるようです。

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

新しい `input` 要素に `type` 属性がありません。 開始タグのタグ名の後にスペースがあることを確認してください。

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

新しい `input` 要素には `type` 属性があり、値が `checkbox` に設定されている必要があります。 値が設定されていないか、誤字脱字があります。 属性値は引用符で囲むことを忘れないでください。

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

新しい `input` 要素の `type` 属性は `checkbox` に設定されていますが、属性の値は常に引用符で囲むことが推奨されています。

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

`Loving` というテキストはチェックボックスの右側に配置してください。 要素とテキストの間には必ずスペースを入れてください。 テキストが設定されていないか、誤字脱字があります。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

