---
id: 5ef9b03c81a63668521804e9
title: ステップ 64
challengeType: 0
dashedName: step-64
---

# --description--

`freeCodeCamp.org` というテキストをアンカー (`a`) 要素で囲んでリンクにしましょう。 `href` 属性は `https://www.freecodecamp.org` に設定してください。

# --hints--

アンカー (`a`) 要素は `footer` 要素内にネストされている必要があります。 アンカー (`a`) 要素の開始タグと終了タグが追加されていることを確認してください。

```js
assert($('footer > p > a').length);
```

アンカー (`a`) 要素には終了タグが必要です。 終了タグは `<` の直後に `/` があります。

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert(aElemClosingTags && aElemClosingTags.length === 3);
```

アンカー (`a`) 要素は値が `https://www.freecodecamp.org` に設定されている `href` 属性を持つ必要があります。 属性または値が設定されていないか、誤字脱字があります。

```js
const nestedAnchor = $('footer > p > a')[0];
assert(nestedAnchor.getAttribute('href').toLowerCase() === 'https://www.freecodecamp.org');
```

リンクのテキストは `freeCodeCamp.org` でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
const nestedAnchor = $('footer > p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'freecodecamp.org'
);
```

アンカー (`a`) 要素をネストした後は、ブラウザーに表示される `p` 要素の中身は `No Copyright - freeCodeCamp.org` だけになるはずです。 `p` 要素およびネストされたアンカー要素の、テキスト、スペース、句読点を再確認してください。

```js
const pText = $('footer > p')[0].innerText.toLowerCase().replace(/\s+/g, ' ');
assert(pText.match(/^no copyright - freecodecamp.org$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```

