---
id: 5ef9b03c81a63668521804e5
title: ステップ 60
challengeType: 0
dashedName: step-60
---

# --description--

チェックボックス、あるいはラジオボタンをデフォルトで選択状態にするには、`checked` 属性を追加する必要があります。 `checked` 属性には、値の設定は必要ありません。 その代わり、`input` 要素に `checked` という単語だけを追加します。他の属性との間には必ずスペースを入れるようにしてください。

1 つ目のラジオボタンと 1 つ目のチェックボックスがデフォルトで選択されるように設定してください。

# --hints--

2 つのラジオボタンと 3 つのチェックボックスは、対応する `fieldset` 要素にネストされたままにしてください。

```js
assert(
  $('input[type="radio"]').length === 2 &&
    $('fieldset > input[type="checkbox"]').length === 3
);
```

1 つ目のラジオボタンに `checked` 属性が不足しています。

```js
assert($('input[type="radio"]')[0].hasAttribute('checked'));
```

2 つ目のラジオボタンには `checked` 属性を設定してはいけません。

```js
assert(!$('input[type="radio"]')[1].hasAttribute('checked'));
```

1 つ目のチェックボックスに `checked` 属性が不足しています。

```js
assert($('input[type="checkbox"]')[0].hasAttribute('checked'));
```

2 つ目のチェックボックスには `checked` 属性を設定してはいけません。

```js
assert(!$('input[type="checkbox"]')[1].hasAttribute('checked'));
```

3 つ目のチェックボックスには `checked` 属性を設定してはいけません。

```js
assert(!$('input[type="checkbox"]')[2].hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

