---
id: 5ef9b03c81a63668521804d8
title: ステップ 37
challengeType: 0
dashedName: step-37
---

# --description--

`input` 要素は、ウェブフォームからデータを収集する方法をいくつか用意しています。 `input` 要素は `img` 要素と同じような<dfn>自己終了</dfn>タグであり、終了タグは必要ありません。

`form` 要素内に `input` 要素をネストしてください。

# --hints--

`form` 要素には開始タグおよび終了タグが正しい順序で必要です。 必要なタグの片方、または両方が欠けているか、順番が誤っている可能性があります。

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

`form` 要素の開始タグには `action` 属性のみを設定してください。 その他に入力したものがあれば削除してください。

```js
assert([...document.querySelector('form').attributes].length < 2);
```

`input` 要素を 1 つ作成する必要があります。 構文を確認してください。

```js
assert(document.querySelector('input'));
```

`input` 要素には開始タグが必要ですが、終了タグは必要ありません。

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

`input` 要素は `form` 要素内にネストされている必要があります。

```js
assert(document.querySelector('form > input'));
```

`form` には `input` 要素だけが含まれるようにしてください。 `form` 要素のタグの間にある他の HTML 要素やテキストは削除してください。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">

        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

