---
id: 5efb23e70dc218d6c85f89b1
title: ステップ 37
challengeType: 0
dashedName: step-37
---

# --description--

`type` 属性を使用して、さまざまな入力欄が作成できます。 パスワード入力欄、リセットボタン、あるいはユーザーが自分のコンピューターからファイルを選択できるフォームコントロールなどを簡単に作成できます。

ユーザーによるテキスト入力を受け付けるテキストフィールドを作成するために、`input` 要素に `type` 属性を追加し、値 `text` を設定してください。

# --hints--

`input` 要素が削除されているか、構文が正しくありません。 属性を追加した場合には、属性値が引用符で囲まれていることを確認してください。

```js
assert($('input').length);
```

`form` 要素には `input` 要素だけが含まれるようにしてください。 `form` 要素のタグの間にある余分な HTML 要素やテキストは削除してください。

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

`input` 要素の `type` 属性が `text` に設定されていません。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
assert($('input')[0].hasAttribute('type'));
```

`input` 要素には `type` 属性があり、値が `text` に設定されている必要があります。 値が設定されていないか、誤字脱字があります。 属性値は引用符で囲むことを忘れないでください。

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

`input` 要素の `type` 属性は `text` に設定されていますが、属性の値は常に引用符で囲むことが推奨されています。

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

