---
id: 5efc4f528d6a74d05e68af74
title: ステップ 56
challengeType: 0
dashedName: step-56
---

# --description--

`input` 要素のテキストと要素自体を関連付ける方法はもう一つあります。 テキストを `label` 要素にネストして、`input` 要素の `id` 属性と同じ値を持つ `for` 属性を追加することです。

Associate the text `Loving` with the checkbox by nesting only the text `Loving` in a `label` element and giving it an appropriate `for` attribute.

# --hints--

チェックボックスが削除されていないことを確認してください。

```js
assert($('input[type="checkbox"]')[0]);
```

変更前と変わらず、チェックボックスが `id` 属性を持ち、値が `loving` に設定されている必要があります。  期待される値は `--fcc-expected--` ですが、見つかった値は `--fcc-actual--` でした。

```js
assert.equal($('input[type="checkbox"]')[0].id, 'loving');
```

チェックボックスの右側に、`Loving` というテキストがそのまま置かれていてはいけません。 `label` 要素に囲まれている必要があります。 2 つの要素の間には必ずスペースを入れてください。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  !checkboxInputElem.nextSibling.nodeValue
    .replace(/\s+/g, ' ')
    .match(/ Loving/i)
);
```

テキスト `Loving` をネストするための、新しい `label` 要素を追加する必要があります。 開始タグと終了タグが両方あることを確認してください。

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

新しい `label` 要素はチェックボックスの右側に配置してください。 2 つの要素の間には必ずスペースを入れてください。

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

新しい `label` 要素に `for` 属性がありません。 開始タグのタグ名の後にスペースがあることを確認してください。

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

新しい `label` 要素には `for` 属性があり、値が `loving` に設定されている必要があります。  期待される値は `--fcc-expected--` ですが、見つかった値は `--fcc-actual--` でした。

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert.equal(labelElem.getAttribute('for'), 'loving');
```

テキスト `Loving` は新しい `label` 要素内にネストされている必要があります。 テキストが設定されていないか、誤字脱字があります。

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

