---
id: 5efc575c8d6a74d05e68af77
title: ステップ 58
challengeType: 0
dashedName: step-58
---

# --description--

一つ前のチェックボックスの後に最後のチェックボックスを追加し、`id` 属性値を `energetic` に設定してください。 `name` 属性は一つ前のチェックボックスと同じにしてください。

更に、新しいチェックボックスの右側に、`label` 要素を追加してテキスト `Energetic` を設定してください。 `label` 要素は新しいチェックボックスに関連付けてください。

# --hints--

新しいチェックボックスを追加する必要があります。

```js
assert($('input[type="checkbox"]').length === 3);
```

新しいチェックボックスは `id` 属性の値が `energetic`、`name` 属性の値が `personality` に設定されている必要があります。 開始タグのタグ名の後にスペースがあることと、すべての属性名の前にスペースがあることを確認してください。

```js
const checkboxes = [...$('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'energetic' &&
      checkbox.getAttribute('name') === 'personality'
  )
);
```

新しいチェックボックスは 2 つ目のチェックボックスの後に置いてください。 順番に誤りがあるようです。

```js
const checkboxes = [...$('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('lazy') < checkboxes.indexOf('energetic'));
```

新しいチェックボックスの右側に、テキスト `Energetic` を持つ `label` 要素が必要です。

```js
const nextElementSibling = $('input[type="checkbox"]')[2].nextElementSibling;
assert(
  nextElementSibling.nodeName === 'LABEL' &&
    nextElementSibling.innerText.replace(/\s+/g, '').match(/^Energetic$/i)
);
```

新しい `label` の `for` 属性は、新しいチェックボックスの `id` 属性と同じ値が設定されている必要があります。 値が設定されていないか、誤字脱字があります。

```js
assert(
  $('input[type="checkbox"]')[2].nextElementSibling.getAttribute('for') ===
    'energetic'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality"> <label for="lazy">Lazy</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

