---
id: 5f05a1d8e233dff4a68508d8
title: ステップ 47
challengeType: 0
dashedName: step-47
---

# --description--

最初のラジオボタンの下に別のラジオボタンを作成しましょう。 ボタンを `label` 要素の中にネストし、`label` テキストとして `Outdoor` を設定します。 ラジオボタンに `id` 属性を付け、値を `outdoor` としてください。

# --hints--

元の `Indoor` ラジオボタンの `input` のみ `id` を `indoor` に設定してください。 必要に応じて元の HTML に戻すためにステップをやり直すことができます。

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

`Indoor` ラジオボタンに変更を加えないでください。 必要に応じて元の HTML に戻すためにステップをやり直すことができます。

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

新しい `input` 要素を 1 つだけ追加して新しい `label` 要素内にネストする必要があります。 新しい `label` の開始タグと終了タグの両方があるか確認してください。

```js
assert(document.querySelectorAll('label input').length === 2);
```

開始 `label` タグに属性を追加しないでください。

```js
assert(code.match(/<label\s*>/g).length === 2);
```

新しい `input` には `id` 属性があり、値が `outdoor` に設定されている必要があります。

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

新しい `input` には `type` 属性があり、値が `radio` に設定されている必要があります。

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

`label` 内でネストされた `input` 以外に新しい要素を追加してはいけません。

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

新しい `input` の左側にテキストを入れないでください。

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

新しいラジオボタンのラベルテキストは正確に `Outdoor` にしてください。 大文字小文字も合わせてください。

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

新しいラジオボタンとラベルは `Indoor` ラジオボタンとラベルのすぐ下／後に置いてください。 これらの間には他のタグを入れないでください。

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

