---
id: 62dabe2ef403a12d5d295273
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

`p` 要素内のテキスト `cat photos` をリンクに変換したので、`p` 要素の下の 2 つ目のリンクは必要ありません。 `p` 要素の下にあるアンカー要素全体を削除してください。

# --hints--

コードには段落 (`p`) 要素が 1 つだけ必要です。

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

コードにはアンカー (`a`) 要素が 1 つだけ必要です。

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

前のステップで作成した要素には変更を加えないでください。

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

`p` 要素の下にテキスト `cat photos` があってはいけません。

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
