---
id: 62dabe2ef403a12d5d295273
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

`p` 要素内のテキスト `cat photos` をリンクに変換したので、`p` 要素の下の 2 つ目のリンクは必要ありません。 `p` 要素の下にあるアンカー要素全体を削除してください。

# --hints--

コードにはアンカー (`a`) 要素が 1 つだけ必要です。

```js
assert(document.querySelectorAll('a').length === 1);
```

変更前と変わらず、`p` 要素内にアンカー(`a`) 要素が含まれている必要があります。

```js
assert($('p > a').length);
```

リンクの `href` の値は `https://freecatphotoapp.com` でなければなりません。 `href` の値が設定されていないか、誤字脱字があります。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

リンクのテキストは `cat photos` でなければなりません。 テキストが設定されていないか、誤字脱字があります。

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

ブラウザーに表示される `p` 要素の中身は `See more cat photos in our gallery.` だけのはずです。`p` 要素およびネストされたアンカー要素の、テキスト、スペース、句読点を再確認してください。

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

`p` 要素の下にテキスト `cat photos` があってはいけません。

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
