---
id: 61fd7a160ed17960e971f28b
title: Hatua ya 16
challengeType: 0
dashedName: step-16
---

# --description--

Ndani ya kipengele cha nne cha `tr`, ongeza kipengee cha `th` chenye maandishi `Total Assets`. Funga maandishi `Assets` katika kipengele cha `span` kwa sifa ya `class` iliyowekwa kuwa `sr-only`.

Kufuatia hilo, ongeza vipengele vitatu vya `td` vyenye maandishi yafuatayo (kwa mpangilio): `$579`, `$736`, `$809`. Kipe kipengele cha tatu cha `td` sifa ya `class` iliyowekwa kuwa `current`.

# --hints--

Kipengele chako cha nne cha `tr` kinapaswa kuwa na kipengele cha `th`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th'));
```

Kipengele chako cha `th` kinapaswa kuwa na maandishi `Total Assets`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th')?.innerText === 'Total Assets');
```

Unapaswa kufunga maandishi `Assets` katika kipengele cha `span`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th > span')?.textContent === 'Assets');
```

Kipengele chako cha `span` kinafaa kuwa na sifa ya `class` iliyowekwa kuwa `sr-only`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelector('th > span')?.classList?.contains('sr-only'));
```

Unapaswa kuwa na vipengele vitatu vya `td`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td').length === 3);
```

Kipengele chako cha kwanza cha `td` kinapaswa kuwa na maandishi `$579`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[0]?.textContent === '$579');
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maandishi `$736`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[1]?.textContent === '$736');
```

Kipengele chako cha tatu cha `td` kinapaswa kuwa na maandishi `$809`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[2]?.textContent === '$809');
```

Your third `td` element should have the `class` attribute set to `current`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[3]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
