---
id: 61fd70336ebb3e4f62ee81ba
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

Majedwali ya HTML hutumia kipengele cha `caption` kuelezea jedwali linahusu nini. Kipengele cha `caption` kinapaswa kuwa mtoto wa kwanza katika `table`, lakini kinaweza kuwekwa kwa sifa ya CSS ya `caption-side`.

Ongeza kipengele cha `caption` kwa `table` yako ya kwanza, na uipe maandishi `Assets`.

# --hints--

Kipengele chako cha kwanza cha `table` kinapaswa kuwa na kipengele cha `caption`.

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

Kipengele chako cha `caption` kinapaswa kuwa na maandishi `Assets`.

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
