---
id: 61fd77f7ad2aeb5ae34d07d6
title: Hatua ya 14
challengeType: 0
dashedName: step-14
---

# --description--

Katika kipengele chako cha pili cha `tr`, ongeza kipengele cha `th` chenye maandishi `Checking Our primary transactional account.`. Funga maandishi hayo, isipokuwa `Checking`, katika kipengele cha `span` chenye sifa ya `class` iliyowekwa `description`.

Kufuatia hilo, ongeza vipengele vitatu vya `td` vyenye maandishi yafuatayo (kwa mpangilio): `$54`, `$56`, `$53`. Kipe kipengele cha tatu cha `td` sifa ya `class` iliyowekwa kuwa `current`.

# --hints--

`tr` yako ya pili inapaswa kuwa na kipengele cha `th`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th'));
```

Kipengele chako cha `th` kinapaswa kuwa na maandishi `Checking Our primary transactional account.`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th')?.innerText === 'Checking Our primary transactional account.');
```

Unapaswa kufunga maandishi `Our primary transactional account.` katika kipengele cha `span`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.textContent === 'Our primary transactional account.');
```

Kipengele chako cha `span` kinafaa kuwa na sifa ya `class` iliyowekwa kuwa `description`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.classList?.contains('description'));
```

Unapaswa kuwa na vipengele vitatu vya `td`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td').length === 3);
```

Kipengele chako cha kwanza cha `td` kinapaswa kuwa na maandishi `$54`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[0]?.textContent === '$54');
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maandishi `$56`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[1]?.textContent === '$56');
```

Kipengele chako cha tatu cha `td` kinapaswa kuwa na maandishi `$53`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.textContent === '$53');
```

Kipengele chako cha tatu cha `td` kinapaswa kuwa na `class` iliyowekwa kuwa `current`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
