---
id: 61fd6343fb42013d99bcd7f3
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele cha `body`, ongeza kipengele cha `section` ndani ya kipengele cha `main`.

# --hints--

Kipengele cha kwanza ndani ya kipengee cha `body` kinapaswa kuwa kipengele cha `main` chako.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

Unapaswa kuwa na kipengele cha `section` ndani ya kipengele chako cha `main`.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
