---
id: 61fd6cc9475a784b7776233e
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Chini ya kipengele chako cha `div` kilichopo, ongeza kipengee kipya cha `div` chenye `class` iliyowekwa kuwa `table-wrap`. Hiki kitakuwa chombo (container) cha meza zako.

Ndani ya hayo, ongeza vipengele vitatu vya `table`. Utakuwa unatumia CSS kuweka mtindo kwa hizi katika jedwali moja, lakini kutumia majedwali tofauti kutasaidia wasomaji wa skrini kuelewa mtiririko wa hati.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Kipengee chako kipya cha `div` kinapaswa kuwa na `class` iliyowekwa kuwa `table-wrap`.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

Kipengele chako cha `.table-wrap` kinafaa kuja baada ya `div` yako iliyopo.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

Kipengele chako cha `.table-wrap` kinapaswa kuwa na vipengele vitatu vya `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
