---
id: 61fd66c687e610436494c6f1
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Ndani ya kipengele cha `section`, ongeza kipengele cha `h1` na kipengele cha `span`.

# --hints--

Kipengele chako cha `section` kinapaswa kuwa na kipengele cha `h1`.

```js
assert(document.querySelector('section')?.children?.[0]?.localName === 'h1');
```

Kipengele chako cha `h1` kinapaswa kuwa na kipengele cha `span`.

```js
assert(document.querySelector('h1')?.children?.[0]?.localName === 'span');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
