---
id: 61fd719788899952e67692b9
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Vipengele vya `thead` na `tbody` hutumika kuonyesha ni sehemu gani ya jedwali lako ni kichwa, na ni sehemu gani inayo data msingi au maudhui.

Ongeza `thead` na `tbody` kwenye `table` yako ya kwanza, chini ya kipengele cha `caption`.

# --hints--

Kipengele chako cha kwanza cha `table` kinapaswa kuwa na kipengele cha `thead`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Kipengele chako cha kwanza cha `table` kinapaswa kuwa na kipengele cha `tbody`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

Kipengele chako cha `thead` kinapaswa kuwa mara moja chini ya kipengele chako cha `caption`.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

Kipengele chako cha `tbody` kinapaswa kuwa mara moja chini ya kipengele chako cha `thead`.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
