---
id: 61fd6b7c83dbf54a08cf0498
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Ndani ya kipengele chako cha `div`, ongeza vipengele vitatu vya `span`. Ipe kila mmoja yao sifa ya `class` iliyowekwa kuwa `year`, na uongeze maandishi yafuatayo (kwa mpangilio): `2019`, `2020`, na `2021`.

# --hints--

Kipengele chako cha `div` kinafaa kuwa na vipengele vitatu vya `span`.

```js
assert(document.querySelector('div')?.children?.length === 3);
```

Kipengele chako cha `span` kinafaa kuwa na sifa ya `class` iliyowekwa kuwa `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

`span` yako ya kwanza inapaswa kuwa na maandishi `2019`.

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

`span` yako ya pili inapaswa kuwa na maandishi `2020`.

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

`span` yako ya tatu inapaswa kuwa na maandishi `2021`.

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
