---
id: 5f331e55dfab7a896e53c3a1
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

El elemento `title` es uno de los varios elementos que proporcionan información adicional no visible en la página web, pero que es útil para los motores de búsqueda o en cómo es mostrada la página.

Dentro del elemento `head`, anida un elemento `meta` con un atributo `charset` con el valor `utf-8` para decirle al navegador cómo se codificarán los caracteres para la página. Ten en cuenta que los elementos `meta` se cierran automáticamente.

# --hints--

Debes tener una etiqueta `meta`.

```js
assert(code.match(/<meta.*?\/?>/is));
```

Tu etiqueta `meta` debe tener un atributo `charset`.

```js
assert(code.match(/<meta charset=/i));
```

Tu atributo `charset` debe tener el valor `utf-8`.

```js
assert(code.match(/charset=('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
