---
id: 5f332b23c2045fb843337579
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

To let visitors know the cafe was founded in 2020, add a `p` element below the `h1` element with the text `Est. 2020`.

# --hints--

Debes tener una etiqueta de apertura `<p>`.

```js
assert(code.match(/<p>/i));
```

Debes tener una etiqueta de cierre `</p>`.

```js
assert(code.match(/<\/p>/i));
```

No debes cambiar el elemento `h1` existente. Asegúrate de que no borras la etiqueta de cierre.

```js
assert($('h1').length === 1);
```

Tu elemento `p` debe estar debajo del elemento `h1`.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

Your `p` element should have the text `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

