---
id: 5f3313e74582ad9d063e3a38
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Añade un elemento `head` dentro del elemento `html`, para que puedas añadir un elemento `title`. El texto del elemento `title` deber ser `Cafe Menu`.

# --hints--

Debes tener una etiqueta `<head>` de apertura.

```js
assert(code.match(/<head>/i));
```

Debes tener una etiqueta `</head>` de cierre.

```js
assert(code.match(/<head>/i));
```

Debes tener una etiqueta `<title>` de apertura.

```js
assert(code.match(/<title>/i));
```

Debes tener una etiqueta `</title>` de cierre.

```js
assert(code.match(/<\/title>/i));
```

Tu elemento `title` debe estar anidado dentro del elemento `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Tu elemento `title` debe tener el texto `Cafe Menu`. Talvez debas revisar tu ortografía.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
