---
id: 5f33071498eb2472b87ddee4
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Como has aprendido en los pasos anteriores del proyecto Cat Photo App, hay una estructura básica necesaria para comenzar a construir tu página web.

Añade la etiqueta `<!DOCTYPE html>` y un elemento `html` con un atributo `lang` con el valor `en`.

# --hints--

Debes tener la declaración `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Debes tener una etiqueta `<html>` de apertura con un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Debes tener una etiqueta `<html>` de cierre. Recuerda que las etiquetas de cierre tienen una `/` después del corchete `<` de apertura.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
