---
id: 5f33294a6af5e9188dbdb8f3
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

El nombre de la cafetería es `CAMPER CAFE`. Añade un elemento `h1` dentro tu elemento `body`. Dale el nombre de la cafetería en mayúsculas para hacer que resalte.

# --hints--

Debes tener una etiqueta `<h1>` de apertura.

```js
assert(code.match(/<h1>/i));
```

Debes tener una etiqueta `</h1>` de cierre.

```js
assert(code.match(/<\/h1>/i));
```

No debes modificar el elemento `body` existente.

```js
assert($('body').length === 1);
```

Tu elemento `h1` debe estar anidado dentro del elemento `body`.

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

Tu elemento `h1` debe tener el texto `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
