---
id: 5f332b23c2045fb843337579
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Ya que el elemento `p` añadido en el paso anterior, provee información extra sobre la cafetería, anida el elemento `h1` y el elemento `p` dentro de un elemento `header`.

# --hints--

Debes tener una etiqueta `<header>` de apertura.

```js
assert(code.match(/<header>/i));
```

Debes tener una etiqueta `</header>` de cierre.

```js
assert(code.match(/<\/header>/i));
```

Tu elemento `h1` debe estar anidado dentro del elemento `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

Tu elemento `p` debe estar anidado dentro del elemento `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
</html>
```

