---
id: 5f332a88dc25a0fd25c7687a
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Para que los visitantes sepan que la cafetería fue fundada en el 2020, añade un elemento `p` debajo del elemento `h1` con el texto `Est. 2020`.

# --hints--

Debes tener una etiqueta `<p>` de apertura.

```js
assert(code.match(/<p>/i));
```

Debes tener una etiqueta `</p>` de cierre.

```js
assert(code.match(/<\/p>/i));
```

No debes modificar el elemento `h1` existente. Asegúrese de que no eliminaste tu etiqueta de cierre.

```js
assert($('h1').length === 1);
```

Tu elemento `p` debe estar debajo del elemento `h1`.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

Tu elemento `p` debe tener el texto `Est. 2020`.

```js
assert(code.match(/<p>Est. 2020<\/p>/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
--fcc-editable-region--
  </body>
</html>
```

