---
id: 5f33071498eb2472b87ddee4
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

En este proyecto aprenderás los conceptos básicos de CSS (Hojas de Estilo en Cascada) construyendo un menú de café. CSS es el idioma utilizado para dar estilo a un documento HTML. Describe cómo se deben mostrar los elementos HTML en la pantalla.

Como has aprendido en los pasos anteriores del proyecto Cat Photo App, hay una estructura básica necesaria para comenzar a construir tu página web. Cada documento HTML debe tener una declaración `DOCTYPE` y un elemento `html`. El `DOCTYPE` le indica al navegador qué versión del código HTML está en el documento. Y el elemento `html` representa el elemento raíz que contiene todos los demás elementos.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Añade la etiqueta `<!DOCTYPE html>` y un elemento `html` con un atributo `lang` con el valor `en`.

# --hints--

Tu Codigo debe tener una declaración `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Deberías tener un `<html lang="en">` de apertura con el idioma configurado en inglés.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Debes tener una etiqueta `</html>` de cierre.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
