---
id: 5f33310c1851c6c4da013250
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Es hora de añadir contenido al menú. Añade un elemento `main` debajo del elemento `header` existente. Eventualmente, contendrá información sobre los precios del café y postres ofrecidos por la cafetería.

# --hints--

Debes tener una etiqueta `<main>` de apertura.

```js
assert(code.match(/<main>/i));
```

Debes tener una etiqueta `</main>` de cierre.

```js
assert(code.match(/<\/main>/i));
```

No debes modificar el elemento `header`. Asegúrate de no eliminar accidentalmente tu etiqueta de cierre.

```js
assert($('header').length === 1);
```

Tu etiqueta `main` debe estar después de la etiqueta `header`.

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
