---
id: 5f331e55dfab7a896e53c3a1
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

El elemento `title` es uno de los varios elementos que proporcionan información adicional no visible en la página web, pero que es útil para los motores de búsqueda o en cómo es mostrada la página.

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8` to tell the browser how to encode characters for the page.

# --hints--

Debes tener una etiqueta `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Tu etiqueta `meta` debe tener un atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Tu atributo `charset` debe tener el valor `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
