---
id: 5f344fbc22624a2976425065
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Crea un elemento `h2` en el elemento `section` y dale el texto `Coffee`.

# --hints--

Debes tener una etiqueta `<h2>` de apertura.

```js
assert(code.match(/<h2\s*>/i));
```

Debes tener una etiqueta `</h2>` de cierre.

```js
assert(code.match(/<\/h2\s*>/i));
```

No debes modificar el elemento `section` existente. Asegúrate de que no eliminaste tu etiqueta de cierre.

```js
assert($('section').length === 1);
```

Tu elemento `h2` debe estar anidado dentro de tu elemento `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Tu elemento `h2` debe tener el texto `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
