---
id: 5f34a1fd611d003edeafd681
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Ese fondo color café hace difícil leer el texto. Cambia el color de fondo (background color) del elemento `body` para que sea `burlywood` para que tenga algo de color, pero se pueda leer el texto.

# --hints--

Debes darle el valor `burlywood` a la propiedad `background-color`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'burlywood');
assert(hasBackground);
```

Tu elemento `body` debe tener un color de fondo `burlywood`.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'burlywood');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: brown;
}
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
```
