---
id: 5f3477cb303c5cb61b43aa9b
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

El texto está centrado nuevamente, por lo tanto el vínculo al archivo CSS está funcionando. Añade otro estilo al archivo que cambie la propiedad `background-color` a `brown` para el elemento `body`.

# --hints--

Debes usar un selector `body`.

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

Debes establecer la propiedad `background-color` como `brown`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

Tu elemento `body` debe tener un color de fondo (background color) con el valor `brown` (café).

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

