---
id: 5f3477cb2e27333b1ab2b955
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Ahora tienes que vincular el archivo `styles.css` para que los estilos se apliquen de nuevo. Anida el elemento de cierre automático `link` dentro del elemento `head`. Dale un atributo `rel` con el valor `stylesheet` y un atributo `href` con el valor `styles.css`.

# --hints--

Tu código debe tener un elemento `link`.

```js
const link = document.querySelector('link');
assert(link);
```

No debes modificar el elemento `head` existente. Asegúrate de no haber eliminado la etiqueta de cierre.

```js
assert($('head').length === 1);
```

Tu elemento `link` debe ser un elemento de cierre automático.

```js
assert(code.match(/<link[\w\W\s]+\/?>/i));
```

Tu elemento `link` debe estar dentro del elemento `head`.

```js
const link = document.querySelector('head > link');
assert(link);
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
    <main>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
