---
id: 5f3477ae9675db8bb7655b30
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

En el paso anterior, utilizaste un <dfn>selector de tipo</dfn> para darle estilo al elemento `h1`. Ahora centra los elementos `h2` y `p` añadiendo un selector de tipo para cada elemento.

# --hints--

No debes modificar el elemento `h1` existente.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Debes agregar un nuevo selector `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Debes agregar un nuevo selector `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Tu elemento `h1` debe tener una propiedad, `text-align` el valor `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Tu elemento `h2` deben una propiedad `text-align` con el valor `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Tu elemento `p` debe tener una propiedad `text-align` con el valor `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
    <main>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

