---
id: 5f344fc1520b6719f2e35605
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Habrá dos secciones en el menú, una para cafés (coffes) y otra para postres (desserts). Añade un elemento `section` dentro del elemento `main` para que tengas un lugar donde poner todos los cafés disponibles.

# --hints--

Debes tener una etiqueta `<section>` de apertura.

```js
assert(code.match(/<section\s*>/i));
```

Debes tener una etiqueta `</section>` de cierre.

```js
assert(code.match(/<\/section\s*>/i));
```

No debes modificar el elemento `main` existente. Asegúrate de que no eliminaste la etiqueta de cierre.

```js
assert($('main').length === 1);
```

Tu elemento `section` debe estar anidado dentro de tu elemento `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
