---
id: 5f3477cbcb6ba47918c1da92
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Para que el estilo de la página se vea similar tanto en móvil como en ordenador de escritorio o portátil, necesitas añadir un elemento `meta` con un atributo especial `content`.

Añade lo siguiente dentro del elemento `head`:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Tu código debe tener dos elementos `meta`.

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

Tu elemento `meta` debe tener un atributo `name` con un valor de `viewport`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

Tu elemento `meta` debe tener un atributo `content` con un valor de `width=device-width, initial-scale=1.0`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
