---
id: 5f3477ae34c1239cafe128be
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

Ahora tienes tres selectores de tipo con exactamente el mismo estilo. Puede añadir el mismo grupo de estilos a varios elementos creando una lista de selectores. Cada selector se separa con comas, de esta forma:

```css
selector1, selector2 {
  property: value;
}
```

Elimina los tres selectores de tipo existentes y reemplázalos con una sola lista de selectores, que centre el texto de los elementos `h1`, `h2` y `p`.

# --hints--

Debes usar un solo selector de tipo para los tres elementos, `h1, h2, p`. Asegúrate de usar ese orden.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Debes tener un solo selector en el elemento `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

