---
id: 5f3477ae8466a9a3d2cc953c
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Ahora que tienes el código CSS en el archivo `styles.css`, puedes proseguir a eliminar el elemento `style` y todo su contenido. Una vez eliminado, el texto que estaba centrado regresara hacia la izquierda.

# --hints--

No debe tener ninguna etiqueta `style` en tu código.

```js
assert(!code.match(/style/i));
```

No debes tener ningún selector CSS en tu archivo HTML.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

