---
id: 5f3477ae9675db8bb7655b30
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

En el paso anterior, utilizaste un <dfn>selector de tipo</dfn> para darle estilo al elemento `h1`. Centra los elementos `h2` y `p` añadiendo un nuevo selector de tipo para cada uno en el elemento `style` existente.

# --hints--

No debes modificar el elemento `h1` existente.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

No debes añadir una nueva etiqueta `style`. Añade las nuevas reglas CSS a la etiqueta `style` existente.

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

Debes añadir un nuevo selector `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Debes añadir un nuevo selector `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Tu elemento `h1` debe tener una propiedad `text-align` establecida a `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Tu elemento `h2` debe tener una propiedad `text-align` establecida a `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Tu elemento `p` debe tener una propiedad `text-align` establecida a `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

