---
id: 5f344f9c805cd193c33d829c
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Puedes darle estilo a un elemento, especificándolo en el elemento `style` y añadiendo una propiedad para el elemento, de esta forma:

```css
element {
 property: value;
}
```

Centra tu elemento `h1` dándole a su propiedad `text-align` al valor `center`.

# --hints--

Debes tener un selector `h1` en el elemento `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

Tu propiedad `text-align` debe tener el valor `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

El selector `h1` debe darle a la propiedad `text-align` el valor `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

