---
id: 5f356ed69db0a491745e2bb6
title: Paso 27
challengeType: 0
dashedName: step-27
---

# --description--

Para aplicar el estilo de clase (class) al elemento `div`, elimina el atributo `id` y añade un atributo `class` a la etiqueta de apertura del elemento `div`. Asegúrate de darle el valor `menu` al atributo class.

# --hints--

El elemento `div` aún debería ser visible. Asegúrate de que no has modificado incorrectamente la etiqueta `<div>`.

```js
assert($('div').length === 1);
```

Tu elemento `div` debe tener un atributo class (clase) con el valor `menu`.

```js
assert($('div').attr('class').includes('menu'));
```

Tu elemento `div` ya no debe tener el id `menu`.

```js
assert(!$('div#menu').length);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

