---
id: 5f356ed6199b0cdef1d2be8f
title: Paso 26
challengeType: 0
dashedName: step-26
---

# --description--

Hasta ahora has estado usando selectores de tipo y id para darle estilo a los elementos. Sin embargo, es más común usar un selector diferente para dar estilo a tus elementos.

Un selector de clase (<dfn>class selector</dfn>) se define con un nombre y un punto directamente en frente, así:

```css
.class-name {
  styles
}
```

Cambia el selector existente `#menu` a un selector de clase, reemplazando `#menu` con una clase llamada `.menu`.

# --hints--

Debes tener un selector de clase `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Debes tener un selector `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

