---
id: 5f35e5c44359872a137bd98f
title: Paso 28
challengeType: 0
dashedName: step-28
---

# --description--

Ya que el producto principal a la venta en la cafetería es el café, podrías utilizar una imagen de granos de café como imagen de fondo.

Elimina el comentario y su contenido dentro del selector de tipo `body`. Ahora añade una propiedad `background-image` y dale el valor `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Debes eliminar el comentario con la propiedad `background-color`.

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

Tu selector `body` no debe tener ningún comentario.

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

La propiedad `background-image` debe tener el valor `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

Tu elemento `body` debe tener una imagen de granos de café como imagen de fondo.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

