---
id: 5f356ed63c7807a4f1e6d054
title: Paso 21
challengeType: 0
dashedName: step-21
---

# --description--

El objetivo ahora, es hacer que él `div` no ocupe todo el ancho de la página. La propiedad CSS `width`, es perfecta para esto.

Puedes usar el selector `id` para apuntar a un elemento específico con un atributo `id`. Un selector `id` se define colocando el símbolo de numeral `#` directamente delante del valor del `id`. Por ejemplo, si un elemento tiene un `id` con el valor `cat`, entonces se apuntaría a ese elemento así:

```css 
#cat {
  width: 250px;
}
```

Usa el selector `#menu` para darle a tu elemento, un ancho (width) de `300px`.

# --hints--

Debes tener un selector `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

La propiedad `width` debe tener un valor dé `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

Tu elemento `div` debe tener un ancho (width) de 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

