---
id: 5f35e5c4321f818cdc4bed30
title: Paso 29
challengeType: 0
dashedName: step-29
---

# --description--

Se ve bien. Es hora de empezar a añadir algunos elementos al menú. Añade un elemento `article` vacío debajo del encabezado `Coffee`. Contendrá el sabor y precio de cada café que estás ofreciendo.

# --hints--

Debes tener una etiqueta `<article>` de apertura.

```js
assert(code.match(/<article>/i));
```

Debes tener una etiqueta `</article>` de cierre.

```js
assert(code.match(/<\/article>/i));
```

No debes modificar el elemento `h2` existente.

```js
assert($('h2').length === 1);
```

Tu elemento `article` debe estar después del elemento `h2`.

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

