---
id: 5f3c866de7a5b784048f94b1
title: Paso 36
challengeType: 0
dashedName: step-36
---

# --description--

Esto es más o menos lo que queremos lograr, pero sería mejor si el sabor y el precio estuvieran en la misma línea. Los elementos `p` son elementos <dfn>block-level</dfn> (elementos de bloque), esto quiere decir que ocupan todo el ancho de su elemento padre.

Para colocarlos en la misma línea, necesitas aplicar un estilo a los elementos `p` y que así se comporten más como elementos <dfn>inline</dfn> (en línea). Añade un atributo `class` (clase) con el valor `item` al primer elemento `article` debajo del título `Coffee`.

# --hints--

Debes darle una clase (class) `item` a tu elemento `article`.

```js
assert(code.match(/<article\s*class=('|")item\1\s*>/i))
```

Solo debes tener un elemento con una clase (class) `item`.

```js
assert($('.item').length === 1);
```

Tu primer elemento `article` debe tener una clase (class) `item`.

```js
assert($('article')[0].className === 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
