---
id: 5f3c866daec9a49519871816
title: Paso 30
challengeType: 0
dashedName: step-30
---

# --description--

Los elementos `article` comúnmente contienen múltiples elementos que tienen información relacionada. En este caso, contendrá un sabor de café y un precio para ese sabor. Anida dos elementos `p` dentro de tu elemento `article`. El texto del primero será `French Vanilla`, y el texto del segundo será `3.00`.

# --hints--

No debes modificar el elemento `article` existente.

```js
assert($('article').length === 1);
```

Tu elemento `article` debe tener dos elementos `p`.

```js
assert($('article').children('p').length === 2);
```

Tu primer elemento `p` debe tener el texto `French Vanilla`.

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

Tu segundo elemento `p` debe tener el texto `3.00`.

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
