---
id: 5f3c866d28d7ad0de6470505
title: Paso 32
challengeType: 0
dashedName: step-32
---

# --description--

Por ahora, los sabores (flavors) y precios (prices) están uno encima del otro y se centran con su respectivo elemento `p`. Sería mejor si el sabor estuviera en la izquierda y el precio estuviera en la derecha.

Al elemento `p` con el texto `French Vanilla` agrégale un atributo class con el valor `flavor`.

# --hints--

Debes agregar un atributo class con el valor `flavor` a tu elemento `p`.

```js
assert(code.match(/<p\s*class=('|")flavor\1\s*>/i));
```

Solo debes tener un elemento con el atributo class y el valor `flavor`.

```js
assert($('.flavor').length === 1);
```

El atributo class con el valor `flavor` debe estar en el elemento `p` con el texto `French Vanilla`.

```js
assert($('.flavor')[0].innerText.match(/French Vanilla/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            <p>French Vanilla</p>
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
