---
id: 5f7b87422a560036fd03ccff
title: Paso 50
challengeType: 0
dashedName: step-50
---

# --description--

A los dos elementos `p` que acabas de añadir, añade un atributo `class` con el valor `dessert` al primer elemento `p` y un atributo `class` con el valor `price` al segundo elemento `p`.

# --hints--

Debes tener un elemento `p` con la clase `dessert`.

```js
assert($('.dessert').length === 1);
```

Tu elemento `p` con el texto `Donut` debe tener una clase `dessert`.

```js
assert($('.dessert')[0].innerText.match(/donut/i));
```

Tu elemento `p` con el texto `1.50` debe tener una clase `price`.

```js
assert($('.price').last().text().match(/1\.50/));
```

No debe haber ningún espacio entre tus elementos `p`.

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
--fcc-editable-region--
            <p>Donut</p><p>1.50</p>
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

