---
id: 61967e74a8e3690ab6292daa
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Apunte al elemento `body` para establecer el `background` en un degradado lineal inclinado 45 grados en el sentido de las agujas del reloj, comenzando en `rgb(118, 201, 255)` y terminando en `rgb(247, 255, 222)`.

# --hints--

Debe utilizar el selector de elementos `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Debe usar la propiedad `background` en el selector `body`.

```js
assert.isTrue(new __helpers.CSSHelp(document).isPropertyUsed('background'));
```

Debe establecer `background` en `linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222))`.

```js
assert.include(['linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))', 'rgba(0,0,0,0)linear-gradient(45deg,rgb(118,201,255),rgb(247,255,222))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
