---
id: 619665c9abd72906f3ad30f9
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Estarás construyendo un feliz pingüino Flappy y explorando aún más las transformaciones y animaciones CSS en el proceso.

Comience con su plantilla HTML básica. Incluya la declaración `DOCTYPE`, el elemento `html` con un idioma establecido en inglés, las etiquetas `meta` apropiadas, un `head`, `body` y elemento `title`. Además, vincule su hoja de estilo a la página.

# --hints--

El código debe tener una declaración `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

El código debe tener un elemento `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

El elemento `html` debe tener etiqueta de apertura con el atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

El elemento `html` debe tener etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/gi));
```

El código debe tener un elemento `head` dentro del elemento `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

El código debe tener un elemento `body` dentro del elemento `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

El elemento `head` debe estar antes del elemento `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Debe haber dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Uno de los elementos `meta` debe tener el atributo `name` establecido a `viewport` y el atributo `content` establecido a `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

El otro elemento `meta` debe tener el atributo `charset` establecido a `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

El código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

El elemento `title` debe contener algún texto.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

El código debe tener un elemento `link`.

```js
assert.exists(document.querySelector('link'));
```

El elemento `link` debe estar dentro del elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

El elemento `link` debe tener el atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

El elemento `link` debe tener el atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
