---
id: 60b69a66b6ddb80858c5159b
title: Крок 37
challengeType: 0
dashedName: step-37
---

# --description--

Час використати іконки FontAwesome.

Елемент `i` використовують для ідіоматичного тексту або тексту, який відокремлений від «нормального» текстового вмісту. Це може бути _курсивний_ текст, як-от наукові терміни, або іконки, подібні до тих, що надає FontAwesome.

В елемент `#white-paper` додайте чотири елементи `i`. Надайте їм `class` зі значенням `fas fa-music`.

Завдяки цьому спеціальному класу FontAwesome визначає, яку іконку завантажувати. `fas` вказує категорію іконки (тут FontAwesome Solid), а `fa-music` вибирає конкретну іконку.

# --hints--

Ви повинні мати чотири нових елементи `i` в межах елемента `#white-paper`.

```js
assert(document.querySelectorAll('#white-paper > i')?.length === 4);
```

Всі ваші елементи `i` повинні мати `class` зі значенням `fas fa-music`.

```js
const icons = document.querySelectorAll('#white-paper > i');
for (const icon of icons) {
  assert(icon.classList?.contains('fas'));
  assert(icon.classList?.contains('fa-music'));
};
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
      <div id="black-character">
        <div id="black-hat"></div>
        <div id="gray-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="white-paper">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

#white-hat {
  width: 0;
  height: 0;
  border-style: solid;
  border-width: 0 120px 140px 180px;
  border-top-color: transparent;
  border-right-color: transparent;
  border-bottom-color: GhostWhite;
  border-left-color: transparent;
  position: absolute;
  top: -140px;
  left: 0;
}

#black-mask {
  width: 100%;
  height: 50px;
  background-color: rgb(45, 31, 19);
  position: absolute;
  top: 0;
  left: 0;
  z-index: 1;
}

#gray-instrument {
  width: 15%;
  height: 40%;
  background-color: rgb(167, 162, 117);
  position: absolute;
  top: 50px;
  left: 125px;
  z-index: 1;
}

.black-dot {
  width: 10px;
  height: 10px;
  background-color: rgb(45, 31, 19);
  border-radius: 50%;
  display: block;
  margin: auto;
  margin-top: 65%;
}

#tan-table {
  width: 450px;
  height: 140px;
  background-color: #D2691E;
  position: absolute;
  top: 275px;
  left: 15px;
  z-index: 1;
}
```
