---
id: 60b69a66b6ddb80858c5157f
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Як правило, HTML відображається зверху вниз. Елементи у верхній частині коду розташовані у верхній частині сторінки. Однак часто потрібно перемістити елементи в інші позиції. Це можна зробити з допомогою атрибута `position`.

Встановіть атрибут `position` для елемента `back-wall` на `absolute`. Позиція `absolute` вилучає елемент з низхідного документообігу та дозволяє налаштувати його відносно свого контейнера.

Коли елемент розташовується вручну, ви можете зміщувати його макет з допомогою `top`, `left`, `right` та `bottom`. Встановіть `back-wall` властивості `top` зі значенням `0` та `left` зі значенням `0`.


# --hints--

Ваш селектор `#back-wall` повинен мати властивість `position` зі значенням `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.position === 'absolute');
```

Ваш селектор `#back-wall` повинен мати властивість `top` зі значенням `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.top === '0px');
```

Ваш селектор `#back-wall` повинен мати властивість `left` зі значенням `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.left === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
