---
id: 60b69a66b6ddb80858c51579
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

В свій елемент `head` додайте теґ `meta` з атрибутом `charset` зі значенням `utf-8`. Також додайте елемент `title` з текстом `Picasso Painting`.

# --hints--

Ви повинні додати рівно один елемент `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Елемент `meta` повинен мати атрибут `charset`.

```js
assert(document.querySelector('meta')?.getAttribute('charset'));
```

Ваш атрибут `charset` повинен мати значення `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Ви повинні додати рівно один елемент `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Ваш елемент `title` повинен мати текст `Picasso Painting`. Зауважте, що написання та вживання великої літери мають значення.

```js
assert(document.querySelector('title')?.innerText === 'Picasso Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
  --fcc-editable-region--

  --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
