---
id: 60b69a66b6ddb80858c51580
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Властивість `z-index` використовують, щоб створити «шари» для елементів HTML. Якщо ви знайомі з інструментами редагування зображень, ви, можливо, раніше працювали з шарами. Це схоже поняття.

Елементи з вищим значенням `z-index` виглядатимуть накладеними на елементи з нижчим значенням `z-index`. Це можна поєднати з позиціюванням в попередньому уроці для створення унікальних ефектів.

Оскільки елемент `back-wall` повинен з’явитися «за» іншими елементами, які ви створюватимете, надайте елементу `back-wall` властивість `z-index` зі значенням `-1`.

# --hints--

Ваш селектор `#back-wall` повинен мати властивість `z-index` зі значенням `-1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.zIndex === '-1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
