---
id: 60b69a66b6ddb80858c5157a
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

FontAwesome – це бібліотека іконок на основі SVG, багато з яких у вільному доступі для користування. Ви використовуватимете деякі зі цих іконок в проєкті, тому вам потрібно буде зв’язати зовнішню таблицю стилів зі своїм HTML.

Додайте елемент `link`, що має `rel` зі значенням `stylesheet` та `href` зі значенням `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

# --hints--

Ви повинні мати два елементи `link`.

```js
assert(document.querySelectorAll('link').length === 2);
```

Ваш елемент `link` повинен мати `rel` зі значенням `stylesheet`.

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('rel') === 'stylesheet');
```

Ваш елемент `link` повинен мати `href` зі значенням `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
assert(document.querySelectorAll('link')?.[1]?.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
