---
id: 60b69a66b6ddb80858c51583
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Створіть чотири елементи `div` в своєму елементі `offwhite-character`. Надайте цим елементам `div` такі значення `id`, за порядком: `white-hat`, `black-mask`, `gray-instrument`, `tan-table`.

# --hints--

Ви повинні додати чотири елементи `div` всередину свого елемента `.offwhite-character`.

```js
assert(document.querySelectorAll('#offwhite-character div').length === 4);
```

Ваш перший новий елемент `div` повинен мати `id` зі значенням `white-hat`.

```js
assert(document.querySelectorAll('#offwhite-character div')[0]?.getAttribute('id') === 'white-hat');
```

Ваш другий новий елемент `div` повинен мати `id` зі значенням `black-mask`.

```js
assert(document.querySelectorAll('#offwhite-character div')[1]?.getAttribute('id') === 'black-mask');
```

Ваш третій новий елемент `div` повинен мати `id` зі значенням `gray-instrument`.

```js
assert(document.querySelectorAll('#offwhite-character div')[2]?.getAttribute('id') === 'gray-instrument');
```

Ваш четвертий новий елемент `div` повинен мати `id` зі значенням `tan-table`.

```js
assert(document.querySelectorAll('#offwhite-character div')[3]?.getAttribute('id') === 'tan-table');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
