---
id: 60b69a66b6ddb80858c51592
title: Крок 28
challengeType: 0
dashedName: step-28
---

# --description--

Використайте селектор класу, щоб створити правило для елементів з класом `black-dot`. Встановіть `width` на `10px`, `height` на `10px` та `background-color` на `rgb(45, 31, 19)`.

# --hints--

Ви повинні мати селектор `.black-dot`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.black-dot'));
```

Ваш селектор `.black-dot` повинен мати властивість `width` зі значенням `10px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.black-dot')?.width === '10px');
```

Ваш селектор `.black-dot` повинен мати властивість `height` зі значенням `10px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.black-dot')?.height === '10px');
```

Ваш селектор `.black-dot` повинен мати властивість `background-color` зі значенням `rgb(45, 31, 19)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.black-dot')?.backgroundColor === 'rgb(45, 31, 19)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

#white-hat {
  width: 0;
  height: 0;
  border-style: solid;
  border-width: 0 120px 140px 180px;
  border-top-color: transparent;
  border-right-color: transparent;
  border-bottom-color: GhostWhite;
  border-left-color: transparent;
  position: absolute;
  top: -140px;
  left: 0;
}

#black-mask {
  width: 100%;
  height: 50px;
  background-color: rgb(45, 31, 19);
  position: absolute;
  top: 0;
  left: 0;
  z-index: 1;
}

#gray-instrument {
  width: 15%;
  height: 40%;
  background-color: rgb(167, 162, 117);
  position: absolute;
  top: 50px;
  left: 125px;
  z-index: 1;
}

--fcc-editable-region--

--fcc-editable-region--

```
