---
id: 60b69a66b6ddb80858c51584
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Цій фігурі потрібні очі. Створіть два елементи `div` в елементі `#black-mask`. Надайте їм класи `eyes left` та `eyes right`, за таким порядком.

# --hints--

Ви повинні створити два елементи `div` в елементі `#black-mask`.

```js
assert(document.querySelectorAll('#black-mask div').length === 2);
```

Ваш перший новий `div` повинен мати класи `eyes` та `left`.

```js
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[0]?.classList.contains('left'));
```

Ваш другий новий `div` повинен мати класи `eyes` та `right`.

```js
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('eyes'));
assert(document.querySelectorAll('#black-mask div')[1]?.classList.contains('right'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          --fcc-editable-region--

          --fcc-editable-region--
        </div>
        <div id="gray-instrument"></div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
