---
id: 60b80da8676fb3227967a731
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Прив’яжіть свій CSS-файл зараз, навіть якщо ви ще не написали CSS.

Додайте елемент `link`, що має `rel` зі значенням `stylesheet` та `href` зі значенням `styles.css`.

# --hints--


Ваш код повинен мати елемент `link`.

```js
assert.match(code, /<link/)
```

Ви повинні мати один самозакривний елемент `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
