---
id: 64b163c20e59cbd4a64940b0
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Create a function to add two numbers together.

# --hints--

Adding 3 and 5 should return 8.

```js
({ test: () => assert.equal(__userGlobals.get("add")(3,5), 8) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def add(a, b):
    return a + b
```
