---
id: 61764c602bee6974e7790f35
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

In school, you might have learned that red, yellow, and blue are primary colors, and learned how to create new colors by mixing those. However, this is an outdated model.

These days, there are two main color models: the <dfn>additive</dfn> RGB (red, green, blue) model used in electronic devices, and the <dfn>subtractive</dfn> CMYK (cyan, magenta, yellow, black) model used in print. In this project you'll work with the RGB model.

First, add the class `one` next to the already existing `marker` class into the first marker `div` element.

HTML elements can be assigned multiple classes by listing the classes in the class attribute, with a blank space to separate them.

Imagine you want to add another class `exampleClass2` into a div with an existing `exampleClass1`:

```html
<div class="exampleClass1">
  ...content
</div>
```

You would simply add a blank space followed by the next class name. In this case we added `exampleClass2`:

```html
<div class="exampleClass1 exampleClass2">
  ...content
</div>
```

And so on if you want to add multiple classes:

```html
<div class="exampleClass1 exampleClass2 exampleClass3">
  ...content
</div>
```

Do not forget to add a `blank space` between each class, otherwise html will read it as a single class, in the underneath example `exampleClass1exampleClass2` is a single class:

```html 
<div class="exampleClass1exampleClass2">
  ...content
</div>
```

# --hints--

You should add the class `one` to the first marker `div` element.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Your first marker `div` should have the classes `marker` and `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
