---
id: 66f12a88741aeb16b9246c59
title: Build a Weather App
challengeType: 14 
dashedName: lab-weather-app
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `button` element with an `id` of `get-forecast`.

1. You should have a `select` element with an `option` element with a `value` attribute equal to an empty string, and `option` elements that have text as follow (and a `value` that is the same but all lowercase):
   - New York
   - Los Angeles
   - Chicago
   - Paris
   - Tokyo
   - London

1. If no city is selected, pressing the button should do nothing.

1. If a city is selected, elements to show the weather should appear:

   - You should have an `img` element with the id `weather-icon` for displaying the weather icon.

   - You should have an element with the id `main-temperature` for displaying the main temperature.

   - You should have an element with the id `feels-like` for displaying what the temperature feels like.

   - You should have an element with the id `humidity` for displaying the amount of humidity in air.

   - You should have an element with the id `wind` element for displaying the wind speed.

   - You should have an element with the id `wind-gust` element for displaying the wind gust.

   - You should have an element with the id `weather-main` element for displaying the main weather type.

   - You should have an element with the id `location` element for displaying the current location.

1. You should have an asynchronous function named `getWeather` that fetches the weather information from the `https://weather-proxy.freecodecamp.rocks/api/city/<CITY>` API and returns it.

1. The `getWeather` asynchronous function should accept a city as argument.

1. You should handle any errors that occur within the `getWeather` function and log them to the console.

1. You should have a `showWeather` function that accepts a city as parameter.

1. The `showWeather` function should call the `getWeather` function to retrieve the weather data for a certain city.

1. If the `getWeather` function had an error and returns a falsy value, the `showWeather` function should only show an alert that says `Something went wrong, please try again later`.

1. If the data from `getWeather` are usable, the `showWeather` function should display the weather data in the corresponding elements.

# --hints--

You should have a `button` element with an `id` of `get-forecast`.

```js
assert.exists(document.querySelector('button#get-forecast'));
```

You should have a `select` element.

```js
assert.exists(document.querySelector('select'));
```

Inside the `select` element the first child should be an `option` element with an empty `value` attribute.

```js
assert.exists(document.querySelector('select > option[value=""]'));
assert.strictEqual(document.querySelector('select > option:first-child').value, "");
```

Inside the `select` element there should be 6 `option` elements, one for each of the following cities: Paris, London, Tokyo, Los Angeles, Chicago, New York.

```js
["Paris", "London", "Tokyo", "Los Angeles", "Chicago", "New York"].forEach(city => {
    const el = document.querySelector(`select > option[value="${city.toLowerCase()}"]`);
    assert.exists(el);
    assert.strictEqual(el.innerText.trim(), city)
});
```

You should have an `img` element with the id `weather-icon` for displaying the weather icon.

```js
async () => {
    document.querySelector('select').value = "chicago";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("img#weather-icon"));
}
```

You should have an element with the id `main-temperature` for displaying the main temperature.

```js
async () => {
    document.querySelector('select').value = "london";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("#main-temperature"));
}
```

You should have an element with the id `feels-like` for displaying what the temperature feels like.

```js
async () => {
    document.querySelector('select').value = "london";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("#feels-like"));
}
```

You should have an element with the id `humidity` for displaying the amount of humidity in air.

```js
async () => {
    document.querySelector('select').value = "london";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("#humidity"));
}
```

You should have an element with the id `wind` element for displaying the wind speed.

```js
async () => {
    document.querySelector('select').value = "new york";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("#wind"));
}
```

You should have an element with the id `wind-gust` element for displaying the wind gust.

```js
async () => {
    document.querySelector('select').value = "new york";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("#wind-gust"));
}
```

You should have an element with the id `weather-main` element for displaying the main weather type.

```js
async () => {
    document.querySelector('select').value = "new york";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
    assert.exists(document.querySelector("#weather-main"));
}
```

You should have an element with the id `location` element for displaying the current location.

```js
async () => {
    document.querySelector('select').value = "new york";
    document.querySelector('#get-forecast').click();
       await new Promise(resolve => setTimeout(resolve, 4000));
       assert.exists(document.querySelector("#location"));
}
```

You should have a `clickHandler` function.

```js
assert.isFunction(clickHandler);
```

You should have a `showWeather` function.

```js
assert.isFunction(showWeather);
```

You should have a `getWeather` function.

```js
assert.isFunction(getWeather);
```

The `getWeather` function should accepts 2 arguments, a latitude and a longitude, and return the json from the Weather API.

```js
async () => {
    try {
        const city = "chicago";
        const url = `https://weather-proxy.freecodecamp.rocks/api/city/${city}`;

        const outputFunction = await getWeather(city);

        const json = await fetch(url).then(data => data.json());
        assert.deepEqual(outputFunction, json)

    } catch (err) {
        throw new Error(err);
    }
}
```

The `showWeather` function should call the `getWeather` function to get the weather data.

```js
async () => {
    try {
        let flag = false;
        const temp = getWeather;
        getWeather = async (location) => {
            flag = true;
            return await fetch(`https://weather-proxy.freecodecamp.rocks/api/city/${location}`).then(data => data.json())
        }
        await showWeather("london");
        getWeather = temp;
        assert.isTrue(flag)
    } catch (err) {
        throw new Error(err);
    }
}
```

The `showWeather` function should display the data from the API in the respective HTML elements. If the value from the API is `undefined`, you should show `N/A`.

```js
async () => {
    // function that construct an object with the id-value pairs that we expect in the page from an object

    const helper = (wobj) => ({
        "weather-icon": wobj.weather[0].icon,
        "main-temperature": wobj.main.temp || 'N/A',
        "feels-like": wobj.main.feels_like || 'N/A',
        humidity: wobj.main.humidity || 'N/A',
        wind: wobj.wind.speed || 'N/A',
        "wind-gust": wobj.wind.gust || 'N/A',
        "weather-main": wobj.weather[0].main || 'N/A',
        location: wobj.name || 'N/A'
    })

    document.querySelector('select').value = "london";
    document.querySelector('#get-forecast').click();


    // fetch the expected values from the API to confront
    const body = await fetch(`https://weather-proxy.freecodecamp.rocks/api/city/london`).then(
        data => data.json()
    )

    await new Promise(resolve => setTimeout(resolve, 4000));

    // construct the object from the data from the API
    const extractedData = helper(body);

    // check that all things are in place
    for (const id in extractedData) {
        if (id === "weather-icon") {
            assert.equal(document.querySelector('#weather-icon').src, extractedData[id]);
        } else {
            assert.include(document.querySelector(`#${id}`).innerText, extractedData[id]);
        }
    }
}
```

If there is an error, your `getWeather` function should log the error to the console.

```js
const testArr = [];
const temp1 = fetch;
const temp2 = console.log;
async () => {
  try {
    console.log = obj => {testArr.push(obj.toString())};
    fetch = source => {throw new Error("This is a test error");}
    await getWeather();
    assert.deepEqual(testArr, ["Error: This is a test error"]);
  } finally {
    fetch = temp1;
    console.log = temp2;
  }
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Weather App</title>
</head>

<body>

</body>

</html>

```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Weather App</title>
    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Lato&display=swap" rel="stylesheet">
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="container">
        <button id="get-forecast">Get forecast for</button>
        <select>
            <option value="">-- select a city --</option>
            <option value="new york">New York</option>
            <option value="chicago">Chicago</option>
            <option value="paris">Paris</option>
            <option value="los angeles">Los Angeles</option>
            <option value="tokyo">Tokyo</option>
            <option value="london">London</option>
        </select>
    </div>
    <div class="container hidden">
        <div class="card">
            <span class="card-header" id="location"></span>
            <div class="card-body">
                <div class="weather-icon-ditails">
                    <div class="icon-temp">
                        <img id="weather-icon" src="" alt="">
                        <span id="weather-main"></span>
                    </div>
                </div>
                <div class="weather-texts">
                    <span id="main-temperature" class="weather-text"></span>
                    <span id="feels-like" class="weather-text"></span>
                    <span id="humidity" class="weather-text"></span>
                    <span id="wind" class="weather-text"></span>
                    <span id="wind-gust" class="weather-text"></span>
                </div>
            </div>
        </div>
    </div>
    <script src="script.js"></script>
</body>
</html>

```

```css
body {
    height: 100vh;
    display: flex;
    flex-direction: column;
    justify-content: center;
    gap: 4rem;
    background-color: #f5f6f7;
}

.container {
    font-family: Lato;
    display: flex;
    justify-content: center;
    align-items: center;
}

.card {
    display: flex;
    justify-content: space-between;
    flex-direction: column;

    border-radius: 10px;
    box-shadow: 5px 5px 10px 0px rgba(0, 0, 0, 0.18);
}

.card-body {
    display: flex;
}

.card-header {
    display: flex;
    justify-content: center;
    border-bottom: 0.5px solid;
    padding: 20px;
    font-size: larger;
    font-weight: 200;
}

#weather-main {
    padding: 20px;
    font-size: larger;
    font-weight: 200;
}

.weather-texts {
    display: flex;
    flex-direction: column;
    justify-content: end;
    align-items: flex-start;
    padding: 50px;
    padding-left: 20px;
}

.weather-icon-ditails {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
    padding: 50px;
    padding-right: 20px;
}

.icon-temp {
    display: flex;
    justify-content: center;
    align-items: center;
}

.hidden {
    display: none;
}

.weather-text {
    padding: 5px;
}

```

```js
const getForecastButton = document.querySelector("#get-forecast");
const dropdown = document.querySelector('select');
const container = document.querySelectorAll(".container");

getForecastButton.addEventListener("click", clickHandler);

function clickHandler() {
    const city = dropdown.value;
    if (city) {
        showWeather(city)
    }
}

async function showWeather(location) {
    const json = await getWeather(location)

    if (!json) {
        alert("Something went wrong, please try again later")
        return;
    }

    container[1].classList.remove("hidden");

    document.getElementById("weather-icon").src = json.weather[0].icon;

    document.getElementById("weather-icon").alt = `${json.weather[0].description} weather icon`;

    document.getElementById("main-temperature").textContent = `Temperature: ${json.main.temp || 'N/A'}${String.fromCharCode(176)} C`

    document.getElementById("feels-like").textContent = `Feels like: ${json.main.feels_like || 'N/A'}${String.fromCharCode(176)} C`

    document.getElementById("humidity").textContent = `Humidity: ${json.main.humidity || 'N/A'} %`;

    document.getElementById("wind").textContent = `Wind: ${json.wind.speed || 'N/A'} metre/sec`;

    document.getElementById("wind-gust").textContent = `Wind Gust: ${json.wind.gust || 'N/A'} metre/sec`;

    document.getElementById("weather-main").textContent = json.weather[0].main || 'N/A';

    document.getElementById("location").textContent = `${(json.weather[0].description || 'N/A').toUpperCase()} in ${(json.name || 'N/A').toUpperCase()}`;


}


async function getWeather(location) {
    const url = `https://weather-proxy.freecodecamp.rocks/api/city/${location}`;

    try {
        const response = await fetch(url);
        if (!response.ok) {
            throw new Error(`Response status: ${response.status}`);
        }
        return await response.json();
    } catch (error) {
        console.error(error.message);
    }
}

```
