---
id: 66f12a88741aeb16b9246c59
title: Build a Weather App
challengeType: 14 
dashedName: lab-weather-app
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have an `img` element with the id `weather-icon` for displaying the weather icon.

1. You should have an element with the id `main-temperature` for displaying the main temperature.

1. You should have an element with the id `feels-like` for displaying what the temperature feels like.

1. You should have an element with the id `humidity` for displaying the amount of humidity in air.

1. You should have an element with the id `wind` element for displaying the wind speed.

1. You should have an element with the id `wind-gust` element for displaying the wind gust.

1. You should have an element with the id `weather-main` element for displaying the main weather type.

1. You should have an element with the id `location` element for displaying the current location.

1. Your app should automatically request the user's location when the page loads, and inform them if `geolocation` is not supported by their browser. If supported, then you should get the user's current geographic position. Otherwise display an `alert` with the text `Geolocation is not available` to the user .

1. You should have an asynchronous function named `getWeather` to fetch the weather information from the `https://weather-proxy.freecodecamp.rocks/api/current?lat=<LATITUDE>&lon=<LONGITUDE>` API.

1. The `getWeather`asynchronous function should accept the geolocation position object you got from the user when the page was loaded for the first time as input.

1. You should handle any errors that occur within the `getWeather` function, whether during or after the `geolocation` request.

# --hints--

You should have an `img` element with the id `weather-icon` for displaying the weather icon.

```js
assert.exists(document.querySelector("img#weather-icon"));
```

You should have an element with the id `main-temperature` for displaying the main temperature.

```js
assert.exists(document.querySelector("#main-temperature"));
```

You should have an element with the id `feels-like` for displaying what the temperature feels like.

```js
assert.exists(document.querySelector("#feels-like"));
```

You should have an element with the id `humidity` for displaying the amount of humidity in air.

```js
assert.exists(document.querySelector("#humidity"))
```

You should have an element with the id `wind` element for displaying the wind speed.

```js
assert.exists(document.querySelector("#wind"));
```

You should have an element with the id `wind-gust` element for displaying the wind gust.

```js
assert.exists(document.querySelector("#wind-gust"));
```

You should have an element with the id `weather-main` element for displaying the main weather type.

```js
assert.exists(document.querySelector("#weather-main"));
```

You should have an element with the id `location` element for displaying the current location.

```js
assert.exists(document.querySelector("#location"));
```

You should have a `clickHandler` function.

```js
assert.isFunction(clickHandler);
```

You should have a `checkPermissions` function.

```js
assert.isFunction(checkPermissions);
```

You should have a `showWeather` function.

```js
assert.isFunction(showWeather);
```

You should have a `getWeather` function.

```js
assert.isFunction(getWeather);
```

The `getWeather` function should accepts 2 arguments, a latitude and a longitude, and return the json from the Weather API.

```js
async () => {
    try {
        const lat = 40.7128;
        const lon = 74.0060;
        const url = `https://weather-proxy.freecodecamp.rocks/api/current?lat=${lat}&lon=${lon}`;

        const outputFunction = await getWeather(lat, lon);

        const json = await fetch(url).then(data => data.json());
        assert.deepEqual(outputFunction, json)

    } catch (err) {
        throw new Error(err);
    }
}

```

The `checkPermissions` function should use the Permissions API and return `"denied"`, `"prompt"` or `"granted"` depending on the permissions state.

```js
async () => {
    try {
        const output = await checkPermissions()
        assert.oneOf(output, ["denied", "prompt", "granted"]);
        assert.fail("Find out if I can mock permissions API");
    } catch (err) {
        throw new Error(err);
    }
}
```

If `geolocation` is not supported by the user browser the `clickHandler` function should display a `p` element with the text `Geolocation is not available` to the user.

```js
async () => {
    try {
        const temp1 = checkPermissions;
        const temp2 = showWeather;
        checkPermissions = () => "denied";
        showWeather = () => {};
        await clickHandler();
        checkPermissions = temp1;
        showWeather = temp2;

        assert.isTrue([...document.querySelectorAll('p')].some(p => p.innerText.includes('Geolocation is not available')));
    } catch (err) {
        throw new Error(err);
    }
}
```

The `showWeather` function should call the `getWeather` function to get the weather data.

```js
async () => {
    try {
            let flag = false;
        	const temp = getWeather;
            getWeather = async (lat, lon) => {
                flag = true;
                return await fetch(`https://weather-proxy.freecodecamp.rocks/api/current?lat=${lat}&lon=${lon}`).then(data => data.json())
            }
            await showWeather({coords: {latitude: 44.2793984, longitude: 11.9144448}});
            getWeather = temp;
            assert.isTrue(flag)
    } catch (err) {
        throw new Error(err);
    }
}
```

You should handle any errors that occur within the `getWeather` function, whether during or after the `geolocation` request.

```js
assert(false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Weather App</title>
</head>

<body>

</body>

</html>

```

```css

```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Weather App</title>
    <link rel="preconnect" href="https://fonts.googleapis.com">
    <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin>
    <link href="https://fonts.googleapis.com/css2?family=Lato&display=swap" rel="stylesheet">
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="container">
        <button id="get-forecast">Get forecast for my location</button>
    </div>
    <div class="container hidden">
        <div class="card">
            <span class="card-header" id="location"></span>
            <div class="card-body">
                <div class="weather-icon-ditails">
                    <div class="icon-temp">
                        <img id="weather-icon" src="" alt="">
                        <span id="weather-main"></span>
                    </div>
                </div>
                <div class="weather-texts">
                    <span id="main-temperature" class="weather-text"></span>
                    <span id="feels-like" class="weather-text"></span>
                    <span id="humidity" class="weather-text"></span>
                    <span id="wind" class="weather-text"></span>
                    <span id="wind-gust" class="weather-text"></span>
                </div>
            </div>
        </div>
    </div>
    <script src="script.js"></script>
</body>
</html>

```

```css
body {
    height: 100vh;
    display: flex;
    justify-content: center;
    background-color: #f5f6f7;
}

.container {
    font-family: Lato;
    display: flex;
    justify-content: center;
    align-items: center;
}

.card {
    display: flex;
    justify-content: space-between;
    flex-direction: column;

    border-radius: 10px;
    box-shadow: 5px 5px 10px 0px rgba(0, 0, 0, 0.18);
}

.card-body {
    display: flex;
}

.card-header {
    display: flex;
    justify-content: center;
    border-bottom: 0.5px solid;
    padding: 20px;
    font-size: larger;
    font-weight: 200;
}

#weather-main {
    padding: 20px;
    font-size: larger;
    font-weight: 200;
}

.weather-texts {
    display: flex;
    flex-direction: column;
    justify-content: end;
    align-items: flex-start;
    padding: 50px;
    padding-left: 20px;
}

.weather-icon-ditails {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
    padding: 50px;
    padding-right: 20px;
}

.icon-temp {
    display: flex;
    justify-content: center;
    align-items: center;
}

.hidden {
    display: none;
}

.weather-text {
    padding: 5px;
}

```

```js
const getForecastButton = document.querySelector("#get-forecast");
const container = document.querySelectorAll(".container");

getForecastButton.addEventListener("click", () => {
  clickHandler();
});

async function clickHandler() {
  const state = await checkPermissions();
  if (["granted", "prompt"].includes(state)) {
    container[0].classList.add("hidden");
    container[1].classList.remove("hidden");
    navigator.geolocation.getCurrentPosition(showWeather)
  } else {
    container[0].innerHTML += "<p>Geolocation is not available. You can enable it in the settings of the browser</p>"
  }

}

async function showWeather(location) {
  const json = await getWeather(location.coords.latitude, location.coords.longitude)

  document.getElementById("weather-icon").src = json.weather[0].icon;

  document.getElementById("weather-icon").alt = `${json.weather[0].description} weather icon`;

  document.getElementById("main-temperature").textContent = `Temperature: ${json.main.temp}${String.fromCharCode(176)} C`

  document.getElementById("feels-like").textContent = `Feels like: ${json.main.feels_like}${String.fromCharCode(176)} C`

  document.getElementById("humidity").textContent = `Humidity: ${json.main.humidity} %`;

  document.getElementById("wind").textContent = `Wind: ${json.wind.speed} metre/sec`;

  document.getElementById("wind-gust").textContent = `Wind Gust: ${json.wind.gust} metre/sec`;

  document.getElementById("weather-main").textContent = json.weather[0].main;

  document.getElementById("location").textContent = `${json.weather[0].description.toUpperCase()} in ${json.name.toUpperCase()}`;


}
async function checkPermissions() {
  return await navigator.permissions.query({ name: "geolocation" }).then(({ state }) => state);
}

async function getWeather(lat, lon) {
  const url = `https://weather-proxy.freecodecamp.rocks/api/current?lat=${lat}&lon=${lon}`;

  try {
    const response = await fetch(url);
    if (!response.ok) {
      throw new Error(`Response status: ${response.status}`);
    }
    return await response.json();
  } catch (error) {
    console.error(error.message);
  }
}

```
