---
id: 66f12a88741aeb16b9246c59
title: Build a Weather App
challengeType: 14 
dashedName: lab-weather-app
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have an `img` element with the id `weather-icon` for displaying the weather icon.

1. You should have an element with the id `main-temperature` for displaying the main temperature.

1. You should have an element with the id `feels-like` for displaying what the temperature feels like.

1. You should have an element with the id `humidity` for displaying the amount of humidity in air.

1. You should have an element with the id `wind` element for displaying the wind speed.

1. You should have an element with the id `wind-gust` element for displaying the wind gust.

1. You should have an element with the id `weather-main` element for displaying the main weather type.

1. You should have an element with the id `location` element for displaying the current location.

1. Your app should automatically request the user's location when the page loads, and inform them if `geolocation` is not supported by their browser. If supported, then you should get the user's current geographic position. Otherwise display an `alert` with the text `Geolocation is not available` to the user .

1. You should have an asynchronous function named `getWeather` to fetch the weather information from the `https://weather-proxy.freecodecamp.rocks/api/current?lat=<LATITUDE>&lon=<LONGITUDE>` API.

1. The `getWeather`asynchronous function should accept the geolocation position object you got from the user when the page was loaded for the first time as input.

1. You should handle any errors that occur within the `getWeather` function, whether during or after the `geolocation` request.

# --hints--

You should have an `img` element with the id `weather-icon` for displaying the weather icon.

```js
assert.exists(document.querySelector("img#weather-icon"));
```

You should have an element with the id `main-temperature` for displaying the main temperature.

```js
assert.exists(document.querySelector("#main-temperature"));
```

You should have an element with the id `feels-like` for displaying what the temperature feels like.

```js
assert.exists(document.querySelector("#feels-like"));
```

You should have an element with the id `humidity` for displaying the amount of humidity in air.

```js
assert.exists(document.querySelector("#humidity"))
```

You should have an element with the id `wind` element for displaying the wind speed.

```js
assert.exists(document.querySelector("#wind"));
```

You should have an element with the id `wind-gust` element for displaying the wind gust.

```js
assert.exists(document.querySelector("#wind-gust"));
```

You should have an element with the id `weather-main` element for displaying the main weather type.

```js
assert.exists(document.querySelector("#weather-main"));
```

You should have an element with the id `location` element for displaying the current location.

```js
assert.exists(document.querySelector("#location"));
```

If `geolocation` is not supported by the user browser you should display an `alert` with the text `Geolocation is not available` to the user.

```js
Object.defineProperty(navigator, "geolocation", {});

let alertMessage;
window.alert = message => (alertMessage = message);
alert = message => (alertMessage = message);

document.dispatchEvent(new Event("DOMContentLoaded"));

assert.equal(alertMessage, "Geolocation is not available");
```

You should have an asynchronous function named `getWeather`.

```js
assert.isFunction(getWeather);

```

When geolocation is available, the `getWeather`asynchronous function should accept the geolocation position object you got from the user when the page was loaded for the first time as input and use the data to populate the page.

```js
// function that construct an object with the id-value pairs that we expect in the page from an object
/*
const helper = (wobj) => ({"weather-icon": wobj.weather[0].icon, "main-temperature": wobj.main.temp, "feels-like": wobj.main.feels_like, humidity: wobj.main.humidity, wind: wobj.wind.speed, "wind-gust": wobj.wind.gust, "weather-main": wobj.weather[0].main, location: wobj.name})

// give coords to `geolocation`
const coordObj = {coords: {latitude: 22, longitude: 14}};
Object.defineProperty(navigator, "geolocation", coordObj);
// navigator.geolocation = coordObj;

// dispatch event, getWeather should run
document.dispatchEvent(new Event("DOMContentLoaded"));

// fetch the expected values from the API to confront
const {latitude, longitude} = coordObj.coords;
fetch(`https://weather-proxy.freecodecamp.rocks/api/current?lat=${latitude}&lon=${longitude}`).then(
    data => data.json()
).then(
    body => {
      
      // construct the object from the data from the API
        const extractedData = helper(body);
      
      // check that all things are in place
        for (const id in extractedData) {
            if (id === "weather-icon") {
                assert.equal(document.querySelector('#weather-icon').src, extractedData[id]);
            } else {
                assert.includes(document.querySelector(`#${id}`).innerText, extractedData[id]);
            }
        }
    }
) */

```

You should handle any errors that occur within the `getWeather` function, whether during or after the `geolocation` request.

```js
assert(false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Weather App</title>
</head>

<body>

</body>

</html>

```

```css

```

```js

```

# --solutions--

```html
<!doctype html>
<html lang="en">
    <head>
        <meta charset="UTF-8" />
        <meta name="viewport" content="width=device-width, initial-scale=1.0" />
        <title>Weather App</title>
        <link rel="preconnect" href="https://fonts.googleapis.com" />
        <link rel="preconnect" href="https://fonts.gstatic.com" crossorigin />
        <link
            href="https://fonts.googleapis.com/css2?family=Lato&display=swap"
            rel="stylesheet"
        />
        <link rel="stylesheet" href="styles.css" />
    </head>

    <body>
        <div class="container">
            <div class="card">
                <span class="card-header" id="location"></span>
                <div class="card-body">
                    <div class="weather-icon-ditails">
                        <div class="icon-temp">
                            <img id="weather-icon" src="" alt="" />
                            <span id="main-temperature">{{template}}&deg;</span>
                        </div>
                        <div>
                            <span id="feels-like"
                                >Feels like {{template}}&deg;</span
                            >
                        </div>
                    </div>
                    <div class="weather-texts">
                        <span id="humidity" class="weather-text"
                            >Humidity:
                        </span>
                        <span id="wind" class="weather-text">Wind: </span>
                        <span id="wind-gust" class="weather-text"
                            >Wind Gust:
                        </span>
                    </div>
                </div>
                <div class="card-footer">
                    <span id="weather-main"></span>
                </div>
            </div>
        </div>
        <script src="script.js"></script>
    </body>
</html>

```

```css
body {
    height: 100vh;
    display: flex;
    justify-content: center;
    background-color: #f5f6f7;
}

.container {
    font-family: Lato;
    display: flex;
    justify-content: center;
    align-items: center;
}

.card {
    display: flex;
    justify-content: space-between;
    flex-direction: column;

    border-radius: 10px;
    box-shadow: 5px 5px 10px 0px rgba(0, 0, 0, 0.18);
}

.card-body {
    display: flex;
}

.card-header {
    display: flex;
    justify-content: center;
    border-bottom: 0.5px solid;
    padding: 20px;
    font-size: larger;
    font-weight: 200;
}

.card-footer {
    display: flex;
    border-top: 0.5px solid;
    padding: 20px;
    font-size: larger;
    font-weight: 200;
}

.weather-texts {
    display: flex;
    flex-direction: column;
    justify-content: end;
    align-items: flex-start;
    padding: 50px;
    padding-left: 20px;
}

.weather-icon-ditails {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
    padding: 50px;
    padding-right: 20px;
}

.icon-temp {
    display: flex;
    justify-content: center;
    align-items: center;
}

```

```js
document.addEventListener("DOMContentLoaded", () => {
  if (navigator?.geolocation?.hasOwnProperty("getCurrentLocation")) {
    navigator.geolocation.getCurrentPosition(getWeather);
  } else {
    alert("Geolocation is not available");
  }
});

async function getWeather(location) {
  const url = `https://weather-proxy.freecodecamp.rocks/api/current?lat=${location.coords.latitude}&lon=${location.coords.longitude}`;

  try {
    const response = await fetch(url);
    if (!response.ok) {
      throw new Error(`Response status: ${response.status}`);
    }

    const json = await response.json();

    document.getElementById("weather-icon").src = json.weather[0].icon;

    document.getElementById("main-temperature").textContent = document
      .getElementById("main-temperature")
      .textContent.replace("{{template}}", json.main.temp);
    document.getElementById("feels-like").textContent = document
      .getElementById("feels-like")
      .textContent.replace("{{template}}", json.main.feels_like);

    document.getElementById("humidity").textContent += json.main.humidity;

    document.getElementById("wind").textContent += json.wind.speed;
    document.getElementById("wind-gust").textContent += json.wind.gust;

    document.getElementById("weather-main").textContent += json.weather[0].main;

    document.getElementById("location").textContent =
      `${json.weather[0].description.toUpperCase()} in ${json.name.toUpperCase()}`;
  } catch (error) {
    console.error(error.message);
  }
}

```
