---
id: 69a013a86c97dac3784b2636
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now, the log should show something like `{ head: { element: 42, next: null }, length: 1 }` in the console. 

However, this is not how you would handle a case where the added node is the first element in the list, that is, the head.

To do that properly, use an if statement to check if the list is empty by calling the `isEmpty` function with the `list`. If that's true, then move your `list.head = node` line into the body of the `if` statement. 

# --hints--

You should have an `if` statement with the condition `isEmpty(list)`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'if(isEmpty(list)){');
```

You should move your `list.head = node` line into the `if` statement.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'if(isEmpty(list)){list.head=node');
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = {
    element: element,
    next: null
  };

--fcc-editable-region--
  list.head = node;
--fcc-editable-region--
  list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
