---
id: 69a030fbe487f185f55b8627
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Now, you will work on a function to remove a node from the list.

Create a `remove` function with `list` and `element` parameters.

# --hints--

You should create a `remove` function.

```js
assert.isFunction(remove)
```

Your `remove` function should have `list` and `element` parameters.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.remove.parameters, 2);
assert.isTrue(explorer.allFunctions.remove.parameters[0].matches("list"));
assert.isTrue(explorer.allFunctions.remove.parameters[1].matches("element"));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
    return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

// Step 13: remove function
function remove(list, element) {
    let current = list.head;
    let previous = null;

    while (current !== null) {
        if (current.element === element) {

            // if head node
            if (previous === null) {
                list.head = current.next;
            } else {
                previous.next = current.next;
            }

            list.length--;
            return;
        }

        previous = current;
        current = current.next;
    }
}

const myList = initList();

console.log(isEmpty(myList));

add(myList, 42);
add(myList, 43);
add(myList, 44);

console.log(myList);
console.log(isEmpty(myList));

// optional test
remove(myList, 43);
console.log(myList);
```
