---
id: 69a045fd236d0d973a59f103
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

If `previous` is `null`, this means the element to be removed is at the head of the list, so you should update the `head` to point to the next node in order to remove the first element.

To do that, add an `else` clause. Inside it set `current.next` to `list.head`.

# --hints--

You should add an `else` clause to your `if` statement.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

You should set `current.next` to `list.head` inside the `else` clause.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
