---
id: 69a00eaa43458e618331873c
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now you will get into working on the function that will add a node to the list.

Create an `add` function with a `list` and `element` parameter.

# --hints--

You should create an `add` function.

```js
assert.isFunction(add)
```

Your `add` function should have a `list` and `element` parameter.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
