---
id: 69a00a051b9a18694d00c16b
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you'll implement a linked list.

To begin, create a function named `initList`. Inside it, return an object with a `head` property set to `null`, and a `length` property of `0`.

This is the function you will call to initialize a linked list. The `head` represents the first element in the list, initially, it is `null` because there's no element in the list. The `length` tracks the number of nodes (elements) in the list, which you will increment or decrement when adding or removing a node from the list.

# --hints--

You should create an `initList` function.

```js
assert.isFunction(initList)
```

Your `initList` function should return an object.

```js
assert.isObject(initList());
```

Your `initList` function should return an object with a `head` set to `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Your `initList` function should return an object with a `length` set to `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
