---
id: 69a030fbe487f185f55b8627
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Now, you will work on a function to remove a node from the list.

Create a `remove` function with `list` and `element` parameters.

# --hints--

You should create a `remove` function.

```js
assert.isFunction(remove)
```

Your `remove` function should have `list` and `element` parameters.

```js
const removeParams = __helpers.getFunctionParams(remove.toString());
const removeParamsNames = removeParams.map(param => param.name)
assert.include(removeParamsNames, "list");
assert.include(removeParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
