---
id: 69a03aaa1e6fe8152788ce6a
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

Inside the `remove` function, use `let` to create a `previous` variable set to `null`, and `let` again to create a `current` variable set to `list.head`. These are the variables you will use in the loop that will traverse the list.

`current` tracks the node currently being checked in the list, starting at the head and moving forward one node at a time. `previous`, on the other hand, keeps track of the node right before the current one.

This is important when removing a node because you must update the next value of the previous node so it points to the node after the one being removed, in order to keep the list connected.

# --hints--

You should use `let` to create a `previous` variable set to `null` inside your `remove` function.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

You should use `let` to create a `current` variable set to `list.head` inside your `remove` function.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
