---
id: 69a00eaa43458e618331873c
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now you will work on the function that will add a node to the list.

Create an `add` function with `list` and `element` parameters.

# --hints--

You should create an `add` function.

```js
assert.isFunction(add)
```

Your `add` function should have `list` and `element` parameters.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.add.parameters;
assert.equal(parameters[0].toString(), "list");
assert.equal(parameters[1].toString(), "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
