---
id: 6626a060c4006f793e10cb33
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

現在，老師需要你幫助將學生們的分數轉換成一個字母評級。

完善 `getGrade` 函數， 這個函數以一個數字類型的 `score` 作爲參數。 你的函數應該根據分數返回一個字符串，代表字母評級。

這裏是分數段和它們所對應的字母評級：

| 分數段       | 字母評級    |
| --------- | ------- |
| `100`     | `"A++"` |
| `90 - 99` | `"A"`   |
| `80 - 89` | `"B"`   |
| `70 - 79` | `"C"`   |
| `60 - 69` | `"D"`   |
| `0 - 59`  | `"F"`   |


**提示：**

- 記得你學習過一些條件語句（`if`、`else if`、`else`）。
- 記得你學習過一些比較運算符（`>`、`<`、`>=`、`<=`、`===`）。

# --hints--

如果分數是 `100`，你的 `getGrade` 函數應該返回 `"A++"`。

```js
assert.strictEqual(getGrade(100), "A++");
```

如果分數是 `94`，你的 `getGrade` 函數應該返回 `"A"`。

```js
assert.strictEqual(getGrade(94), "A");
```

如果分數在 `80` 到 `89` 之間，你的 `getGrade` 函數應該返回 `"B"`。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

如果分數是 `78`，你的 `getGrade` 函數應該返回 `"C"`。

```js
assert.strictEqual(getGrade(75), "C");
```

如果分數在 `60` 到 `69` 之間，你的 `getGrade` 函數應該返回 `"D"`。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

如果分數是 `57`，你的 `getGrade` 函數應該返回 `"F"`。

```js
assert.strictEqual(getGrade(57), "F");
```

如果分數在 `0` 到 `59` 之間，你的 `getGrade` 函數應該返回 `"F"`。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
