---
id: 6626b4c58c027d86478ff5eb
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

老師對你目前所做的方案感到非常開心。 但是現在他們希望有一種簡單的方法來檢查學生的成績是否合格。 只要不是 `"F"` 評級，就是成績合格。

完善 `hasPassingGrade` 函數，這個函數把學生的分數作爲參數。 如果這個學生的成績合格，那麼你的函數應該返回 `true`，否則返回 `false`。

**提示：**

- 使用 `getGrade` 函數來獲得學生的評級。 然後檢查這個評級是否及格。

# --hints--

你的 `hasPassingGrade` 函數應該返回一個布爾值。

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

如果評級是 `"A"`，你的 `hasPassingGrade` 函數應該返回 `true`。

```js
assert.isTrue(hasPassingGrade(100));
```

如果評級是 `"F"`，你的 `hasPassingGrade` 函數應該返回 `false`。

```js
assert.isFalse(hasPassingGrade(53));
```

如果評級合格，你的 `hasPassingGrade` 函數應該返回 `true`。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {

}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
