---
id: 6626b8dcf5057f896f948440
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在，教師們已經掌握了他們所需要的所有信息，他們希望能夠將結果告訴學生。

完善 `studentMsg` 函數，這個函數以 `totalScores` 和 `studentScore` 作爲參數。 這個函數應該返回一個字符串，代表給學生的信息。

如果學生通過了該課程，字符串應遵循這種格式：

```js
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

如果學生沒通過該課程，字符串應該遵循這種格式：

```js
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

將 `average-goes-here` 替換爲總分的平均值。 將 `grade-goes-here` 替換爲學生的成績。

**提示：**

- 使用 `getAverage` 函數來獲得班級平均分。
- 調用 `getGrade` 函數來獲得學生的評級。
- 用字符串連接符（`+`）來創建發給學生的信息。
- 請留意信息中的標點符號和空格。

# --hints--

當調用 `studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` 時，你的函數應該返回以下信息：`"Class average: 71.7. Your grade: F. You failed the course."`。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

當調用 `studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` 時，你的函數應該返回以下信息：`"Class average: 50.8. Your grade: A++. You passed the course."`。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

你的 `studentMsg` 函數應該基於學生的成績和班級平均分來返回正確的信息。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
