---
id: 662693f82c91a66be46c881b
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

一位老師已經完成了學生的考試評級，需要你幫忙計算全班的平均分。

完成 `getAverage` 函數，該函數接收考試分數數組並返回平均分數。

平均分是通過累加所有得分併除以得分的數量來計算的。

```js
average = sum of all scores / total number of scores
```

這裏有一些寫好的函數調用的例子，便於你來測試你自己的代碼。

**提示：**

- 你可以使用循環來累加 `scores` 數組以計算總分。
- 你可以使用 `length` 屬性來獲取得分的數量。

# --hints--

你的 `getAverage` 函數應該返回一個數字。

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` 應該返回 `71.7`。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` 應該返回 `85.4`。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` 應該返回 `92.4`。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

你的 `getAverage` 函數應該返回考試分數的平均分。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
