---
id: 6568997f94c673a68b035b60
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--
Next, you would join the items in a row using `'|'` between each item.
For that, create a generator expression that iterates over each element (`item`) in the iterable `part`.

# --hints--
You should use the for loop to iterate over each element in the iterable `part`.

```js
assert.match(code, /for\s+item\s+in\s+part/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
            row_list = []
--fcc-editable-region--
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                
--fcc-editable-region--
```
