---
id: 6569df1d6fb83d22623b38c5
title: Step 52
challengeType: 20
dashedName: step-52
---

# --description--

The next step is to check if the specified number (`num`) is already present in the current cell of the 3x3 square.

Inside the inner `for` loop, create an `if` statement that checks if the current cell in `self.board` is equal to `num`.

# --hints--

You should have `if self.board[row_no][col_no] == num` within `valid_in_square`.

```js
const tCode = code.replace(/\r/g, '');
const innerFor = __helpers.python.getBlock(tCode, "for col_no in range(col_start, col_start + 3)");
const {block_body} = innerFor;
assert.match(block_body, /if\s+self\.board\s*\[\s*row_no\s*\]\s*\[\s*col_no\s*\]\s*==\s*num\s*\:/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\nв•”в•ђв•ђв•ђ{"в•¤в•ђв•ђв•ђ"*2}{"в•¦в•ђв•ђв•ђ"}{"в•¤в•ђв•ђв•ђ"*2}{"в•¦в•ђв•ђв•ђ"}{"в•¤в•ђв•ђв•ђ"*2}в•—\n'
        middle_lines = f'в•џв”Ђв”Ђв”Ђ{"в”јв”Ђв”Ђв”Ђ"*2}{"в•«в”Ђв”Ђв”Ђ"}{"в”јв”Ђв”Ђв”Ђ"*2}{"в•«в”Ђв”Ђв”Ђ"}{"в”јв”Ђв”Ђв”Ђ"*2}в•ў\n'
        lower_lines = f'в•љв•ђв•ђв•ђ{"в•§в•ђв•ђв•ђ"*2}{"в•©в•ђв•ђв•ђ"}{"в•§в•ђв•ђв•ђ"*2}{"в•©в•ђв•ђв•ђ"}{"в•§в•ђв•ђв•ђ"*2}в•ќ\n'
        board_string = upper_lines
        for index, line in enumerate(self.board):
            row_list = []
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                row_square = '|'.join(str(item) for item in part)
                row_list.extend(row_square)
                if square_no != 3:
                    row_list.append('в•‘')

            row = f'в•‘ {" ".join(row_list)} в•‘\n'
            row_empty = row.replace('0', ' ')
            board_string += row_empty

            if index < 8:
                if index % 3 == 2:
                    board_string += f'в• в•ђв•ђв•ђ{"в•Єв•ђв•ђв•ђ"*2}{"в•¬в•ђв•ђв•ђ"}{"в•Єв•ђв•ђв•ђ"*2}{"в•¬в•ђв•ђв•ђ"}{"в•Єв•ђв•ђв•ђ"*2}в•Ј\n'
                else:
                    board_string += middle_lines
            else:
                board_string += lower_lines

        return board_string

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(
            self.board[row][col] != num
            for row in range(9)
        )

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start=(col // 3) * 3
        for row_no in range(row_start, row_start + 3):
--fcc-editable-region--
            for col_no in range(col_start, col_start + 3):
                
--fcc-editable-region--
```
