---
id: 659422d6875503e53218abb2
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

In the existing `for each`, convert each element `item` to a string. This is necessary because when joining all elements, they should be strings.

# --hints--

You should use the `str` function to convert each element to a string.

```js
assert.match(code, /str\(item\)/);
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def __str__(self):
        upper_lines = f'\n╔═══{"╤═══"*2}{"╦═══"}{"╤═══"*2}{"╦═══"}{"╤═══"*2}╗\n'
        middle_lines = f'╟───{"┼───"*2}{"╫───"}{"┼───"*2}{"╫───"}{"┼───"*2}╢\n'
        lower_lines = f'╚═══{"╧═══"*2}{"╩═══"}{"╧═══"*2}{"╩═══"}{"╧═══"*2}╝\n'
        board_string = upper_lines

        for index, line in enumerate(self.board):
            row_list = []
--fcc-editable-region--
            for square_no, part in enumerate([line[:3], line[3:6], line[6:]], start=1):
                for item in part
--fcc-editable-region--
```
