---
id: 66069167b3307b2f4067b22b
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

A new instance of a class is created by using the function notation, which involves appending a pair of parentheses to the class name.

Outside the class definition, create an instance of the `Board` class and assign it to a variable named `gameboard`.

# --hints--

You should declare a variable `gameboard` outside the `Board` class.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Your `gameboard` variable should have the value of `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
