---
id: 64dc90e837670a3ba26a1e61
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер, коли типом є `list` відповідно до вимог, ви можете видалити виклик `print()`.

# --hints--

Видаліть `print(type(rods['A']))` з коду.

```js
({ test: () => assert.match(code, /\}(?!\s+print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")A\1\s*\]\s*\)\s*\))/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
