---
id: 64dc963d866fff3dd0329993
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Надрукуйте змінну, яку ви оголосили в попередньому кроці, і позмінюйте кількість дисків, щоб побачити, як швидко зростає необхідна мінімальна кількість ходів.

# --hints--

Надрукуйте змінну `number_of_moves`.

```js
({ test: () => assert.match(code, /print\s*\(\s*number_of_moves\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move():
    pass
```
