---
id: 64dc92a9718fb53ca3d1bad7
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Ханойську вежу можна розв’язати за 2<sup>n</sup> - 1 ходів, де n — це кількість дисків. Оголосіть змінну під назвою `number_of_moves` і призначте до неї загальну кількість ходів.

Оператором піднесення до степеня в Python є `**`.

# --hints--

Оголосіть змінну під назвою `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

Значення `number_of_moves` має бути виразом для обчислення кількості ходів.

```js
({ test: () => {
    assert(__pyodide.runPython(`        
        a = __locals.get('NUMBER_OF_DISKS')
        __locals.get('number_of_moves') == 2**a -1        
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
