---
id: 64dc9004f658183afa85f4f1
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now check the data type of your `'A'` key by passing it to the `type()` function and print the result on the terminal.

# --hints--

You should pass your `'A'` key to the `type()` function.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

You should print the type of `rods['A']` using the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
