---
id: 64dc916138dbe03bdfe6cec6
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Мета Ханойської вежі — перемістити всі диски на останній стрижень. Для цього необхідно дотримуватися трьох простих правил:

1. Можна переміщати лише диски з самого верху
2. Можна переміщати лише один диск за раз
3. Більші диски не можна класти поверх менших

Під наявним кодом оголосіть порожню функцію під назвою `move`. Пізніше ви будете використовувати цю функцію, щоб переміщувати диски між стрижнями. Наразі, щоб уникнути помилок, використайте ключове слово `pass` в тілі функції.

# --hints--

Оголосіть порожню функцію під назвою `move`. Не забудьте використати ключове слово `pass` в межах тіла функції з правильним відступом.

```js
({ test: () => {
        assert(__pyodide.runPython(`
        import inspect
        inspect.isfunction(__locals.get('move'))
    `))
    } 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
