---
id: 64dc8ea01436383a88256d7f
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Головоломка починається з того, що на першому стрижні зібрані диски в спадному порядку (найменший диск розташований зверху, а найбільший — внизу). Заповніть перший список числами, що представляють різні розміри дисків.

Замість того, щоб додавати елементи до першого списку вручну, створіть множину чисел зі зворотним відліком від `3` до `1` за допомогою функції `range()` та призначте цю множину до `rods['A']`. Цифра `3` представляє найбільший диск внизу, а `1` представляє найменший диск зверху.

Синтаксис такий: `range(x, y, h)`, де `x` — початкове ціле число (включно), `y` — останнє ціле число (виключно), а `h` — це різниця між числом та наступним числом множини.

# --hints--

Використайте функцію `range()`, щоб призначити множину чисел до `rods['A']`. Синтаксис для виклику функції `range()` такий: `range(x, y, h)`.

```js
({ test: () => assert(__pyodide.runPython(`
  a = __locals.get('rods')
  type(a['A']) is range
  `))
})
```

Використайте функцію `range()`, щоб призначити множину чисел від `3` до `1` до `rods['A']`.

```js
({ test: () => assert(__pyodide.runPython(`
 a = __locals.get('rods')
 a['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
