---
id: 64dc90949ff85e3b37be40d0
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Функція `range()` повертає незмінну множину чисел. Як бачите, типом даних `rods['A']` є `range`, але потрібно, щоб це був список.

Для цього передайте виклик `range()` до функції `list()`.

# --hints--

Передайте виклик `range()` до функції `list()`.

```js
({ test: () => assert(__pyodide.runPython(`
  a = __locals.get('rods')
  a['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
