---
id: 64dc922df2919e3c38cead61
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

У верхній частині коду оголосіть змінну під назвою `NUMBER_OF_DISKS`, щоб зберегти кількість дисків і присвойте до неї значення `3`. Потім замініть перший аргумент, переданий до функції `range()`, на нову змінну.

# --hints--

Оголосіть змінну під назвою `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

Змінна `NUMBER_OF_DISKS` повинна мати значення `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Змініть значення першого аргументу функції `range()` на змінну `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
