---
id: 64dc8b1faf9c60304ca8a9b6
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви розгадаєте математичну головоломку, відому як Ханойська вежа. Вона складається з трьох стрижнів та декількох дисків різного діаметра.

Мета цієї головоломки — перемістити диски з першого стрижня на третій, дотримуючись певних правил, які обмежують розміщення більшого диска на меншому.

Для початку створіть порожній словник під назвою `rods`, щоб представити стрижні.

# --hints--

Ви повинні мати змінну під назвою `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Змінна `rods` має бути порожнім словником.

```js
({ test: () => assert(__pyodide.runPython(`
    rods = __locals.get("rods")
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
