---
id: 64de1469fccce61940f498c4
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Словник `rods` представлятиме три стрижні з їхніми дисками. Надайте йому рядки `'A'`, `'B'` та `'C'` як ключі, значення яких встановіть на порожній список.

# --hints--

Словник `rods` повинен мати ключ `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` має бути порожнім списком.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Словник `rods` повинен мати ключ `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` має бути порожнім списком.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Словник `rods` повинен мати ключ `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` має бути порожнім списком.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
