---
id: 64df47b32b92301a815d5ef8
title: Крок 49
challengeType: 20
dashedName: step-49
---

# --description--

Спочатку рекурсивний виклик, який ви щойно додали, справляється з переміщенням `n - 1` дисків до другого стрижня.

Через це аргумент `target` відповідає другому стрижню, а аргумент `auxiliary` — третьому стрижню. Майте на увазі, що вони продовжуватимуть мінятися місцями під час рекурсії.

Виправте порядок аргументів, помінявши місцями `target` та `auxiliary` в рекурсивному виклику.

# --hints--

Змініть порядок аргументів у виклику `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)

        # move the nth disk from source to target
        rods[target].append(rods[source].pop())

        # display our progress
        print(rods, '\n')

--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
