---
id: 656879a66338b2a461d5d307
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Explain why the odd digits are summated in the Luhn algorithm -->

The Luhn algorithm is used to validate credit card numbers. The algorithm is as follows:

1. From the rightmost digit, which is the check digit, moving left, double the value of every second digit; if the product of this doubling operation is greater than 9 (e.g., 8 × 2 = 16), then sum the digits of the products (e.g., 16: 1 + 6 = 7, 18: 1 + 8 = 9) or alternatively subtract 9 from the product (e.g., 16: 16 - 9 = 7, 18: 18 - 9 = 9).
2. Take the sum of all the digits.
3. If the total modulo 10 is equal to 0 (if the total ends in zero) then the number is valid according to the Luhn formula; else it is not valid.
4. Assume an example of an account number "7992739871" that will have a check digit added, making it of the form 7992739871x:

```markdown
Account number      7  9  9  2  7  3  9  8  7  1  x
Double every other  7 18  9  4  7  6  9 16  7  2  x
Sum 2-char digits   7  9  9  4  7  6  9  7  7  2  x
```


Replace the `pass` statement with a variable `sum_of_odd_digits` and a value of `0`.

# --hints--

You should have `sum_of_odd_digits = 0` within the `verify_card_number` function.

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
    const { function_body } = verify_card_number;

    assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
  }
})
```

You should not have `pass` within the `verify_card_number` function.

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
    const { function_body } = verify_card_number;

    assert.notMatch(function_body, /pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
  pass

--fcc-editable-region--
def main():
  card_number = '4111-1111-4555-1142'
  card_translation = str.maketrans({'-': '', ' ': ''})
  translated_card_number = card_number.translate(card_translation)

  print(translated_card_number)

  verify_card_number(translated_card_number)

main()
```
