---
id: 656b47dc2cf39e37025dc033
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Reverse the order of the digits in the first four digits of `card_number`, by using a slice with a step of `-1`.

# --hints--

You should have `card_number_reversed = card_number[0:4:-1]` within the `verify_card_number` function.

```js
({
  test: () => {
    const transformedCode = e.code.original["main.py"].replace(/\r/g, "");
    const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
    const { function_body } = verify_card_number;

    assert.match(function_body, /card_number_reversed\s*=\s*card_number\[0:4:-1\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
  sum_of_odd_digits = 0
  card_number_reversed = card_number[0:4:2]
  print(card_number_reversed)
  
--fcc-editable-region--
def main():
  card_number = '4111-1111-4555-1142'
  card_translation = str.maketrans({'-': '', ' ': ''})
  translated_card_number = card_number.translate(card_translation)

  print(translated_card_number)

  verify_card_number(translated_card_number)

main()
```
