---
id: 68e41084bb275539f72de9fe
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will practice working with the `replace()`, `replaceAll()` and `repeat()` methods by building a string transformer app. 

To start, create a variable called `originalString` and assign it the value of `"I love cats."`.

# --hints--

You should have a variable called `originalString`.

```js
assert.exists(originalString);
```

Your `originalString` variable should be a string.

```js
assert.isString(originalString);
```

You should assign the string `"I love cats."` to the `originalString` variable.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
