---
id: 686478b5616b5804b0600e4d
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Right now, anytime the user edits a note and clicks outside of the note, the `currentContent` variable is being updated. But what about situations where the user focuses on the note and decides to leave without making any edits?

Well, in this situation, you can add a condition to handle that use case. 

Above your `currentContent = newContent;` line, add an `if` statement to check if `currentContent` is equal to `newContent`. If so, return.

# --hints--

You should add an `if` statement to check if `currentContent` is equal to `newContent`. If so, return.

```js
assert.match(code, /note\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*note\.innerHTML[^}]*if\s*\(currentContent\s*===\s*newContent\)\s*return[^}]*currentContent\s*=\s*newContent[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const note = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

note.addEventListener("blur", () => {
  const newContent = note.innerHTML;
  --fcc-editable-region--

  currentContent = newContent;
  console.log(currentContent);
  --fcc-editable-region--
  statusEl.textContent = "Note saved successfully!";
});

window.addEventListener("DOMContentLoaded", () => {
  currentContent = note.textContent;
});
```
