---
id: 6564aee9c077774ea49c3faf
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

The merge sort algorithm performs three actions:

- Divide an unsorted sequence of items into sub-parts
- Sort the items in the sub-parts
- Merge the sorted sub-parts

The above happens recursively until the sub-parts are merged into the complete sorted sequence. Let's start by dividing the sequence.

First, replace the `pass` keyword with a variable `middle_point`. Use the integer division operator (`//`) to divide the length of the `array` list by 2 and assign the result to your new `middle_point` variable. Remember to indent your code.

# --hints--

You should not have `pass` in the body of the function.

```js
assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
```

You should declare a `middle_point` variable and assign `len(array) // 2` to it.

```js
assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
