---
id: 6569b4e0bd29d17d4b53b16c
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

Now, you are going to replicate the same `while` loop logic for `right_part`.

Create a `while` loop that runs when `right_array_index` is less than `len(right_part)`. Within the `while` loop, assign `right_part[right_array_index]` to `array[sorted_index]`. At the end, increment `right_array_index` and `sorted_index` by `1`.

# --hints--

Your `while` loop should check if `right_array_index` is less than `len(right_part)`.

```js
assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`))
```

You should assign `right_part[right_array_index]` to `array[sorted_index]` in the body of the loop. 

```js
assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
```

You should increment `right_array_index` by `1` within the `while` loop. Note that the increment should be done after assigning the right part element to the sorted array.

```js
assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `)) 
```

You should use the increment the current value of `sorted_index` by `1` within the `while` loop. Note that the increment should be done after assigning the right part element to the sorted array.

```js
assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `)) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
