---
id: 5d5a813321b9e3db6c106a46
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript is a powerful language which allows you to build websites that are interactive.

To get started, create your standard HTML boilerplate with a `DOCTYPE`, `html`, `head`, and `body`, then add a `meta` tag for the `charset`. Add a `title` element and use the text `RPG - Dragon Repeller` for it. Include a `link` tag for your stylesheet to link the `styles.css` file.

Finally, create a `div` element with `id` set to `game` within your `body`.

# --hints--

Your code should contain the `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE/gi));
```

You should include a space after the `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

You should define the document type to be `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Your `html` element should have an opening tag. Vergiss nicht das `lang`-Attribut.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Your `html` element should have a closing tag.

```js
assert(code.match(/<\/html\s*>/));
```

Your `DOCTYPE` declaration should be at the beginning of your HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

You should have an opening `head` tag.

```js
assert(code.match(/<head\s*>/i));
```

You should have a closing `head` tag.

```js
assert(code.match(/<\/head\s*>/i));
```

You should have an opening `body` tag.

```js
assert(code.match(/<body\s*>/i));
```

You should have a closing `body` tag.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister bzw. gleichrangig sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Dein Code sollte über ein `meta`-Element verfügen.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Dein `meta`-Element sollte ein `charset`-Attribut mit dem Wert `UTF-8` enthalten.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Dein Code sollte über ein `title`-Element verfügen.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein Code sollte über ein `link`-Element verfügen.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Du solltest über ein `div`-Element verfügen.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Dein `div`-Element sollte über ein `id`-Attribut mit dem Wert `game` verfügen.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Dein `div`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
