---
id: 5d5a813321b9e3db6c106a46
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript ist eine leistungsstarke Sprache, die es dir ermöglicht Webseiten zu erstellen, die interaktiv sind.

*Note*: For all remaining projects in this curriculum, you will need a basic level of knowledge in HTML and CSS. If you are new to HTML and CSS, please go through the [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Um loszulegen, erstelle deinen Standard-HTML-Boilerplate-Code mit `DOCTYPE`, `html`, `head` und `body`. Füge dann einen `meta`-Tag für das `charset` hinzu. Füge ein `title`-Element hinzu und verwende dafür den Text `RPG - Dragon Repeller`. Füge einen `link`-Tag für dein Stylesheet hinzu, um die `styles.css`-Datei zu verknüpfen.

Erstelle dann innerhalb deines `body` ein `div`-Element mit `game` als `id`.

# --hints--

Dein Code sollte über die `DOCTYPE`-Referenz verfügen.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest nach der `DOCTYPE`-Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag haben. Vergiss nicht das `lang`-Attribut.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister bzw. gleichrangig sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Dein Code sollte über ein `meta`-Element verfügen.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Dein `meta`-Element sollte ein `charset`-Attribut mit dem Wert `UTF-8` enthalten.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Dein Code sollte über ein `title`-Element verfügen.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Du solltest ein schließendes `title`-Tag haben.

```js
assert(code.match(/<\/title\s*>/i));
```

Dein Code sollte über ein `link`-Element verfügen.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Du solltest über ein `div`-Element verfügen.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Dein `div`-Element sollte über ein `id`-Attribut mit dem Wert `game` verfügen.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Dein `div`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
