---
id: 62a1166ed9a56d439c0770e7
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Erstelle drei `span`-Elemente innerhalb deines `#stats`-Elements. Gib jedem von ihnen die Klasse `stat`. Gib dann dem ersten den Text `XP: 0`, dem zweiten den Text `Health: 100` und dem dritten den Text `Gold: 50`.

# --hints--

Du solltest drei `span`-Elemente innerhalb des `#stats`-Elements haben.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Du solltest den drei neuen `span`-Elementen eine Klasse von `stat` geben.

```js
assert.exists(document.querySelectorAll('#stats > .stat')?.length, 3);
```

Dein erstes `.stat`-Element sollte den bereitgestellten Text `XP: 0` haben.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Dein neues `.stat`-Element sollte den bereitgestellten Text `Health: 100` haben.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Dein drittes `.stat`-Element sollte den bereitgestellten Text `Gold: 50` haben.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
