---
id: 62a255dae245b52317da824a
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Eines der leistungsfähigsten Werkzeuge ist deine Entwicklerkonsole. Abhängig von deinem Browser kann dies durch Drücken von `F12` oder `Ctrl+Shift+I` geöffnet werden. Auf dem Mac kannst du `Option + ⌘ + C` drücken und "Konsole" auswählen. Du kannst auch auf die Schaltfläche "Konsole" oberhalb des Vorschaufensters klicken, um unsere integrierte Konsole zu sehen.

Die Entwicklerkonsole enthält Fehler, die durch deinen Code verursacht werden, aber du kannst sie auch verwenden, um die Werte von Variablen in deinem Code zu sehen, was bei der Fehlersuche hilfreich ist.

Füge eine `console.log("Hello World");`-Zeile zwischen deinen `script`-Tags hinzu. Klicke dann auf die "Konsole"-Schaltfläche, um die Konsole zu öffnen. Du solltest den Text `Hello World` sehen.

Beachte, dass die Zeile mit einem Semikolon endet. In JavaScript ist es üblich, Codezeilen mit Semikolons zu beenden.

# --hints--

Du solltest eine `console.log("Hello World");`-Zeile in deinen Code einfügen. Vergesse das Semikolon nicht.

```js
assert.match(code, /console\.log\("Hello World"\);/);
```

Deine `console.log("Hello World");` Zeile sollte sich zwischen deinen `script` Tags befinden.

```js
assert.match(code, /<script>\s*console\.log\("Hello World"\);\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```
