---
id: 62a255dae245b52317da824a
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Eines der leistungsfähigsten Werkzeuge ist deine Entwicklerkonsole. Abhängig von deinem Browser kann dies durch Drücken von `F12` oder `Ctrl+Shift+I` geöffnet werden. Auf dem Mac kannst du `Option + ⌘ + C` drücken und "Konsole" auswählen. Du kannst auch auf die Schaltfläche "Konsole" oberhalb des Vorschaufensters klicken, um unsere integrierte Konsole zu sehen.

Die Entwicklerkonsole enthält Fehler, die durch deinen Code verursacht werden, aber du kannst sie auch verwenden, um die Werte von Variablen in deinem Code zu sehen, was bei der Fehlersuche hilfreich ist.

Füge eine `console.log("Hello World");`-Zeile zwischen deinen `script`-Tags hinzu. Klicke dann auf die "Konsole"-Schaltfläche, um die Konsole zu öffnen. You should see the text `"Hello World"`.


# --hints--

Du solltest eine `console.log("Hello World");`-Zeile in deinen Code einfügen.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

Deine `console.log("Hello World");` Zeile sollte sich zwischen deinen `script` Tags befinden.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
