---
id: 62a257659d0d1e2456f24ba2
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Bevor du mit dem Schreiben deines Projektcodes beginnst, solltest du ihn in seine eigene Datei verschieben, um Ordnung zu halten.

Entferne deine `console.log("Hello World");`-Zeile. Gib dann deinem nun leeren `script`-Element ein `src`-Attribut mit dem Wert `./script.js`.

# --hints--

Dein Code sollte keine Zeile mit `console.log("Hello World");` enthalten.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Dein `script`-Element sollte über ein `src`-Attribut mit dem Wert `./script.js` verfügen.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
