---
id: 62a2509ba163e020bb9d84ea
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Nun kannst du beginnen, dein JavaScript zu schreiben. Beginne, indem du ein `script`-Element erstellst. Dieses Element wird verwendet, um JavaScript in deine HTML-Datei zu laden.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Du solltest ein `script`-Element haben.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

Dein `script`-Element sollte ein öffnendes Tag haben.

```js
assert.match(code, /<script\s*>/i);
```

Dein `script`-Element sollte ein schließendes Tag haben.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
