---
id: 62a23d1c5f1c93161f3582ae
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Ähnlich wie dein `#stats`-Element benötigt dein `#monsterStats`-Element zwei `span`-Elemente. Gib ihnen die Klasse `stat` und gib dem ersten Element den Text `Monster Name:` und dem zweiten den Text `Health:`. After the text in each, add a `strong` element with an empty nested `span` element.

# --hints--

Dein `#monsterStats`-Element sollte zwei `span`-Elemente haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

Deine neuen `span`-Elemente sollten beide einen `class`-Wert von `stat` haben.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Dein erstes `span`-Element sollte den Text `Monster Name:` haben. Achte darauf, dass die Abstände korrekt sind.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].innerText, 'Monster Name: ');
```

Dein zweites `span`-Element sollte den Text `Health:` haben. Achte darauf, dass die Abstände korrekt sind.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[1].textContent, 'Health: ');
```

Dein erstes `span`-Element sollte ein `strong`-Element mit einem leeren eingebetteten `span`-Element enthalten.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Dein zweites `span`-Element sollte ein `strong`-Element mit einem leeren eingebetteten `span`-Element enthalten.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Deine `strong` und `span`-Elemente sollten nach dem Text kommen.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name: <strong>/);
assert.match(spans[1].innerHTML, /Health: <strong>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
--fcc-editable-region--
        <div id="monsterStats">

        </div>
--fcc-editable-region--
        <div id="text"></div>
    </div>
</body>
</html>
```
