---
id: 62a3c4a0e52767482c5202d4
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

Setze mit der gleichen Syntax die `onclick`-Eigenschaften von `button2` und `button3` auf `goCave` bzw. auf `fightDragon`.

Öffne anschließend Ihre Konsole und versuche die Schaltflächen deines Projekts anzuklicken.

# --hints--

You should assign the `goStore` function reference to `button1.onclick`. Make sure not to call the function.

```js
assert.match(code, /button1\.onclick\s*=\s*goStore\s*;?\s*$/m);
```

Du solltest die Punktnotation verwenden, um auf die `onclick`-Eigenschaft von `button2` zuzugreifen.

```js
assert.match(code, /button2\.onclick/);
```

Du solltest nicht `let` oder `const` verwenden, um `button2.onclick` zuzuweisen.

```js
assert.notMatch(code, /(let|const)\s+button2\.onclick/);
```

You should assign the `goCave` function reference to `button2.onclick`. Make sure not to call the function.

```js
assert.match(code, /button2\.onclick\s*=\s*goCave\s*;?\s*$/m);
```

Du solltest die Punktnotation verwenden, um auf die `onclick`-Eigenschaft von `button3` zuzugreifen.

```js
assert.match(code, /button3\.onclick/);
```

Du solltest nicht `let` oder `const` verwenden, um `button3.onclick` zuzuweisen.

```js
assert.notMatch(code, /(let|const)\s+button3\.onclick/);
```

You should assign the `fightDragon` function reference to `button3.onclick`. Make sure not to call the function.

```js
assert.match(code, /button3\.onclick\s*=\s*fightDragon\s*;?\s*$/m);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

--fcc-editable-region--
// initialize buttons
button1.onclick = goStore;

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
