---
id: 62a3a75d8466a12e009eff76
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Initialisiere eine andere Variable namens `health` mit einem Wert von `100` und eine Variable namens `gold` mit einem Wert von `50`.

# --hints--

Du solltest `let` verwenden, um eine Variable namens `health` zu deklarieren.

```js
assert.match(code, /let\s+health/);
```

Du solltest eine Variable namens `health` mit einem Wert von `100` initialisieren.

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

Du solltest `let` verwenden, um eine Variable namens `gold` zu deklarieren.

```js
assert.match(code, /let\s+gold/);
```

Du solltest eine Variable namens `gold` mit einem Wert von `50` initialisieren.

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` sollte einen Wert von `100` haben.

```js
assert.equal(health, 100);
```

`gold` sollte einen Wert von `50` haben.

```js
assert.equal(gold, 50);
```

`xp` sollte einen Wert von `0` haben.

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
