---
id: 62a3b49686792938718b90d3
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Die Variablen, die du zugewiesen hast, hatten alle Werte, die Zahlen sind. JavaScript hat mehrere verschiedene Datentypen. Das nächste Element, das du verwenden wirst, ist der <dfn>String</dfn>. Strings werden verwendet, um Dinge wie Wörter oder Text zu speichern. Zeichenketten werden mit doppelten Anführungszeichen, einfachen Anführungszeichen oder Backticks umgeben. Hier ist ein Beispiel für die Deklaration einer Variablen mit einem String:

```js
let developer = "Naomi";
```

Weise der `inventory`-Variable den Wert `stick` zu.

# --hints--

Du solltest `inventory` auf die Zeichenfolge `stick` setzen.

```js
assert.equal(inventory, "stick");
```

Du solltest deine `inventory`-Variable mit dem String `stick` initialisieren.

```js
assert.match(code, /let\s+inventory\s*=\s*('|"|`)stick\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
