---
id: 62a3a0a3c0a4b32915d26a6e
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Entferne deine `console.log("Hello World");` Zeile, um mit dem Schreiben deines Projektcodes zu beginnen.

In JavaScript wird eine <dfn>variable</dfn> verwendet, um einen Wert zu halten. Um eine Variable zu verwenden, musst du sie zuerst <dfn>deklarieren</dfn>. Um zum Beispiel eine Variable namens `camperbot` zu deklarieren, würdest du folgendes schreiben:

```js
let camperbot;
```

Das `let`-Schlüsselwort gibt JavaScript an, dass du eine Variable deklarierst. Deklariere eine Variable namens `xp`.

# --hints--

Du solltest keine `console.log("Hello World");`-Zeile in deinem Code haben.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Du solltest mit dem `let`-Schlüsselwort deine Variable deklarieren.

```js
assert.match(code, /let/);
```

Du solltest eine Variable namens `xp` deklarieren.

```js
assert.match(code, /let\s+xp/);
```

Du solltest deiner Variablen keinen Wert zuweisen.

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

Vergiss das Semikolon am Ende der Zeile nicht.

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
