---
id: 62a3cfc8328d3351b95d4f61
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

Die `innerText`-Eigenschaft steuert den Text, der in einem HTML-Element angezeigt wird. Zum Beispiel:

```html
<p id="info">Demo content</p> 
```

```js
const info = document.querySelector("#info"); 
info.innerText = "Hello World"; 
```

The following example would change the text of the `p` element from `Demo content` to `Hello World`.

Wenn ein Spieler auf die Schaltfläche `Go to store` klickt, möchtest du die Schaltflächen und den Text ändern. Entferne den Code innerhalb der `goStore`-Funktion und füge eine Zeile hinzu, die den Text von `button1` so aktualisiert, dass er `"Buy 10 health (10 gold)"` sagt.

# --hints--

Du solltest keine `console.log("Going to store.");`-Zeile in deinem Code haben.

```js
assert.notMatch(code, /console\.log\(\s*('|")Going to store\.\1\s*\)/);
```

Du solltest die Punktnotation verwenden, um auf die `innerText`-Eigenschaft von `button1` zuzugreifen.

```js
assert.match(code, /button1\.innerText/);
```

Du solltest `let` oder `const` verwenden, um auf die `innerText`-Eigenschaft von `button1` zuzugreifen.

```js
assert.notMatch(code, /(let|const)\s+button1\.innerText/);
```

Du solltest die `innerText`-Eigenschaft von `button1` so aktualisieren, dass sie `Buy 10 health (10 gold)` ergibt.

```js
assert.match(code, /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

Du solltest die `innerText`-Eigenschaft innerhalb deiner `goStore`-Funktion aktualisieren.

```js
assert.match(goStore.toString(), /button1\.innerText\s*=\s*('|")Buy 10 health \(10 gold\)\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector('#button1');
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
const text = document.querySelector("#text");
const xpText = document.querySelector("#xpText");
const healthText = document.querySelector("#healthText");
const goldText = document.querySelector("#goldText");
const monsterStats = document.querySelector("#monsterStats");
const monsterName = document.querySelector("#monsterName");
const monsterHealthText = document.querySelector("#monsterHealth");

// initialize buttons
button1.onclick = goStore;
button2.onclick = goCave;
button3.onclick = fightDragon;

--fcc-editable-region--
function goStore() {
  console.log("Going to store.");
}
--fcc-editable-region--

function goCave() {
  console.log("Going to cave.");
}

function fightDragon() {
  console.log("Fighting dragon.");
}
```
