---
id: 62a3bec30ea7f941412512dc
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

Erstelle, genau wie du es mit den Tasten getan hast, Variablen für die folgenden `id`'s und verwende `querySelector()`, um ihnen das Element als Wert zuzuweisen:

`text`, `xpText`, `healthText`, `goldText`, `monsterStats` und `monsterName`.

Denke daran diese mit dem `const`-Schlüsselwort zu deklarieren und die Variablen so zu benennen, dass sie mit den `id`'s übereinstimmen.

# --hints--

Du solltest eine `text`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+text/);
```

Deine `text`-Variable sollte den Wert deines `#text`-Elements haben.

```js
assert.deepEqual(text, document.querySelector('#text'));
```

Du solltest mit `const` eine `xpText`-Variable deklarieren.

```js
assert.match(code, /const\s+xpText/);
```

Deine `xpText`-Variable sollte den Wert deines `#xpText`-Elements haben.

```js
assert.deepEqual(xpText, document.querySelector('#xpText'));
```

Du solltest eine `healthText`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+healthText/);
```

Deine `healthText`-Variable sollte den Wert des Elements `#healthText` enthalten.

```js
assert.deepEqual(healthText, document.querySelector('#healthText'));
```

Du solltest eine `goldText`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+goldText/);
```

Deine `goldText`-Variable sollte den Wert deines `#goldText`-Elements enthalten.

```js
assert.deepEqual(goldText, document.querySelector('#goldText'));
```

Du solltest eine `monsterStats`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+monsterStats/);
```

Deine `monsterStats`-Variable sollte den Wert deines `#monsterStats`-Elements haben.

```js
assert.deepEqual(monsterStats, document.querySelector('#monsterStats'));
```

Du solltest eine `monsterName`-Variable mit `const` deklarieren.

```js
assert.match(code, /const\s+monsterName/);
```

Deine `monsterName`-Variable sollte den Wert deines `#monsterName`-Elements haben.

```js
assert.deepEqual(monsterName, document.querySelector('#monsterName'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  display: none;
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");

--fcc-editable-region--
```
