---
id: 62a3b506dbaead396f58a701
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Das Inventar des Spielers in deinem Spiel wird mehrere Gegenstände enthalten können. Du wirst einen Datentyp verwenden müssen, der dies tun kann. Ein <dfn>array</dfn> kann verwendet werden, um mehrere Werte zu speichern. Zum Beispiel:

```js
let order = ["first", "second", "third"];
```

Dies ist ein Array, das drei Werte enthält. Beachte, wie die Werte durch Kommas getrennt werden.

Change your `inventory` variable to be an array with the strings `"stick"`, `"dagger"`, and `"sword"`.

# --hints--

Deine `inventory`-Variable sollte ein Array sein.

```js
assert.isArray(inventory);
```

Deine `inventory`-Variable sollte drei Werte haben.

```js
assert.lengthOf(inventory, 3);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

Your `inventory` variable should include the string `"dagger"`.

```js
assert.include(inventory, "dagger");
```

Your `inventory` variable should include the string `"sword"`.

```js
assert.include(inventory, "sword");
```

Die Werte deiner `inventory`-Variable sollten sich in der richtigen Reihenfolge befinden.

```js
assert.deepEqual(inventory, ["stick", "dagger", "sword"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = "stick";
--fcc-editable-region--
```
