---
id: 62a3b41c9494f937560640ab
title: Schitt 11
challengeType: 0
dashedName: step-11
---

# --description--

Deklariere zwei weitere Variablen namens `monsterHealth` und `inventory`, initialisiere sie aber nicht.

# --hints--

Du solltest `let` verwenden, um deine `monsterHealth` Variable zu deklarieren.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Du solltest camelCase verwenden, um deine `monsterHealth`-Variable zu benennen.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` sollte keinen Wert haben.

```js
assert.isUndefined(monsterHealth);
```

Du solltest deiner `monsterHealth`-Variablen keinen Wert zuweisen. Denke an dein Semikolon.

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

Du solltest `let` verwenden, um deine `inventory`-Variable zu deklarieren.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` sollte keinen Wert haben.

```js
assert.isUndefined(inventory);
```

Du solltest deiner `inventory`-Variablen keinen Wert zuweisen. Denke an dein Semikolon.

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
