---
id: 62a3a7e4f1060e2fc5ffb34b
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Erstelle eine weitere Variable namens `currentWeapon` und setze sie auf `0`.

Wenn ein Variablenname aus mehreren Wörtern besteht, ist es in JavaScript üblich, die sogenannte <dfn>camelCase</dfn> zu verwenden. Das erste Wort wird kleingeschrieben und der erste Buchstabe jedes folgenden Wortes wird großgeschrieben.

```js
let thisIsCamelCase;
```

# --hints--

Du solltest `let` verwenden, um eine Variable namens `currentWeapon` zu deklarieren.

```js
assert.match(code, /let\s+currentWeapon/i);
```

Du solltest camelCase verwenden, um deine Variable zu benennen.

```js
assert.match(code, /currentWeapon/);
```

Deine `currentWeapon`-Variable sollte auf `0` gesetzt sein.

```js
assert.equal(currentWeapon, 0);
```

Du solltest deine Variable auf `0` initialisieren.

```js
assert.match(code, /let\s+currentWeapon\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
