---
id: 62a3a488b24fb32b91155d56
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Variablen kann ein Wert <dfn>zugewiesen</dfn> werden. Wenn du dies bei der Deklaration tust, nennt man das <dfn>Initialisierung</dfn>. Zum Beispiel:

```js
let age = 32;
```

This would initialize the `age` variable with a value of `32`, a number.

Initialisiere deine `xp`-Variable mit dem Wert `0`, einer Zahl.

# --hints--

`xp` sollte den Wert `0` besitzen.

```js
assert.equal(xp, 0);
```

Du solltest die `xp`-Variable mit `0` initialisieren. Vergiss das Semikolon am Ende der Zeile nicht.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp;
--fcc-editable-region--
```
