---
id: 62fc20387ef88d1d1998aac5
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Verwende `querySelector()`, um die anderen beiden `button`-Elemente zu erhalten, idem du ihre `id`s: `button2` und `button3` verwendest. Speichere sie in Variablen namens `button2` und `button3`. Denke daran, `const` zu verwenden.

# --hints--

Du solltest mit `const` eine `button2`-Variable deklarieren.

```js
assert.match(code, /const\s+button2/);
```

Deine `button2` Variable sollte den Wert deines `"#button2"` Elements haben.

```js
assert.deepEqual(button2, document.querySelector('#button2'));
```

Du solltest mit `const` eine `button3`-Variable deklarieren.

```js
assert.match(code, /const\s+button3/);
```

Deine `button3` Variable sollte den Wert deines `"#button3"` Elements haben.

```js
assert.deepEqual(button3, document.querySelector('#button3'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--
const button1 = document.querySelector("#button1");

--fcc-editable-region--
```
