---
id: 646c48df8674cf2b91020ecc
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

You have been provided with a basic boilerplate. Begin by adding a `link` element within your `head` element. For that `link` element, set the `rel` attribute to `stylesheet` and the `href` to `./styles.css`.

# --hints--

You should have a `link` element.

```js
assert.exists(document.querySelector("head>link"));
```

Your `link` element should have a `rel` and `href` attribute set to `stylesheet` and `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert(href === "./styles.css" || href === "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
