---
id: 61329b210dac0b08047fd6ab
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Продовжуючи з елементами `meta`: дефініція `viewport` розповідає браузеру, як зобразити сторінку. Це одна з речей, що покращує візуальну доступність на мобільних пристроях та удосконалює _SEO_ (пошукову оптимізацію сайту).

Додайте здатність `viewport` з атрибутом `content`, деталізуючи `width` та `initial-scale` сторінки.

# --hints--

Ви повинні створити інший елемент `meta` в `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Ви повинні надати `meta` атрибут `name` зі значенням `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Ви повинні надати `meta` атрибут `content` зі значенням `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
