---
id: 61329d52e5010e08d9b9d66b
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Інший важливий елемент `meta` для доступності та SEO – це дефініція `description`. Значення атрибуту `content` використовується пошуковими системами для надання опису вашої сторінки.

Додайте елемент `meta` з атрибутом `name`, встановленим на `description`, та надайте йому придатний атрибут `content`.

# --hints--

Ви повинні додати новий елемент `meta` до `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Ви повинні надати `meta` атрибут `name` зі значенням `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Ви повинні надати `meta` атрибут `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Значення атрибуту `content` не повинне перевищувати 165 символів. _Це максимальна довжина опису Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
