---
id: 613297a923965e0703b64796
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Можливо, ви вже знайомі з елементом `meta`: його використовують для визначення інформації про сторінку, такої як заголовок, опис, ключові слова та автор.

Надайте своїй сторінці елемент `meta` з відповідним значенням `charset`.

Атрибут `charset` визначає кодування символів сторінки і сьогодні `UTF-8` є єдиним кодуванням, яке підтримується більшістю браузерів.

# --hints--

Ви повинні створити елемент `meta` в межах елемента `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Теґу `meta` надайте `charset` зі значенням `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
