---
id: 6133acc353338c0bba9cb553
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Останній елемент `title` в `head` корисний для читачів екрану, щоб зрозуміти вміст сторінки. Крім того, це важлива частина _SEO_.

Надайте своїй сторінці описовий та чіткий `title`.

# --hints--

Ви повинні додати елемент `title` до `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Не створюйте `title` більшим за 60 символів.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Намагайтесь зробити свій елемент `title` більш описовим. _Підказка: принаймні 15 символів_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
