---
id: 613e2546d0594208229ada50
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

В межах `header` надайте контекст сторінки шляхом вкладення по одному елементу `img`, `h1` та `nav`.

`img` повинен вказувати на `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` та мати `id` зі значенням `logo`.

`h1` повинен містити текст `HTML/CSS Quiz`.

# --hints--

Ви повинні додати елемент `img` до елемента `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Ви повинні додати елемент `h1` до елемента `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Ви повинні додати елемент `nav` до елемента `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Ви повинні розмістити елементи `img`, `h1` та `nav` в такому ж порядку.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Ви повинні надати елементу `img` атрибут `src` зі значенням `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Ви повинні надати елементу `img` атрибут `id` зі значенням `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Ви повинні надати елементу `h1` текст `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
