---
id: 6141fabd6f75610664e908fd
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Оскільки це вікторина, вам знадобиться форма для користувачів, щоб відправляти відповіді. Ви можете самостійно відокремити вміст в межах форми, використовуючи елементи `section`.

В межах елементу `main` створіть форму з трьома вкладеними елементами `section`. Потім зробіть так, щоб форма відправляла на `https://freecodecamp.org/practice-project/accessibility-quiz`, використовуючи правильний метод.

# --hints--

Ви повинні вкласти елемент `form` в межах елемента `main`.

```js
assert.exists(document.querySelector('main > form'));
```

Ви повинні вкласти три елементи `section` в межах елемента `form`.

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

Ви повинні надати елементу `form` атрибут `action`.

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

Ви повинні дати атрибуту `action` значення `https://freecodecamp.org/practice-project/accessibility-quiz`.

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

Ви повинні надати елементу `form` атрибут `method`.

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

Ви повинні надати елементу `form` атрибут `method` зі значенням `post`.

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
