---
id: 614202874ca576084fca625f
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Кожній ролі `region` потрібна мітка, яка допомагає користувачам читача екрана зрозуміти призначення регіону. Одним зі способів додавання мітки є додавання заголовкового елемента всередині регіону та посилання на нього за допомогою атрибуту `aria-labelledby`.

Додайте наступні атрибути `aria-labelledby` до елементів `section`:

- `student-info`
- `html-questions`
- `css-questions`

Потім в межах кожного елемента `section` вкладіть один елемент `h2` з `id`, що відповідає атрибуту `aria-labelledby`. Надайте кожному `h2` пригожий текстовий вміст.

# --hints--

Ви повинні надати першому елементу `section` атрибут `aria-labelledby` зі значенням `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

Ви повинні надати другому елементу `section` атрибут `aria-labelledby` зі значенням `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

Ви повинні надати третьому елементу `section` атрибут `aria-labelledby` зі значенням `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

Ви повинні вкласти один елемент `h2` в межах першого елемента `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

Ви повинні вкласти один елемент `h2` в межах другого елемента `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

Ви повинні вкласти один елемент `h2` в межах третього елемента `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

Ви повинні надати першому елементу `h2` атрибут `id` зі значенням `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

Ви повинні надати другому елементу `h2` атрибут `id` зі значенням `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

Ви повинні надати третьому елементу `h2` атрибут `id` зі значенням `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

Ви повинні надати першому елементу `h2` пригожий текстовий вміст. _Підказка: я б обрав `Student Info`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[0]?.textContent?.length, 1);
```

Ви повинні надати другому елементу `h2` пригожий текстовий вміст. _Підказка: я б обрав `HTML`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[1]?.textContent?.length, 1);
```

Ви повинні надати третьому елементу `h2` пригожий текстовий вміст. _Підказка: я б обрав `CSS`_

```js
assert.isAtLeast(document.querySelectorAll('h2')?.[2]?.textContent?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
