---
id: 6143639d5eddc7094161648c
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

Важливо пов'язати кожен `input` з відповідним елементом `label`. Це надає користувачам з допоміжними технологіями візуальне посилання на введення.

Це роблять завдяки наданню `label` атрибута `for`, який містить `id` теґа `input`.

Ця секція візьме ім'я, адресу електронної пошти та дату народження учня. Надайте елементам `label` відповідні атрибути `for`, а також текстовий вміст. Потім прив'яжіть елементи `input` до відповідних елементів `label`.

# --hints--

Ви повинні надати першому елементу `label` відповідний атрибут `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

Ви повинні надати другому елементу `label` відповідний атрибут `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

Ви повинні надати третьому елементу `label` відповідний атрибут `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

Ви повинні надати першому елементу `label` відповідний текстовий вміст.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.textContent?.length, 1);
```

Ви повинні надати другому елементу `label` відповідний текстовий вміст.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.textContent?.length, 1);
```

Ви повинні надати третьому елементу `label` відповідний текстовий вміст.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.textContent?.length, 1);
```

Ви повинні надати першому елементу `input` атрибут `id`, що відповідає атрибуту `for` першого `label`.

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

Ви повинні надати другому елементу `input` атрибут `id`, що відповідає атрибуту `for` другого `label`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

Ви повинні надати третьому елементу `input` атрибут `id`, що відповідає атрибуту `for` третього `label`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

Не використовуйте той самий атрибут `id` для більш ніж одного елемента `input`.

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

Не використовуйте той самий атрибут `for` для більш ніж одного елемента `label`.

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```
