---
id: 614ccc21ea91ef1736b9b578
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Вітаємо в першій частині вікторини «Доступність». Оскільки ви стаєте досвідченим HTML та CSS розробником, ми розпочали з базового шаблонного коду.

Розпочніть подорож по доступності з надання атрибута `lang` елементу `html`. Це допоможе читачам екрану з ідентифікацією мови сторінки.

# --hints--

Ви повинні надати елементу `html` атрибут `lang`. _Підказка: ви можете використати значення `en` для англійської._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
