---
id: 686253f1b051998ad4904e3e
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

You will now create the main function that accepts three parameters: the adjacency matrix, the starting node, and an optional target node.

Create a function named `shortest_path` that takes in three parameters: `matrix`, `start_node`, and `target_node`. Assign `None` as the default value for `target_node`. 


Add a `pass` statement inside the function body for now.

# --hints--

You should create a function named `shortest_path`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_function("shortest_path")`))
  }
})
```

The function should take in three parameters: `matrix`, `start_node`, and `target_node=None`.

```js
({
  test: () => {
    assert(runPython(`
      func = _Node(_code).find_function("shortest_path")
      func.has_args(["matrix", "start_node", "target_node"]) and func.args[2].has_default("None")
    `))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

--fcc-editable-region--

--fcc-editable-region--

```
