---
id: 686251fb49ebc489ba60ccb5
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

You will need to create a 2D list to represent the adjacency matrix of the graph. This matrix will be used to represent the weights of the edges between nodes in the graph.

Create a variable named `adj_matrix` and assign it a 2D list representing the graph with the following weights:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

You should create a variable named `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` should be a 2D list containing the provided values.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
