---
id: 686b78bc13d518e1267448e4
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

In addition to tracking distances, you also need to keep track of the actual paths taken to reach each node.

You'll create a list named `paths` where each entry stores the path taken to reach that node. Initially, each node's path will just contain itself.

List comprehensions provides a concise way to create lists. For example:

```py
[x * 2 for x in range(3)]`
```

Create a variable named `paths` and initialize it using a list comprehension that creates a list containing `[node_no]` for each `node_no` in `range(n)`.

# --hints--

You should have a variable called `paths`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("paths")`))
  }
})
```

You should initialize `paths` using a list comprehension.

```js
({
  test: () => {
    assert(runPython(`
func = _Node(_code).find_function("shortest_path").find_variable("paths")
func and "for" in str(func)
    `))
  }
})
```

Your list comprehension should create a nested list with `[node_no]` as the element.

```js
({
  test: () => {
    assert(runPython(`
func = _Node(_code).find_function("shortest_path").find_variable("paths")
"[node_no]" in str(func)
    `))
  }
})
```

You should have a `for` loop inside the list comprehension that iterates `node_no` over `range(n)`.

```js
({
  test: () => {
    assert(runPython(`
func = _Node(_code).find_function("shortest_path").find_variable("paths")
"range(n)" in str(func)
    `))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  
--fcc-editable-region--
```
