---
id: 6870acd39a70e73368b33db6
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

In this step, you will add a loop that will run once for each node in the graph. This loop will allow the algorithm to update distances and paths over multiple passes.

Create a `for` loop that runs `n` times. Use `_` as the loop variable since we don't need to use the iteration value. Add `pass` as the body of the loop for now.

# --hints--

You should create a `for` loop that iterates over `range(n)`.

```js
({
  test: () => {
    assert(runPython(`"for _ in range(n):" in _code`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

--fcc-editable-region--
```
