---
id: 68760464eb9c2e79cc913fac
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Next, you need to prepare for selecting the next node to process. For this step, you will create two variables inside the loop you just created. 

One variable will hold the smallest distance found so far in the current loop, and the other will store the index of the node that has this smallest distance.

Inside the `for` loop, create variables `min_distance` and `current` and set them to `INF` and `-1`, respectively.

# --hints--

You should create a variable named `min_distance` and set it to `INF`.

```js
({
  test: () => {
    assert(runPython(`"min_distance = INF" in _code`))
  }
})
```

You should create a variable named `current` and set it to `-1`.

```js
({
  test: () => {
    assert(runPython(`"current = -1" in _code`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

--fcc-editable-region--
  for _ in range(n):
    pass
    
--fcc-editable-region--
```
