---
id: 68777b7d316cf739269d207b
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

If the neighbor is reachable and unvisited, you can now calculate what the total distance would be to reach that neighbor through the current node.

Inside the `if` block, create a variable named `new_distance` and assign it the sum of `distances[current]` (the shortest distance to the current node) and `distance` (the distance from the current node to the neighbor).

# --hints--

You should create a variable named `new_distance`.

```js
({
  test: () => {
    assert(runPython(`"new_distance" in _code`))
  }
})
```

You should assign `new_distance` the sum of `distances[current]` and `distance`.

```js
({
  test: () => {
    assert(runPython(`"new_distance = distances[current] + distance" in _code`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

--fcc-editable-region--
    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        pass
        
--fcc-editable-region--
```
