---
id: 68777816e95d4936f129819f
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

Before trying to update the distance to a neighbor, you need to verify that the neighbor is both reachable and unvisited.

Inside the `for node_no in range(n):` loop, add an `if` statement that checks:

- The `distance` is not equal to `INF` (meaning there's an edge between the nodes)
- The neighbor `node_no` has not been visited yet

Add `pass` as a placeholder inside the conditional block for now.

# --hints--

You should add an `if` statement inside the loop.

```js
({
  test: () => {
    assert(runPython(`"if" in _code`))
  }
})
```

You should check if `distance` is not equal to `INF`.

```js
({
  test: () => {
    assert(runPython(`("distance != INF") in _code`))
  }
})
```

You should check if the neighbor node has not been visited using `not visited[node_no]`.

```js
({
  test: () => {
    assert(runPython(`("not visited[node_no]") in _code`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

--fcc-editable-region--
    for node_no in range(n):
      distance = matrix[current][node_no]
      
--fcc-editable-region--
```
