---
id: 687776efa7d0d135e694a570
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

Inside the loop that goes through the neighbors, you’ll need to know how far the current node is from each one.

Remove the `pass` statement from the loop you added in the previous step.

Add a variable `distance` and set it to `matrix[current][node_no]`. This will give you the distance from the current node to the neighbor node.

# --hints--

You should not have a `pass` statement inside the loop anymore.

```js
({
  test: () => {
    assert(!runPython(`"pass" in _code`))
  }
})
```

You should have a variable `distance`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("distance")`))
  }
})
```

Your variable should be set to `matrix[current][node_no]`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).find_variable("distance").is_assigned("matrix[current][node_no]")`))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

--fcc-editable-region--
    for node_no in range(n):
      pass

--fcc-editable-region--
```
