---
id: 687b3ef0c2fe185b2abc4654
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

Once the algorithm has finished running, you need to decide which node(s) to display results for.

If a specific `target_node` was provided, you’ll only show the distance and path to that node. Otherwise, you’ll show results for all nodes.

Create a variable named `targets` and assign it a list containing only the `target_node` if a target was provided. Otherwise, assign it a list of all node indices using `range(n)`.

# --hints--

You should create a variable named `targets`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("targets")`))
  }
})
```

You should assign `targets` to a list containing only the `target_node`.

```js
({
  test: () => {
    assert(runPython(`
      any(
        "targets" in assign.targets and "target_node" in assign.value
        for assign in _Node(_code).find_assignments()
      )
    `))
  }
})
```

You should assign `targets` to a list of all node indices if no target was provided.

```js
({
  test: () => {
    assert(runPython(`
      any(
        "range(n)" in assign.value
        for assign in _Node(_code).find_assignments("targets")
      )
    `))
  }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
     [0, 5, 3, INF, 11, INF],
     [5, 0, 1, INF, INF, 2],
     [3, 1, 0, 1, 5, INF],
     [INF, INF, 1, 0, 9, 3],
     [11, INF, 5, 9, 0, INF],
     [INF, 2, INF, 3, INF, 0],
 ]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        new_distance = distances[current] + distance
        if new_distance < distances[node_no]:
          distances[node_no] = new_distance
          paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--

--fcc-editable-region--
```
