---
id: 687b45e16f38ad5f4e4bd799
title: Step 23
challengeType: 20
dashedName: step-23
---

# --description--

Now that you've determined a node should be displayed, you need to format its path so it can be printed clearly.

Inside the loop after the `if` statement, create a variable called `path`. Assign it a string that joins all the node numbers in `paths[node_no]` using `" -> "` as the separator.

Use `str()` to convert each node number to a string before joining them.

# --hints--

You should create a variable called `path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("path")
      `)
      );
    }
});
```

You should use `" -> ".join()` to join the path elements.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].find_variable("path").is_equivalent('path = " -> ".join(str(n) for n in paths[node_no])')
      `)
      );
    }
});
```

You should convert each element to a string using `str()` in the generator expression.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("path")
"str(n)" in str(path_var)
      `)
      );
    }
});
```

You should use a `for` loop to iterate over `paths[node_no]`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("path")
"for n in paths[node_no]" in str(path_var)
      `)
      );
    }
});
```

You should assign the joined string to the variable `path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].find_variable("path").is_equivalent('path = " -> ".join(str(n) for n in paths[node_no])')
      `)
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
      min_distance = INF
      current = -1
      for node_no in range(n):
        if not visited[node_no] and distances[node_no] < min_distance:
          min_distance = distances[node_no]
          current = node_no

      if current == -1:
        break
      visited[current] = True

      for node_no in range(n):
        distance = matrix[current][node_no]
        if distance != INF and not visited[node_no]:
          new_distance = distances[current] + distance
          if new_distance < distances[node_no]:
            distances[node_no] = new_distance
            paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
      if node_no == start_node or distances[node_no] == INF:
        continue

--fcc-editable-region--
```
