---
id: 687b4bd2e4c7c962ec08310c
title: Step 30
challengeType: 20
dashedName: step-30
---

# --description--

Now you will demonstrate that your function works. At the end of your program, call your function `shortest_path` with the `adj_matrix` and a `start_node` of your choice as well as a `target_node` as arguments.

You can use the `start_node` as `0` and the `target_node` as `5`.

With that, the shortest path algorithm is complete!

# --hints--

You should call the `shortest_path` function.

```js
({
  test: () => {
  assert(runPython(`"shortest_path(" in _code`))
 }
})
```

You should call `shortest_path` with `adj_matrix` as the first argument.

```js
({
  test: () => {
  assert(runPython(`"shortest_path(adj_matrix" in _code`))
 }
})
```

You should call `shortest_path` with `0` as the start node.

```js
({
  test: () => {
  assert(runPython(`"shortest_path(adj_matrix, 0" in _code`))
 }
})
```

You should call `shortest_path` with `5` as the target node.

```js
({
  test: () => {
  assert(runPython(`"shortest_path(adj_matrix, 0, 5" in _code`))
 }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        new_distance = distances[current] + distance
        if new_distance < distances[node_no]:
          distances[node_no] = new_distance
          paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
  targets = [target_node] if target_node is not None else range(n)
  for node_no in targets:
    if node_no == start_node or distances[node_no] == INF:
      continue
    path = " -> ".join(str(n) for n in paths[node_no])
    print(f"\n{start_node}-{node_no} distance: {distances[node_no]}\nPath: {path}")
  
  return distances, paths


--fcc-editable-region--
```

# --solutions--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        new_distance = distances[current] + distance
        if new_distance < distances[node_no]:
          distances[node_no] = new_distance
          paths[node_no] = paths[current] + [node_no]

  targets = [target_node] if target_node is not None else range(n)
  for node_no in targets:
    if node_no == start_node or distances[node_no] == INF:
      continue
    path = " -> ".join(str(n) for n in paths[node_no])
    print(f"\n{start_node}-{node_no} distance: {distances[node_no]}\nPath: {path}")
  
  return distances, paths

shortest_path(adj_matrix, 0, 5)
```
