---
id: 687b426c3dabfe5cf80fad80
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Now that you've defined which nodes you want to display results for, you need to loop through them.

Add a `for` loop that iterates through each `node_no` in the `targets` list. Add `pass` as a placeholder inside the loop body.

# --hints--

You should create a `for` loop that uses `node_no` to iterate over `targets`.

```js
({
    test: () => {
      runPython(`
        loop = _Node(_code).find_function("shortest_path").find_for_loops()[1]
        assert loop.find_for_iter().is_equivalent("targets")
        assert loop.find_for_vars().is_equivalent("node_no")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no  

        if current == -1:  
            break  
        visited[current] = True  

        for node_no in range(n):  
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance  
                if new_distance < distances[node_no]:  
                    distances[node_no] = new_distance  
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)

--fcc-editable-region--
```
