---
id: 687b48ccb46c5e6126d30bc6
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--

Finally, print the results for the current node.

Your output should show two things, the distance from the `start_node` to `node_no` and the path taken to get there.

Use an `f-string` to format the output to show both the distance and the full path. For example:

```py
print(f"\n{begin_node}-{node_num} length: {lengths[node_num]}\nPath: {pathway}")
```

# --hints--

You should have a `print` statement.

```js
({
  test: () => {
  assert(runPython(`"print(" in _code`))
 }
})
```

You should use an `f-string` to format the output.

```js
({
  test: () => {
  assert(runPython(`"print(f" in _code`))
 }
})
```

You should include `start_node` output.

```js
({
  test: () => {
  assert(runPython(`"{start_node}" in _code`))
 }
})
```

You should include `node_no` output.

```js
({
  test: () => {
  assert(runPython(`"{node_no}" in _code`))
 }
})
```

You should print the distance using `distances[node_no]`.

```js
({
  test: () => {
  assert(runPython(`"distances[node_no]" in _code`))
 }
})
```

You should print the path taken to get to `node_no`.

```js
({
  test: () => {
  assert(runPython(`"path" in _code`))
 }
})
```

# --seed--

## --seed-contents--

```py
INF = float("inf")
adj_matrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
  n = len(matrix)
  distances = [INF] * n
  distances[start_node] = 0
  paths = [[node_no] for node_no in range(n)]
  visited = [False] * n

  for _ in range(n):
    min_distance = INF
    current = -1
    for node_no in range(n):
      if not visited[node_no] and distances[node_no] < min_distance:
        min_distance = distances[node_no]
        current = node_no

    if current == -1:
      break
    visited[current] = True

    for node_no in range(n):
      distance = matrix[current][node_no]
      if distance != INF and not visited[node_no]:
        new_distance = distances[current] + distance
        if new_distance < distances[node_no]:
          distances[node_no] = new_distance
          paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
  targets = [target_node] if target_node is not None else range(n)
  for node_no in targets:
    if node_no == start_node or distances[node_no] == INF:
      continue
    path = " -> ".join(str(n) for n in paths[node_no])
    
--fcc-editable-region--
```
