---
id: 68e4c0ff9edcc92771d7549a
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Right now, the `distances` list only has one element. However, you need one distance value for each node in the graph.

In Python, you can multiply a list by an integer to repeat it. For example, `[0] * 3` creates `[0, 0, 0]`.

Update the `distances` list to contain `n` copies of `INF` by multiplying `[INF]` by `n`.

# --hints--

You should update `distances` to contain `n` copies of `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
