---
id: 641da42481d90c4314c99e94
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Die `fetch()`-Methode gibt ein `Promise` zurück, das ein Platzhalterobjekt ist, das entweder erfüllt wird, wenn deine Anfrage erfolgreich ist, oder abgelehnt wird, wenn deine Anfrage nicht erfolgreich ist.

Wenn das `Promise` erfüllt ist, wird es in ein `Response`-Objekt aufgelöst, und du kannst die `.then()`-Methode verwenden, um auf die `Response` zuzugreifen.

Hier siehst du, wie du `.then()` mit der `fetch()`-Methode verketten kannst:

```js
fetch("sample-url-goes-here")
  .then((res) => res)
```

Verkette die `.then()`-Methode mit deinem `fetch`-Aufruf. Füge innerhalb der `.then()`-Methode eine Callback-Funktion mit `res` als Parameter hinzu und protokolliere dann `res` auf der Konsole, damit du das `Response`-Objekt sehen kannst. Öffne deine Konsole und erweitere das `Response`-Objekt, um zu sehen, was es enthält.

Nochmals: Beende den Code noch nicht mit einem Semikolon.

# --hints--

Du solltest die `fetch()`-Methode verwenden, um eine `GET`-Anfrage an `https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json` zu stellen.

```js
assert.match(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)/)
```

Du solltest die `.then()`-Methode mit deiner `fetch()`-Syntax verketten

```js
assert.match(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)\s*\.then\(/)
```

Du solltest deiner `.then()`-Methode einen `res`-Parameter hinzufügen.

```js
assert.match(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)\s*\.then\(\s*(\(\s*res\s*\)|res)/)
```

Deine `.then()`-Methode sollte eine Arrow-Funktion-Syntax verwenden.

```js
assert.match(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)\s*\.then\(\s*(\(\s*res\s*\)|res)\s*=>\s*/)
```

Du solltest den `res`-Parameter auf der Konsole protokollieren, damit du das `Response`-Objekt sehen kannst.

```js
assert.match(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)\s*\.then\(\s*(\(\s*res\s*\)|res)\s*=>\s*console\.log\(\s*res\s*\)\s*\)/)
```

Nach deinem `.then()` sollte kein Semikolon stehen.

```js
assert.notMatch(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)\s*\.then\(\s*(\(\s*res\s*\)|res)\s*=>\s*console\.log\(\s*res\s*\)\s*\)\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--
fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')

--fcc-editable-region--
```
