---
id: 641da3c6b6fbd742bff6ee40
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Die <dfn>Fetch API</dfn> ist ein eingebautes JavaScript-Interface mit dem man Netzwerk-Anfragen an einen Serven stellen kann. Sie verf端gt 端ber eine `fetch()`-Methode, mit der du `GET`-, `POST`-, `PUT`- oder `PATCH`-Anfragen stellen kannst. In diesem Projekt wirst du eine `GET`-Anfrage an eine URL f端r eine JSON-Datei mit Informationen 端ber Autoren auf freeCodeCamp-News stellen.

Hiermit kannst du eine `GET`-Anfrage mit der `fetch()`-Methode stellen:

```js
fetch("url-goes-here")
```

Stelle eine `GET`-Anfrage an diese URL: `https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json`. Beende deinen Code noch nicht mit einem Semikolon.

# --hints--

Du solltest die `fetch()`-Methode verwenden, um eine `GET`-Anfrage an `https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json` zu stellen.

```js
assert.match(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)/)
```

Nach deinem `fetch`-Aufruf sollte kein Semikolon stehen.

```js
assert.notMatch(code, /fetch\(\s*('|"|`)https:\/\/cdn\.freecodecamp\.org\/curriculum\/news\-author\-page\/authors\.json\1\s*\)\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--
```
