---
id: 6ac09f8420b4b67e075d2cd2
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

step 1 instructions

# --hints--

Test 1

```js
({ test: () => assert(runPython(`<python expression that returns truthy>`)) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
