---
id: 66d9af3897e7d75a895b72c2
title: Build the nth Fibonacci number generator
challengeType: 14
dashedName: lab-nth-fibonacci-number-generator

---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should define an array named `sequence` in the global scope. It should contain the number `0` and `1`.
1. You should create a function named `fibonacci`.
1. The `fibonacci` function should take a single parameter, which is a positive integer.
1. A call to `fibonacci(n)` should use a dynamic programming approach to compute and return the `n`-th number from the Fibonacci sequence, in which each number is the sum of the two preceding numbers.
1. Each computed number at the position `n` in the Fibonacci sequence should be stored at the index `n - 1` of the `sequence` array.

# --hints--

You should have a `sequence` array.

```js
assert.isArray(sequence);
```

The `sequence` array should contain `0` and `1` at definition.

```js
assert.match(__helpers.removeJSComments(code), /(let|const)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]\s*;?/)
```

You should create a `fibonacci` function.

```js
assert.isFunction(fibonacci);
```

THe `fibonacci` function should take a single parameter.

```js

```

If the `fibonacci` function receives an argument that is not a positive integer it should return `"Invalid input"`.

```js

```

`fibonacci(1)` should return `0`.

```js

```

`fibonacci(32)` should return `1346269`.

```js
assert.strictEqual(fibonacci(32), 1346269)

```

`fibonacci(5)` should return `3`.

```js

```

After `fibonacci(10)` is called, `sequence` should contain 10 numbers.

```js
fibonacci(10)
assert.lengthOf(sequence, 10)
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const sequence = [0, 1];

const fibonacci = (n) => {
    if (n <= sequence.length) {
        return sequence[n - 1];
    }
    for (let i = sequence.length; i <= n; i++) {
        sequence.push(sequence[i - 1] + sequence[i - 2]);
    }
    return sequence[n - 1]
}

```
