---
id: 60eebd07ea685b0e777b5583
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the Registration Form workshop! 

Start by adding the `!DOCTYPE html` at the top of the document so the browser knows what type of document it's reading.

# --hints--

Your code should contain the `DOCTYPE` reference.

```js
assert.match(code, /<!DOCTYPE/gi);
```

You should include a space after the `DOCTYPE` reference.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

You should define the document type to be `html`.

```js
assert.match(code, /html/gi);
```

You should close the `DOCTYPE` with a `>` after the type.

```js
assert.match(code, /html\s*>/gi);
```

Your `DOCTYPE` should be at the beginning of your HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
