---
id: 686fbc534f9befa91a24c51c
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Now that you've grouped the Personal and Business account options using the `name` attribute, it's time to add a `value` attribute to each radio button. This ensures the selected option is properly submitted with the form. Set the value of Personal to `personal` and Business to `business`.

# --hints--

The `first` radio button should have a `value` attribute set to `personal`.

```js
const inputs = document.querySelectorAll('input[type="radio"]');
assert.equal(inputs[0].value, 'personal');
```

The `second` radio button should have a `value` attribute set to `business`.

```js
const inputs = document.querySelectorAll('input[type="radio"]');
assert.equal(inputs[1].value, 'business');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required></label>
      </fieldset>

      --fcc-editable-region--
      <fieldset>
        <label><input type="radio" name="account-type"> Personal</label>
        <label><input type="radio" name="account-type"> Business</label>
      </fieldset>
      --fcc-editable-region--

      <fieldset></fieldset>
      <input type="submit" value="Submit">
    </form>
  </body>
</html>

```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
