---
id: 66faa50000000000000000b9
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--
Pull the door object out of the same object layer. The door is a 16 × 16
rectangle on the right wall — read it with `findObject` again, then create
a Phaser zone with a *static* physics body so overlap detection works:

```js
const doorObj = map.findObject("objects", (o) => o.type === "door");
this.door = this.add.zone(doorObj.x, doorObj.y, 16, 16).setOrigin(0, 0);
this.physics.add.existing(this.door, true);
```

The second argument to `physics.add.existing(obj, true)` is the
`isStatic` flag — passing `true` gives the zone a static body that does
not move and stays where you place it.

# --hints--
You should call `map.findObject("objects", (o) => o.type === "door")` and store it in `doorObj`.

```js
assert.match(code, /const\s+doorObj\s*=\s*map\.findObject\(\s*["']objects["']\s*,\s*\(?\s*o\s*\)?\s*=>\s*o\.type\s*===\s*["']door["']\s*\)/);
```

You should create a 16 × 16 zone at `doorObj.x, doorObj.y` and attach a static body.

```js
assert.match(code, /this\.door\s*=\s*this\.add\.zone\(\s*doorObj\.x\s*,\s*doorObj\.y\s*,\s*16\s*,\s*16\s*\)[\s\S]{0,40}?setOrigin\(\s*0\s*,\s*0\s*\)/);
assert.match(code, /this\.physics\.add\.existing\(\s*this\.door\s*,\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "room1",
      "/curriculum-assets/phaser/tilemaps/room-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "room1" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.player = this.add.rectangle(40, 60, 12, 12, 0x66ddff);
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    walls.setCollisionByProperty({ collides: true });
    this.physics.add.collider(this.player, walls);
    this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.startFollow(this.player);
    this.debugGfx = this.add.graphics().setDepth(50).setVisible(false);
    this.input.keyboard.on("keydown-G", () => {
      this.debugGfx.clear();
      if (!this.debugGfx.visible) {
        this.debugGfx.setVisible(true);
        walls.renderDebug(this.debugGfx, {
          tileColor: null,
          collidingTileColor: new Phaser.Display.Color(255, 80, 80, 80),
          faceColor: new Phaser.Display.Color(255, 255, 80, 100)
        });
      } else {
        this.debugGfx.setVisible(false);
      }
    });
    const spawn = map.findObject("objects", (o) => o.name === "spawn");
    this.player.x = spawn.x;
    this.player.y = spawn.y;
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const speed = 90;
    this.player.body.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.body.setVelocityX(-speed);
    } else if (this.cursors.right.isDown) {
      this.player.body.setVelocityX(speed);
    }
    if (this.cursors.up.isDown) {
      this.player.body.setVelocityY(-speed);
    } else if (this.cursors.down.isDown) {
      this.player.body.setVelocityY(speed);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
