---
id: 66faa50000000000000000af
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--
Drop a placeholder player into the room. Use a `Rectangle` game object —
no sprite sheet required — and give it an arcade physics body so it can
move and collide:

```js
this.player = this.add.rectangle(40, 60, 12, 12, 0x66ddff);
this.physics.add.existing(this.player);
this.player.body.setCollideWorldBounds(true);
this.cursors = this.input.keyboard.createCursorKeys();
```

`physics.add.existing(target)` attaches a *dynamic* arcade body to a game
object you already created. The cursor-key handle goes on `this.cursors`
so `update` can read it.

# --hints--
You should create the player as a `12 × 12` rectangle at `(40, 60)` and attach a physics body.

```js
assert.match(code, /this\.player\s*=\s*this\.add\.rectangle\(\s*40\s*,\s*60\s*,\s*12\s*,\s*12\s*,\s*0x66ddff\s*\)/);
assert.match(code, /this\.physics\.add\.existing\(\s*this\.player\s*\)/);
```

You should call `this.player.body.setCollideWorldBounds(true)`.

```js
assert.match(code, /this\.player\.body\.setCollideWorldBounds\(\s*true\s*\)/);
```

You should grab cursor keys as `this.cursors`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "room1",
      "/curriculum-assets/phaser/tilemaps/room-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "room1" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
