---
id: 66faa50000000000000000aa
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
Switch to `create`. Pull the parsed JSON out of the cache and turn it into
a `Tilemap` object you can render against:

```js
const map = this.make.tilemap({ key: "room1" });
```

Nothing renders yet — `make.tilemap` only builds the data structure. You
will reference `map` for the rest of `create`.

# --hints--
You should declare `const map = this.make.tilemap({ key: "room1" })` in `create`.

```js
assert.match(code, /const\s+map\s*=\s*this\.make\.tilemap\(\s*\{\s*key\s*:\s*["']room1["']\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "room1",
      "/curriculum-assets/phaser/tilemaps/room-01.json"
    );
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
