---
id: 66faa50000000000000000ae
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--
Pin each layer to an explicit depth so you can drop a player or HUD between
them later without re-numbering everything:

```js
floor.setDepth(0);
walls.setDepth(10);
```

Spaced gaps (`0`, `10`, `20`, `30`) leave room for decorations and the
player to slot in cleanly later in the workshop.

# --hints--
You should call `floor.setDepth(0)`.

```js
assert.match(code, /floor\.setDepth\(\s*0\s*\)/);
```

You should call `walls.setDepth(10)`.

```js
assert.match(code, /walls\.setDepth\(\s*10\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "room1",
      "/curriculum-assets/phaser/tilemaps/room-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "room1" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
