---
id: 66faa50000000000000000b0
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Drive the player from `update`. Top-down dungeons use 4-directional
movement: zero the velocity each frame, then re-apply it on whichever
arrow keys are held:

```js
const speed = 90;
this.player.body.setVelocity(0);
if (this.cursors.left.isDown) {
  this.player.body.setVelocityX(-speed);
} else if (this.cursors.right.isDown) {
  this.player.body.setVelocityX(speed);
}
if (this.cursors.up.isDown) {
  this.player.body.setVelocityY(-speed);
} else if (this.cursors.down.isDown) {
  this.player.body.setVelocityY(speed);
}
```

Clearing velocity first means releasing a key stops motion immediately —
no drift to fight with later. The player walks around but currently passes
straight through walls; that is the next fix.

# --hints--
You should zero `this.player.body.setVelocity(0)` each frame.

```js
assert.match(code, /this\.player\.body\.setVelocity\(\s*0\s*\)/);
```

You should drive horizontal velocity from cursor left / right.

```js
assert.match(code, /cursors\.left\.isDown[\s\S]{0,120}?setVelocityX\(/);
assert.match(code, /cursors\.right\.isDown[\s\S]{0,120}?setVelocityX\(/);
```

You should drive vertical velocity from cursor up / down.

```js
assert.match(code, /cursors\.up\.isDown[\s\S]{0,120}?setVelocityY\(/);
assert.match(code, /cursors\.down\.isDown[\s\S]{0,120}?setVelocityY\(/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "room1",
      "/curriculum-assets/phaser/tilemaps/room-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "room1" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.player = this.add.rectangle(40, 60, 12, 12, 0x66ddff);
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.cursors = this.input.keyboard.createCursorKeys();
  }

  update(time, delta) {
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
