---
id: 66faa50000000000000000a9
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Now load the Tiled map JSON. After the existing `load.image` call, queue
the map JSON with `this.load.tilemapTiledJSON`, using the cache key
`"room1"`:

```js
this.load.tilemapTiledJSON(
  "room1",
  "/curriculum-assets/phaser/tilemaps/room-01.json"
);
```

`tilemapTiledJSON` is a dedicated loader that parses the JSON immediately
when it lands in the cache, so by the time `create` runs the parsed map is
ready to wire up.

# --hints--
You should call `this.load.tilemapTiledJSON("room1", "/curriculum-assets/phaser/tilemaps/room-01.json")`.

```js
assert.match(code, /this\.load\.tilemapTiledJSON\(\s*["']room1["']\s*,\s*["']\/curriculum-assets\/phaser\/tilemaps\/room-01\.json["']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
