---
id: 66faa50000000000000000be
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--
Surface the score and the room name as a HUD that follows the camera.
Read the room name from the map's *map-level* `properties` array
(set in Tiled under Map > Map Properties), and render the HUD with a
zero scroll factor so it pins to the screen:

```js
this.roomName = map.properties.find((p) => p.name === "name").value;
this.hud = this.add
  .text(4, 4, "Score: 0  " + this.roomName, {
    fontSize: "8px",
    color: "#ffffff",
    backgroundColor: "rgba(0,0,0,0.55)"
  })
  .setScrollFactor(0)
  .setDepth(60);
```

`map.properties` is the array of Tiled custom properties. Each entry has
`{ name, type, value }`. The room is named `"Stone Chamber"` in the JSON.

# --hints--
You should read the room name from `map.properties` into `this.roomName`.

```js
assert.match(code, /this\.roomName\s*=\s*map\.properties\.find\(\s*\(?\s*p\s*\)?\s*=>\s*p\.name\s*===\s*["']name["']\s*\)\.value/);
```

You should render `this.hud` text pinned to the screen with `setScrollFactor(0)`.

```js
assert.match(code, /this\.hud\s*=\s*this\.add\s*\.\s*text\([\s\S]{0,200}?setScrollFactor\(\s*0\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>First Tile Room with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
  image-rendering: pixelated;
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "tiles",
      "/curriculum-assets/phaser/tilemaps/tileset-dungeon.png"
    );
    this.load.tilemapTiledJSON(
      "room1",
      "/curriculum-assets/phaser/tilemaps/room-01.json"
    );
  }

  create() {
    const map = this.make.tilemap({ key: "room1" });
    const tileset = map.addTilesetImage("dungeon", "tiles");
    const floor = map.createLayer("floor", tileset, 0, 0);
    const walls = map.createLayer("walls", tileset, 0, 0);
    floor.setDepth(0);
    walls.setDepth(10);
    this.player = this.add.rectangle(40, 60, 12, 12, 0x66ddff);
    this.physics.add.existing(this.player);
    this.player.body.setCollideWorldBounds(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    walls.setCollisionByProperty({ collides: true });
    this.physics.add.collider(this.player, walls);
    this.physics.world.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.setBounds(0, 0, map.widthInPixels, map.heightInPixels);
    this.cameras.main.startFollow(this.player);
    this.debugGfx = this.add.graphics().setDepth(50).setVisible(false);
    this.input.keyboard.on("keydown-G", () => {
      this.debugGfx.clear();
      if (!this.debugGfx.visible) {
        this.debugGfx.setVisible(true);
        walls.renderDebug(this.debugGfx, {
          tileColor: null,
          collidingTileColor: new Phaser.Display.Color(255, 80, 80, 80),
          faceColor: new Phaser.Display.Color(255, 255, 80, 100)
        });
      } else {
        this.debugGfx.setVisible(false);
      }
    });
    const spawn = map.findObject("objects", (o) => o.name === "spawn");
    this.player.x = spawn.x;
    this.player.y = spawn.y;
    const doorObj = map.findObject("objects", (o) => o.type === "door");
    this.door = this.add.zone(doorObj.x, doorObj.y, 16, 16).setOrigin(0, 0);
    this.physics.add.existing(this.door, true);
    this.physics.add.overlap(this.player, this.door, () => {
      console.log("next room");
    });
    this.coins = this.physics.add.staticGroup();
    map.filterObjects("objects", (o) => o.type === "coin").forEach((c) => {
      const coin = this.add.rectangle(c.x + 8, c.y + 8, 6, 6, 0xffe066);
      this.physics.add.existing(coin, true);
      this.coins.add(coin);
    });
--fcc-editable-region--
    this.score = 0;
--fcc-editable-region--
    this.physics.add.overlap(this.player, this.coins, (player, coin) => {
      coin.destroy();
      this.score += 10;
      this.hud.setText("Score: " + this.score + "  " + this.roomName);
    });
  }

  update(time, delta) {
    const speed = 90;
    this.player.body.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.body.setVelocityX(-speed);
    } else if (this.cursors.right.isDown) {
      this.player.body.setVelocityX(speed);
    }
    if (this.cursors.up.isDown) {
      this.player.body.setVelocityY(-speed);
    } else if (this.cursors.down.isDown) {
      this.player.body.setVelocityY(speed);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 240,
  height: 160,
  parent: "game-container",
  backgroundColor: "#0b0b18",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
