---
id: cf1111c1c11feddfaeb3bdef
title: JavaScript で 2 つの数値を足し算する
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` は数値データを表す JavaScript のデータ型です。

では、JavaScript で 2 つの数字の足し算をしてみましょう。

JavaScript では `+` 記号を加算演算子として 2 つの数字の間に置いて使用します。

**例:**

```js
const myVar = 5 + 10;
```

`myVar` の値はいま `15` になりました。

# --instructions--

`0` を、和が `20` と等しくなるように変更してください。

# --hints--

`sum` は `20` と等しくなければいけません。

```js
assert(sum === 20);
```

`+` 演算子を使用してください。

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
