---
id: 56533eb9ac21ba0edf2244a9
title: 代入演算子で変数を初期化する
challengeType: 1
videoUrl: 'https://scrimba.com/c/cWJ4Bfb'
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

変数は宣言と同じ行の中で、初期値を使用して<dfn>初期化</dfn>するのが一般的です。

```js
var myVar = 0;
```

上記は `myVar` という名前の新しい変数を作成し、初期値 `0` を代入します。

# --instructions--

変数 `a` を `var` で定義し、値 `9` で初期化してください。

# --hints--

`a` を値 `9` で初期化する必要があります。

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(code));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
