---
id: cf1231c1c11feddfaeb5bdef
title: JavaScript で 2 つの数値を掛け算する
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

数値どうしの掛け算をすることもできます。

JavaScript では 2 つの数値の掛け算に `*` 記号を使用します。

**例**

```js
const myVar = 13 * 13;
```

`myVar` の値は `169` になります。

# --instructions--

`0` を変更して、掛け算の積が `80` と等しくなるようにしてください。

# --hints--

変数 `product` は 80 に等しくなる必要があります。

```js
assert(product === 80);
```

`*` 演算子を使用する必要があります。

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'product = '+z;})(product);
```

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
