---
id: 66ea7fc048a6e26e33c598fe
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now, it is time to generate a random number using the `Math.random()` method.

Create a variable called `randomNum` and assign it the value of the result of calling the `Math.random()` method.

Then, log the `randomNum` variable to the console.

Try adding a space in the code to re-run the bot and see different random numbers that are generated.

# --hints--

You should have a variable called `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

You should assign the result of calling the `Math.random()` method to the `randomNum` variable.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

You should log the `randomNum` variable to the console.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
