---
id: 66eb488c0129e192542f0865
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

In the previous lecture videos, you learned how to work with the `Math.ceil()` method like this:

```js
const price = 10.01;
Math.ceil(price); // 11
```

The `Math.ceil()` method rounds the value up to the nearest whole integer.

Create a variable called `numRoundedUp` and assign it the result of rounding the floating point number `3.2` up to the nearest whole integer.

Then, log the `numRoundedUp` variable to the console to see the result.

# --hints--

You should have a `numRoundedUp` variable.

```js
assert.isNotNull(numRoundedUp);
```

The `numRoundedUp` variable should be a number.

```js
assert.isNumber(numRoundedUp);
```

Your `numRoundedUp` variable should equal the result of rounding the floating point number `3.2` down to the nearest whole integer.

```js
assert.equal(numRoundedUp, 4);
```

You should not hardcode the value of `4` for the `numRoundedUp` variable.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

You should log the `numRoundedUp` variable to the console.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
