---
id: 66eb480d996af5918d7c9c86
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Now, it is time to review how the `Math.ceil()` method works.

Start by logging the string `"The Math.ceil() method rounds the value up to the nearest whole integer."` to the console.

# --hints--

You should log the message `"The Math.ceil() method rounds the value up to the nearest whole integer."`

```js
assert.match(code, /console\.log\((['"`])The\s+Math\.ceil\(\)\s+method\s+rounds\s+the\s+value\s+up\s+to\s+the\s+nearest\s+whole\s+integer\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);
--fcc-editable-region--

--fcc-editable-region--
```
