---
id: 65578d4fc3afc3b8f554c882
title: Крок 36
challengeType: 20
dashedName: step-36
---

# --description--

`min()` також приймає аргумент лише за ключовим словом. Передавши функцію як додатковий аргумент до `min()`, можна змінити те, як порівнюються елементи списку.

Результатом рядка, який ви щойно написали в попередньому кроці, є вузол, який стоїть першим за алфавітним порядком. Замість нього потрібно вибрати невідвіданий вузол, який має найменшу відстань від початкового вузла.

Передайте `key=distances.get` як другий аргумент до виклику `min()`. Таким чином порівняння відбуватиметься залежно від значення кожного елемента списку `unvisited` в словнику `distances`.

# --hints--

Передайте `key=distances.get` як другий аргумент до виклику `min()`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
