---
id: 65577791ad8c26a7705e2919
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

За допомогою розуміння словника можна створити словник, починаючи з наявного словника:

```py
{key: val for key in dict}
```

Вам потрібно відстежувати шляхи між початковим вузлом та всіма іншими вузлами.

Створіть змінну `paths` після змінної `distances` та призначте до неї словник з всіма ключами з `graph`. Призначте порожній список до кожного ключа та використайте розуміння словника, щоб створити словник.

# --hints--

Ви повинні мати змінну `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

Змінна `paths` має використати синтаксис розуміння словника, щоб призначити порожній список до кожного вузла в графі.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}

    print(f'Unvisited: {unvisited}\nDistances: {distances}')

shortest_path(my_graph, 'A')
--fcc-editable-region--
```
