---
id: 65578d0f6c78a0b868a43b9c
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

Перше, що потрібно зробити в межах циклу `while` — визначити поточний вузол, який потрібно відвідати. Для цього можна скористатися функцією `min()`. Вона повертає найменший елемент з ітерованого об’єкта, переданого як аргумент.

Видаліть `pass`, потім створіть змінну під назвою `current` і призначте до неї `min(unvisited)`.

# --hints--

Створіть змінну `current` в циклі `while`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=/ms));
  }
})
```

Призначте `min(unvisited)` до змінної `current`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
