---
id: 6557746aad2844a0cd864e12
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Алгоритм досліджує граф, але в той самий час він повинен відстежувати найкоротшу відстань між початковим вузлом та іншими вузлами, яка відома на цю мить.

Створіть нову змінну під назвою `distances` перед циклом `for` та призначте до неї порожній словник.

# --hints--

Ви повинні мати змінну `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=.*(?=^\1for.*:)/ms));
  }
})
```

Змінна `distances` має бути порожнім словником.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)distances\s*=\s*\{\s*\}.*(?=^\1for.*:)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
