---
id: 6557924d47c325bf27afbe51
title: Крок 48
challengeType: 20
dashedName: step-48
---

# --description--

Алгоритм завершено, але вихідні дані можна покращити. Крім того, ви можете надати функції додатковий аргумент, щоб повернути лише шлях між двома вузлами.

Додайте `target` як третій параметр до оголошення функції та надайте йому значення за замовчуванням у вигляді порожнього рядка.

# --hints--

Функція має приймати три параметри: `graph`, `start` та `target`. Порядок має значення.

```js
({ test: () => assert.match(code, /^def\s+shortest_path\s*\(\s*graph\s*,\s*start\s*,\s*target\s*=?\s*.*\s*\)\s*:/m) })
```

Параметр `target` повинен мати значення за замовчуванням у вигляді порожнього рядка.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    foo = __locals.get("shortest_path")
    sig = str(inspect.signature(foo))
    sig == "(graph, start, target='')"
  `))
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
--fcc-editable-region--
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)

    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

shortest_path(my_graph, 'A')

```
