---
id: 65578cee7f2cb8b80127cce2
title: Крок 34
challengeType: 20
dashedName: step-34
---

# --description--

Перед викликом `print` створіть цикл `while`, який виконується, поки `unvisited` не порожній. Використайте ключове слово `pass`, щоб заповнити тіло циклу.

# --hints--

Ви повинні мати цикл `while`, який виконується, поки `unvisited` не порожній. Не забудьте ключове слово `pass`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1pass(?=\s*^\1print)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')
--fcc-editable-region--
```
