---
id: 655773f8b8b5db9fc6d0ae76
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Щоб відстежувати відвідані вузли, вам потрібен список усіх вузлів у графі. Після відвідування вузла його буде видалено з цього списку.

Тепер замініть ключове слово `pass` на змінну під назвою `unvisited` та призначте до неї порожній список.

# --hints--

Ви повинні мати змінну під назвою `unvisited`.

```js
({ test: () =>
  {
    const shortest_path = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest_path;    
    assert(function_body.match(/unvisited\s*=/));
  }
})
```

Змінна `unvisited` має бути порожнім списком.

```js
({ test: () =>
  {
    const shortest_path = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest_path;    
    assert(function_body.match(/unvisited\s*=\s*\[\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
