---
id: 6557712d77ce2d9bd7e63afd
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Додайте ще один вузол, приєднаний до `B`, до графа та назвіть його `C`.

Змініть наявний словник, щоб представити це розташування. Використайте список, щоб представити декілька з’єднань вузла `B`.

# --hints--

Словник повинен мати 3 ключі під назвами `A`, `B` та `C`.

```js
({ test: () => assert(__pyodide.runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Значенням `my_graph["A"]` має бути вузол `B`.

```js
({ test: () => assert(__pyodide.runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph["B"]` має бути списком.

```js
({ test: () => assert(__pyodide.runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Значення `my_graph["B"]` має бути списком, що містить два інші вузли.

```js
({ test: () => assert(__pyodide.runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Значення `my_graph["C"]` має бути приєднаним вузлом.

```js
({ test: () => assert(__pyodide.runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
