---
id: 655791e6cf5e03be3de73451
title: Крок 45
challengeType: 20
dashedName: step-45
---

# --description--

Якщо ви спробуєте розкоментувати виклик функції, у вас не вийде. Вам потрібно виправити пару помилок. Перша помилка трапляється через те, що ви намагаєтесь використати у вкладеній інструкції `if` те, чого може не існувати в словнику `paths`. Отже, перед доступом до `paths[node][-1]` потрібно переконатися, що `paths[node]` не порожній.

Додайте додаткову умову до вкладеної інструкції `if`, щоб переконатися, що `paths[node]` не порожній, перш ніж отримати доступ до `paths[node][-1]`.

# --hints--

Додайте `paths[node]` як першу умову до вкладеної інструкції `if`. Використайте оператор `and`, щоб об’єднати умови.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)

    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
