---
id: 6557712d77ce2d9bd7e63afd
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Додайте ще один вузол, приєднаний до `B`, до графа та назвіть його `C`.

Змініть наявний словник, щоб представити таке розташування: додайте ще один ключ `'C'` до `my_graph` та надайте йому значення рядка `'B'`.

Також змініть значення наявного ключа `'B'` на список `['A', 'C']`, щоб представити декілька зв’язків вузла `'B'`.

# --hints--

Словник повинен мати 3 ключі під назвами `'A'`, `'B'` та `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Значенням `my_graph['A']` має бути рядок `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` має бути списком.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Значенням `my_graph['B']` має бути список, що містить `'A'` та `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Значенням `my_graph['C']` має бути рядок `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
