---
id: 65576ff7888f9e96f52a4be1
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Словники зберігають дані у вигляді пар *ключ*-*значення*. Ключ відокремлюється від відповідного значення двокрапкою. Кожна пара ключ-значення відокремлюється від іншої пари комою:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Додайте нову пару ключ-значення до словника. Use the string `'species'` as the key, and the string `'guinea pig'` as the value.

# --hints--

You should have a `'species'` key with the value `'guinea pig'` inside your `copper` dictionary.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Your `copper` dictionary should have a single key-value pair.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
