---
id: 6557709b0aee699a6a00528c
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Наприклад, граф можна використати, щоб представити в просторі точки `A` та `B`, з’єднані шляхом. Такий граф буде складатися з двох вузлів, з’єднаних ребром.

Замініть наявний ключ `'species'` на рядки `'A'`. Потім замініть відповідне значення на рядок `'B'`, щоб продемонструвати зв’язок між вузлами `'A'` та `'B'`.

# --hints--

Словник повинен мати ключ `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ключ `'A'` повинен мати значення `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Словник повинен мати два ключі.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
