---
id: 655790d113d14dbb727eaf41
title: Крок 40
challengeType: 20
dashedName: step-40
---

# --description--

Як тільки в словнику `distances` встановлено відстань до вузла, потрібно відстежувати шлях до цього вузла. Якщо відстань для вузла в обробленому кортежі оновлено, останнім елементом на шляху є сам вузол.

Вкладіть ще одну інструкцію `if` в межах умовної інструкції, яка запускатиметься, якщо останній елемент `paths[node]` дорівнює `node`.

# --hints--

Створіть вкладену інструкцію `if`, яка перевірятиме, чи `paths[node][-1]` дорівнює `node`. Не забудьте використати `pass`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1(\s{4})if\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\2\2pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')
```
