---
id: 655773f8b8b5db9fc6d0ae76
title: Крок 25
challengeType: 20
dashedName: step-25
---

# --description--

Щоб відстежувати відвідані вузли, вам потрібен список усіх вузлів у графі. Після відвідування вузла його буде видалено з цього списку.

Тепер замініть ключове слово `pass` на змінну під назвою `unvisited` та призначте до неї порожній список.

# --hints--

Ви повинні мати змінну під назвою `unvisited` в межах функції `shortest_path`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

Призначте порожній список до змінної `unvisited`. Не забудьте видалити `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
