---
id: 65577a17564ce8a8e06c1460
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Розуміння словників також підтримує умовний синтаксис `if`/`else`:

```py
{key: val_1 if condition else val_2 for key in dict}
```

Використайте розуміння словника, щоб створити словник на основі `graph` та призначити його до змінної `distances`. Надайте ключу значення 0, якщо вузол дорівнює початковому вузлу; в іншому випадку надайте нескінченне значення. Щоб досягти останнього, використайте `float('inf')`.

# --hints--

Використайте синтаксис розуміння словника, щоб надати значення змінній `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')

shortest_path(my_graph, 'A')
--fcc-editable-region--
```
