---
id: 6557743527cb92a06417ea97
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Створіть цикл `for`, щоб ітерувати над графом та додайте всі вузли до списку `unvisited`.

# --hints--

Створіть цикл `for`, щоб ітерувати над `graph` в межах функції `shortest_path`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for\s+(\w+)\s+in\s+graph\s*:/m));
  }
})
```

Додайте всі вузли до `unvisited` в межах циклу `for`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for\s+(\w+)\s+in\s+graph\s*:\s*^\1\1unvisited\.append\s*\(\s*\2\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
--fcc-editable-region--
```
