---
id: 655774955b097ea14897db12
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Відстань від початкового вузла дорівнює нулю, оскільки алгоритм починає оцінювання саме звідти.

Після того, як ви додали `node` до `unvisited` в циклі, створіть інструкцію `if`, яка виконуватиметься, якщо вузол дорівнює початковому вузлу. Потім призначте `0` до цього вузла в межах словника `distances`.

# --hints--

Створіть інструкцію `if`, яка виконується, коли `node` дорівнює `start`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for.*:.*^\1\1if\s+node\s*==\s*start\s*:/ms));
  }
})
```

Призначте `0` до вузла в словнику `distances` в межах нової інструкції `if`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)for.*:.*^\1\1if\s+node\s*==\s*start\s*:\s*^\1\1\1distances\s*\[\s*node\s*\]\s*=\s*0/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
