---
id: 6557709b0aee699a6a00528c
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер замініть наявні ключі на рядки `A` та `B` — по одному для кожного вузла. Потім замініть кожне значення на рядок, який представляє вузол, приєднаний до ключа.

# --hints--

Словник повинен мати ключ `A`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ключ `A` повинен мати значення `B`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Словник повинен мати ключ `B`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

Ключ `B` повинен мати значення `A`.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Словник повинен мати два ключі.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
