---
id: 655773b0591c5f9f4045883e
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Алгоритм почнеться з вказаного вузла. Потім він досліджуватиме граф, щоб знайти найкоротший шлях між початковим вузлом, або *джерелом*, та всіма іншими вузлами.

Для цього функції потрібні два параметри: `graph` та `start`. Додайте їх до оголошення функції.

# --hints--

Функція має приймати `graph` та `start` як параметри (в такому ж порядку).

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    f = __locals.get("shortest_path")
    sig = str(inspect.signature(f))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
