---
id: 655771d889132f9ccd341060
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

Граф називається *зваженим*, якщо його ребра пов’язані з вагою, яка представляє відстань, час або інше кількісне значення.

У цьому випадку вагою буде відстань між вузлами або точками в просторі. Щоб представити зважений граф, можна змінити словник, використовуючи список кортежів для кожного значення.

Першим елементом в кортежі буде приєднаний вузол, а другим елементом буде ціле число, яке вказує відстань.

Modify `my_graph['A']` into a list of tuples, considering that the `A-B` distance is `3` and the `A-D` distance is `1`.

# --hints--

`my_graph["A"]` should be a list containing the tuples `('B', 3)` and `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` should have 4 keys named `'A'`, `'B'`, `'C'`, and `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
