---
id: 65576ff7888f9e96f52a4be1
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Словники зберігають дані у вигляді пар *ключ*-*значення*. Ключ відокремлюється від відповідного значення двокрапкою. Кожна пара ключ-значення відокремлюється від іншої пари комою:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Додайте нову пару ключ-значення до словника. Використайте рядок `species` як ключ, а рядок `guinea pig` як значення.

# --hints--

Додайте нову пару ключ-значення до словника `copper`.

```js
({ test: () => assert(__pyodide.runPython(`
    copper = __locals.get("copper")
    len(copper) == 1
  `))
})
```

Ви повинні мати ключ `species` зі значенням `guinea pig` в межах словника `copper`.

```js
({ test: () => assert(__pyodide.runPython(`
    copper = __locals.get("copper")
    copper == {"species": "guinea pig"}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
