---
id: 65577236b056379d5dbc7000
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Зараз ви почнете розробляти алгоритм для обчислення найкоротшого шляху між вузлами в новому графі.

Оголосіть порожню функцію під назвою `shortest_path` та не забудьте ключове слово `pass`.

# --hints--

Ви повинні мати функцію `shortest_path`.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    f = __locals.get("shortest_path")
    inspect.isfunction(f)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
