---
id: 6559d70c5161b16ff1d6530d
title: Крок 49
challengeType: 20
dashedName: step-49
---

# --description--

Python забезпечує стислий спосіб запису умовних інструкцій `if`/`else` за допомогою тернарного синтаксису:

```py
val_1 if condition else val_2
```

Наведений вище вираз обчислюється як `val_1`, якщо `condition` має значення `true`, а в іншому випадку — як `val_2`.

Видаліть виклик `print` та створіть зміну під назвою `targets_to_print` після циклу `while`. Використайте тернарний синтаксис, щоб призначити до неї `[target]`, якщо `target` оцінюється як `true`, або `graph`, якщо навпаки.


# --hints--

Видаліть виклик `print`.

```js
({ test: () => assert.isFalse( /print\s*\(\s*f("|')Unvisited:\s*\{\s*unvisited\s*\}\\nDistances:\s\{\s*distances\s*\}\\nPaths:\s\{\s*paths\s*\}\1\s*\)/.test(code)) })
```

Створіть змінну під назвою `targets_to_print` після циклу `while`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=/ms));
  }
})
```

Використайте тернарний синтаксис, щоб призначити `[target]` до змінної `targets_to_print`, якщо `target` оцінюється як `true`, або `graph`, якщо навпаки.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\.remove\(\s*current\s*\).*^\s{4}targets_to_print\s*=\s*\[\s*target\s*\]\s+if\s+target\s+else\s+graph/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start, target = ''):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)

    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current][:]
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

shortest_path(my_graph, 'A')
--fcc-editable-region--
```
