---
id: 65661b72d6745ebec6a96923
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Таким самим чином змініть два інших списки, враховуючи, що відстань `C-D` становить `7`.

# --hints--

`my_graph["C"]` має бути списком кортежів.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    type(graph["C"]) is list and all(type(i) is tuple for i in graph["C"])
  `))
})
```

`my_graph["C"]` має бути списком кортежів, де першим елементом кортежу є приєднаний вузол, а другим — відстань.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    tuples = [("B", 4), ("D", 7)]
    len(graph["C"]) == 2 and all(t in graph["C"] for t in tuples) 
  `))
})
```

`my_graph["D"]` має бути списком кортежів.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    type(graph["D"]) is list and all(type(i) is tuple for i in graph["D"])
  `))
})
```

`my_graph["D"]` має бути списком кортежів, де першим елементом кортежу є приєднаний вузол, а другим — відстань.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    tuples = [("A", 1), ("C", 7)]
    len(graph["D"]) == 2 and all(t in graph["D"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
