---
id: 65661b72d6745ebec6a96923
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Таким самим чином змініть два інших списки, враховуючи, що відстань `C-D` становить `7`.

# --hints--

`my_graph['C']` має бути списком, що містить кортежі `('B', 4)` та `('D', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 4), ("D", 7)]
    len(my_graph["C"]) == 2 and all(t in my_graph["C"] for t in tuples) 
  `))
})
```

`my_graph['D']` має бути списком, що містить кортежі `('A', 1)` та `('C', 7)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 1), ("C", 7)]
    len(my_graph["D"]) == 2 and all(t in my_graph["D"] for t in tuples)
  `))
})
```

`my_graph` повинен мати 4 ключі під назвами `'A'`, `'B'`, `'C'` та `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

`my_graph['A']` має бути списком, що містить кортежі `('B', 3)` та `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph['B']` має бути списком, що містить кортежі `('A', 3)` та `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}
--fcc-editable-region--
```
