---
id: 6566195b0a021bb660b2b4b1
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Тепер змініть `my_graph["B"]` на список кортежів. Відстань `B-C` становить `4`.

# --hints--

`my_graph["B"]` має бути списком кортежів.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    type(graph["B"]) is list and all(type(i) is tuple for i in graph["B"])
  `))
})
```

`my_graph["B"]` має бути списком кортежів, де першим елементом кортежу є приєднаний вузол, а другим — відстань.

```js
({ test: () => assert(__pyodide.runPython(`
    graph = __locals.get("my_graph")
    tuples = [("A", 3), ("C", 4)]
    len(graph["B"]) == 2 and all(t in graph["B"] for t in tuples)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
