---
id: 6578b57361f2f132a02e2a18
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ви можете отримати доступ до даних в словнику через ключі:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

After your dictionary, follow the example above to access the `'species'` key of `copper` and print the result.

# --hints--

You should not modify the assignment of your dictionary.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

You should use `copper['species']` to access the value of the `'species'` key.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Викличте `print()`, передавши `copper['species']` як аргумент.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
