---
id: 65789506b30453080f77470c
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Ви вже стикалися з різними типами даних:

- незмінні типи даних (наприклад, цілі числа, рядки, кортежі та булеві значення);
- Mutable data types, such as lists, and dictionaries.

A dictionary is identified by a pair of curly braces, `{}`.

Для початку створіть змінну під назвою `copper` та призначте до неї порожній словник за допомогою пари фігурних дужок, так само як ви б створили порожній список за допомогою пари квадратних дужок.

# --hints--

Ви повинні мати змінну під назвою `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Змінна `copper` повинна мати значення порожнього словника. Використайте для цього пару фігурних дужок.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
