---
id: 6578b57361f2f132a02e2a18
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ви можете отримати доступ до даних в словнику через ключі:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack`
}

my_dict['name'] # 'Michael'
```

Дотримайтесь прикладу вище після словника, щоб отримати доступ до ключа `species` в `copper` та надрукуйте результат.

# --hints--

Не змінюйте словник.

```js
({ test: () => assert(__pyodide.runPython(`
    copper = __locals.get("copper")
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Використайте `copper['species']`, щоб отримати доступ до значення ключа `species`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Викличте `print()`, передавши `copper['species']` як аргумент.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
