---
id: 6578b13757611e2825beb8a5
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Ключі в межах словника мають бути унікальними, і вони можуть бути лише незмінними типами даних. Це означає, що ви не можете використовувати список або інший словник як ключі.

Додайте ще один ключ `age` до словника та надайте йому значення цілого числа `2`.

# --hints--

Додайте нову пару ключ-значення до словника `copper`.

```js
({ test: () => assert(__pyodide.runPython(`
    copper = __locals.get("copper")
    len(copper) == 2
  `))
})
```

Ви повинні мати ключ `age` зі значенням `2` одразу після `'species': 'guinea pig'` в межах словника `copper`. Не забудьте про кому.

```js
({ test: () => assert(__pyodide.runPython(`
    copper = __locals.get("copper")
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
