---
id: 65797670e0c0d016f17e7660
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер надрукуйте `copper` в нижній частині коду.

# --hints--

Надрукуйте `copper` у нижній частині коду.

```js
assert.match(code, /print\s*\(\s*copper\s*\)/)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

--fcc-editable-region--
```
