---
id: 6579717f0920131304286804
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Щоб додати нову пару ключ-значення після оголошення словника, можна вказати ключ (так само як при отриманні наявного ключа) і встановити значення нового ключа за допомогою оператора присвоєння:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Видаліть виклик `print()`. Then, after the `copper` declaration, add the key `'food'` to your dictionary and set its value to `'hay'`.

# --hints--

Код не повинен містити `print(copper['age'])`.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

You should add the key `'food'` to `copper` after declaring the dictionary.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

You should set `copper['food']` to `'hay'` after declaring the dictionary.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
