---
id: 6579ca0923cfa7162089d2f0
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Такий самий синтаксис можна використовувати, щоб змінити значення наявного ключа.

Отримайте доступ до ключа `species` одразу перед викликом `print()` та перепризначте його значення до `Cavia porcellus`.

# --hints--

Перепризначте `copper['species']` до рядка `Cavia porcellus` перед викликом `print()`.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2.*^print\s*\(\s*copper\s*\)/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
