---
id: 6579cbab9825b8170974c69a
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Ви можете видалити пару ключ-значення зі словника за допомогою ключового слова `del`. Синтаксис такий:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack`
}

del my_dict['occupation']
```

Використайте ключове слово `del` одразу перед викликом `print()`, щоб видалити ключ `age` та його значення з `copper`.

# --hints--

Використайте ключове слово `del`, щоб видалити `copper['age']` перед викликом `print()`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^print\s*\(\s*copper\s*\)/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
