---
id: 6579cd5f6dd62c189e53ddbb
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Ви пригадали основні аспекти словників, тому можете перейти до створення алгоритму найкоротшого шляху.

Видаліть кожен рядок коду після оголошення словника `copper`.

# --hints--

Видаліть рядки коду після оголошення словника.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
