---
id: 6579cd5f6dd62c189e53ddbb
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер ви розумієте основні аспекти словників, тому можете перейти до створення алгоритму найкоротшого шляху.

Видаліть кожен рядок коду після оголошення словника `copper`.

# --hints--

Видаліть рядки після оголошення словника.

```js
({ test: () => assert.isFalse( /copper\s*\[.*?\]|del|print\s*\(.*?\)/ms.test(code)) })
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']
print(copper)
--fcc-editable-region--
```
