---
id: 6579dd49fa8a8e1fd06b85a9
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Графи — це структури даних, що представляють зв’язки між парами елементів. Ці елементи, які називаються *вузлами*, можуть бути реальними об’єктами, сутностями, точками в просторі тощо. З’єднання між вузлами називаються *ребрами*.

Наприклад, граф можна використати, щоб представити в просторі точки `A` та `B`, з’єднані шляхом. Такий граф буде складатися з двох вузлів, з’єднаних ребром.

Перейменуйте словник `copper` на `my_graph`. Це представлятиме граф для тестування алгоритму.

# --hints--

Перейменуйте словник `copper` на `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Змінна `my_graph` має бути словником.

```js
({ test: () => assert(__pyodide.runPython(`
    my_graph = __locals.get("my_graph")
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
