---
id: 6579717f0920131304286804
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Щоб додати нову пару ключ-значення після оголошення словника, можна вказати ключ (так само як при отриманні наявного ключа) і встановити значення нового ключа за допомогою оператора присвоєння:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack`
}

my_dict['country'] = 'Canada'
```

Видаліть виклик `print()`. Потім, після оголошення `copper`, додайте ключ `food` до словника та встановіть його значення на `hay`.

# --hints--

Код не повинен містити `print(copper['age'])`.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Додайте ключ `food` до `copper` після оголошення словника.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Встановіть `copper['food']` на `hay` після оголошення словника.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
