---
id: 65796fac81f983127558f3f4
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер змініть наявний виклик `print()`, щоб надрукувати значення ключа `age`.


# --hints--

Використайте `copper['age']`, щоб отримати доступ до значення ключа `age`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Викличте `print()`, передавши `copper['age']` як аргумент.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
