---
id: 6579dd49fa8a8e1fd06b85a9
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Графи — це структури даних, що представляють зв’язки між парами елементів. Ці елементи, які називаються *вузлами*, можуть бути реальними об’єктами, сутностями, точками в просторі тощо. З’єднання між вузлами називаються *ребрами*.

Here's a visual representation of a graph:

<img class="img-responsive center-block" alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px;" />


Перейменуйте словник `copper` на `my_graph`. Це представлятиме граф для тестування алгоритму.

# --hints--

Перейменуйте словник `copper` на `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

Змінна `my_graph` має бути словником.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
