---
id: 65affe602a2a828b0382a058
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Замініть ключ `'age'` на рядок `'B'` та встановіть його значення на рядок `'A'`, щоб продемонструвати зв’язок між вузлами в обох напрямках.

# --hints--

Словник повинен мати ключ `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

Ключ `'B'` повинен мати значення `'A'`.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Словник повинен мати ключ `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

Ключ `'A'` повинен мати значення `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Словник повинен мати два ключі.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
