---
id: 65b7cf140d34058d7ea8935f
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Якщо ви хочете ітерувати над значеннями ключів словника, можна використати метод `.values()`.

Змініть цикл `for`, щоб він ітерував над `copper.values()`, а не `copper` та гляньте на результат.

# --hints--

Змініть цикл `for`, щоб ітерувати над `copper.values()`, а не `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
