---
id: 65b7d3319ebec69b983fb91b
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Як можна побачити з виводу, `.items()` створює структури даних, які зберігають пари ключ-значення в окремих кортежах. Щоб ітерувати над елементами в цих кортежах, ви можете додати другу змінну циклу:

```py
for i, j in dict.items():
    print(i, j)
```

Змініть цикл `for`, щоб він приймав дві змінні циклу та надрукував їх в межах тіла циклу.

# --hints--

Змініть цикл `for`, щоб він використовував дві змінні.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Надрукуйте обидві змінні циклу в межах циклу `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
