---
id: 6894c1b715b8d18a18ef3203
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you'll get some practice with Python's classes and objects by building a musical instrument inventory.

To define a class in Python, use the `class` keyword followed by the class name and a colon.

```python
class ClassName:
    # class contents go here
```

Remember that in Python, class names typically use PascalCase, which means each word starts with an uppercase letter with no spaces between words.

Start by creating a class called `MusicalInstrument`. Add `pass` as the body of the class for now.

# --hints--

You should create a class called `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
