---
id: 60f030d388cb74067cf291c3
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

見出しの下で、段落要素内に次のテキストを使用してユーザーに登録を促しましょう。

```md
Please fill out this form with the required information
```

# --hints--

`body` 内に `p` 要素を追加する必要があります。

```js
assert.exists(document.querySelector('body > p'));
```

`h1` の下に `p` 要素を追加する必要があります。

```js
assert.exists(document.querySelector('h1 + p'));
```

`p` 要素に `Please fill out this form with the required information` というテキストを入れる必要があります。

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
