---
id: 60f0286404aefb0562a4fdf9
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`head` 内に、`title` 要素と `meta` 要素を追加してください。 プロジェクトのタイトルは `Registration Form` とし、`meta` 要素には `charset` 属性と値 `UTF-8` を設定してください。

# --hints--

コードには `title` 要素が 1 つ必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 要素は `head` 要素内にある必要があります。

```js
assert.exists(document.querySelector('head > title'));
```

プロジェクトには `Registration Form` というタイトルが必要です。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

タイトルについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
