---
id: 60f02e7361b68405e27b62a5
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`body` 内に、`Registration Form` というテキストを持つ `h1` を追加して、コンテンツの見出しを提供しましょう。

# --hints--

`body` 内に `h1` を追加する必要があります。

```js
assert.exists(document.querySelector('body > h1'));
```

`h1` に `Registration Form` というテキストを入れる必要があります。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
