---
id: 60f1a9cbd23023082e149fee
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

良くなりました。 では、`body` の `background-color` を `#1b1b32` に変更して、目に優しい背景にしましょう。 そして、テキストが見えるようにするため、`color` を `#f5f6f7` に変更してください。

# --hints--

`body` 要素セレクター内に `background-color` を追加する必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

`background-color` の値を `#1b1b32` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

`body` 要素セレクター内に `color` を追加する必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

`color` の値を `#f5f6f7` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
