---
id: 60f1922fcbd2410527b3bd89
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

プロジェクトの見た目を整えるために、CSS を追加しましょう。 まず、`body` の `width` を `100%` に、`height` を `100vh` に設定してください。

# --hints--

`body` 要素セレクターを使用する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` の `width` を `100%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

`body` の `height` を `100vh` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
