---
id: 60f1a5e2d2c23707a4f9a660
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

では `body` について、一部のブラウザで追加されるデフォルトの `margin` を `0` に設定して、水平方向のスクロールバーを取り除きましょう。

# --hints--

`body` 要素セレクター内に `margin` を追加する必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

`margin` の値を `0` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
