---
id: 60f5dc35c07ac1078f140916
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

`label` 要素に以下のテキストを追加してください:

- `Enter Your First Name:`
- `Enter Your Last Name:`
- `Enter Your Email:`
- `Create a New Password:`

# --hints--

1 つ目の `label` には、テキスト `Enter Your First Name:` が必要です。

```js
assert.match(document.querySelector('label')?.innerHTML, /Enter Your First Name:/i);
```

2 つ目の `label` には、テキスト `Enter Your Last Name:` が必要です。

```js
assert.match(document.querySelector('fieldset > label:nth-child(2)')?.innerHTML, /Enter Your Last Name:/i);
```

3 つ目の `label` には、テキスト `Enter Your Email:` が必要です。

```js
assert.match(document.querySelector('fieldset > label:nth-child(3)')?.innerHTML, /Enter Your Email:/i);
```

4 つ目の `label` には、テキスト `Create a New Password:` が必要です。

```js
assert.match(document.querySelector('fieldset > label:nth-child(4)')?.innerHTML, /Create a New Password:/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label></label>
        <label></label>
        <label></label>
        <label></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
