---
id: 60f5d2776c854e069560fbe6
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

1 つ目の `fieldset` には、名前、メールアドレス、パスワードを入れます。 まずは 1 つ目の `fieldset` に `label` 要素を 4 つ追加してください。

# --hints--

`fieldset` 要素は全部で 3 つ必要です。

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

すべての `fieldset` 要素が `form` 要素の直接の子要素でなければなりません。

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

4 つの `label` 要素はすべて、最初の `fieldset` 要素の直接の子要素でなければなりません。 各 `label` 要素が終了タグで閉じられているか確認してください。

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

最後の 2 つの `fieldset` 要素は空でなければなりません。

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

各 `label` 要素が終了タグで閉じられているか確認してください。

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
