---
id: 60f80e0081e0f2052ae5b505
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

期待するデータの種類をブラウザーに伝えるために、フォーム要素の `type` 属性を指定することが重要です。 `type` が指定されていない場合、ブラウザーはデフォルトで `text` を使用します。

最初の 2 つの `input` 要素の `type` 属性を `text` に設定し、3 つ目の `type` 属性は `email`、4 つ目の `type` 属性は `password` に設定してください。

`email` タイプは、`@` とドメイン内の `.` が含まれるメールアドレスのみを受け付けます。 `password` タイプは入力内容を隠し、また、サイトが HTTPS を使用していない場合に警告を表示します。

# --hints--

1 つ目の `input` 要素の `type` 属性を `text` に設定する必要があります。

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

2 つ目の `input` 要素の `type` 属性を `text` に設定する必要があります。

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

3 つ目の `input` 要素の `type` 属性を `email` に設定する必要があります。

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

4 つ目の `input` 要素の `type` 属性を `password` に設定する必要があります。

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
