---
id: 60f805f813eaf2049bc2ceea
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

各 `label` 内に、`input` 要素を 1 つずつネストしてください。 各 `input` は `label` のテキストの後に追加し、コロンの後にスペースを入れるようにしてください。

# --hints--

`fieldset` 要素に `input` 要素を 4 つ追加する必要があります。

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

`input` 要素は `label` 要素内にネストする必要があります。

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

`label` のテキスト `Enter Your First Name:` の後に 1 つ目の `input` を追加し、コロンの後にはスペースを入れてください。

```js
assert.equal(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), 'Enter Your First Name: <input>');
```

`label` のテキスト `Enter Your Last Name:` の後に 2 つ目の `input` を追加し、コロンの後にはスペースを入れてください。

```js
assert.equal(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), 'Enter Your Last Name: <input>');
```

`label` のテキスト `Enter Your Email:` の後に 3 つ目の `input` を追加し、コロンの後にはスペースを入れてください。

```js
assert.equal(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), 'Enter Your Email: <input>');
```

`label` のテキスト `Create a New Password:` の後に 4 つ目の `input` を追加し、コロンの後にはスペースを入れてください。

```js
assert.equal(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), 'Create a New Password: <input>');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
