---
id: 60f81167d0d4910809f88945
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

`type` が `submit` に設定された最初の `input` 要素は、自動的に、一番近い親要素の `form` を送信するよう設定されます。

フォームの送信を処理するため、最後の `fieldset` 要素の後に `input` 要素を追加し、`type` 属性を `submit` に、`value` 属性を `Submit` に設定してください。

# --hints--

最後の `fieldset` 要素の後に `input` 要素を追加する必要があります。

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

`input` 要素の `type` 属性を `submit` に設定する必要があります。

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

`input` 要素の `value` 属性を `Submit` に設定する必要があります。

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>

--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
