---
id: 60f803d5241e6a0433a523a1
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--
`rem` という単位は「root (ルート要素) の `em`」を表し、`html` 要素のフォントサイズに比例します。

`label` 要素はデフォルトではインラインなので、すべて同じ行に並んで表示されてテキストが読み辛くなります。 別の行に表示させるには、`label` 要素に `display: block` を追加してください。また、それぞれ離れて表示させるために、`margin` を `0.5rem 0` に設定してください。

# --hints--

`label` セレクターを使用する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

`display` プロパティの値を `block` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

`margin` プロパティの値を `0.5rem 0` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
