---
id: 60f84ec41116b209c280ba91
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

`type="password"` を設定した場合、`pattern` 属性を使用して、有効なパスワードと見なされるためにマッチすべき正規表現を定義できます。

パスワードの `input` 要素に `pattern` 属性を追加して、入力値が `[a-z0-5]{8,}` にマッチしなければならないように設定してください。

上記は、8 文字以上の小文字のアルファベットまたは `0` ～ `5` の数字にマッチする正規表現です。 次に、`minlength` 属性を削除して入力欄の動作を試してみてください。

# --hints--

パスワードの `input` 要素に `pattern` 属性を追加してください。

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

`pattern` 属性の値を `[a-z0-5]{8,}` に設定する必要があります。

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

パスワードの `input` 要素から `minlength` 属性を削除する必要があります。

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
