---
id: 60fabf0dd4959805dbae09e6
title: ステップ 32
challengeType: 0
dashedName: step-32
---

# --description--

1 つ目の後にもう 1 つの `label` を追加し、テキストは `Input your age (years):` としてください。 次に、`type` が `number` の `input` をネストしてください。

次に、13 歳未満のユーザーが登録できないように、`input` に `min` 属性を追加して値を `13` に設定してください。 また、120 歳より年上のユーザーが登録することはおそらくないでしょう。`max` 属性を追加して、値を `120` に設定してください。

これで、範囲外の値でフォームを送信しようとすると警告が表示され、フォームが送信されないようになりました。 試してみてください。

# --hints--

3 つ目の `fieldset` 内で、既存の `label` の後に `label` を追加する必要があります。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

`label` に `Input your age (years):` というテキストを入れる必要があります。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

`label` 内に、`type` が `number` の `input` を追加する必要があります。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

`input` の `min` 属性の値を `13` に設定する必要があります。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

`input` の `max` 属性の値を `120` に設定する必要があります。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
