---
id: 60fab9f17fa294054b74228c
title: ステップ 31
challengeType: 0
dashedName: step-31
---

# --description--

では最後の `fieldset` に移ります。 ユーザーがプロフィール画像をアップロードできるようにしたい場合はどうしたらいいでしょうか？

これは `input` 要素のタイプ `file` で実現できます。 `label` を追加して、その中に `Upload a profile picture:` というテキストと、ファイルのアップロードを受け付ける `input` をネストしてください。

# --hints--

`Upload a profile picture:` というテキストを持つ `label` 要素を追加する必要があります。

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label')?.innerText, /Upload a profile picture:/i);
```

`input` 要素を `label` 要素内にネストする必要があります。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label > input'));
```

`input` 要素の `type` を `file` に設定する必要があります。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label > input')?.type, 'file');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
