---
id: 60fad1cafcde010995e15306
title: ステップ 41
challengeType: 0
dashedName: step-41
---

# --description--

フォームの送信においては、それぞれの送信可能な要素に `name` 属性を与えることが便利かつ良い方法とされています。 フォームの送信データ内で要素を識別するためにこの属性が使用されます。

2 つの `radio` の input 要素を除く、送信可能な各要素に、ユニークな `name` 属性を任意に設定してください。

# --hints--

You should give the `input` expecting a first name a `name` attribute. _例えば `first-name` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

ラストネームの入力欄の `input` に、`name` 属性が必要です。 _例えば `last-name` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

`email` 欄に `name` 属性が必要です。 _例えば `email` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

`password` 欄に `name` 属性が必要です。 _例えば `password` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

`checkbox` 欄に `name` 属性が必要です。 _例えば `terms` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

`file` 欄に `name` 属性が必要です。 _例えば `file` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

`number` 欄に `name` 属性が必要です。 _例えば `age` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

`select` 要素に `name` 属性が必要です。 _例えば `referrer` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

`textarea` 要素に `name` 属性が必要です。 _例えば `bio` などが適切でしょう。_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

`option` 要素には `name` 属性を設定しないでください。

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

`label` 要素には `name` 属性を設定しないでください。

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

`fieldset` 要素には `name` 属性を設定しないでください。

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
