---
id: 62cc5b1779e4d313466f73c5
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`head` 要素内に、自己終了要素の `link` 要素を追加しましょう。 `rel` 属性の値は `stylesheet`、`href` 属性の値は `styles.css` に設定してください。

# --hints--

コード内に `link` 要素が 1 つ必要です。

```js
assert.exists(document.querySelector('link'));
```

既存の `head` タグを変更しないでください。 終了タグを削除していないか確認してください。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

自己終了要素 `link` が 1 つ必要です。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
