---
id: 62ff8e998d3e7eae14d6ae3b
title: ステップ 30
challengeType: 0
dashedName: step-30
---

# --description--

ユーザーが利用規約 (terms and conditions) を読んだことの確認が必要です。

`label` 要素を追加しましょう。 新しく作成した `label` 要素の中に、`input` 要素を追加して `type` 属性を `checkbox` に設定してください。 この `input` 要素を `required` にして、ユーザーが利用規約に同意しなければサインアップできないようにしてください。

アクセシビリティのために、それぞれの要素に `id` 属性と `for` 属性を追加して、値を `terms-and-conditions` にしてください。

# --hints--

3 つ目の `fieldset` 要素の後に `label` を追加してください。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

`label` 要素内に `input` 要素を追加してください。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

`input` 要素に `type` 属性を追加して、値を `checkbox` にする必要があります。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

`input` 要素に `required` 属性を追加する必要があります。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

`input` 要素には `terms-and-conditions` という `id` が必要です。

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

`label` 要素には `terms-and-conditions` という値の `for` 属性が必要です。

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
