---
id: 63541ef4f96cd82e8e6c788a
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`method` 属性は、フォームのデータを `action` 属性で指定された URL へ「どうやって」送信するかを指定します。 フォームのデータは、`GET` リクエストで URL パラメーターとして送信したり (`method="get"` を指定した場合)、`POST` リクエストでリクエストボディのデータとして送信したり (`method="post"` を指定した場合) できます。

フォームのデータを `POST` リクエストで送信するように `method` 属性を設定してください。

# --hints--

新しい `form` 要素は追加しないでください。

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

`form` 要素には `method` 属性が必要です。

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

`method` 属性は `post` に設定してください。

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
