---
id: 61696ef7ac756c829f9e4048
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

このプロジェクトでは外部 CSS ファイルを使用してページのスタイルを設定します。 `styles.css` ファイルは既に作成してあります。 ですがそれを使えるようにするには、ページにリンクする必要があります。

Nest a `link` element within the `head` element. その `rel` 属性を `stylesheet` に、`href` 属性を `styles.css` に設定してください。

# --hints--

自己終了要素 `link` が 1 つ必要です。

```js
assert(document.querySelectorAll('link').length === 1);
```

`link` 要素は `head` 要素内にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素には `rel` 属性があり、値が `stylesheet` に設定されている必要があります。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
