---
id: 61695ab9f6ffe951c16d03dd
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`html` 要素内に `head` 要素を 1 つネストしてください。 `head` 要素のすぐ後に、`body` 要素を 1 つ追加してください。

# --hints--

`head` の開始タグが 1 つ必要です。

```js
assert(code.match(/<head\s*>/i));
```

`head` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/head\s*>/i));
```

`body` の開始タグが 1 つ必要です。

```js
assert(code.match(/<body\s*>/i));
```

`body` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` の開始タグは 1 つだけ必要です。

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

`head` の終了タグは 1 つだけ必要です。

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

`body` の開始タグは 1 つだけ必要です。

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

`body` の終了タグは 1 つだけ必要です。

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

`head` 要素は空にしてください。

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

`body` 要素は空にしてください。

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

`body` 要素は `head` 要素の後に置く必要があります。

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
