---
id: 616968c2c94c8071b349c146
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

これでコンテンツをページに追加する準備ができました。

`body` 要素内に `h1` 要素をネストし、テキストを `CSS Color Markers` としてください。

# --hints--

コードには `h1` 要素が 1 つ必要です。

```js
const title = document.querySelector('h1');
assert.exists(title);
```

`h1` の開始タグが 1 つ必要です。

```js
assert(code.match(/<h1\s*>/i));
```

`h1` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/h1\s*>/i));
```

`h1` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

`h1` 要素は `CSS Color Markers` というテキストをもつ必要があります。

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

`h1` のテキストについて、大文字小文字の区別とスペルに気をつけましょう。

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
