---
id: 61695197ac34f0407e339882
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

これまでのプロジェクトで見てきたように、ウェブページは `DOCTYPE html` 宣言で始め、その後に続けて `html` 要素を置くようにしましょう。

文書の先頭に `DOCTYPE html` 宣言を追加し、その後に `html` 要素を追加してください。 `html` 要素に `lang` 属性を設定し、値は `en` にしてください。

# --hints--

コードに `DOCTYPE` 参照を入れる必要があります。

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 参照の後にスペースを入れる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/gi));
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
