---
id: 616965351e74d4689eb6de30
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

ウェブページには、自己終了要素 `meta` を複数追加できます。 各 `meta` 要素は、他の HTML 要素では表現できない、そのページ自体に関する情報を追加します。

自己終了要素である `meta` 要素をもう 1 つ、`head` 要素内に追加してください。 `name` 属性を `viewport` に、`content` 属性を `width=device-width, initial-scale=1.0` に設定してください。これにより、あらゆるデバイスでページが同じように表示されます。

# --hints--

2 つの `meta` 要素が必要です。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

新しい `meta` 要素は自己終了要素である必要があります。

```js
assert(code.match(/<\/meta>/i) === null);
```

新しい `meta` 要素の `name` 属性には `viewport` を、`content` 属性には `width=device-width, initial-scale=1.0` を設定する必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
