---
id: 616971b3cd990186b66c99a1
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

外部 CSS ファイルが設定されたので、ページのスタイリングを始めることができます。

ヒントとして、段落要素を選択して右揃えにする方法は次の通りです:

```css
p {
  text-align: right;
}
```

`h1` 要素を選択する新しい CSS ルールを作成し、その `text-align` プロパティを `center` に設定してください。

# --hints--

`h1` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

`h1` 用の CSS ルールの `text-align` プロパティを `center` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
