---
id: 61695d1fbc003856628bf561
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

ブラウザーにページの文字エンコーディングについて指示するため、`charset` を `utf-8` に設定します。 `utf-8` とは、あらゆる言語のほぼすべての文字を含む世界共通の文字セットです。

`head` 要素内に `meta` 要素をネストし、`charset` 属性を `utf-8` に設定してください。 `meta` 要素は自己終了要素なので、終了タグは必要ないことを覚えておきましょう。

# --hints--

`meta` 要素が 1 つ必要です。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

`meta` 要素は自己終了要素である必要があります。

```js
assert(code.match(/<\/meta>/i) === null);
```

`meta` 要素は `charset` 属性を `utf-8` に設定する必要があります。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
