---
id: 616d4a84b756d9c4b8255093
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

It's time to add some color to the marker. 要素に色をつける方法の 1 つに、`black` や `cyan`、`yellow` のような<dfn>色キーワード</dfn>を使う方法があることを思い出してください。

ヒントとして、クラス `freecodecamp` を選択する方法は次の通りです:

```css
.freecodecamp {

}
```

クラス `marker` を選択する新しい CSS ルールを作成し、`background-color` プロパティを `red` に設定してください。

**Note:** You will not see any changes after adding the CSS.

# --hints--

You should create a class selector to target the `marker` class.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Your `.marker` CSS rule should have a `background-color` property set to `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
