---
id: 616d55bd160a95e22453a081
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

これで 1 つのマーカーに色を付けて中央に配置できたので、今度は別のマーカーを追加しましょう。

`container` の `div` 内に、`div` 要素をさらに 2 つ追加して、それぞれにクラス `marker` を設定してください。

# --hints--

1 つ目の新しい `div` 要素の開始タグが必要です。

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

1 つ目の新しい `div` 要素の終了タグが必要です。

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

2 つ目の新しい `div` 要素の開始タグが必要です。

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

2 つ目の新しい `div` 要素の終了タグが必要です。

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

新しく追加した `div` 要素は、`container` クラスの `div` の中にある必要があります。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

新しい `div` 要素には `marker` のクラスが必要です。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
