---
id: 616d525007b8c5d8b3308b1c
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

マーカーがページの中央にあった方が見た目が良くなりそうです。 そうするには `margin` <dfn>一括指定プロパティ</dfn>を使うのが簡単です。

一つ前のプロジェクトでは、`margin-top` や `margin-left` などのプロパティを使って要素のマージン領域を個別に設定しました。 `margin` 一括指定プロパティを使うと、複数のマージン領域を同時に設定することが簡単にできます。

マーカーをページの中央に配置するために、`margin` プロパティを `auto` に設定してください。 これにより `margin-top`、`margin-right`、`margin-bottom`、および `margin-left` のすべてが `auto` に設定されます。

# --hints--

`.marker` CSS ルールの `margin` プロパティを `auto` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
