---
id: 616d3a67ccf800ad94ec89ae
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

ここで、最終的にカラーマーカーの形になるようスタイルを適用する要素をいくつか追加します。

First, within the `body` element, add a `div` element and set its `class` attribute to `container`. `div` 要素は `h1` 要素の下に置いてください。

# --hints--

`div` の開始タグが 1 つ必要です。

```js
assert(code.match(/<div.*>/i));
```

`div` の終了タグが 1 つ必要です。

```js
assert(code.match(/<\/div\s*>/i));
```

`div` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

`div` 要素は `h1` 要素の後に置く必要があります。

```js
assert.exists(document.querySelector('h1 + div'));
```

`div` 要素の `class` 属性を `container` に設定する必要があります。

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
