---
id: 61764c602bee6974e7790f35
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

マーカーに異なる色を付けるため、それぞれに固有のクラスを追加します。 `class` 属性にスペースで区切って並べることで、1 つの要素に複数のクラスを追加することができます。 たとえば、次の例は `animal` と `dog` の両方のクラスを `div` 要素に追加します。

```html
<div class="animal dog">
```

1 つの HTML 要素に複数のクラスを追加する場合、最初のクラスのスタイルは後のクラスで上書きされる可能性があります。

はじめに、クラス `one` を最初のマーカーの `div` 要素に追加してください。


# --hints--

クラス `one` を最初のマーカーの `div` 要素に追加してください。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

最初のマーカーの `div` には `marker` と `one` というクラスが必要です。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
