---
id: 61764c602bee6974e7790f35
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

学校で、赤、黄、青が原色であることと、それを混ぜて新しい色を作る方法を習ったことがあるかもしれません。 しかし、これは旧式の理論です。

現代では、主に 2 つのカラーモデルがあります。電子機器内で使用される<dfn>加法</dfn> RGB (赤、緑、青) モデルと、印刷で使用される<dfn>減法</dfn> CMYK (シアン、マゼンタ、黄、黒) モデルです。 このプロジェクトでは、RGB モデルを扱います。

まず、1 つ目のマーカーの `div` 要素に、`one` という `class` を追加してください。

```html
<div class="exampleClass1 exampleClass2">
  ...content
</div>
```

# --hints--

クラス `one` を最初のマーカーの `div` 要素に追加してください。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

最初のマーカーの `div` には `marker` と `one` というクラスが必要です。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
