---
id: 6176519636a76b810ab1219a
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

クラス `two` を選択する CSS ルールを作成し、その `background-color` プロパティを `green` に設定してください。

また、それとは別にクラス `three` を選択する CSS ルールを作成し、その `background-color` プロパティを `blue` に設定してください。

# --hints--

クラスセレクターを使用してクラス `two` を選択してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

`.two` CSS ルールの `background-color` プロパティを `green` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

クラスセレクターを使用してクラス `three` を選択してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

`.three` CSS ルールの `background-color` プロパティを `blue` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
