---
id: 61764c602bee6974e7790f35
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

マーカーに異なる色を付けるには、それぞれに固有のクラスを追加する必要があります。 複数のクラスを `class` 属性に並べ、スペースで区切ることで、要素に追加することができます。 たとえば、次の例は `animal` と `dog` の両方のクラスを `div` 要素に追加します。

```html
<div class="animal dog">
```

はじめに、クラス `one` を最初のマーカーの `div` 要素に追加してください。


# --hints--

クラス `one` を最初のマーカーの `div` 要素に追加してください。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

最初のマーカーの `div` には `marker` と `one` というクラスが必要です。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
