---
id: 617bd6ec666b1da2587e4e37
title: ステップ 34
challengeType: 0
dashedName: step-34
---

# --description--

三次色には他にもシャルトルーズグリーン (緑 + 黄、黄色がかった緑色)、アジュール (青 + シアン、明るい青)、ローズ (赤 + マゼンタ、薔薇色) の 3 つがあります。

シャルトルーズグリーンを作成するには、`.one` ルール内の `rgb` 関数を変更して赤を `127` に、緑を最大値にしてください。

アジュールを得るには、`.two` ルール内の `rgb` 関数を変更して緑を `127` に、青を最大値にしてください。

そしてローズ (明るいピンクと呼ばれることもあります) を得るには、`.three` ルール内の `rgb` 関数を変更して青を `127` に、赤を最大値にしてください。

# --hints--

`.one` CSS ルールの `background-color` プロパティを `rgb(127, 255, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

`.two` CSS ルールの `background-color` プロパティを `rgb(0, 127, 255)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

`.three` CSS ルールの `background-color` プロパティを `rgb(255, 0, 127)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
