---
id: 617b8b38f32bf2080a140675
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

関数とは、入力を受け取って特定の処理を実行するコードのまとまりです。 CSS の `rgb` 関数は、赤、緑、青の値 (<dfn>引数</dfn>) を受け取り、1 つの色を作成します:

```css
rgb(red, green, blue);
```

赤、緑、青の各値は `0` から `255` の数値です。 `0` はその色が 0% であり、黒であることを意味します。 `255` はその色が 100% であることを意味します。

`.one` CSS ルール内で、色キーワード `red` を `rgb` 関数に置き換えてください。 その `rgb` 関数に対し、赤の値を `255` に、緑の値を `0` に、青の値を `0` に設定してください。

# --hints--

`.one` CSS ルールの `background-color` プロパティの設定には、色キーワード `red` を使わないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

`.one` CSS ルールの `background-color` プロパティを `rgb(255, 0, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
