---
id: 617b9ad735109e217284e095
title: ステップ 27
challengeType: 0
dashedName: step-27
---

# --description--

これで原色の RGB カラーができたので、それを混合しましょう。

`.container` ルール内の `rgb` 関数に対し、赤、緑、青の値を最大の `255` に設定してください。

# --hints--

`.container` CSS ルールの `background-color` プロパティを `rgb(255, 255, 255)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
