---
id: 617bb5624a75e86463b7e638
title: ステップ 28
challengeType: 0
dashedName: step-28
---

# --description--

<dfn>二次色</dfn>とは原色を混合して得られる色のことです。 前のステップで赤、緑、青の値を変えている途中で、二次色に気づいたかもしれません。

1 つ目の二次色として黄色を作成するために、`.one` CSS ルール内の `rgb` 関数を変更して純粋な赤と純粋な緑を混ぜ合わせてください。

# --hints--

`.one` CSS ルールの `background-color` プロパティを `rgb(255, 255, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
