---
id: 617b97abd9f3f61d1590b815
title: ステップ 26
challengeType: 0
dashedName: step-26
---

# --description--

加法 RGB カラーモデルにおける<dfn>原色</dfn>とは、混合された時に純粋な白を作る色のことです。 純粋な白を作るには、それぞれの色が最大強度である必要があります。

色を混ぜ合わせる前に、緑のマーカーを純粋な緑に戻しましょう。 `.two` CSS ルール内の `rgb` 関数について、緑を最大値の `255` に戻してください。

# --hints--

`.two` CSS ルールの `background-color` プロパティを `rgb(0, 255, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
