---
id: 617b65579ce424bf5f02ca73
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

先ほど RGB カラーモデルは加法混合であることを学びました。 これは色が黒から始まり、異なる量の赤、緑、青が加えられることによって変化することを意味します。

これを簡単に確認する方法が CSS の `rgb` 関数です。

クラス `container` を選択する新しい CSS ルールを作成し、その `background-color` を `rgb(0, 0, 0)` で黒に設定してください。

# --hints--

クラスセレクタ―を使用してクラス `container` を選択してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

`.container` CSS ルールの `background-color` プロパティを `rgb(0, 0, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
