---
id: 617bc3386dc7d07d6469bf20
title: ステップ 32
challengeType: 0
dashedName: step-32
---

# --description--

オレンジを作るには、`rgb` 値の赤の強度を増やし、緑の強度を減らさなければならないことに注目してください。 これは、オレンジが赤と黄色を混ぜ合わせたものであり、カラーホイール上でその 2 つの色の間にあるためです。

三次色のスプリンググリーン (黄緑) を作成するには、シアンと緑を混ぜ合わせます。 `.two` CSS ルール内の `rgb` 関数を変更して、緑を最大値に、青を `127` に設定してください。

# --hints--

`.two` CSS ルールの `background-color` プロパティを `rgb(0, 255, 127)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
