---
id: 617b8e0d93a8d10d9a90e720
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

マーカーの `background-color` がまだ赤であることに注目してみましょう。 `rgb` 関数の赤の値を最大の `255` (100%) に設定し、緑と青の両方の値を `0` に設定しているので、このようになります。

では `rgb` 関数を使用して他の色を設定してみましょう。

`.two` CSS ルール内で、`rgb` 関数を使用して `background-color` の緑の値を最大に設定し、他の値は `0` にしてください。 また `.three` CSS ルール内で、`rgb` 関数を使用して `background-color` の青の値を最大に設定し、他の値は `0` にしてください。

# --hints--

`.two` CSS ルールの `background-color` プロパティの設定には、色キーワード `green` を使わないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

`.two` CSS ルールの `background-color` プロパティを `rgb(0, 255, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

`.three` CSS ルールの `background-color` プロパティの設定には、色キーワード `blue` を使わないでください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

`.three` CSS ルールの `background-color` プロパティを `rgb(0, 0, 255)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
