---
id: 617bbb6b97a83f6d1f7d6e38
title: ステップ 31
challengeType: 0
dashedName: step-31
---

# --description--

二次色に慣れてきたところで、<dfn>三次色</dfn>の作り方を学びましょう。 三次色は、原色を隣接する二次色の 1 つと混合することで得られます。

三次色のオレンジを作成するために、`.one` CSS ルール内の `rgb` 関数を変更して赤を最大値に、緑を `127` に設定してください。

# --hints--

`.one` CSS ルールの `background-color` プロパティを `rgb(255, 127, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
