---
id: 618a0927005553b74bfa05ae
title: ステップ 36
challengeType: 0
dashedName: step-36
---

# --description--

カラーホイールとは、同じような色 (<dfn>色相</dfn>) が互いに近くに、異なる色は離れて配置された円のことです。 例えば、純粋な赤はローズとオレンジの色相の間に位置します。

カラーホイール上で互いに反対の位置にある 2 つの色は<dfn>補色</dfn>と呼ばれています。 2 つの補色を混ぜ合わせると、灰色ができます。 しかしそれらを並べて配置すると、視覚的なコントラストが強くなり明るく見えます。

`.one` CSS ルールの `rgb` 関数内で、赤の値を最大の `255` に設定して純粋な赤を作成してください。 `.two` CSS ルールの `rgb` 関数内では、緑と青の値を最大の `255` に設定してシアンを作成してください。

# --hints--

`.one` CSS ルールの `background-color` プロパティを `rgb(255, 0, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

`.two` CSS ルールの `background-color` プロパティを `rgb(0, 255, 255)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
