---
id: 618a132676346ac9f7fd59dd
title: ステップ 39
challengeType: 0
dashedName: step-39
---

# --description--

自然に中央の赤色が目を引くことに気がつきましたか？ サイトをデザインする際に、この効果を利用して重要な見出し、ボタン、あるいはリンクに注意を引き付けることができます。

補色の他にも重要な色の組み合わせがありますが、それらはもう少し後で学びます。

For now, use the `rgb` function in the `.two` CSS rule to set the `background-color` to black.

# --hints--

`.two` CSS ルールの `background-color` プロパティを `rgb(0, 0, 0)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
