---
id: 619b761916dac02643940022
title: ステップ 46
challengeType: 0
dashedName: step-46
---

# --description--

CSS で要素に色を適用する方法として、<dfn>16 進数</dfn> (hexadecimal、ヘキサ) の値がよく使われます。 16 進数と聞くとわかりにくいですが、実のところ RGB 値の別の形式に過ぎません。

16 進数の色の値は `#` 記号から始まり、0-9 および A-F から 6 つの文字を取ります。 最初の二文字のペアが赤を、2 つ目のペアは緑を、3 つ目のペアは青を表します。 たとえば `#4B5320` のように記述します。

In the `.green` class selector, set the `background-color` property to a hex color code with the values `00` for red, `FF` for green, and `00` blue.

# --hints--

`.green` CSS ルールの `background-color` プロパティを `#00FF00` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
