---
id: 619b74fa777a2b2473c94f82
title: ステップ 45
challengeType: 0
dashedName: step-45
---

# --description--

CSS クラスセレクター `.two` を変更して、新しい `green` クラスを選択してください。 また `.three` セレクターを変更して新しい `blue` クラスを選択してください。

# --hints--

コード内に `.two` クラスセレクターがないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

クラスセレクターを使用してクラス `green` を選択してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

`.green` CSS ルールの `background-color` プロパティは `rgb(0, 0, 0)` に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

コード内に `.three` クラスセレクターがないようにしてください。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

クラスセレクターを使用してクラス `blue` を選択してください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

`.blue` CSS ルールの `background-color` プロパティは `rgb(0, 0, 0)` に設定されている必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
