---
id: 619b7fd56aa2253778bcf5f7
title: ステップ 48
challengeType: 0
dashedName: step-48
---

# --description--

<dfn>HSL</dfn> カラーモデル、または色相 (hue)、彩度 (saturation)、および明度 (lightness) とは、色を表現するもう一つの方法です。

The CSS hsl function accepts 3 values: a number from 0 to 360 for hue, a percentage from 0 to 100 for saturation, and a percentage from 0 to 100 for lightness.

カラーホイールを想像すると、赤の色相は 0 度、緑は 120 度、青は 240 度に位置します。

Saturation is the intensity of a color from 0%, or gray, to 100% for pure color. You must add the percent sign `%` to the saturation and lightness values.

明度は色の明るさで、0% の完全な黒から 100% の完全な白までを表し、50% で中間色となります。

`.blue` CSS ルール内で、`hsl` 関数を使用して `background-color` プロパティを変更し、純粋な青にしてください。 色相を `240` に、彩度を `100%` に、明度を `50%` に設定してください。

# --hints--

`.blue` CSS ルールの `background-color` プロパティを `hsl(240, 100%, 50%)` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
