---
id: 61a49d15bdbb5e57cc6fd280
title: ステップ 54
challengeType: 0
dashedName: step-54
---

# --description--

色経由点 (color-stops) を利用すると、グラデーション軸に沿って色がどう配置されるかを微調整できます。 色経由点は、`linear-gradient` 関数内で色指定に続けて `px` などの長さの単位、またはパーセンテージで指定します。

例えば下記の赤から黒のグラデーションでは、赤から黒への変化はグラデーション軸に沿って 90% の位置で発生し、スペースのほとんどを赤が占めることになります。

```css
linear-gradient(90deg, red 90%, black);
```

`linear-gradient` 関数内で、`75%` の色経由点を 1 つ目の 赤色の引数の後に追加してください。 他の色の引数には色経由点を追加しないでください。

# --hints--

`.red` CSS ルールの `background` プロパティを `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))` に設定する必要があります。

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
