---
id: 61a489b8579e87364b2d2cdb
title: ステップ 49
challengeType: 0
dashedName: step-49
---

# --description--

CSS で均一な色を設定する方法をいくつか学んできましたが、色の移り変わり、<dfn>グラデーション</dfn>を要素に適用することもできます。

グラデーションとは、ある色を別の色へ変化させることです。 CSS の `linear-gradient` 関数を使用して、線に沿った変化の方向と使用する色を制御できます。

ひとつ忘れてはならないことがあります。`linear-gradient` 関数は実際は `image` 要素を作り出すため、画像を値として受け付ける `background` プロパティと通常はペアで使われます。

`.red` CSS ルール内の `background-color` プロパティを `background` に変更してください。

# --hints--

`.red` CSS ルールには `background` プロパティと値 `rgb(255, 0, 0)` が必要です。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
