---
id: 61a493ead935c148d2b76312
title: ステップ 50
challengeType: 0
dashedName: step-50
---

# --description--

`linear-gradient` 関数は非常に柔軟性があります。このチュートリアルで使用する基本的な構文はこちらです:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` は変化に使用する線の方向です。 `color1` と `color2` は色の引数で、変化の中で使用される色です。 色にはカラーキーワード、16 進数、`rgb`、または `hsl` を含め、どの形式でも使用できます。

では 1 つ目のマーカーに、赤から緑のグラデーションを 90 度の線に沿って適用してみましょう。

最初に、`.red` CSS ルール内で `background` プロパティを `linear-gradient()` に設定し、それに値 `90deg` を `gradientDirection` として渡してください。

# --hints--

`.red` CSS ルールには `background` プロパティと値 `linear-gradient(90deg)` が必要です。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
