---
id: 61a5c906ab73313316e342f0
title: ステップ 65
challengeType: 0
dashedName: step-65
---

# --description--

`linear-gradient` 関数に `gradientDirection` 引数が与えられなかった場合、デフォルトでは色が上から下、つまり 180 度の線に沿って配置されます。

両方の `linear-gradient` 関数から `gradientDirection` 引数を削除して、コードを整理しましょう。

# --hints--

`.red` および `.green` の CSS ルール内の `linear-gradient` 関数から、`gradientDirection` 引数を削除する必要があります。

```js
assert(!code.match(/linear-gradient\s*\(\s*180deg/gi));
```

`.red` CSS ルールの `background` プロパティを `linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))` に設定する必要があります。

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

`.green` CSS ルールの `background` プロパティを `linear-gradient(#55680D, #71F53E, #116C31)` に設定する必要があります。

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
