---
id: 61adc91467b5d59328b9f781
title: ステップ 79
challengeType: 0
dashedName: step-79
---

# --description--

1 つ前のプロジェクトで、境界線 (border) と `border-color` プロパティについて少し学習しました。

すべての HTML 要素は境界線を持ちますが、通常、デフォルト値は `none` に設定されています。 CSS では、要素の境界線のあらゆる外観を制御できます。1 度にすべての側面の境界線を設定することも、1 つずつ設定することもできます。 境界線が見えるようにするには、幅とスタイルを設定する必要があります。

`.sleeve` CSS ルール内に、`border-left-width` プロパティと値 `10px` を追加してください。

# --hints--

`.sleeve` CSS ルールには `border-left-width` プロパティが必要で、その値は `10px` である必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeftWidth === '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}
--fcc-editable-region--

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
