---
id: 61b0a44a6b865738ba49b9fb
title: ステップ 86
challengeType: 0
dashedName: step-86
---

# --description--

最後に各マーカーに少し影を付けて、更にリアルに見えるようにします。

`box-shadow` プロパティを使うと、要素の周りに 1 つまたは複数の影を適用できます。 基本的な構文は次のとおりです。

```css
box-shadow: offsetX offsetY color;
```

`offsetX` と `offsetY` の値は次のように動作します。

* `offsetX` と `offsetY` はどちらも、`px` など CSS で使われる単位の数値を受け付けます。
* 正の `offsetX` 値は影を右に、負の値は左に動かします。
* 正の `offsetY` 値は影を下に、負の値は上に動かします。
* `offsetX` と `offsetY` の両方またはどちらかにゼロ (`0`) の値を指定したい場合、単位は必要ありません。 すべてのブラウザーで、ゼロは位置を動かさないと解釈されます。

影の高さと幅は、その影が適用されている要素の高さと幅によって決まります。 オプションの値 `spreadRadius` を使って影の範囲を広げることもできます。 これについては後ほど詳しく説明します。

赤いマーカーにシンプルな影を付けることから始めましょう。

`.red` CSS ルール内に `box-shadow` プロパティを追加して、値は `offsetX` を `5px`、`offsetY` を `5px`、`color` を `red` にしてください。

# --hints--

`.red` CSS ルールの `box-shadow` 一括指定プロパティの値を `5px 5px red` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
