---
id: 61b093179e7fc020b43b1bb3
title: ステップ 73
challengeType: 0
dashedName: step-73
---

# --description--

<dfn>不透明度 (Opacity)</dfn> は、ある物がどのくらい不透明かを表します。 例えば頑丈な壁は不透明で、光を通しません。 一方飲み物用のグラスは透明で、グラス越しに向こう側が見えます。

CSS の `opacity` プロパティを使って、要素の不透明さ、あるいは透明さを制御できます。 値が `0`、または 0% の時、要素は完全に透明になり、`1.0`、または 100% の時、要素は (デフォルトの状態と同じように) 完全に不透明になります。

`.sleeve` CSS ルール内で、`opacity` プロパティを `0.5` に設定してください。

# --hints--

`.sleeve` CSS ルールの `opacity` プロパティを `0.5` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
