---
id: 61b30995968123ceb6b76167
title: ステップ 88
challengeType: 0
dashedName: step-88
---

# --description--

影のふちがくっきりしていることに注目してください。 これは、`box-shadow` プロパティにはオプションの値 `blurRadius` があるためです。

```css
box-shadow: offsetX offsetY blurRadius color;
```

`blurRadius` の値が含まれない場合、デフォルト値は `0` となり、くっきりしたふちになります。 `blurRadius` の値が大きいほど、ぼかし効果が大きくなります。

`.green` CSS ルール内に `box-shadow` プロパティを追加して、値は `offsetX` を `5px`、`offsetY` を `5px`、`blurRadius` を `5px`、`color` を `green` にしてください。

# --hints--

`.green` CSS ルールの `box-shadow` 一括指定プロパティの値を `5px 5px 5px green` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow === 'green 5px 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

--fcc-editable-region--
.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
