---
id: 5a94fdf869fb03452672e45b
title: 使用 align-items 垂直对齐所有项目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ckzPeUv'
forumTopicId: 301121
---

# --description--

对网格容器使用 `align-items` 属性可以让网格中所有的网格项沿竖直方向对齐。

# --instructions--

请使用 `align-items` 属性将所有网格项移动到单元格的末尾。

# --hints--

class 为 `container` 的元素应具有 `align-items` 属性且属性值应为 `end`。

```js
assert(
  code.match(/.container\s*?{[\s\S]*align-items\s*?:\s*?end\s*?;[\s\S]*}/gi)
);
```

# --solutions--

