---
id: 5a90375238fddaf9a66b5d3b
title: 使用 align-self 垂直对齐项目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzd4fz'
forumTopicId: 301123
---

# --description--

正如能设置网格项沿行轴对齐方式一样，也可以设置网格项沿列轴对齐：你可以对网格项使用`align-self`属性。在上一个挑战中适用于`justify-self`属性的每个值也都适用于`align-self`属性。

# --instructions--

用值`end`使类为`item3`的网格项底端对齐。

# --hints--

`item3`类应该有`align-self`属性且值为`end`。

```js
assert(code.match(/.item3\s*?{[\s\S]*align-self\s*?:\s*?end\s*?;[\s\S]*}/gi));
```

# --solutions--

