---
id: 5a94fdf869fb03452672e45b
title: 使用 align-items 垂直对齐所有项目
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ckzPeUv'
forumTopicId: 301121
---

# --description--

对网格容器使用`align-items`属性可以给网格中所有的网格项设置沿列轴对齐的方式。

# --instructions--

请使用`align-items`属性将所有网格项移动到单元格的末尾。

# --hints--

`container`类应该有`align-items`属性且值为`end`。

```js
assert(
  code.match(/.container\s*?{[\s\S]*align-items\s*?:\s*?end\s*?;[\s\S]*}/gi)
);
```

# --solutions--

