---
id: 5a9036ee38fddaf9a66b5d36
title: 使用 grid-row-gap 创建多行之间的间距
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cPbJ2Cv'
forumTopicId: 301125
---

# --description--

和上个挑战在两列之间添加间距一样，我们还可以用 `grid-row-gap` 设置行间距。

# --instructions--

请为网格中的行添加高度为 `5px` 的间距。

# --hints--

class 为 `container` 的元素应具有 `grid-row-gap` 属性且属性值应为 `5px`。

```js
assert(
  code.match(/.container\s*?{[\s\S]*grid-row-gap\s*?:\s*?5px\s*?;[\s\S]*}/gi)
);
```

# --solutions--

