---
id: 5a858944d96184f06fd60d61
title: 创建你的第一个 CSS 网格
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
---

# --description--

通过将属性`display`的值设为`grid`，使 HTML 元素变为网格容器。通过前面的操作，你可以对该容器使用与 CSS 网格（CSS Grid）相关的属性。

**注意：**  
在 CSS 网格中，父元素称为<dfn>容器（container）</dfn>，它的子元素称为<dfn>项（items）</dfn>。

# --instructions--

将类为`container`的 div 的`display`属性改为`grid`。

# --hints--

`container`类应该有`display`属性且值为`grid`。

```js
assert(code.match(/.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi));
```

# --solutions--

