---
id: 641da5462576784453146ec2
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Agora que você tem os dados que quer, você pode usá-los para preencher a interface do usuário. Os dados obtidos, no entanto, contêm um array de 26 autores. Se você adicionar todos eles à página ao mesmo tempo, poderá levar a um desempenho ruim.

Ao invés disso, você deve adicionar 8 autores de uma vez e ter um botão para adicionar mais 8 até que não haja mais dados para serem exibidos.

Use `let` para criar 2 variáveis chamadas `startingIndex` e `endingIndex` e atribua a elas os valores numéricos `0` e `8`, respectivamente. Além disso, crie uma variável `authorDataArr` com `let` e defina-a como um array vazio.

# --hints--

Você deve usar `let` para declarar uma variável chamada `startingIndex`.

```js
assert.match(code, /let\s+startingIndex/)
```

Você deve definir a variável `startingIndex` como `0`.

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

Você deve usar `let` para declarar uma variável chamada `endingIndex`.

```js
assert.match(code, /let\s+endingIndex/)
```

Você deve definir a variável `endingIndex` como `8`.

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

Você deve usar `let` para declarar uma variável chamada `authorDataArr`.

```js
assert.match(code, /let\s+authorDataArr/)
```

Você deve definir a variável `authorDataArr` como um array vazio (`[]`).

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
