---
id: 641da86294fd9f485d3c2bf0
title: Passo 19
challengeType: 0
dashedName: step-19
---

# --description--

Em seguida, adicione um link para a página do autor no editorial do freeCodeCamp.

Adicione um elemento de âncora com a `class` `author-link`, interpole `url` como o valor para o atributo `href` e defina `target` como `_blank`. Para o texto do elemento de âncora, interpole `author` seguido pelo texto `'s author page`. Por exemplo, `Quincy Larson's author page`.

# --hints--

Você deve criar um elemento de âncora.

```js
assert.exists(document.querySelector('a'));
```

O elemento de âncora deve ter a classe `author-link`.

```js
assert.include(document.querySelector('a')?.className, "author-link");
```

Você deve interpolar `url` como o valor do atributo `href` do elemento de âncora.

```js
assert.equal(document.querySelector('a')?.getAttribute('href'), authorDataArr[0].url);
```

Você deve definir o atributo `target` do elemento de âncora como `_blank`.

```js
assert.equal(document.querySelector('a')?.getAttribute('target'), '_blank');
```

Você deve interpolar `author` seguido pelo texto `'s author page` dentro do elemento de âncora.

```js
assert.equal(document.querySelector('a')?.textContent.trim(), `${authorDataArr[0].author}'s author page`);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

let startingIndex = 0;
let endingIndex = 8;
let authorDataArr = [];

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    authorDataArr = data;
    displayAuthors(authorDataArr.slice(startingIndex, endingIndex));   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });

const displayAuthors = (authors) => {
  authors.forEach( ({ author, image, url, bio }, index) =>
  (authorContainer.innerHTML += `
    <div id="${index}" class="user-card">
      <h2 class="author-name">${author}</h2>
      <img class="user-img" src="${image}" alt="${author} avatar" />
      <p class="bio">${bio}</p>
      --fcc-editable-region--

      --fcc-editable-region--
    </div>
  `)
  )
};
```
