---
id: 61695197ac34f0407e339882
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Wie du in den vorherigen Projekten gesehen hast, sollten Webseiten mit einer `DOCTYPE html`-Deklaration beginnen, gefolgt von einem `html`-Element.

Füge eine `DOCTYPE html`-Deklaration oben im Dokument und ein `html`-Element nach dieser hinzu. Gib dem `html`-Element ein `lang`-Attribut mit `en` als Wert.

# --hints--

Dein Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest ein Leerzeichen nach der `DOCTYPE`-Referenz einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut mit dem Wert `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/gi));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTMLs stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
