---
id: 61695d1fbc003856628bf561
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Das `charset`-Attribut gibt an, wie die Zeichen vom Dokument kodiert werden. `utf-8` (Unicode Transformation Format – 8-bit) ist eine Standard-Kodierung, die für elektronische Kommunikation verwendet wird.

Verschachtele innerhalb des `head`-Elements ein `meta`-Element mit einem `charset` von `"utf-8"`.

# --hints--

Du solltest ein `meta`-Element haben.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Dein `meta`-Element sollte ein leeres Element sein, es sollte kein schließendes Tag `</meta>` haben.

```js
assert.notMatch(code, /<\/meta>/i);
```

Dein `meta`-Element sollte ein `charset`-Attribut haben, das auf den Wert `"utf-8"` gesetzt ist.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Dein `meta`-Element sollte innerhalb des `head`-Elements liegen.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
