---
id: 61695c4aad56f95497c19583
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Beachte, dass das `title`-Element den Suchmaschinen zusätzliche Informationen über die Seite gibt. It also displays the content of that `title` element in two more ways:

* in the title bar when the page is open
* in der Browser-Registerkarte für die Seite, wenn du mit dem Mauszeiger darüber fährst. Auch wenn dieses Tab nicht aktiv ist, wird der `title`-Text angezeigt, sobald du mit der Maus darüber fährst.

Niste innerhalb des `head`-Elements ein `title`-Element mit dem Text `Colored Markers`.

# --hints--

Du solltest ein öffnendes `title`-Tag haben.

```js
assert(code.match(/<title\s*>/i));
```

Du solltest ein schließendes `title`-Tag haben.

```js
assert(code.match(/<\/title\s*>/i));
```

Dein Projekt sollte den Titel `Colored Markers` haben.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Remember, casing and spelling matter for the title text.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
