---
id: 61696ef7ac756c829f9e4048
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

In diesem Projekt wirst du mit einer externen CSS-Datei arbeiten, um die Seite zu gestalten. Wir haben bereits eine `styles.css` für dich erstellt. Aber bevor du sie benutzen kannst, musst du sie mit der Seite verlinken.

Bette ein `link`-Element in das `head`-Element ein. Weise ihm ein `rel`-Attribut mit `stylesheet` und ein `href`-Attribut mit `styles.css` zu.

# --hints--

Du solltest ein selbstschließendes `link`-Element haben.

```js
assert(document.querySelectorAll('link').length === 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Dein `link`-Element sollte ein `rel`-Atribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
