---
id: 616965351e74d4689eb6de30
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Du kannst mehrere selbst schließende `meta`-Elemente auf einer Webseite haben. Jedes `meta`-Element fügt Informationen über die Seite hinzu, die nicht durch andere HTML-Elemente wiedergegeben werden kann.

Füge ein weiteres selbst schließendes `meta`-Element in den `head` ein. Gib ihm ein `name`-Attribut mit dem Wert `viewport` und ein `content`-Attribut mit dem Wert `width=device-width, initial-scale=1.0`, so dass deine Seite auf allen Geräten gleich aussieht.

# --hints--

Du solltest zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Dein neues `meta`-Element sollte ein selbst schließendes Element sein.

```js
assert(code.match(/<\/meta>/i) === null);
```

Dein neues `meta`-Element sollte ein `name`-Attribut auf `viewport` und ein `content`-Attribut auf `width=device-width, initial-scale=1.0` gesetzt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
