---
id: 61695ab9f6ffe951c16d03dd
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Bette innerhalb des `html`-Elements ein `head`-Element ein. Füge nach dem `head`-Element ein `body`-Element hinzu.

# --hints--

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein schließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein schließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Es sollte nur ein öffnendes `head`-Tag vorhanden sein.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Es sollte nur ein schließendes `head`-Tag geben.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Es sollte nur ein öffnendes `body`-Tag geben.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Es sollte nur ein schließendes `body`-Tag geben.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Dein `head`-Element sollte leer sein.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Dein `body`-Element sollte leer sein.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Dein `body`-Element sollte nach deinem `head`-Element stehen.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Dein `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Dein `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
