---
id: 61695d1fbc003856628bf561
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Setze das `charset` auf `utf-8`, um den Browsern mitzuteilen, wie man Zeichen auf deiner Seite kodiert. `utf-8` ist ein universelles Zeichen-Set, das fast jedes Zeichen aller menschlichen Sprachen enthält.

Bette innerhalb des `head`-Elements ein `meta`-Element mit dem Attribut `charset` mit `utf-8` als Wert ein. Denke daran, dass `meta`-Elemente selbstabschließend sind und kein abschließendes Tag benötigen.

# --hints--

Du solltest ein `meta`-Element haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

Dein `meta`-Element sollte ein selbstschließendes Element sein.

```js
assert(code.match(/<\/meta>/i) === null);
```

Dein `meta`-Element sollte ein `charset`-Attribut auf `utf-8` gesetzt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
