---
id: 616971b3cd990186b66c99a1
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Da nun deine externe CSS-Datei eingerichtet ist, kannst du mit der Gestaltung der Seite beginnen.

Hier zur Erinnerung, wie man ein Absatz-Element auswählt und es nach rechts ausrichtet:

```css
p {
  text-align: right;
}
```

Erstelle eine neue CSS-Regel, die das `h1`-Element auswählt, und setze ihre `text-align`-Eigenschaft auf `center`.

# --hints--

Du solltest einen `h1`-Selektor verwenden.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Deine `h1`-CSS-Regel sollte eine `text-align`-Eigenschaft von `center` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
