---
id: 61696ef7ac756c829f9e4048
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

In diesem Projekt wirst du mit einer externen CSS-Datei arbeiten, um die Seite zu gestalten. Wir haben bereits eine `styles.css` für dich erstellt. Aber bevor du sie benutzen kannst, musst du sie mit der Seite verlinken.

Bette ein `link`-Element in das `head`-Element ein. Gib ihm ein `rel`-Attribut mit `"stylesheet"` und ein `href`-Attribut mit `"styles.css"`.

# --hints--

Du solltest ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Dein`link`-Element sollte ein leeres Element sein, es sollte kein schließendes Tag besitzen `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Dein`link`-Element sollte ein`rel`-Attribut mit dem Wert `"stylesheet"` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `"styles.css"` enthalten.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Dein`link`-Element sollte innerhalb des `head`-Elements liegen.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
