---
id: 61695197ac34f0407e339882
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

As you've seen in the previous projects, webpages should start with a `DOCTYPE html` declaration, followed by an `html` element.

Füge eine `DOCTYPE html`-Deklaration oben im Dokument und ein `html`-Element nach dieser hinzu. Gib dem `html`-Element ein `lang`-Attribut mit `en` als Wert.

# --hints--

Dein Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest ein Leerzeichen nach der `DOCTYPE`-Referenz einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

You should close the `DOCTYPE` declaration with a `>` after the type.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut mit dem Wert `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/gi));
```

Your `DOCTYPE` declaration should be at the beginning of your HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
