---
id: 616968c2c94c8071b349c146
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Jetzt kannst du damit beginnen, der Seite Inhalte hinzuzufügen.

Niste innerhalb des `body` ein `h1`-Element mit dem Text `CSS Color Markers`.

# --hints--

Dein Code sollte ein `h1`-Element haben.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Du solltest ein öffnendes `h1`-Tag haben.

```js
assert(code.match(/<h1\s*>/i));
```

Du solltest ein schließendes `h1`-Tag haben.

```js
assert(code.match(/<\/h1\s*>/i));
```

Dein `h1`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Dein `h1`-Element sollte den Text `CSS Color Markers` haben.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Denke daran, dass die Schreibweise und Rechtschreibung für den `h1`-Text wichtig sind.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
