---
id: 616d3a67ccf800ad94ec89ae
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Du wirst nun einige Elemente hinzufügen, die du eventuell als Farbmarkierungen gestalten kannst.

Füge zuerst ein `div`-Element innerhalb des `body`-Elements hinzu und setze sein `class`-Attribut auf `container`. Stelle sicher, dass das `div` unterhalb des `h1`-Elements steht.

# --hints--

Du solltest ein öffnendes `div`-Tag haben.

```js
assert(code.match(/<div.*>/i));
```

Du solltest ein schließendes `div`-Tag haben.

```js
assert(code.match(/<\/div\s*>/i));
```

Dein `div`-Element sollte sich innerhalb deines `body`-Elements befinden.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Dein `div`-Element sollte nach dem `h1`-Element stehen.

```js
assert.exists(document.querySelector('h1 + div'));
```

Dein `div`-Element sollte ein `class`-Attribut auf `container` gesetzt haben.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
