---
id: 616d50b93ba424d6282c99cf
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Die Hintergrundfarbe wurde angewendet, aber da das `div`-Markierungselement keinen Inhalt besitzt, hat es standardmäßig keine Höhe.

Setze in deiner `.marker`-CSS-Regel die `height`-Eigenschaft auf `25px` und die `width`-Eigenschaft auf `200px`

# --hints--

Deine `.marker`-CSS-Regel sollte eine `width`-Eigenschaft auf `200px` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

Deine `.marker`-CSS-Regel sollte eine `height`-Eigenschaft auf `25px` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
