---
id: 616d4a84b756d9c4b8255093
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Es ist an der Zeit, der Markierung etwas Farbe zu verleihen. Denke daran, dass die Verwendung eines <dfn>color keyword</dfn>, wie `black`, `cyan`, oder `yellow` eine Möglichkeit ist, um einem Element Farbe hinzuzufügen.

Zur Erinnerung: So kannst du die Klasse `freecodecamp` bestimmen:

```css
.freecodecamp {

}
```

Erstelle eine neue CSS-Regel, die die Klasse `marker` bestimmt, und setze ihre `background-color`-Eigenschaft auf `red`.

**Hinweis:** Nach dem Hinzufügen des CSS wirst du keine Änderungen sehen.

# --hints--

Du solltest einen Klassenselektor erstellen, um die `marker`-Klasse zu bestimmen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Deine `.marker`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `red` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
