---
id: 616d55bd160a95e22453a081
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Now that you've got one marker centered with color, it's time to add the other markers.

Füge im `container` `div` zwei weitere `div`-Elemente hinzu und weise jedem die Klasse `marker` zu.

# --hints--

Your first new `div` element should have an opening tag.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

Dein erstes neues `div`-Element sollte ein abschließendes Tag haben.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Dein zweites neues `div`-Element sollte ein öffnendes Tag haben.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

Dein zweites neues `div`-Element sollte ein abschließendes Tag haben.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Deine neuen `div`-Elemente sollten sich innerhalb des `div` mit der Klasse `container` befinden.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

Deine neuen `div`-Elemente sollten beide eine Klasse von `marker` haben.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
