---
id: 616d55bd160a95e22453a081
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Now that you've got one marker centered with color, it's time to add the other markers.

In the `container` `div`, add two more `div` elements and give them each a class of `marker`.

# --hints--

Your first new `div` element should have an opening tag.

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

Your first new `div` element should have a closing tag.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Your second new `div` element should have an opening tag.

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

Your second new `div` element should have a closing tag.

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Your new `div` elements should be within the `div` with the class `container`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

Your new `div` elements should both have a class of `marker`.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
