---
id: 616d595270d902f0e7086e18
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Während du drei separate `div`-Markierungselemente hast, sehen sie wie ein einziges großes Rechteck aus. Du solltes etwas Abstand zwischen ihnen hinzufügen, um es einfacher zu machen, jedes Element zu erkennen.

When the shorthand `margin` property has two values, it sets `margin-top` and `margin-bottom` to the first value, and `margin-left` and `margin-right` to the second value.

Setze in deiner `.marker`-CSS-Regel die `margin`-Eigenschaft auf `10px auto`.

# --hints--

Deine `.marker`-CSS-Regel sollte eine `margin`-Eigenschaft auf `10px auto` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
