---
id: 61764f7e5240eb7ccc7f6a0a
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Erstelle dann eine neue CSS-Regel, die die Klasse `one` auswählt und ihre `background-color`-Eigenschaft auf `red` einstellt.

# --hints--

Du solltest einen Klassen-Selektor verwenden, um die Klasse `one` anzuvisieren.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one'));
```

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `red` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
