---
id: 61764c602bee6974e7790f35
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Um den Markierungen verschiedene Farben zu geben, musst du jeder Markierung eine eindeutige Klasse hinzufügen. Mehrere Klassen können zu einem Element hinzugefügt werden, indem man diese in dem `class`-Attribut auflistet und sie mit einem Leerzeichen trennt. Dies fügt also zum Beispiel `animal`- und `dog`-Klassen zu einem `div`-Element hinzu.

```html
<div class="animal dog">
```

Wenn du einem HTML-Element mehrere Klassen hinzufügst, können die Stile der ersten Klassen, die du aufzählst, von späteren Klassen überschrieben werden.

Füge zunächst die Klasse `one` zum ersten `div`-Markierungselement hinzu.


# --hints--

Du solltest die Klasse `one` zum ersten `div`-Markierungselement hinzufügen.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Deine erste Markierung `div` sollte die Klassen `marker` und `one` haben.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
