---
id: 6176519636a76b810ab1219a
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Erstelle eine CSS-Regel, die die Klasse `two` bestimmt und setze seine `background-color`-Eigenschaft auf `green`.

Erstelle ebenfalls eine getrennte CSS-Regel, die die Klasse `three` bestimmt und setzte seine `background-color` auf `blue`.

# --hints--

Du solltest einen Klassen-Selektor verwenden, um die Klasse `two` zu bestimmen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

Deine `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `green` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

Du solltest einen Klassen-Selektor verwenden, um die Klasse `three` zu bestimmen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

Deine `.three`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `blue` gestetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
