---
id: 617b8e0d93a8d10d9a90e720
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Beachte, dass die `background-color` für deine Markierung noch immer rot ist. Das liegt daran, dass du den Rot-Wert der `rgb`-Funktion auf den Maximalwert von `255`, also 100% Rot, und die Grün- und Blau-Werte auf `0` gesetzt hast.

Verwende jetzt die `rgb`-Funktion, um die anderen Farben festzulegen.

In der CSS-Regel `.two` verwendest du die Funktion `rgb`, um die `background-color` auf den Maximalwert für Grün und `0` für die anderen Werte zu setzen. Verwende in der `.three`-CSS-Regel die `rgb`-Funktion, um die `background-color` auf den maximalen Wert für Blau und `0` für andere Werte zu setzen.

# --hints--

Deine `.two`-CSS-Regel sollte nicht das `green`-Farbschlüsselwort verwenden, um die `background-color`-Eigenschaft festzulegen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

Deine `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(0, 255, 0)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

Deine `.three`-CSS-Regel sollte nicht das `blue`-Farbschlüsselwort verwenden, um die `background-color`-Eigenschaft festzulegen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

Deine `.three`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(0, 0, 255)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
