---
id: 617b65579ce424bf5f02ca73
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Es gibt zwei Hauptfarbenmodelle: Das <dfn>additive</dfn> RGB (Rot, Grün, Blau) Modell für elektronische Geräte und das <dfn>subtractive</dfn> CMYK (Cyan, Magenta, Gelb, Schwarz) Modell für Druck.

In diesem Projekt wirst du mit dem RGB-Modell arbeiten. Das bedeutet, dass die Farben als Schwarz beginnen und sich bei der Zugabe verschiedener Rot-, Grün- und Blauanteile verändern. Eine einfache Möglichkeit, dies zu sehen, ist die `rgb`-CSS-Funktion.

Create a new CSS rule that targets the class `container` and set its `background-color` to black with `rgb(0, 0, 0)`.

# --hints--

You should use a class selector to target the class `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Deine `.container`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 0, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
