---
id: 617bbb6b97a83f6d1f7d6e38
title: Schritt 31
challengeType: 0
dashedName: step-31
---

# --description--

Da du nun mit Sekundärfarben vertraut bist, wirst du lernen, wie du <dfn>tertiary colors</dfn> erstellen kannst. Tertiäre Farben werden durch das Kombinieren einer Primärfarbe mit einer benachbarten Sekundärfarbe erstellt.

Um die tertiäre Farbe Orange zu erstellen, aktualisiere die `rgb`-Funktion in der `.one`-CSS-Regel, so dass rot den maximalen Wert hat und setze grün auf `127`.

# --hints--

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(255, 127, 0)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
