---
id: 617b8b38f32bf2080a140675
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

A function is a piece of code that can take an input and perform a specific action. Die `rgb`-CSS-Funktion nimmt Werte oder <dfn>arguments</dfn> für Rot, Grün und Blau an und erzeugt eine Farbe:

```css
rgb(red, green, blue);
```

Jeder rote, grüne und blaue Wert ist eine Zahl von `0` bis `255`. `0` bedeutet, dass es 0% dieser Farbe gibt und es schwarz ist. `255` bedeutet, dass es 100% dieser Farbe gibt.

Ersetze in der `.one`-CSS-Regel das Schlüsselwort der Farbe `red` mit der `rgb`-Funktion. Setze für die `rgb`-Funktion den Wert für Rot auf `255`, den Wert für Grün auf `0` und den Wert für Blau auf `0`.

# --hints--

Deine `.one`-CSS-Regel sollte nicht das `red`-Farbschlüsselwort verwenden, um die `background-color`-Eigenschaft festzulegen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(255, 0, 0)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
