---
id: 617bd6ec666b1da2587e4e37
title: Schritt 34
challengeType: 0
dashedName: step-34
---

# --description--

Es gibt drei weitere Tertiärfarben: Charteusegrün (Grün + Gelb), Azurblau (Blau + Cyan) und Rosa (Rot + Magenta).

Um Chartreusegrün herzustellen, aktualisiere die `rgb`-Funktion in der `.one`-CSS-Regel, so dass Rot auf `127` und Grün auf den maximalen Wert gesetzt ist.

Aktualisiere für Azurblau die `rgb`-Funktion in der `.two`-CSS-Regel, so dass Grün auf `127` und Blau auf den maximalen Wert gesetzt ist.

Und aktualisiere für Rosa, das manchmal als Pink bezeichnet wird, die `rgb`-Funktion in der `.three`-CSS-Regel, so dass Blau auf `127` und Rot auf den maximalen Wert gesetzt ist.

# --hints--

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(127, 255, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

Deine `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 127, 255)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

Deine `.three`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(255, 0, 127)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
