---
id: 618a0927005553b74bfa05ae
title: Schritt 36
challengeType: 0
dashedName: step-36
---

# --description--

A color wheel is a circle where similar colors, or <dfn>hues</dfn>, are near each other, and different ones are further apart. For example, pure red is between the hues rose and orange.

Two colors that are opposite from each other on the color wheel are called <dfn>complementary colors</dfn>. If two complementary colors are combined, they produce gray. But when they are placed side-by-side, these colors produce strong visual contrast and appear brighter.

Setze in der `rgb`-Funktion f端r die `.one`-CSS-Regel den roten Wert auf das Maximum von `255`, um ein reines Rot zu erzeugen. Setze in der `rgb`-Funktion f端r die `.two`-CSS-Regel die Werte f端r Gr端n und Blau auf das Maximum von `255`, um Cyan zu erzeugen.

# --hints--

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(255, 0, 0)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

Deine `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(0, 255, 255)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
