---
id: 618a00ed1ca871a2b3aca0cb
title: Schritt 35
challengeType: 0
dashedName: step-35
---

# --description--

Nun, da du alle primären, sekundären und tertiären Farben auf einem Farbrad durchgegangen bist, wird es einfacher sein, andere Farbtheoriekonzepte und deren Auswirkungen auf das Design zu verstehen.

Passe zuerst in den CSS-Regeln `.one`, `.two` und `.three` die Werte in der `rgb`-Funktion an, so dass die `background-color` jedes Elements auf reines Schwarz gesetzt ist. Denke daran, dass die `rgb`-Funktion das additive Farbmodell verwendet, in dem Farben als Schwarz beginnen und sich ändern, wenn die Werte von Rot, Grün, und Blau steigen.

# --hints--

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 0, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

Deine `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 0, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

Deine `.three`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 0, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
