---
id: 618a0927005553b74bfa05ae
title: Schritt 36
challengeType: 0
dashedName: step-36
---

# --description--

Ein Farbkreis ist ein Kreis, in dem ähnliche Farben oder <dfn>Farbtöne</dfn> nahe beieinander liegen und unterschiedliche Farben weiter voneinander entfernt sind. Zum Beispiel ist pures rot zwischen den Farbtönen rosa und orange.

Zwei Farben, die sich im Farbkreis gegenüber stehen, werden <dfn>Komplementärfarben</dfn> genannt. Wenn zwei Komplementärfarben kombiniert werden, erzeugen sie die Farbe grau. Wenn sie aber nebeneinander platziert werden, erzeugen diese Farben einen starken visuellen Kontrast und wirken heller.

Setze in der `rgb`-Funktion für die `.one`-CSS-Regel den roten Wert auf das Maximum von `255`, um ein reines Rot zu erzeugen. Setze in der `rgb`-Funktion für die `.two`-CSS-Regel die Werte für Grün und Blau auf das Maximum von `255`, um Cyan zu erzeugen.

# --hints--

Deine `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(255, 0, 0)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

Deine `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft von `rgb(0, 255, 255)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
