---
id: 619b7fd56aa2253778bcf5f7
title: Schritt 48
challengeType: 0
dashedName: step-48
---

# --description--

Das <dfn>HSL</dfn>-Farbmodell, oder Farbton, Sättigung und Helligkeit, ist ein weiterer Weg Farben darzustellen.

Die hsl-CSS-Funktion nimmt 3 Werte auf: Eine Zahl von 0 bis 360 für einen Farbton, einen Prozentsatz von 0 bis 100 für die Sättigung und einen Prozentsatz von 0 bis 100 für Helligkeit.

Wenn du dir einen Farbkreis vorstellst ist der Farbton Rot auf 0 Grad, Grün auf 120 Grad und Blau auf 240 Grad.

Sättigung ist die Intensität einer Farbe von 0%, grau, bis zu 100% für die reine Farbe. Du musst das Prozentzeichen `%` zu den Sättigungs- und Helligkeitswerten hinzufügen.

Die Helligkeit gibt an, wie hell eine Farbe erscheint. Sie reicht von 0 %, also komplett schwarz, bis 100 %, also komplett weiß, wobei 50 % neutral ist.

Verwende in der `.blue`-CSS-Regel die `hsl`-Funktion, um die `background-color`-Eigenschaft auf reines Blau zu ändern. Setze den Farbton auf `240`, die Sättigung auf `100%` und die Helligkeit auf `50%`.

# --hints--

Deine `.blue`-CSS-Regel sollte eine `background-color`-Eigenschaft von `hsl(240, 100%, 50%)` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
