---
id: 619b74fa777a2b2473c94f82
title: Schritt 45
challengeType: 0
dashedName: step-45
---

# --description--

Aktualisiere den CSS-Klassen-Selektor `.two`, sodass er die neue `green`-Klasse auswählt. Aktualisiere auch den `.three`-Klassenselektor, sodass er die neue `blue`-Klasse auswählt.

# --hints--

Dein Code sollte nicht länger einen `.two`-Klassen-Selektor haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

Du solltest einen Klassenselektor verwenden, um die Klasse `green` auszuwählen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

Deine `.green`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 0, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Dein Code sollte nicht länger einen `.three`-Klassen-Selektor haben.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

Du solltest einen Klassenselektor verwenden, um die Klasse `blue` auszuwählen.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Deine `.blue`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgb(0, 0, 0)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
