---
id: 61a489b8579e87364b2d2cdb
title: Schritt 49
challengeType: 0
dashedName: step-49
---

# --description--

Du hast ein paar Möglichkeiten gelernt, flache Farben in CSS zu setzen, aber du kannst auch einen Farbübergang oder <dfn>gradient</dfn> auf ein Element verwenden.

Farbverlauf bedeutet, wenn eine Farbe in eine andere übergeht. Die `linear-gradient`-CSS-Funktion lässt dich die Richtung des Übergangs entlang einer Linie kontrollieren und welche Farben verwendet werden sollen.

Eine Sache, die man sich merken sollte, ist, dass die `linear-gradient` Funktion tatsächlich ein `image`-Element erstellt und normalerweise mit der `background`-Eigenschaft gepaart wird, die ein Bild als einen Wert annehmen kann.

Ändere in der `.red` CSS-Regel die `background-color`-Eigenschaft in `background`.

# --hints--

Deine `.red`-CSS-Regel sollte eine `background`-Eigenschaft mit dem Wert `rgb(255, 0, 0)` haben.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
