---
id: 61a493ead935c148d2b76312
title: Schritt 50
challengeType: 0
dashedName: step-50
---

# --description--

Die `linear-gradient`-Funktion ist sehr flexibel - hier die grundlegende Syntax, die du in diesem Tutorial verwenden wirst:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` ist die Richtung der Linie, die für diesen Übergang verwendet wird. `color1` und `color2` sind Farbargumente, die im Übergang selbst verwendet werden. Diese können jede Art von Farbe sein, einschließlich Farbschlüsselwörter, hex, `rgb`, oder `hsl`.

Nun wirst du einen rot-grünen Farbverlauf entlang einer 90°-Linie auf der ersten Makierung anwenden.

Setze zuerst in die `.red`-CSS-Regel, die `background`-Eigenschaft auf `linear-gradient()` und gebe ihr den Wert `90deg` als `gradientDirection`.

# --hints--

Deine `.red`-CSS-Regel sollte eine `background`-Eigenschaft mit dem Wert `linear-gradient(90deg)` haben.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
