---
id: 61a49d15bdbb5e57cc6fd280
title: Schritt 54
challengeType: 0
dashedName: step-54
---

# --description--

Farbregler erlauben dir festzulegen, wo genau die Farben entlang der Farbverlaufslinie platziert werden. Sie sind eine Längeneinheit wie `px` oder Prozentsätze, die einer Farbe in der `linear-gradient`-Funktion folgen.

In diesem rot-schwarzen Farbverlauf zum Beispiel findet der Übergang von Rot zu Schwarz am 90%-Punkt entlang der Farbverlaufslinie statt, so dass Rot den größten Teil des verfügbaren Platzes einnimmt:

```css
linear-gradient(90deg, red 90%, black);
```

Füge in die `linear-gradient`-Funktion einen `75%`-Farbstopp nach dem ersten roten Farbargument hinzu. Füge den anderen Farbargumenten keinen Farbstopp hinzu.

# --hints--

Deine `.red` CSS-Regel sollte eine `background`-Eigenschaft von `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))` haben.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
