---
id: 61a5be9833e7dc178de2af10
title: Schritt 63
challengeType: 0
dashedName: step-63
---

# --description--

Das sieht besser aus, aber der untere Rand der grünen Markierung muss dunkler sein, um ein wenig mehr Dimension zu kreieren.

Füge in der gleichen `linear-gradient`-Funktion einen Hex-Farbcode mit den Werten `11` für Rot, `6C` für Grün und `31` für Blau als drittes Farbargument hinzu.

# --hints--

Deine `.green` CSS-Regel sollte eine `background`-Eigenschaft auf `linear-gradient(180deg, #55680D, #71F53E, #116C31)` gesetzt haben.

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background: linear-gradient(180deg, #55680D, #71F53E);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
