---
id: 61b093179e7fc020b43b1bb3
title: Schritt 73
challengeType: 0
dashedName: step-73
---

# --description--

<dfn>Opacity</dfn> beschreibt, wie blickdicht oder undurchsichtig etwas ist. Eine massive Wand ist zum Beispiel undurchsichtig und kein Licht kann hindurchdringen. Ein Trinkglas dagegen ist viel transparenter und du kannst durch das Glas hindurch auf die andere Seite schauen.

Mit der `opacity`-CSS-Eigenschaft kannst du kontrollieren, wie blickdicht oder transparent ein Element ist. Mit dem Wert `0`, oder 0%, wird das Element vollständig transparent und bei `1.0`, oder 100%, wird das Element komplett blickdicht, wie es standardmäßig der Fall ist.

Setze in der `.sleeve`-CSS-Regel die `opacity`-Eigenschaft auf `0.5`.

# --hints--

Deine `.sleeve`-CSS-Regel sollte eine `opacity`-Eigenschaft mit dem Wert `0.5` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
