---
id: 61b093329e7fc020b43b1bb5
title: Schritt 75
challengeType: 0
dashedName: step-75
---

# --description--

Du bist bereits mit der Verwendung der `rgb`-Funktion beim Festlegen von Farben vertraut. Um einen Alphakanal zu einer `rgb`- Farbe hinzuzufügen, verwende stattdessen die `rgba`-Funktion.

Die `rgba`-Funktion funktioniert genau wie die `rgb`-Funktion, aber sie nimmt für den Alpha-Kanal eine weitere Zahl von `0` bis `1.0`:

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

Du kannst auch einen Alphakanal mit `hsl`- und `hex`-Farben verwenden. Du wirst bald sehen, wie man das macht.

Verwende die `rgba`-Funktion in der `.sleeve`-Regel, um die `background-color`-Eigenschaft auf reines Weiß mit 50% Deckkraft zu setzen.

# --hints--

Deine `.sleeve`-CSS-Regel sollte eine `background-color`-Eigenschaft auf `rgba(255, 255, 255, 0.5)` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor === 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
