---
id: 61b30995968123ceb6b76167
title: Schritt 88
challengeType: 0
dashedName: step-88
---

# --description--

Beachte, dass die Kanten des Schattens scharf sind. Das liegt daran, dass es einen optionalen `blurRadius`-Wert für die `box-shadow`-Eigenschaft gibt:

```css
box-shadow: offsetX offsetY blurRadius color;
```

Wenn ein `blurRadius`-Wert nicht enthalten ist, ist er standardmäßig `0` und erzeugt scharfe Kanten. Je höher der Wert von `blurRadius`, desto größer ist der Verwischungseffekt.

Füge in der `.green`-CSS-Regel die `box-shadow`-Eigenschaft mit den Werten `5px` für `offsetX`, `5px` für `offsetY`, `5px` für `blurRadius` und `green` für `color` hinzu.

# --hints--

Deine `.green`-CSS-Regel sollte eine zusammenfassende `box-shadow`-Eigenschaft mit dem Wert `5px 5px 5px green` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.boxShadow === 'green 5px 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

--fcc-editable-region--
.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
