---
id: 68e9a3e82da294fb6395b647
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

It is time to begin sprucing up the project with some CSS. You will begin by creating a `body` attribute selector within your CSS.

Next, you will need to set the `display` with a value of `flex`. This will begin the setup of your flexbox.

Then you will need to set `justify-content` and `align-items` to `center`. This completes the vertical and horizontal setup using flexbox properties.

Lastly, give the `body` a `height` of `100vh` and a background color set to `#f0f0f0`.

# --hints--

You should use the `body` element selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

You should add the `display` property within the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.display);
```

You should have a `display` with a value of `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.display, "flex");
```

You should add the `justify-content` within the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.justifyContent);
```

You should have `justify-content` with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.justifyContent, "center");
```

You should add the `align-items` to the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.alignItems);
```

You should have `align-items` with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.alignItems, "center");
```

You should add the `height` property within the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

You should have a `height` property with the value of `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

You should add the `background-color` property within the `body` element selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

You should have a `background-color` property with a value of `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
