---
id: 68e9a6bf1a74721a4e43e217
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Now you are going to remove the default checkbox that is applied by browsers. Inside of your `input[type='checkbox']` selector you are going to add `appearance` with a value of `none`.

Before the `appearance` property was standardised, you would need to use vendor prefixes like `-webkit-appearance` or `-moz-appearance`. However, in modern web browsers these are no longer necessary.

Setting the `appearance` property to `none` will clear the appearance that other browsers apply to checkboxes, allowing you to show the style you want.

Now, since the checkbox is no longer visible you will set a `border` with `2px` `solid` and a hex code of `#f1be32`.

# --hints--

You should add an `appearance` property to your `input[type="checkbox]` selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

You should have an `appearance` property with a value of `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

You should add a `border` property to your `input[type="checkbox]` selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

You should have a `border` property with a value of `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

label {
  display: flex;
  align-items: center;
  gap: 6px;
  margin: 8px 0;
  cursor: pointer;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--

--fcc-editable-region--
}
```
