---
id: 68e9a834dfef752dc35f6f90
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now that your checkboxes change color when the user clicks them, you are going to apply some styling to make it a little more obvious that the checkbox has been selected.

First, you are going to create a new attribute selector of `input[type="checkbox"]` with the previous pseudo-class selector of `:checked` and add the pseudo-element selector of `::after`.
This will specifically target what happens after the user has checked the box.

Next, you are going add `content` with the value of `"✓"`.

Set the `display` to `block` and `text-align` to `center`.

After this, you'll want to set the `font-weight` for the check mark to be easier to see to `bold`. This will increase the visibility of the checkmark.

Now that this is easier to see, you'll want to change the `color` to a value of `white`. This is going to change the color of the checkmark within the checkbox when it is checked.

Lastly, you are going to alter the `line-height` to a value of `20px`.

# --hints--

You should have a combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

You should add `content` to your combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

You should have `content` with a value of `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

You should add `display` to your combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

You should have `display` with a value of `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

You should add `text-align` to your combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

You should have `text-align` with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

You should add `font-weight` to your combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight);
```

You sould have `font-weight` with a value of `bold`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.fontWeight, "bold");
```

You should add `color` to your combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color);
```

You should have `color` with a value of `white`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.color, ["white", "#ffffff", "rgb(255,255,255)"]);
```

You should add `line-height` to your combined attribute selector with a pseudo-class selector and pseudo-element selector of `input[type="checkbox"]:checked::after`.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight);
```

You should have `line-height` with a value of `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.lineHeight, "20px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

label {
  display: flex;
  align-items: center;
  gap: 6px;
  margin: 8px 0;
  cursor: pointer;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  -webkit-appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

label {
  display: flex;
  align-items: center;
  gap: 6px;
  margin: 8px 0;
  cursor: pointer;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type='checkbox']:checked::after {
  content: '✓';
  display: block;
  text-align: center;
  font-weight: bold;
  color: white;
  line-height: 20px;
}
```

