---
id: 68e9a4a47d8a9802cba886b7
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now that the body has been styled, you are going to create a container card for Game Settings panel. To start, you will need to create a class selector for `settings-card`.

This will be where all of your formating for the container will go. You will need to set the `max-width` to `250px`. This will define the overall size of our container.

Next, you will want to set your `padding` to `20px` so that your content has space between it and the border of the container.

After this, you are going to create a rounded edge by setting your `border-radius` to `10px`.

Then you will establish a `box-shadow` with the values of `0 2px 6px rgba(0,0,0,0.2)`. This will create a subtle "lifted" look that will create depth for the container.

# --hints--

You should have a class selector in your css for `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

You should add `max-width` to the `settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

You should have a `max-width` with a value of `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

You should add `padding` to the `settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

You should have a `padding` with a value of `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

You should add `border-radius` to your `settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

You should have a `border-radius` with a value of `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

You should add `box-shadow` to your `settings-card` class selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

You should have a `box-shadow` with the value of `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

--fcc-editable-region--

--fcc-editable-region--
```
