---
id: 68e9a54dd97fd9091e18acce
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

This is where you'll be creating spacing between the checkboxes and the labels. You'll begin by using the selector for `label` in your CSS.

Continue using Flexbox by setting the `display` property of these elements to `flex` and aligning their content vertically in the center by setting the `align-items` property to `center`.

# --hints--

You should have a type selector in your CSS for `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

You should add a `display` property to your `label` selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

You should have a `display` property with a value of `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "flex");
```

You should add the `align-items` property to your `label` selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.alignItems);
```

You should have an `align-items` property with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.alignItems, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

--fcc-editable-region--

--fcc-editable-region--
```
