---
id: 68e9a54dd97fd9091e18acce
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

This is where you'll be creating spacing between the checkboxes and the labels. You'll begin by using the attribute selector for `label` in your CSS.

You should set the `display` for these to `flex` and set `align-items` to `center`.

Now you need to create some space between each `label` and `checkbox` for readability. You will do this by setting a `gap` of `6px` and applying a `margin` of `8px 0`.

Lastly you are going to set a cursor effect for when the mouse is hovering over the `label` by setting `cursor` to `pointer`.

# --hints--

You should have an attribute selector in your CSS of `label`

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

You should add `display` to your `label` attribute selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

You should have `display` with a value of `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "flex");
```

You should add `align-items` to your `label` attribute selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.alignItems);
```

You should have `align-items` with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.alignItems, "center");
```

You should add `gap` to your `label` attribute selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.gap);
```

You should have `gap` with a value of `6px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.gap, "6px");
```

You should add `margin` to your `label` attribute selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

You should have a `margin` with a value of `8px 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px 0px");
```

You should add `cursor` to your `label` attribute selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

You should have `cursor` with a value of `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

--fcc-editable-region--

--fcc-editable-region--
```
