---
id: 68ee86e3ee07b01626a392cc
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

In this step you will assign your newly created class `settings-card` to your `div`.

# --hints--

Your `div` should have a class of `settings-card`.

```js
assert.strictEqual(document.querySelector("div")?.className, "settings-card");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <div>
--fcc-editable-region--
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}
```
