---
id: 68f15606be62da0fd918d547
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Set the `align-items` property to `center` to align the flex items vertically in the middle of the flex container.

The `align-items` property controls how the flex items are positioned along the cross axis, which runs perpendicular to the main axis (vertically by default).

# --hints--

You should add the `align-items` to the `body` type selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.alignItems);
```

You should have `align-items` with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.alignItems, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
--fcc-editable-region--

--fcc-editable-region--
}
```
