---
id: 68f1560aff1214103ab6b59b
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Set the body to have a `height` property with a value of `100vh` and a `background-color` property with a value of `#f0f0f0`.

The `height` of `100vh` makes the `body` take up the full height of the browser viewport, while the light gray background color provides a subtle base for the page.

# --hints--

You should add the `height` property within the `body` type selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

You should have a `height` property with the value of `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

You should add the `background-color` property within the `body` type selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

You should have a `background-color` property with a value of `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
--fcc-editable-region--

--fcc-editable-region--
}
```
