---
id: 68f155d9b02a3d0e2a6a797c
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Flexbox (short for Flexible Box Layout) is a CSS layout system that makes it easier to arrange and align items within a container. It’s especially useful for creating responsive layouts because it can automatically adjust the size and spacing of elements to fit different screen sizes.

In this lesson, you’ll get a quick look at how Flexbox works, and you’ll explore it in greater detail in a future module.

To start your Flexbox setup will need to set the `display` with a value of `flex`.

# --hints--

You should add the `display` property within the `body` type selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.display);
```

You should have a `display` with a value of `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.display, "flex");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--

--fcc-editable-region--
}
```
