---
id: 68f15dea11f2952e494913cc
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Set the `gap` property to `6px` to create space between each `label` and `checkbox` for readability.

The `gap` property controls the spacing between flex items along both the main and cross axes.

# --hints--

You should add the `gap` property to your `label` selector.

```js
assert.isOk(new __helpers.CSSHelp(document).getStyle("label")?.gap);
```

You should have a `gap` property with a value of `6px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.gap, "6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

label {
  display: flex;
  align-items: center;
--fcc-editable-region--

--fcc-editable-region--
}
```
