---
id: 68f15ef0446464325ce1aa45
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Set a `margin` property of `8px 0` to add vertical spacing between elements. This will create consistent spacing above and below each element, helping to separate content and improve readability.

Next, set the `cursor` property to `pointer` on the `label` elements. This will change the cursor to a hand icon when a user hovers over a `label`, signaling that the element is clickable and improving the overall user experience.

# --hints--

You should add the `margin` property to your `label` selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

You should have a `margin` property with a value of `8px 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px 0px");
```

You should add a `cursor` property to your `label` selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

You should have a `cursor` property with a value of `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
}

label {
  display: flex;
  align-items: center;
  gap: 6px;
--fcc-editable-region--

--fcc-editable-region--
}
```
