---
id: 68f15604a5a93d0f8a599ac4
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Set the `justify-content` property to `center`.

The `justify-content` property controls how items are positioned along the main axis of a flex container — which, by default, runs horizontally. Setting it to `center` will evenly distribute space so that all flex items are aligned at the center of the container.

# --hints--

You should add the `justify-content` property within the `body` type selector.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.justifyContent);
```

You should have `justify-content` property with a value of `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.justifyContent, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h2>Game Settings</h2>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Hepatic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  display: flex;
--fcc-editable-region--

--fcc-editable-region--
}
```
