---
id: 5d822fd413a79914d39e98cd
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

CSS では、アスタリスク (*) ですべての要素を選択できます。 すべての要素に境界線を追加するため、`*` セレクターを使用して `border` を `1px solid black` に設定してください。 これは要素の位置とサイズを視覚化するのに役立つ技です。 この設定は後ほど削除します。

# --hints--

`*` セレクターを使用する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

`border` プロパティを使用してすべての要素のスタイルを設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

すべての要素が `1px solid black` の境界線を持つ必要があります。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
