---
id: 5d822fd413a79914d39e98d3
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

`.bb1` コンテナー内に `div` 要素を 4 つネストしてください。 それぞれのクラスを順に `bb1a`、`bb1b`、`bb1c`、`bb1d` としてください。 このビルは 4 つのパーツから成ります。

# --hints--

新しい `div` 要素を 4 つ作成する必要があります。

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

新しい `div` 要素の 1 つに `bb1a` というクラスを設定する必要があります。

```js
assert.exists(document.querySelector('div.bb1a'));
```

新しい `div` 要素の 1 つに `bb1b` というクラスを設定する必要があります。

```js
assert.exists(document.querySelector('div.bb1b'));
```

新しい `div` 要素の 1 つに `bb1c` というクラスを設定する必要があります。

```js
assert.exists(document.querySelector('div.bb1c'));
```

新しい `div` 要素の 1 つに `bb1d` というクラスを設定する必要があります。

```js
assert.exists(document.querySelector('div.bb1d'));
```

新しい `div` 要素を正しい順番で配置する必要があります。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

新しい `div` 要素は `.bb1` 要素の中に入れる必要があります。

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

