---
id: 5d822fd413a79914d39e98de
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

`--building-color1` 変数の下に、新しい変数を作成してください。 新しい変数の名前は `--building-color2`、値は `#66cc99` としてください。 次にその変数を `.bb2` の `background-color` として設定してください。

# --hints--

`--building-color2` という新しいカスタムプロパティ (変数) を定義する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('--building-color2'));
```

`.bb1` クラス内で、`--building-color2` に値 `#66cc99` を代入してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color2').trim(), '#66cc99');
```

`.bb2` の `background-color` を設定する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb2')?.backgroundColor);
```

`background-color` の指定には変数 `--building-color2` を使用してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb2')?.getPropVal('background-color', true), 'var(--building-color2)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
}
--fcc-editable-region--
.bb3 {
  width: 10%;
  height: 55%;
}

.bb4 {
  width: 11%;
  height: 58%;
}

```
