---
id: 5d822fd413a79914d39e98da
title: ステップ 18
challengeType: 0
dashedName: step-18
---

# --description--

1 つ目のビルはだいぶ良くなりました。 では `.background-buildings` コンテナーの中に新しい `div` 要素を 3 つ追加して、そのクラスを順に `bb2`、`bb3`、`bb4` に設定してください。 これが残り 3 つの奥のビルになります。

# --hints--

クラスが `bb2` の `div` を作成してください。

```js
assert.exists(document.querySelector('div.bb2'));
```

クラスが `bb3` の `div` を作成してください。

```js
assert.exists(document.querySelector('div.bb3'));
```

クラスが `bb4` の `div` を作成してください。

```js
assert.exists(document.querySelector('div.bb4'));
```

新しい `div` 要素を 3 つ作成してください。

```js
assert.equal(document.querySelectorAll('div')?.length, 9);
```

これらの `div` 要素は `.background-buildings` 要素の中に入れる必要があります。

```js
assert.equal(document.querySelector('div.background-buildings')?.children?.length, 4);
```

要素を正しい順番で配置してください。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1','bb2') && __t('div.bb2','bb3') && __t('div.bb3','bb4'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

```

