---
id: 5d822fd413a79914d39e98d0
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`body` 内に、`background-buildings` というクラスを持つ `div` 要素を作成してください。 これがビル群のコンテナーになります。

# --hints--

`div` 要素を作成してください。

```js
assert.exists(document.querySelector('div'));
```

`div` 要素は `body` の中にある必要があります。

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

`div` 要素には `background-buildings` のクラスが必要です。

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

