---
id: 5d822fd413a79914d39e98cf
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`body` (ページに表示されている内側のボックス) が見えます。その周りのボックスは `html` 要素です。 `body` をビューポートいっぱいに広げるため、`height` を `100vh` に設定してください。 また、`body` からデフォルトの `margin` を取り除くために `margin` を `0` に設定してください。 最後に、ビューポートの外にはみ出す要素があった場合に表示されるスクロールバーを隠すため、`overflow` プロパティを `hidden` に設定してください。

# --hints--

`body` セレクターを使用してください。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` の `height` を `100vh` に設定する必要があります。

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` の `margin` を `0` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body` の `overflow` プロパティを `hidden` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

