---
id: 5d822fd413a79914d39e98e7
title: ステップ 31
challengeType: 0
dashedName: step-31
---

# --description--

6 つの新しい要素について、`width` と `height` の値を以下の通り設定してください。`.fb1` は `10%` と `60%` に、`.fb2` は `10%` と `40%` に、`.fb3` は `10%` と `35%` に、`.fb4` は `8%` と `45%` に、`.fb5` は `10%` と `33%` に、`.fb6` は `9%` と `38%` に、それぞれ設定してください。

# --hints--

`.fb1` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

`.fb1` セレクターに `10%` の `width` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

`.fb1` セレクターに `60%` の `height` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

`.fb2` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

`.fb2` セレクターに `10%` の `width` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

`.fb2` セレクターに `40%` の `height` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

`.fb3` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

`.fb3` セレクターに `10%` の `width` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

`.fb3` セレクターに `35%` の `height` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

`.fb4` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

`.fb4` セレクターに `8%` の `width` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

`.fb4` セレクターに `45%` の `height` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

`.fb5` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

`.fb5` セレクターに `10%` の `width` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

`.fb5` セレクターに `33%` の `height` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

`.fb6` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

`.fb6` セレクターに `9%` の `width` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

`.fb6` セレクターに `38%` の `height` を設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--

```

