---
id: 5d822fd413a79914d39e98d2
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

奥のビル群のコンテナー内に、クラス `bb1` を持つ `div` をネストしてください。 次に `styles.css` ファイルを開いて、`.bb1` の `width` を `10%` に、`height` を `70%` に設定してください。 「bb」は「background building」の略で、これを 1 つ目のビルにします。

# --hints--

新しい `div` 要素を作成してください。

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

新しい `div` にクラス `bb1` を設定してください。

```js
assert.exists(document.querySelector('div.bb1'));
```

要素のスタイル設定にクラスセレクター `.bb1` を使用する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

`.bb1` 要素の `width` を `10%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

`.bb1` 要素の `height` を `70%` に設定する必要があります。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

