---
id: 5d822fd413a79914d39e98ca
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`DOCTYPE` の下に `html` の開始タグと終了タグを追加して、コードを書き始める場所を用意しましょう。 言語は英語に設定してください。

# --hints--

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 要素が開始タグを持ち、`lang` 属性が `en` という値に設定されている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 要素には終了タグが必要です。

```js
assert(code.match(/<\/html\s*>/));
```

`html` タグは正しい順序になっている必要があります。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

`html` 要素は 1 つだけにしてください。

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
