---
id: 56533eb9ac21ba0edf2244a9
title: Initializing Variables with the Assignment Operator
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

It is common to <dfn>initialize</dfn> a variable to an initial value in the same line as it is declared.

```js
var myVar = 0;
```

Creates a new variable called `myVar` and assigns it an initial value of `0`.

### Real-world Example
Think of it like naming a pet and identifying it at the same time. For example, `var myPet = "Buddy"` is like saying, "I have a new cat, and I'll name it Buddy."

### Step-by-step Breakdown
Combining declaration and assignment (`var a = 9`) does two things at once:
1. **Declaration**: It creates the variable `a`.
2. **Assignment**: It immediately puts the number `9` inside of it.


# --instructions--

Define a variable `a` with `var` and initialize it to a value of `9`.

# --hints--

You should initialize `a` to a value of `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
