---
id: 56533eb9ac21ba0edf2244a8
title: Storing Values with the Assignment Operator
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

In JavaScript, you can store a value in a variable with the <dfn>assignment</dfn> operator (`=`).

```js
myVariable = 5;
```

This assigns the `Number` value `5` to `myVariable`.

If there are any calculations to the right of the `=` operator, those are performed before the value is assigned to the variable on the left of the operator.

```js
var myVar;
myVar = 5;
```

First, this code creates a variable named `myVar`. Then, the code assigns `5` to `myVar`. Now, if `myVar` appears again in the code, the program will treat it as if it is `5`.

### Real-world Example
Think of a variable as a labeled storage box. Using the assignment operator (`=`) is like putting something into that box. For example, `myBox = "Book"` means you have placed a "Book" inside the box labeled `myBox`.

### Step-by-step Breakdown
1. **Declaration (`var myVar;`)**: You tell the computer, "I need a box labeled `myVar`." 
2. **Assignment (`myVar = 5;`)**: You take the number `5` and put it inside the `myVar` box.

### Simple Use Case
If you wanted to keep track of a player's score, you might start with:
```js
var score;
score = 0;
```
Later, when the player earns points, you can update the value in that same box:
```js
score = 10;
```


# --instructions--

Assign the value `7` to variable `a`.

# --hints--

You should not change code above the specified comment.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` should have a value of 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
