---
id: 60a3e3396c7b40068ad6996a
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

到目前为止，你应该熟悉 HTML 页面应具有的基本元素。

添加 `DOCTYPE` 声明、语言设置为英语的 `html` 元素、 `head` 和 `body` 元素。

# --hints--

代码应该有一个 `<!DOCTYPE html>` 声明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你应该有一个值为 `en` 的 `lang` 属性的 `<html>` 开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

代码应该在 `html` 元素中包含一个 `head` 元素。

```js
assert(document.querySelectorAll('head').length === 1);
```

代码应该在 `html` 元素中包含一个 `body` 元素。

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 元素应该在你的 `body` 元素之前。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
