---
id: 60a3e3396c7b40068ad6997e
title: 步骤 21
challengeType: 0
dashedName: step-21
---

# --description--

现在 `.one` 水平居中，但它的上边距推过画布并进入框架的边框，将整个画布向下移动 20 像素。

将 `1px` 的 `padding` 添加到 `.canvas` 元素中，以使 `.one` 元素具有坚实的可推力。

# --hints--

应该将 `padding` 属性设置为 `1px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

`.canvas` 元素的 `padding` 值应为 `1px`。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
