---
id: 60a3e3396c7b40068ad6997c
title: 步骤 19
challengeType: 0
dashedName: step-19
---

# --description--

将 `.one` 的 `background-color` 设置为 `#efb762`。

# --hints--

应该将 `background-color` 属性设置为 `#efb762`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(239, 183, 98)');
assert(hasBackground)
```

`.one` 元素的 `background-color` 值应为 `#efb762`。

```js
const oneBackground = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('background-color');
assert(oneBackground === 'rgb(239, 183, 98)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
