---
id: 60a3e3396c7b40068ad6996c
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在 CSS 盒子模型中，每个 HTML 元素都被视为一个有四个区域的盒子。

想象一下，你从最喜欢的在线零售商那里收到一个盒子——内容是盒子里的物品，或者在我们的例子中，是标题、段落或图像元素。

更改 `<img>` 中的 `src` 属性，从 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 修改为 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`。

# --hints--

`img` 元素的 `src` 应该是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

