---
id: 60a3e3396c7b40068ad6996d
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

内容被称为填充的空间包围，类似于气泡包装将项目与其周围的盒子分开的方式。

将边框想象成运送物品的纸板箱。

将 `src` 属性更改为 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

`img` 元素的 `src` 应该是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
