---
id: 60a3e3396c7b40068ad69977
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

使用填充来调整元素内的间距。

在 `.frame` 中，使用 `padding` 简写属性来增加 `.frame` 和 `.canvas` 元素之间的空间到 `50px`。 简写方式将增加元素边框和画布的顶部、底部、左侧和右侧的空间。

# --hints--

应该将 `padding` 属性设置为 `50px`。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

`.frame` 元素的 `padding` 值应为 `50px`。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
