---
id: 60a3e3396c7b40068ad6997f
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

在画布的顶部、底部、左侧和右侧添加 1 个像素的填充将其尺寸更改为 502 像素 x 602 像素。

将 `padding` 属性替换为 `overflow` 设置为 `hidden` - 将画布更改回其原始尺寸。

# --hints--

应该从 `.canvas` 选择器中删除 `padding` 属性。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

应该将 `overflow` 属性设置为 `hidden`。

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

`.canvas` 元素的 `overflow` 值应为 `hidden`。

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
