---
id: 60a3e3396c7b40068ad69971
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

CSS 的时间到了。

即使你的 `<div>` 没有文本，它仍然被视为包含内容的框。 编写一个使用 `.canvas` 类选择器的 CSS 规则，并将其 `width` 设置为 500 像素。 这是一个 CSS 规则，将类 `card` 的宽度设置为 300 像素：

```css
.card {
  width: 300px;
}
```

# --hints--

代码应该有一个 `.canvas` 选择器。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

应该将 `width` 属性设置为 `500px`。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

`.canvas` 元素的 `width` 应为 `500px`。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
