---
id: 60a3e3396c7b40068ad6996b
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `head` 元素中，添加一个 `meta` 标记，将 `charset` 设置为 `UTF-8`，以及一个 `title` 元素，其值为 `Rothko Painting`。

在 `body` 元素中，添加一个 `img` 元素，其中 `src` 为 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`。
# --hints--

代码应该有一个 `meta` 标记。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` 标记应将 `charset` 设置为 `UTF-8`。

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

代码应该有一个 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` 应该是 `Rothko Painting`。

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

代码应该有一个 `img` 元素。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 元素的 `src` 应该是 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
