---
id: 60a3e3396c7b40068ad69972
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

将值为 `600px` 的 `height` 属性添加到您的 `.canvas` 规则中。

# --hints--

应该将 `height` 属性设置为 `600px`。

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

`.canvas` 元素的 `height` 应为 `600px`。

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
