---
id: 60a3e3396c7b40068ad6996e
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

边距是框外的区域，可用于控制其他框或元素之间的间距。

此处底部元素具有较大的上边距，将其推到页面下方。

现在你了解了 CSS 盒子模型，让我们开始 Rothko 绘画。

移除 `<img>` 元素。

# --hints--

代码中不应包含 `img` 元素。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
