---
id: 60a3e3396c7b40068ad69979
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

在 `.canvas` 元素内添加一个新的 `div` 元素。

为新的 `div` 赋予 `class` 属性，其值为 `one`。 这是你的第一个矩形。

# --hints--

应该创建一个新的 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 3);
```

应该将新的 `div` 元素嵌套在 `.canvas` 元素中。

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

新的 `div` 应该有一个 `class` 属性，其值为 `one`。

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
