---
id: 60a3e3396c7b40068ad6998b
title: 步骤 33
challengeType: 0
dashedName: step-33
---

# --description--

将边距推向一个方向会很有帮助。

在这种情况下，`.one` 元素的下边距将 `.two` 向下推 20 个像素。

在 `.two` 选择器中，使用 `margin` 简写属性将上边距设置为 `0`，水平边距设置为 `auto`，和 `20px` 的下边距。 这将删除其上边距，将其水平居中，并将其下边距设置为 20 像素。

# --hints--

应该将 `margin` 属性设置为 `0 auto 20px`。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

`.two` 元素的 `margin` 值应为 `0 auto 20px`。

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
