---
id: 60a3e3396c7b40068ad6996f
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `body` 中添加一个 `div` 元素。

将 `class` 属性设置为等于 `canvas`。 例如，`<div class="my-div">`。

这将作为你绘画的画布。

# --hints--

代码应该有一个 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 1)
```

`div` 元素应该有一个 `class`，其值为 `canvas`。

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
