---
id: 60eebd07ea685b0e777b5583
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

¡Bienvenido al proyecto de formulario de registro! Comience añadiendo la declaración `!DOCTYPE html` en la parte superior del documento para que el navegador sepa qué tipo de documento lee.

# --hints--

Tu código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio después de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debes definir que el tipo de documento sea `html`.

```js
assert(code.match(/html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/html\s*>/gi));
```

Tu declaración de `DOCTYPE` debe estar al principio de tu HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
