---
id: 60f027099a15b00485563dd2
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Abajo del `DOCTYPE`, añade un elemento `html` con un atributo `lang` establecido a `en`, para que tengas espacio para comenzar a poner código.

# --hints--

Tu declaración `DOCTYPE` debe estar al principio de tu HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Tu elemento `html` debe tener una etiqueta de apertura, y un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu etiqueta `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu etiqueta `html` debe estar en el orden correcto.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Solamente debes de tener una etiqueta `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
