---
id: 60f02e7361b68405e27b62a5
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro del `body`, proporciona un título, añadiendo un `h1` con el texto `Registration Form`.

# --hints--

Debes añadir el `h1` dentro del `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Debes añadir al `h1` el texto de `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
