---
id: 60f02e7361b68405e27b62a5
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro del `body`, proporciona un contexto de encabezado para el contenido, agregando un `h1` con el texto `Registration Form`.

# --hints--

Debes añadir el `h1` dentro del `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Debes dar al `h1` el texto `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
