---
id: 60f0286404aefb0562a4fdf9
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Añade un `title` y un elemento `meta` al `head`. Dale a tu proyecto un título de `Registration Form`, y dale un atributo `charset` con el valor `UTF-8` a tu elemento `meta`.

# --hints--

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

El elemento `title` debe estar dentro del elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Tu proyecto debe tener el título `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
