---
id: 60f027c87bc98f050395c139
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

A continuación, agrega las etiquetas de apertura y cierre de las etiquetas `head` y `body` dentro del elemento `html`.

# --hints--

Debes tener una etiqueta `head` de apertura.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta `head` de cierre.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta `body` de apertura.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta `body` de cierre.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

El elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
