---
id: 60f027099a15b00485563dd2
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Añade una etiqueta de apertura y cierre `html` abajo de `DOCTYPE`para que tengas espacio para comenzar a poner código.

# --hints--

Tu declaración `DOCTYPE` debe estar al principio de tu HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Tu `html` debe tener una etiqueta de apertura.

```js
assert(code.match(/<html\s*>/gi));
```

Tu etiqueta `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Tu etiqueta `html` debe estar en el orden correcto.

```js
assert(code.match(/<html\s*>\s*<\/html\s*>/));
```

Solamente debes de tener una etiqueta `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
