---
id: 60f1a5e2d2c23707a4f9a660
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Ahora, deshacerte de la barra de desplazamiento horizontal, al establecer el `body` `margin` por defecto añadido por algunos navegadores a `0`.

# --hints--

Deberías añadir `margin` dentro del `body` selector de elementos.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

Debes dar a `margin` un valor de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
