---
id: 60f1a5e2d2c23707a4f9a660
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Ahora, deshazte de la barra de desplazamiento horizontal, estableciendo el `margin` por defecto del `body` añadido por algunos navegadores a `0`.

# --hints--

Debes agregar `margin` dentro del selector del elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

Debes dar al `margin` un valor de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
