---
id: 60f1a9cbd23023082e149fee
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Eso es mejor. Ahora, haz que el fondo sea fácil de ver, cambiando el `body` `background-color` a `#1b1b32`. Luego, para ver el texto, cambia el `color` a `#f5f6f7`.

# --hints--

Debes agregar el `background-color` dentro del selector del elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Deberás dar al `background-color` un valor de `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Debes agregar el `color` dentro del selector del elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Deberás dar al `color` un valor de `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
