---
id: 60f1922fcbd2410527b3bd89
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--
La unidad `vh` significa viewport height (altura del viewport), y es relativa al 1% de la altura (`height`) del viewport.

Es hora de engalanar el proyecto con un poco de CSS. Comienza estableciendo para `body` la propiedad `width` a `100%`, y la propiedad `height` a `100vh`.

# --hints--

Debes usar el selector para el elemento `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Debes establecer a `100%` la propiedad `width` del elemento `body`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Debes establecer a `100vh` la propiedad `height` del elemento `body`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
