---
id: 60f5c3e399ff1a05629964e4
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Como sugiere el título, está creando un formulario. Así que después del elemento `p` inserta un `form` con un atributo de `action` "tageting" `https://register-demo.freecodecamp.org`.

# --hints--

Tebes añadir un `form` después de el elemento `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Debes añadir el `form`un attributo de `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Debes dar a la `action` un valor de `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}
```
