---
id: 60f5d2776c854e069560fbe6
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Al primer `fieldset`deber치 tener los campos de nombre, email y contrase침a. Comienza a침adiendo cuetro elementos `label`al primer `fieldset`.

# --hints--

Debes a침adir cuatro elementos `label`.

```js
assert.equal(document.querySelectorAll('label')?.length, 4);
```

Debes aladir los elementos `label` al primer `fieldset`.

```js
assert.equal(document.querySelector('fieldset')?.querySelectorAll('label')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}
```
