---
id: 60f5c3e399ff1a05629964e4
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Como sugiere el título, estás creando un formulario. Por lo tanto, después del elemento `p` inserta un `form` con un atributo de `action` apuntando a `https://register-demo.freecodecamp.org`.

# --hints--

Debes agregar un elemento `form` adyacente al elemento `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Debes dar al `form` un atributo `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Debes dar al `action` un valor de `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
