---
id: 60f5d2776c854e069560fbe6
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

El primer `fieldset` contendrá los campos de nombre, email y contraseña. Empieza agregando cuatro elementos `label` al primer `fieldset`.

# --hints--

Debes agregar cuatro elementos `label`.

```js
assert.equal(document.querySelectorAll('label')?.length, 4);
```

Debes agregar los elementos `label` al primer `fieldset`.

```js
assert.equal(document.querySelector('fieldset')?.querySelectorAll('label')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
