---
id: 60f80e0081e0f2052ae5b505
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Especificar el atributo `type` de un elemento de formulario es importante para que el navegador sepa qué tipo de datos debe esperar. Si el `type` no está especificado, el navegador utilizará por defecto `text`.

Da a los dos primeros elementos `input` un atributo `type` de `text`, al tercero un atributo `type` de `email` y al cuarto un atributo `type` de `password`.

El type `email` solo permite correos electrónicos con un `@` y un `.` en el dominio. El type `password` oculta la entrada, y advierte si el sitio no utiliza HTTPS.

# --hints--

Debes dar al primer elemento `input` un atributo `type` de `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Debes dar al segundo elemento `input` un `type` de `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Debes dar al tercer elemento `input` un atributo `type` de `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Debes dar al cuarto elemento `input` un atributo `type` de `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
