---
id: 60f84ec41116b209c280ba91
title: Paso 23
challengeType: 0
dashedName: step-23
---

# --description--

Con el `type="password"` puedes usar el atributo `pattern` para definir una expresión regular que la contraseña debe coincidir para ser considerada válida.

Agrega un atributo `pattern` al elemento `input` de contraseña para requerir que la entrada coincida con: `[a-z0-5]{8,}`

Lo anterior es una expresión regular que coincide con ocho o más letras minúsculas o los dígitos del `0` al `5`. Luego, elimina el atributo `minlength`, y pruébalo.

# --hints--

Debes dar al elemento `input` de contraseña un atributo `pattern`.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

Debes dar al atributo `pattern` un valor de `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

Debes eliminar el atributo `minlength` del elemento `input` de contraseña.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
