---
id: 60f83e7bfc09900959f41e20
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Algunos valores del atributo `type` vienen con una validación de formulario incorporada. Por ejemplo, `type="email"` requiere que el valor sea una dirección de correo electrónico válida.

Agrega validación personalizada al elemento `input` de contraseña, agregando un atributo `minlength` con un valor de `8`. De este modo se evita que se envíen entradas de menos de 8 caracteres.

# --hints--

Debes dar al elemento `input` de contraseña un atributo `minlength`.

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

Debes dar al atributo `minlength` un valor de `8`.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
