---
id: 60f803d5241e6a0433a523a1
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Como los elementos `label` están en línea por defecto, todos ellos se muestran juntos en la misma línea, haciendo difícil leer su texto. Para que aparezcan en líneas separadas, agrega `display: block` al elemento `label`, y agrega un `margin` de `0.5rem 0`, para separarlos entre sí.

# --hints--

Debes usar el selector `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Debes agregar una propiedad `display` con el valor `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Debes agregar una propiedad `margin` con el valor `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
