---
id: 60f805f813eaf2049bc2ceea
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Anida un elemento `input` dentro de cada `label`. Asegúrate de agregar cada `input` después del texto del `label`, e incluye un espacio después de los dos puntos.

# --hints--

Debes agregar cuatro elementos `input` al elemento `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

Debes anidar los elementos `input` dentro de los elementos `label`.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

Debes agregar el primer `input` después del texto del `label` que es `Enter Your First Name:`, e incluye un espacio después de los dos puntos.

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

Debes agregar el segundo `input` después del texto del `label` que es `Enter Your Last Name:`, e incluye un espacio después de los dos puntos.

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

Debes agregar el tercer `input` después del texto del `label` que es `Enter Your Email:`, e incluye un espacio después de los dos puntos.

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

Debes agregar el cuarto `input` después del texto del `label` que es `Create a New Password:`, e incluye un espacio después de los dos puntos.

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
