---
id: 62cc5b1779e4d313466f73c5
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Anida un elemento `link` de cierre automático dentro del elemento `head`. Dale un atributo `rel` con el valor `stylesheet` y un atributo `href` con un valor de `styles.css`.

# --hints--

Tu código debe tener un elemento `link`.

```js
assert.exists(document.querySelector('link'));
```

No debes cambiar tus etiquetas `head` existentes. Asegúrate de no haber eliminado la etiqueta de cierre.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Debes tener un elemento `link` de autocierre.

```js
assert(document.querySelectorAll('link').length === 1);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Tu elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
