---
id: 62ff8b9dab5ac88e4d3d43a3
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

Siguiendo las mejores prácticas de accesibilidad, vincule los elementos `input` y los elementos `label` utilizando el atributo `for`.

Utiliza `first-name`, `last-name`, `email` y `new-password` como valores para los respectivos atributos `id`.

# --hints--

El primer elemento `input` debe tener un `id` de `first-name`.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

El segundo elemento `input` debe tener un `id` de `last-name`.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

El tercer elemento `input` debe tener un `id` de `email`.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

El cuarto elemento `input` debe tener un `id` de `new-password`.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

El primer elemento `label` debe tener un atributo `for` con el valor `first-name`.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

El segundo elemento `label` debe tener un atributo `for` con el valor `last-name`.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

El tercer elemento `label` debe tener un atributo `for` con el valor `email`.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

El cuarto elemento `label` debe tener un atributo `for` con el valor `new-password`.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
