---
id: 63541ef4f96cd82e8e6c788a
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

El atributo `method` especifica cómo enviar datos de formulario a la dirección URL especificada en el atributo `action`. Los datos del formulario se pueden enviar a través de una solicitud `GET` como parámetros de URL (con `method="get"`) o mediante una solicitud `POST` como datos en el cuerpo de la solicitud (con `method="post"`).

Establezca el atributo `method` para enviar los datos del formulario a través de una solicitud `POST`.

# --hints--

No debe agregar un nuevo elemento `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

El elemento `form` debe tener un atributo `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

El atributo `method` debe establecerse en `post`.

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
