---
id: 67c562286b29447da020d407
title: Build an OTP Generator
challengeType: 25
dashedName: build-an-otp-generator
demoType: onClick
---

# --description--

In this lab, you will generate a 6-digit OTP (One-Time Password) and display it to the user. The OTP will expire after 5 seconds, and a new OTP will be generated when the user clicks the `"Generate New OTP"` button.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should use the `useEffect` hook to manage the countdown timer.
2. Your `OTPGenerator` component should return a `div` element with the class name `container`.
3. The `.container` `div` should contain the following elements:

    - An `h1` element with the text `"OTP Generator"`.
    - An `h2` element that displays the OTP when generated. Initially, it should display the message `"Click 'Generate OTP' to get a code"`.
    - A `p` element that displays the message `"Expires in: X seconds"`, where `X` is the number of seconds remaining before the OTP expires. Initially, it should not be displayed.
    - A `button` element with the text `"Generate OTP"`. When clicked, it should generate a new OTP and start a 5 -second countdown.

4. You should display an the following message when the countdown timer reaches `0`: `OTP expired. Click the button to generate a new OTP.`.
5. You should ensure the countdown timer stops automatically once it reaches `0` seconds to prevent unnecessary updates.
6. The generated OTP should be 6 digits long.

# --hints--

You should use the `useState` hook.

```js
assert.match(code, /useState/);
```

You should render a `div` with the class `.container`.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

The `.container` should contain an `h1` element with the text `"OTP Generator"`.

```js
const headingText = document.querySelector('.container h1').textContent;
assert.strictEqual(headingText, 'OTP Generator');
```

The `.container` should contain an `h2` element that displays the OTP when generated. 

```js
assert.exists(document.querySelector(".container h2"));
```

Initially, the `h2` inside `.container` should display the message `"Click 'Generate OTP' to get a code"`.

```js
assert.strictEqual(document.querySelector('.container h2').textContent.trim(), "Click 'Generate OTP' to get a code");
// write test for when the text is replaced with the OTP
```

The `.container` should contain a `p` element that is initially empty.

```js
const paragraph = document.querySelector('.container p');
assert.exists(paragraph);
```

The `.container` should contain a `button` element with the text `"Generate OTP"`. 

```js
const buttonText = document.querySelector('.container button').textContent;
assert.strictEqual(buttonText, 'Generate OTP');
```

When the button is clicked, it should generate a new OTP and start a 5-second countdown.

```js
```


When the generate OTP button is clicked, the `p` element should display the message `"Expires in: X seconds"`, where `X` is the number of seconds remaining before the OTP expires.

```js
```

When the countdown timer reaches `0`, you should display the message `OTP expired. Click the button to generate a new OTP.`.

```js
```


The generated OTP should be 6 digits long.

```js
const theOTP = generateOTP();
console.log(theOTP);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    cursor: pointer;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect } = React;

export const OTPGenerator = () => {
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
  };

  useEffect(() => {
    let timerId;
    if (isCounting && secondsLeft > 0) {
      timerId = setInterval(() => {
        setSecondsLeft(prevSeconds => prevSeconds - 1);
      }, 1000);
    } else if (secondsLeft === 0) {
      setIsCounting(false);
    }

    return () => clearInterval(timerId);
  }, [isCounting, secondsLeft]);

  return (
    <div className="container">
      <h1>OTP Generator</h1>
      <h2>{otp ? otp : "Click 'Generate OTP' to get a code"}</h2>
      {isCounting ? (
        <p>Expires in: {secondsLeft} seconds</p>
      ) : (
        hasGeneratedOtp && <p>OTP expired. Click the button to generate a new OTP.</p>
      )}
      <button onClick={handleGenerateOtp}>Generate OTP</button>
    </div>
  );
};
```
