---
id: 67c562286b29447da020d407
title: Build an OTP Generator
challengeType: 25
dashedName: lab-otp-generator
demoType: onClick
---

# --description--

In this lab, you will generate a 6-digit OTP (One-Time Password) and display it to the user. The OTP will expire after 5 seconds, and a new OTP will be generated when the user clicks the `"Generate New OTP"` button.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. Your `"Generate New OTP"` button should display a new 6-digit OTP in the `h2` element when clicked.
2. When `"Generate New OTP"` button is clicked, it should start a 5-second countdown.
3. Initially, when no OTP has been generated, the `h2` display should prompt the user to click the button with the message `"Click 'Generate OTP' to get a code"` 
4. You should display a countdown timer next to the OTP that shows how many seconds remain before the OTP expires.
5. The timer should begin at `5` seconds and count down to `0` once the OTP is generated.
6. You should display an `"OTP expired"` message when the countdown timer reaches `0`, prompting the user to generate a new OTP.
7. You should ensure the countdown timer stops automatically once it reaches `0` seconds to prevent unnecessary updates.

# --hints--

You should use the `useState` hook.

```js
assert.match(code, /useState/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Color Picker</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    cursor: pointer;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect } = React;

export const OTPGenerator = () => {
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
  };

  useEffect(() => {
    let timerId;
    if (isCounting && secondsLeft > 0) {
      timerId = setInterval(() => {
        setSecondsLeft(prevSeconds => prevSeconds - 1);
      }, 1000);
    } else if (secondsLeft === 0) {
      setIsCounting(false);
    }

    return () => clearInterval(timerId);
  }, [isCounting, secondsLeft]);

  return (
    <div className="container">
      <h1>OTP Generator</h1>
      <h2>{otp ? otp : "Click 'Generate OTP' to get a code"}</h2>
      {isCounting ? (
        <p>Expires in: {secondsLeft} seconds</p>
      ) : (
        hasGeneratedOtp && <p>OTP expired. Click the button to generate a new OTP.</p>
      )}
      <button onClick={handleGenerateOtp}>Generate OTP</button>
    </div>
  );
};
```
