---
id: 642def66e6a60432c9a0371e
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Setze die `className` des `label`-Elements auf `label` und den `textContent` auf den `name`-Parameter.

# --hints--

Du solltest auf die `className`-Eigenschaft des `label`-Elements zugreifen.

```js
assert.match(code, /label\.className/);
```

Du solltest die `className`-Eigenschaft auf den String `label` setzen.

```js
assert.match(code, /label\.className\s*=\s*('|"|`)label\1/);
```

Du solltest auf die `textContent`-Eigenschaft des `label`-Elements zugreifen.

```js
assert.match(code, /label\.textContent/);
```

Du solltest die `textContent`-Eigenschaft auf den Wert von `name` setzen.

```js
assert.match(code, /label\.textContent\s*=\s*name/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");

  }
}
--fcc-editable-region--
```
