---
id: 642ddfdea4200e313f80a4b6
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Denke daran, dass das `document`-Objekt über eine `.createElement()`-Methode verfügt, die dir ermöglicht dynamisch neue HTML-Elemente zu erstellen.

Deklariere in deiner `createLabel`-Funktion eine `label`-Variable und weise ihr ein neues `div`-Element zu.

# --hints--

Du solltest eine `label`-Variable in deiner `createLabel`-Funktion deklarieren.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*(?:const|let|var)\s+label/);
```

Deine `label`-Variable sollte mit `const` deklariert werden.

```js
assert.match(code, /const\s+createLabel\s*=\s*(\(\s*name\s*\)|name)\s*=>\s*\{\s*const\s+label/);
```

Du solltest die `.createElement()`-Methode des `document`-Objekts verwenden.

```js
assert.match(code, /document\.createElement\(/);
```

Du solltest den String `div` an die `.createElement()`-Methode übergeben.

```js
assert.match(code, /document\.createElement\(\s*('|"|`)div\1\s*\)/);
```

Du solltest dein neues `div`-Element `label` zuweisen.

```js
assert.match(code, /const\s+label\s*=\s*document\.createElement\(\s*('|"|`)div\1\s*\)/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--
window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {

  }
}
--fcc-editable-region--
```
