---
id: 6434552bcc0a951a0a99df3b
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Deine `range`-Funktion erwartet Zahlen, aber deine `start`- und `end`-Werte werden Strings sein (genauer gesagt, sind sie einzelne Zeichen wie z. B. `A`).

Wandle deine `start`- und `end`-Werte in deinem `range()`-Aufruf in Zahlen um, indem du die `.charCodeAt()`-Methode auf sie anwendest und die Zahl `0` als Argument an diese Methode 端bergibst.

# --hints--

Du solltest die Methode `.charCodeAt()` verwenden.

```js
assert.match(code, /\.charCodeAt\(/);
```

Du solltest die `.charCodeAt()`-Methode bei `start` aufrufen.

```js
assert.match(code, /start\.charCodeAt\(/);
```

Du solltest `0` an die `.charCodeAt()`-Methode von `start` 端bergeben.

```js
assert.match(code, /start\.charCodeAt\(\s*0\s*\)/);
```

Du solltest die `.charCodeAt()`-Methode f端r `end` aufrufen.

```js
assert.match(code, /end\.charCodeAt\(/);
```

Du solltest `0` an die `.charCodeAt()`-Methode von `end` 端bergeben.

```js
assert.match(code, /end\.charCodeAt\(\s*0\s*\)/);
```

Du solltest die `.charCodeAt()`-Methoden direkt in deinem `range`-Aufruf verwenden.

```js
assert.match(code, /const\s+charRange\s*=\s*\(\s*start\s*,\s*end\s*\)\s*=>\s*range\(\s*start\.charCodeAt\(\s*0\s*\)\s*,\s*end\.charCodeAt\(\s*0\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
--fcc-editable-region--
const charRange = (start, end) => range(start, end);
--fcc-editable-region--

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
}
```
