---
id: 6434750c53db16218f41e6e1
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Denke daran, dass `range()` ein Array zurückgibt, sodass du Arraymethoden direkt mit Funktionsaufrufen verketten kannst.

Rufe `range()` mit `1` und `99` als Argumente auf und verkette die `.forEach()`-Methode. Übergib der `.forEach()`-Methode einen leeren Callback, der `number` als Parameter verwendet.

# --hints--

Du solltest deine `range()`-Funktion aufrufen.

```js
assert.lengthOf(code.match(/range\(/g), 2);
```

Du solltest `1` als erstes Argument an deinen `range()`-Aufruf übergeben.

```js
assert.match(code, /range\(\s*1/);
```

Du solltest `99` als zweites Argument an deinen `range()`-Aufruf übergeben.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)/);
```

Du solltest die `.forEach()`-Methode mit deinem `range()`-Aufruf verketten.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(/);
```

Du solltest eine Callback-Funktion mit Hilfe von Arrow-Syntax an `.forEach()` übergeben.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Deine Callback-Funktion sollte `number` als einzigen Parameter haben.

```js
assert.match(code, /range\(\s*1\s*,\s*99\s*\)\.forEach\(\s*(\(\s*number\s*\)|number)\s*=>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
--fcc-editable-region--

--fcc-editable-region--
}
```
