---
id: 643498328cb52026123e2b91
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Füge das `input`-Element als untergeordnetes Element an dein `container`-Element an.

Du solltest nun in der Lage sein, die Zellen deiner Tabelle zu sehen.

# --hints--

Du solltest die `.appendChild()`-Methode für dein `container`-Element aufrufen.

```js
assert.lengthOf(code.match(/container\.appendChild\(/g), 2);
```

Du solltest dein `input`-Element an die `.appendChild()`-Methode übergeben.

```js
assert.match(code, /container\.appendChild\(\s*input\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
--fcc-editable-region--
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;

    })
--fcc-editable-region--
  })
}
```
