---
id: 6437124c4c03dd4c8fb35d56
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Deklariere eine `isEven`-Funktion, die einen `num`-Parameter verwendet und `true` zurückgibt, wenn die Zahl gerade ist, und andernfalls `false` zurückgibt. Verwende den Modulo-Operator `%`, um festzustellen, ob eine Zahl gerade oder ungerade ist.

# --hints--

Du solltest eine `isEven`-Variable deklarieren.

```js
assert.match(code, /(?:let|const|var)\s+isEven/);
```

Du solltest `const` verwenden, um deine `isEven`-Variable zu deklarieren.

```js
assert.match(code, /const\s+isEven/);
```

Deine `isEven`-Variable sollte eine Funktion sein.

```js
assert.isFunction(isEven);
```

Deine `isEven`-Funktion sollte eine Arrow-Syntax verwenden.

```js
assert.match(code, /const\s+isEven\s*=\s*(\([^)]*\)|[^\s()]+)\s*=>/);
```

Your `isEven` function should have a `num` parameter.

```js
assert.match(code, /const\s+isEven\s*=\s*(\(\s*num\s*\)|num)\s*=>/);
```

Dein `isEven`-Funktion sollte den modulo-Operator `%` verwenden.

```js
assert.match(isEven.toString(),  /%/);
```

Deine `isEven`-Funktion sollte `true` für gerade Zahlen zurückgeben.

```js
assert.isTrue(isEven(2));
assert.isTrue(isEven(1000));
assert.isTrue(isEven(42));
```

Deine `isEven`-Funktion sollte `false` für ungerade Zahlen zurückgeben.

```js
assert.isFalse(isEven(1));
assert.isFalse(isEven(333));
assert.isFalse(isEven(777777777));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
const sum = nums => nums.reduce((acc, el) => acc + el, 0);

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
