---
id: 643715013330824ecaa70442
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Deine n채chste Funktion wird den Medianwert einer Reihe von Zahlen berechnen. Beginne damit, die `median`-Arrow-Funktion zu deklarieren, die einen `nums`-Parameter annimmt.

In the function, declare a `sorted` variable and assign it the value of sorting a copy of the `nums` array.

Du solltest die `slice()`-Methode verwenden, um eine flache Kopie des Arrays zu erstellen.

# --hints--

Du solltest eine `median`-Variable deklarieren.

```js
assert.match(code, /(?:let|const|var)\s+median/);
```

Du solltest `const` verwenden, um deine `median`-Variable zu deklarieren.

```js
assert.match(code, /const\s+median/);
```

Deine `median`-Variable sollte eine Funktion sein.

```js
assert.isFunction(median);
```

Deine `median`-Funktion sollte eine Arrow-Syntax verwenden.

```js
assert.match(code, /const\s+median\s*=\s*\(?/);
```

Your `median` function should have a `nums` parameter.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)/);
```

Deine `median`-Funktion sollte keine implizite R체ckgabe verwenden.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{/);
```

Deine `median`-Funktion sollte eine `sorted`-Variable enthalten.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*(?:let|var|const)\s+sorted/);
```

Du solltest `const` verwenden, um deine `sorted`-Variable zu deklarieren.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted/);
```

You should use `.slice()` to assign a copy of the `nums` array to `sorted`.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)/);
```

Du solltest die `.sort()`-Methode mit deiner `.slice()`-Methode verbinden.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(/);
```

Du solltest eine Callback-Funktion an deine `sort`-Methode 체bergeben, um die Zahlen genau in aufsteigender Reihenfolge zu sortieren. Verwende zur Verdeutlichung eine implizite R체ckgabe.

```js
assert.match(code, /const\s+median\s*=\s*(\(\s*nums\s*\)|nums)\s*=>\s*\{\s*const\s+sorted\s*=\s*nums\.slice\(\s*\)\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>\s*a\s*-\s*b\s*\)\s*\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Functional Programming Spreadsheet</title>
  </head>
  <body>
    <div id="container">
      <div></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
#container {
  display: grid;
  grid-template-columns: 50px repeat(10, 200px);
  grid-template-rows: repeat(11, 30px);
}

.label {
  background-color: lightgray;
  text-align: center;
  vertical-align: middle;
  line-height: 30px;
}
```

```js
const isEven = num => num % 2 === 0;
const sum = nums => nums.reduce((acc, el) => acc + el, 0);
const average = nums => sum(nums) / nums.length;

--fcc-editable-region--

--fcc-editable-region--

const range = (start, end) => Array(end - start + 1).fill(start).map((element, index) => element + index);
const charRange = (start, end) => range(start.charCodeAt(0), end.charCodeAt(0)).map(code => String.fromCharCode(code));

window.onload = () => {
  const container = document.getElementById("container");
  const createLabel = (name) => {
    const label = document.createElement("div");
    label.className = "label";
    label.textContent = name;
    container.appendChild(label);
  }
  const letters = charRange("A", "J");
  letters.forEach(createLabel);
  range(1, 99).forEach(number => {
    createLabel(number);
    letters.forEach(letter => {
      const input = document.createElement("input");
      input.type = "text";
      input.id = letter + number;
      input.ariaLabel = letter + number;
      container.appendChild(input);
    })
  })
}
```
