---
id: 5d822fd413a79914d39e98cd
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在 CSS 中，你可以使用星號定位所有內容。 使用 `*` 選擇器爲所有內容添加邊框，併爲其提供 `1px solid black` 的 `border`。 這是一個有助於可視化元素的位置及其大小的技巧。 你稍後會刪除這個。

# --hints--

你應該使用 `*` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你應該使用 `border` 屬性來設置所有元素的樣式。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

所有元素都應該有一個 `1px solid black` 的邊框。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
