---
id: 5d822fd413a79914d39e98ce
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

還要將 `border-box` 的 `box-sizing` 添加到所有內容。 這將使你添加的邊框不會給你的元素增加任何尺寸。

# --hints--

你應該使用 `box-sizing` 屬性。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

你應該使用現有的 `*` 選擇器。

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

所有元素都應具有 `box-sizing` 爲 `border-box`。

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

