---
id: 5d822fd413a79914d39e98d7
title: 步驟 15
challengeType: 0
dashedName: step-15
---

# --description--

要使用變量，請將變量名放在括號中，並在它們前面加上 `var`，例如：`var(--variable-name)`。 添加變量作爲 `.bb1a` 類的 `background-color` 屬性的值。 你給變量的任何值都將應用於你使用它的任何屬性。 在這個例子中，你的變量的值爲 `#999`。 因此 `#999` 將用作 `background-color` 屬性的值。

# --hints--

`bb1a` 元素的 `background-color` 應該被設置。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

你應該使用 `var(--building-color1)` 來設置 `.bb1a` 元素的 `background-color`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
