---
id: 5d822fd413a79914d39e98da
title: 步驟 18
challengeType: 0
dashedName: step-18
---

# --description--

你的第一棟建築現在看起來不錯。 在 `.background-buildings` 容器中嵌套三個新的 `div` 元素，並按順序爲它們提供 `bb2`、`bb3` 和`bb4` 類。 這將是另外三座作爲背景的建築物。

# --hints--

你應該創建一個具有 `bb2` 類的 `div`。

```js
assert.exists(document.querySelector('div.bb2'));
```

你應該創建一個具有 `bb3` 類的 `div`。

```js
assert.exists(document.querySelector('div.bb3'));
```

你應該創建一個具有 `bb4` 類的 `div`。

```js
assert.exists(document.querySelector('div.bb4'));
```

你應該創建三個新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div')?.length, 9);
```

你應該將這些 `div` 元素放在 `.background-buildings` 元素中。

```js
assert.equal(document.querySelector('div.background-buildings')?.children?.length, 4);
```

你應該以正確的順序放置元素。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1','bb2') && __t('div.bb2','bb3') && __t('div.bb3','bb4'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

```

