---
id: 5d822fd413a79914d39e98d4
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

爲你的建築物的 `width` 和 `height` 屬性賦予這些值：`.bb1a` 的是 `70%` 和 `10%`、`.bb1b` 的是 `80%` 和 `10%`、`.bb1c` 的是 `90%` 和 `10%`、`.bb1d` 的是 `100%` 和 `70%`。 請記住，這些百分比是相對於父級的，並注意高度將加起來爲 100%——垂直填充容器。

# --hints--

你應該使用類選擇器來設置 `.bb1a` 的樣式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

你應該將 `.bb1a` 的 `width` 設置爲 `70%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

你應該將 `.bb1a` 的 `height` 設置爲 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

你應該使用類選擇器來設置 `.bb1b` 的樣式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

你應該將 `.bb1b` 的 `width` 設置爲 `80%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

你應該將 `.bb1b` 的 `height` 設置爲 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

你應該使用類選擇器來設置 `.bb1c` 的樣式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

你應該將 `.bb1c` 的 `width` 設置爲 `90%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

你應該將 `.bb1c` 的 `height` 設置爲 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

你應該使用類選擇器來設置 `.bb1d` 的樣式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

你應該將 `.bb1d` 的 `width` 設置爲 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

你應該將 `.bb1d` 的 `height` 設置爲 `70%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

