---
id: 5d822fd413a79914d39e98ca
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `DOCTYPE` 下方添加開始和結束 `html` 標籤，以便你可以開始放置一些代碼。 確保將語言設置爲英語。

# --hints--

你的 `DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你的 `html` 元素應該有一個開始標籤，並且 `lang` 屬性的值爲 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的 `html` 標籤應該有正確的順序。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

你應該只有一個 `html` 元素。

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
