---
id: 5d822fd413a79914d39e98cf
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

你可以看到 `body` （它是你的頁面最裏面的​​框）；它周圍的框是 `html` 元素。 通過給它一個 `100vh` 的 `height` 來讓你的 `body` 填滿整個視口。 通過將 `margin` 設置爲 `0`，從 `body` 中刪除默認的 `margin`。 最後，將 `overflow` 屬性設置爲 `hidden` 以隱藏任何超出視口時出現的滾動條。

# --hints--

你應該使用 `body` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你的 `body` 的 `height` 應爲 `100vh`。

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

你的 `body` 的 `margin` 應爲 `0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

你的 `body` 的 `overflow` 屬性應爲 `hidden`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

