---
id: 5d822fd413a79914d39e98d3
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

在 `.bb1` 容器中嵌套四個 `div` 元素。 按順序爲它們添加類 `bb1a`、`bb1b`、`bb1c` 和 `bb1d`。 這座建築將有四個部分。

# --hints--

你應該創建四個新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

你應該給其中一個新的 `div` 元素一個 `bb1a` 類。

```js
assert.exists(document.querySelector('div.bb1a'));
```

你應該給其中一個新的 `div` 元素一個 `bb1b` 類。

```js
assert.exists(document.querySelector('div.bb1b'));
```

你應該給其中一個新的 `div` 元素一個 `bb1c` 類。

```js
assert.exists(document.querySelector('div.bb1c'));
```

你應該給其中一個新的 `div` 元素一個 `bb1d` 類。

```js
assert.exists(document.querySelector('div.bb1d'));
```

你應該以正確的順序放置新的 `div` 元素。

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

你應該將新的 `div` 元素放在 `.bb1` 元素中。

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

