---
id: 5d822fd413a79914d39e98cc
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在 `head` 中，嵌套一個 `meta` 元素，其中 `charset` 爲 `UTF-8`，一個標題文本爲 `City Skyline` 的 `title` 元素，以及一個鏈接 `styles.css` 文件的 `link` 元素。

# --hints--

你應該在 `head` 元素中創建一個 `meta` 元素。

```js
assert.exists(document.querySelector('head > meta'));
```

你應該給 `meta` 標籤一個值爲 `UTF-8` 的 `charset`。

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

你的代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 元素應該在 `head` 元素內。

```js
assert.exists(document.querySelector('head > title'));
```

你的項目的標題應爲 `City Skyline`。

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

你的代碼應該有一個 `link` 元素。

```js
assert.match(code, /<link/)
```

你不應更改現有的 `head` 標籤。 確保你沒有刪除結束標籤。

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

你應該有一個自閉合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
