---
id: 5d822fd413a79914d39e98d2
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

在背景建築物容器中嵌套一個具有 `bb1` 類的 `div`。 打開你的 `styles.css` 文件，將 `.bb1` 的 `width` 設置爲 `10%`，`height` 設置爲 `70%`。 “bb” 代表“背景建築”，這將是你的第一棟建築。

# --hints--

你應該創建一個新的 `div` 元素。

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

你應該給新的 `div` 一個 `bb1` 類。

```js
assert.exists(document.querySelector('div.bb1'));
```

你應該使用 `.bb1` 類選擇器來設置元素的樣式。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

你應該將 `.bb1` 元素的 `width` 設置爲 `10%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

你應該將 `.bb1` 元素的 `height` 設置爲 `70%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

