---
id: 5f3313e74582ad9d063e3a38
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. The `title` element's text should be `Cafe Menu`.

# --hints--

You should have a `head` element.

```js
assert.exists(document.querySelector('head'));
```

You should have a `title` element.

```js
assert.exists(document.querySelector('title'));
```

Your `title` element should be nested in your `head` element.

```js
assert.exists(document.querySelector('head > title'));
```

Your `title` element should have the text `Cafe Menu`. You may need to check your spelling.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
