---
id: 5f3326b143638ee1a09ff1e3
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

You should have a `body` element.

```js
assert.exists(document.querySelector('body'));
```

You should not change your `head` element. Make sure you did not delete your closing tag.

```js
assert.exists(document.querySelector('head'));
```

Your `body` element should come after your `head` element.

```js
assert.strictEqual(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

