---
id: 5f356ed69db0a491745e2bb6
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

To apply the class's styling to the `div` element, remove the `id` attribute and add a `class` attribute to the `div` element's opening tag. Make sure to set the class value to `menu`.

# --hints--

Your `div` should still render. Make sure you haven't malformed the `<div>` tag.

```js
assert.lengthOf(document.querySelectorAll('div'),1);
```

Your `div` element should have the `menu` class.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

Your `div` element should no longer have the `menu` id.

```js
assert.lengthOf(document.querySelectorAll('div#menu'),0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

