---
id: 5dc17dc8f86c76b9248c6eb5
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

التعليق يسمح لك بترك الرسائل دون التأثير على العرض في المتصفح. كما يسمح لك بجعل الكود غير نشط. يبدأ التعليق في HTML بـ `<!--`، يحتوي على أي عدد من الأسطر ، وينتهي بـ `-->`. على سبيل المثال، التعليق `<!-- TODO: Remove h1 -->` يحتوي على النص `TODO: Remove h1`.

قم إضافة تعليق أعلى عنصر `p` مع هذا النص:

`TODO: Add link to cat photos`

# --hints--

يجب أن يبدأ تعليقك بـ `<!--`. إنك تفتقد واحدًا أو أكثر من الرموز التي تحدد بداية التعليق.

```js
assert(code.match(/<!--/));
```

ينبغي أن ينتهي تعليقك بـ `-->`.  إنك تفتقد واحدًا أو أكثر من الرموز التي تحدد نهاية التعليق.

```js
assert(code.match(/-->/));
```

لا ينبغي أن يحتوي الكود الخاص بك على رموز opening/closing اضافية للتعليق. لديك `<!--` إضافي أو `-->` معروض في المتصفح.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

يجب أن يحتوي تعليقك على النص `TODO: Add link to cat photos`.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

يجب أن يكون تعليقك فوق عنصر `p`. لقد وضعتهم بترتيب خاطئ.

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

