---
id: 5dc1798ff86c76b9248c6eb3
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

يتم استخدام عناصر العنوان heading من `h1` إلى `h6` للدلالة على أهمية المحتوى أدناه. كلما قل العدد، كلما زادت الأهمية، إذاً عناصر `h2` أقل أهمية من عناصر `h1`. استخدم عنصر `h1` واحد في كل صفحة وضع العناوين (headings) الأقل أهمية تحت العناوين ذات الأهمية أعلى.

تحت عنصر `h1` ، أضف عنصر `h2` مع هذا النص:

`Cat Photos`

# --hints--

يجب أن يحتوي العنصر `h1` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

يجب أن يحتوي العنصر `h1` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/h1\>/));
```

يجب أن يكون لديك عنصر `h1` واحد فقط. احذف الزيادة.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

النص في عنصر `h1` يجب أن يكون 'CatPhotoApp'. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

يجب أن يحتوي العنصر `h2` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

يجب أن يحتوي العنصر `h2` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/h2\>/));
```

النص في عنصر `h2` يجب أن يكون 'Cat Photos'. ضع النص `Cat Photos` فقط بين علامات الافتتاح والإغلاق لـ `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

يجب أن يكون عنصر `h2` الخاص بك تحت عنصر `h1`. عنصر `h1` له أهمية أكبر ويجب أن يكون فوق عنصر `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

