---
id: 5dc174fcf86c76b9248c6eb2
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

عناصر HTML لديها علامات فتح (opening tags) مثل `<h1>` وإغلاق (closing tags) مثل `</h1>`.

تيقن من أن نص العنصر يقع بين سمات أفتاح (open tag) و أغلاق (closing tag).

ابحث عن عنصر `h1` وغيّر نصه إلى:

`CatPhotoApp`

# --hints--

النص `CatPhotoApp` يجب أن يكون موجوداً في الكود. قد تحتاج إلى التحقق من الإملاء.

```js
assert(code.match(/catphotoapp/i));
```

يجب أن يحتوي العنصر `h1` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

يجب أن يحتوي العنصر `h1` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/h1\>/));
```

لديك أكثر من عنصر واحد `h1`. قم بإزالة عنصر `h1` الإضافي.

```js
assert(document.querySelectorAll('h1').length === 1);
```

النص في عنصر `h1` يجب أن يكون `CatPhotoApp`. إما أن تكون قد حذفت النص، أو لديك خطأ إملائي، أو أنه ليس بين علامتي فتح وإغلاق العنصر `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

يبدو أنك تستخدم ملحق (extension) المتصفح الذي يعدل الصفحة. تيقن من إيقاف جميع ملحقات المتصفح.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

