---
id: 5dc17d3bf86c76b9248c6eb4
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

يتم استخدام عنصر `p` لإنشاء فقرة من النصوص على مواقع الويب. قم بإنشاء عنصر `p` تحت عنصر `h2` الخاص بك وقم بإعطائه النص التالي:

`See more cat photos in our gallery.`

# --hints--

عنصرك `p` ينبغي أن يكون له علامة أفتاح (opening tag). Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('p'));
```

يجب أن يحتوي العنصر `p` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/p\>/));
```

يجب أن يكون نص عنصرك `p` بقيمة `See more cat photos in our gallery.` إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/see more cat photos in our gallery\.?$/i));
```

يجب أن يكون عنصر `p` الخاص بك تحت عنصر `h2`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```

