---
id: 5dc23f9bf86c76b9248c6eba
title: الخطوة ٧
challengeType: 0
dashedName: step-7
---

# --description--

يمكنك إضافة صور إلى موقع الويب الخاص بك باستخدام عنصر `img`. `img` عناصر تحتوي على opening tag بدون closing tag. الـ tag لعنصر بدون closing tag معروفة باسم <dfn>self-closing tag</dfn>.

أضف عنصر `img` تحت عنصر `p`. في هذه المرحلة، لن تظهر أي صورة في المتصفح.

# --hints--

يجب أن يحتوي العنصر `img` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('img'));
```

لا يجب أن يحتوي العنصر `img` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(!code.match(/<\/img\>/));
```

يجب أن يكون لديك عنصر `img` واحد فقط. قم بإزالة أي زيادات.

```js
assert(document.querySelectorAll('img').length === 1);
```

يجب أن يكون عنصر `img` الخاص بك تحت عنصر `p`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

