---
id: 5dc24614f86c76b9248c6ebd
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

يمكنك الربط إلى صفحة أخرى باستخدام عنصر (anchor) الرابط (`a`). على سبيل المثال، `<a href='https://freecodecamp.org'></a>` سيحيل علي `freecodecamp.org`.

قم بإضافة عنصر anchor بعد الفقرة التي تربط بـ `https://freecatphotoapp.com`. في هذه المرحلة، لن يظهر الرابط في المعاينة.

# --hints--

يجب أن يحتوي العنصر (`a`) على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('a'));
```

يجب أن يحتوي العنصر (`a`) على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/a\>/));
```

يجب أن يكون عنصر (`a`) الخاص بك تحت عنصر `p`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

عنصر (`a`) الخاص بك ليس لديه سمة `href`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

(`a`) الخاص بك يجب أن يحيل الي `https://freecatphotoapp.com`. إما أنك حذفت الـ URL أو لديك خطأ إملائي.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

على الرغم من أنك قمت بتعيين سمة الرابط ('a') `href` إلى الرابط الصحيح، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

