---
id: 5dc24073f86c76b9248c6ebb
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>attributes</dfn> اي السمات، هي كلمات خاصة تستخدم داخل opening tag لعنصر للتحكم في سلوك العنصر. السمة `src` في عنصر `img` تحدد عنوان URL للصورة (حيث توجد الصورة). مثال على عنصر `img` باستخدام `src` كسمة: `<img src="https://www.example.com/the-image.jpg">`.

داخل عنصر `img` الموجود، أضف خاصية `src` مع هذا الرابط:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `img`. ربما تكون قد قمت بإزالة عنصر `img` أو لم تحيط قيمة سمة `src` بعلامات الاقتباس.

```js
assert(document.querySelector('img'));
```

يجب أن يحتوي عنصر `img` الخاص بك علي سمة `src`. إما أنك حذفت السمة أو لديك خطأ إملائي. تأكد من وجود مسافة بين اسم العنصر واسم السمة.

```js
assert(document.querySelector('img').src);
```

يجب تعيين عنصر `img` الخاص بك سمة `src` بقيمة '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. إما أنك حذفت الـ URL أو لديك خطأ إملائي. حالة (case) عنوان الـ URL مهمة.

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

على الرغم من أنك قمت بتعيين عنصر `img` سمة `src` إلى عنوان الـ URL الصحيح، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

