---
id: 5dc24165f86c76b9248c6ebc
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

يجب أن تحتوي جميع عناصر `img` على سمة `alt`. يتم استخدام نص السمة `alt` لقراء الشاشة (screen readers) لتسهيل الوصول (accessibility) ويتم عرضه إذا فشل تحميل الصورة. على سبيل المثال، `<img src="cat.jpg" alt="A cat">` يحتوي على سمة `alt` مع النص `A cat`.

داخل عنصر `img` ، أضف سمة `alt` مع هذا النص:

`A cute orange cat lying on its back`

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `img`. لقد قمت بإزالة عنصر `img` في خطوة سابقة.

```js
assert(document.querySelector('img'));
```

عنصر `img` الخاص بك لا يحتوي على سمة `alt`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

تم تعيين قيمة السمة `alt` للعنصر `img` بقيمة آخري بخلاف 'A cute orange cat lying on its back'. تأكد من أن قيمة السمة `alt` محاطة بعلامات اقتباس.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

