---
id: 5ddbd81294d8ddc1510a8e56
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

يجب وضع نص الرابط بين علامات فتح وإغلاق عنصر الـ (`a`). على سبيل المثال، `<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` هو رابط مع النص `click here to go to freeCodeCamp.org`.

أضف النص `link to cat pictures` إلى عنصر الرابط. هذا سيصبح نص الرابط.

# --hints--

يجب أن يحتوي العنصر (`a`) على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('a'));
```

يجب أن يحتوي العنصر (`a`) على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/a\>/));
```

يجب أن يكون نص عنصرك (`a`) بقيمة `link to cat pictures`. تيقن من وضع نص الرابط بين opening tag و closing tag لعنصر (`a`).

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

