---
id: 5dfa22d1b521be39a3de7be0
title: الخطوة ١٢
challengeType: 0
dashedName: step-12
---

# --description--

في الخطوة السابقة حولت الكلمات `link to cat pictures` إلى رابط عن طريق وضعها بين علامات أفتاح (opening tag) و إغلاق (closing tag) الرابط (`a`). يمكنك فعل الشيء نفسه للكلمات داخل عنصر، مثل عنصر `p`.

في نص عنصرك `p`، تحويل الكلمات `cat photos` إلى رابط إلى `https://freecatphotoapp.com` عن طريق إضافة الكلمات داخل علامات الأفتاح والإغلاق الرابط (`a`) بين تلك الكلمات.


# --hints--

يجب عليك أدخال عنصر `a` داخل عنصر `p`.

```js
assert($('p > a').length);
```

يجب أن يكون سمة الرابط `href` بقيمة `https://freecatphotoapp.com`. إما أنك حذفت قيمة `href` أو لديك خطأ إملائي.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

يجب أن يكون نص الرابط `cat photos`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

بعد دمج عنصر الرابط (`a`)، محتوى العنصر `p` الوحيد المرئي في المتصفح يجب أن يكون `Click here to view more cat photos.` تحقق مرتين من النص أو المسافات أو علامات الترقيم لكل من `p` وعنصر الرابط المدمج.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

