---
id: 5dfa2407b521be39a3de7be1
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

قم بإضافة سمة `target` مع القيمة `_blank` إلى opening tag العنصر (`a`) حتى يفتح الرابط في علامة تبويب جديدة.

# --hints--

يجب أن يكون للعنصر `p` الخاص بك رابط مدمج (`a`) مع النص `cat photos`. ربما تكون قد حذفته أو لديك خطأ إملائي.

```js
const anchor = $('p > a');
assert(
  anchor.length &&
    anchor[0].innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

عنصر (`a`) الخاص بك ليس لديه سمة `target`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
assert(document.querySelector('a').hasAttribute('target'));
```

قيمة السمة `target` يجب أن تكون `_blank`. إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن قيم السمة يجب أن تكون محاطة بعلامات اقتباس.

```js
assert(document.querySelector('a').getAttribute('target') === '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

