---
id: 5dfa37b9eacea3f48c6300b0
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

بعد عنصر `h3` مع نص `Things cats love:`، أضف عنصر قائمة غير مرتبة (`ul`). لاحظ أنه لن يتم عرض أي شيء في هذه المرحلة.

# --hints--

يجب أن يحتوي العنصر `ul` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('ul'));
```

يجب أن يحتوي العنصر `ul` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/ul>/));
```

عنصر `ul` الخاص بك يجب أن يكون أعلى closing tag عنصر `section` الثاني.

```js
const secondSectionLastElemNode = $('main > section')[1].lastElementChild;
assert(secondSectionLastElemNode.nodeName === 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

