---
id: 5dfb655eeacea3f48c6300b3
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

عنصر `figure` يمثل محتوى قائم بذاته وسيسمح لك بربط صورة مع الشرح.

ادمج الصورة التي أضفتها للتو ضمن عنصر `figure`.

# --hints--

يجب أن يحتوي العنصر `figure` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

يجب أن يحتوي العنصر `figure` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/figure\>/));
```

يجب أن يكون هناك عنصر `figure` أعلى closing tag عنصر الثاني `section`.

```js
assert($('section')[1].lastElementChild.nodeName === 'FIGURE');
```

عنصر `img` للـ lasagna يجب أن يدمج في عنصر `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

