---
id: 5dfb6a35eacea3f48c6300b4
title: الخطوة 24
challengeType: 0
dashedName: step-24
---

# --description--

يتم استخدام عنصر التسمية التوضيحية الشكلية (`figcaption`) لإضافة تسميات توضيحية لوصف الصورة التي بداخل عنصر `figure`. على سبيل المثال، `<figcaption>A cute cat</figcaption>` تضيف الشرح التوضيحي `A cute cat`.

بعد الصورة التي في عنصر `figure` ، أضف عنصر `figcaption` مع تعيين النص إلى:

`Cats love lasagna.`

# --hints--

يجب أن يحتوي العنصر `figcaption` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

يجب أن يحتوي العنصر `figcaption` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

عنصر `figcaption` يجب أن يدمج في عنصر `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

عنصر `img` للـ lasagna يجب أن يدمج في عنصر `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

يجب أن يكون عنصر `figcaption` المدمج في عنصر `figure` تحت عنصر `img`. لقد وضعتهم بترتيب خاطئ.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

يجب أن يكون نص عنصر `figcaption` الخاص بك هو `Cats love lasagna.` إما أنك حذفت النص أو لديك خطأ إملائي.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

