---
id: 5dfb6250eacea3f48c6300b2
title: الخطوة 22
challengeType: 0
dashedName: step-22
---

# --description--

بعد unordered list، أضف صورة جديدة مع قيمة السمة `src` المحددة إلى:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

وقيمة السمة `alt` الخاصة بها إلى:

`A slice of lasagna on a plate.`

# --hints--

يجب أن يكون هناك عنصر `img` مباشرة بعد علامة الإغلاق `</ul>`.

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

الصورة الجديدة لا تحتوي على سمة `alt`. تحقق من وجود مسافة بعد اسم opening tag و/أو أن هناك مسافات قبل جميع أسماء السمات.

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

يجب أن تحتوي الصورة الجديدة على قيمة `alt` من `A slice of lasagna on a plate.` تأكد من أن قيمة السمة `alt` محاطة بعلامات اقتباس.

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

الصورة الجديدة لا تحتوي على سمة `src`. تحقق من وجود مسافة بعد اسم opening tag و/أو أن هناك مسافات قبل جميع أسماء السمات.

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

يجب أن يكون للصورة الجديدة `src`بقيمة `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. تأكد من أن قيمة السمة `src` محاطة بعلامات اقتباس.

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

مع أنّك عيّنت لسمة `src` في عنصر img العنوان URL الصحيح، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

