---
id: 5ef9b03c81a63668521804d4
title: الخطوة 32
challengeType: 0
dashedName: step-32
---

# --description--

يتم استخدام عنصر `strong` للإشارة إلى أن بعض النصوص ذات أهمية كبيرة أو ملحة.

في `figcaption` الذي قمت بإضافته للتو، أشر إلى أن `hate` تتسم بأهمية كبيرة من خلال تغليفها في عنصر `strong`.

# --hints--

يجب أن يحتوي العنصر `strong` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('strong'));
```

يجب أن يحتوى عنصر `strong` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/strong\>/));
```

عنصر `strong` الخاص بك يجب أن يحيط بكلمة `hate` في النص `Cats hate other cats.` إما أنك حذفت النص أو لديك خطأ إملائي.

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.toLowerCase() === 'hate'
);
```

يجب أن يكون نص `figcaption` هو `Cats hate other cats.` تحقق من عدم وجود اي أخطاء إملائية وأن المسافات اللازمة موجودة حول علامات فتح وإغلاق عنصر `strong`.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

