---
id: 5ef9b03c81a63668521804da
title: الخطوة 42
challengeType: 0
dashedName: step-42
---

# --description--

استخدم عنصر `button` لإنشاء زر قابل للنقر. على سبيل المثال ، `<button>Click Here</button>` ينشئ زر مع النص `Click Here`.

أضف عنصر `button` مع النص `Submit` أسفل عنصر `input`. السلوك الافتراضي بالنقر على زر النموذج دون أي سمات يرسل النموذج إلى الموقع المحدد في سمة `action` للنموذج.

# --hints--

يجب أن يحتوي العنصر `button` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('button'));
```

يجب أن يحتوي العنصر `button` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/button\>/));
```

يجب أن يكون نص عنصر `button` هو `Submit`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

يجب أن يكون عنصر `button` الخاص بك تحت عنصر `input`. لقد وضعتهم بترتيب خاطئ.

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>

--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

