---
id: 5ef9b03c81a63668521804d7
title: الخطوة 36
challengeType: 0
dashedName: step-36
---

# --description--

السمة `action` تشير إلى أين يجب إرسال بيانات النموذج. على سبيل المثال ، `<form action="/submit-url"></form>` يخبر المتصفح بأنه يجب إرسال بيانات النموذج إلى المسار `/submit-url`.

قم بإضافة سمة `action` مع القيمة `https://freecatphotoapp.com/submit-cat-photo` إلى عنصر `form`.

# --hints--

عنصر `form` الخاص بك يجب أن يحتوي على علامة فتح وعلامة إغلاق بالترتيب الصحيح. قد تكون اغفلت واحدة من العلامات المطلوبة، أو قد تكون اخطأت في ترتيبهم.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

يجب أن يكون عنصر `form` المدمج في عنصر `section` تحت عنصر `h2`. عنصر `h2` و `form` لديك، بالترتيب الخاطئ.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

يجب أن لا يحتوي عنصر `form` على محتوى. قم بإزالة أي عناصر HTML أو نصوص بين tags العنصر `form`.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

عنصر `form` الخاص بك لا يحتوي على سمة `action`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

عنصر `form` الخاص بك يجب أن يحتوي على سمة `action` بقيمة `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

على الرغم من أنك قمت بتعيين سمة `action` إلى عنوان URL الصحيح، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

