---
id: 5ef9b03c81a63668521804df
title: الخطوة 46
challengeType: 0
dashedName: step-46
---

# --description--

يتم استخدام السمة `id` لتحديد عناصر HTML معينة. يجب أن تكون قيمة كل `id` فريدة من كل قيم `id` الأخرى للصفحة بأكملها.

أضف السمة `id` مع القيمة `indoor` لزر الراديو. عندما يكون للعناصر سمات متعددة، لا يهم ترتيب السمات.

# --hints--

زر الراديو الخاص بك يجب أن يظل موجودا بين علامات الافتتاح والإغلاق الخاصة بعنصر `label`.

```js
const labelChildNodes = [...$('form > label')[0].childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

زر الراديو الخاص بك يجب أن يحتوي على السمة `id`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
assert($('input')[0].hasAttribute('id'));
```

يجب أن يحتوي زرك الراديو على سمة `type`. تحقق من وجود مسافة بعد اسم opening tag و/أو أن هناك مسافات قبل جميع أسماء السمات.

```js
assert($('input')[0].hasAttribute('type'));
```

يجب أن يكون لعنصر الراديو سمة `id` مع القيمة `indoor`. إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن تحيط قيمة السمة بعلامات اقتباس.

```js
assert($('input')[0].id.match(/^indoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

