---
id: 5ef9b03c81a63668521804d0
title: الخطوة 25
challengeType: 0
dashedName: step-25
---

# --description--

قم بتأكيد كلمة `love` في عنصر `figcaption` من خلال تغليفها في عنصر التأكيد (`em`).

# --hints--

يجب أن يحتوي عنصر التأكيد الخاص بك (`em`) علي opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('em'));
```

يجب أن يحتوي عنصر التأكيد الخاص بك (`em`) علي closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/em\>/));
```

إما أنك قمت بحذف عنصر `figcaption` أو أنها تفتقد علامة فتح أو إغلاق.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

يجب أن يحيط عنصر التأكيد الخاص بك (`em`) بالنص `love`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

يجب أن يكون نص `figcaption` هو `Cats love lasagna`. تحقق من عدم وجود اي أخطاء إملائية وأن المسافات اللازمة موجودة حول علامات فتح وإغلاق عنصر `em`.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

