---
id: 5ef9b03c81a63668521804d8
title: الخطوة 37
challengeType: 0
dashedName: step-37
---

# --description--

عنصر `input` يسمح لك بعدة طرق لجمع البيانات من نموذج الويب. مثل عناصر `img` ،عناصر `input` هي <dfn>self-closing</dfn> ولا تحتاج إلى علامات الإغلاق (closing tags).

ادمج عنصر `input` في عنصر `form`.

# --hints--

عنصر `form` الخاص بك يجب أن يحتوي على علامة فتح وعلامة إغلاق بالترتيب الصحيح. قد تكون اغفلت واحدة من العلامات المطلوبة، أو قد تكون اخطأت في ترتيبهم.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

يجب أن تحتوي علامة فتح عنصر `form` الخاص بك على سمة `action` فقط. قم بإزالة أي شيء آخر قد تكون كتبته فيها.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

يجب عليك إنشاء عنصر `input`. تحقق من بناء الجملة (syntax).

```js
assert(document.querySelector('input'));
```

يجب أن يحتوي عنصر `input` الخاص بك على علامة فتح ولكن ليس علامة الإغلاق.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

يجب أن يكون عنصر `input` مدمج داخل عنصر `form`.

```js
assert(document.querySelector('form > input'));
```

يجب أن يحتوي `form` على عنصر `input` فقط. قم بإزالة أي عناصر HTML أو نصوص بين tags العنصر `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">

        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

