---
id: 5ef9b03c81a63668521804d2
title: الخطوة 27
challengeType: 0
dashedName: step-27
---

# --description--

كود القائمة المرتبة (`ol`) مشابه لـ unordered list، لكن قائمة العناصر (list items) في القائمة المطلوبة يتم ترقيمها عند عرضها.

بعد عنصر `h3` الأخير لعنصر `section` الثاني، أضف ordered list مع عناصر القائمة الثلاثة هذه:

`flea treatment` `thunder` `other cats`

# --hints--

يجب أن يحتوي العنصر `ol` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

يجب أن يحتوى عنصر `ol` على closing tag. Closing tags يجب ان تتبع بـ `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/ol>/));
```

عنصر `ol` الخاص بك يجب أن يكون أعلى closing tag عنصر `section` الثاني. لقد وضعتهم بترتيب خاطئ.

```js
assert($('main > section')[1].lastElementChild.nodeName === 'OL');
```

عناصر `li` الثلاثة يجب أن تكون مدمجة داخل عنصر `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

يجب أن يكون لديك ثلاثة عناصر `li` مع النص `flea treatment` و `thunder` و `other cats` باي ترتيب.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

يجب أن يكون لديك عنصر `ol` واحد فقط.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

