---
id: 5ef9b03c81a63668521804db
title: الخطوة 41
challengeType: 0
dashedName: step-41
---

# --description--

لمنع المستخدم من إرسال النموذج الخاص بك عندما تكون المعلومات المطلوبة مفقودة، تحتاج إلى إضافة السمة `required` إلى عنصر `input`. لا توجد حاجة إلى تعيين قيمة للسمة `required`. بدلاً من ذلك، أضف كلمة `required` إلى عنصر `input` ، تأكد من وجود مساحة بينها وبين السمات أخرى.

# --hints--

لقد قمت إما بحذف عنصر `input` الخاص بك أو أن بناء الجملة (syntax) غير صحيح. يجب أن تكون جميع قيم السمات محاطة بعلامات اقتباس.

```js
assert($('input').length);
```

يجب أن يحتوي `form` على عنصر `input` فقط. قم بإزالة أي عناصر HTML إضافية أو نص داخل عنصر `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

يجب أن يحتوي عنصر `input` الخاص بك علي سمة `required`. تذكر، أنت فقط تضيف الكلمة `required` داخل tag العنصر `input`.

```js
assert($('input')[0].hasAttribute('required'));
```

يجب عدم إعطاء قيمة إلى السمة `required`.

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

