---
id: 5ef9b03c81a63668521804d5
title: الخطوة 34
challengeType: 0
dashedName: step-34
---

# --description--

داخل عنصر `section` الثالث، أضف عنصر `h2` مع النص:

`Cat Form`

# --hints--

غير قادر على العثور على عنصر `section` الثالث. ربما قمت بحذفه عن طريق الخطأ أو حذفت علامة الافتتاح أو علامة الإغلاق.

```js
assert(
  document.querySelectorAll('section').length === 3 &&
    code.match(/<\/section>/g).length === 3
);
```

يجب أن يكون للعنصر `h2` الخاص بك علامة فتح وعلامة الإغلاق. قد تكون اغفلت احدي الـ tags المطلوبة.

```js
assert(
  document.querySelectorAll('h2').length >= 3 &&
    code.match(/<\/h2>/g).length >= 3
);
```

يجب عليك إضافة عنصر `h2` واحد فقط. الرجاء إزالة أي زيادات.

```js
assert(document.querySelectorAll('h2').length === 3);
```

يجب أن يكون عنصر `h2` الجديد أعلى closing tag عنصر `section` الثالث.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(thirdSection.lastElementChild.nodeName === 'H2');
```

النص الخاص بك `h2` يجب أن يكون `Cat Form`.

```js
const thirdSection = document.querySelectorAll('section')[2];
assert(
  thirdSection
    .querySelector('h2')
    .innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'cat form'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>
        </figure>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

