---
id: 5ef9b03c81a63668521804e2
title: الخطوة 54
challengeType: 0
dashedName: step-54
---

# --description--

تستخدم النماذج عادة خانات اختيار (checkboxes) للأسئلة التي قد تحتوي على أكثر من إجابة واحدة. على سبيل المثال، إليك مربع إختيار (checkbox) مع خيار `tacos`: `<input type="checkbox"> tacos`.

تحت العنصر `legend` الذي أضفته للتو، قم بإضافة `input` مع السمة `type` المحددة إلى `checkbox` وإعطائه الخيار:

`Loving`

# --hints--

عنصر `input` في مربع الاختيار الخاص بك يجب أن يحتوي على علامة opening tag ليس closing tag.

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

يجب عليك إضافة عنصر input واحد فقط إلى مربع الاختيار (checkbox) الخاص بك. قم بإزالة أي زيادات.

```js
assert($('fieldset > input').length === 1);
```

عنصر `input` الجديد يجب أن يكون تحت عنصر `legend` مع النص `What's your cat's personality?`. لقد وضعتهم بترتيب خاطئ.

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

عنصر `input` الخاص بك لا يحتوي على سمة `type`. تحقق من أن هناك مسافة بعد اسم الـ opening tag.

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

عنصر `input` الجديد الخاص بك يجب أن يحتوي على سمة `type` بقيمة `checkbox`. إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن قيم السمة يجب أن تكون محاطة بعلامات اقتباس.

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

على الرغم من أنك قمت بتعيين عنصر `input` سمة `type` إلى `checkbox`، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

يجب أن يكون النص `Loving` علي يمين خانة الاختيار الخاصة بك مباشرة. تأكد من وجود مسافة بين العنصر والنص. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

