---
id: 5ef9b03c81a63668521804dd
title: الخطوة 45
challengeType: 0
dashedName: step-45
---

# --description--

عناصر `label` تستخدم للمساعدة في ربط النص لعنصر `input` بعنصر الـ `input` نفسه (خاصة للتكنولوجيات المساعدة مثل قارئ الشاشة). على سبيل المثال ،`<label><input type="radio"> cat</label>` يجعلها تجعل النقر على كلمة `cat` يؤدي أيضًا إلى تحديد زر الراديو المقابل.

قم بدمج زر `radio` الخاص بك داخل عنصر `label`.

# --hints--

يجب عليك التأكد من أن زر الراديو لا يزال موجودا.

```js
assert($('input[type="radio"]')[0]);
```

النص `Indoor` يجب أن يكون مباشرة على يمين زر الـ `radio`. تأكد من وجود مسافة بين العنصر والنص. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Indoor/i)
);
```

يجب أن يحتوي العنصر `label` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('label'));
```

يجب أن يحتوي عنصر `label` الخاص بك على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/label\>/));
```

زر الراديو الخاص بك ونصه يجب أن يكونا بين علامات الافتتاح والإغلاق الخاصة بعنصر `label`.

```js
const labelChildNodes = [...$('form > label')[0].childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

