---
id: 5ef9b03c81a63668521804dc
title: الخطوة 44
challengeType: 0
dashedName: step-44
---

# --description--

يمكنك استخدام أزرار الراديو (radio buttons) للأسئلة حيث تريد إجابة واحدة فقط من خيارات متعددة.

هذا مثال على زر الراديو مع خيار `cat`: `<input type="radio"> cat`. تذكر أن عناصر `input` هي <dfn>self-closing</dfn>.

قبل الـ text input، أضف زر radio مع الخيار محدد كالتالي:

`Indoor`

# --hints--

يجب عليك إنشاء عنصر `input` لزر الراديو الخاص بك. تحقق من بناء الجملة (syntax).

```js
assert($('form > input').length >= 2);
```

يجب أن يحتوي عنصر `input` الخاص بك على علامة فتح ولكن ليس علامة الإغلاق.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

يجب عليك فقط إضافة عنصر واحد `input` لزر الراديو الخاص بك. قم بإزالة أي زيادات.

```js
assert($('form > input').length === 2);
```

عنصر `input` الجديد يجب أن يكون فوق `input` الحالي مع سمة `type` بقيمة `text`. لقد وضعتهم بترتيب خاطئ.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

عنصر `input` الخاص بك لا يحتوي على سمة `type`. تحقق من أن هناك مسافة بعد اسم الـ opening tag.

```js
assert($('input')[0].hasAttribute('type'));
```

يجب ألا يحتوي عنصرك `input` على سمة `type`. أزل أي زيادات.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

يجب أن يحتوي عنصرك `input` على سمة `type` بقيمة `radio`. إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن تحيط قيمة السمة بعلامات اقتباس.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

مع أنّك عيّنت عنصر `input` سمة `type` إلى `radio`، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

يجب أن يوجد نص `Indoor` لزر `radio`، بعد الرز وليس قبله.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

يجب أن يكون نص `Indoor` مباشرة على يمين زر `radio`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

