---
id: 5efb23e70dc218d6c85f89b1
title: الخطوة 38
challengeType: 0
dashedName: step-38
---

# --description--

هناك العديد من أنواع المدخلات التي يمكنك إنشاؤها باستخدام السمة `type`. يمكنك بسهولة إنشاء حقل كلمة المرور، أو زر إعادة الضبط، أو السماح للمستخدمين بتحديد ملف من جهاز الكمبيوتر.

قم بإنشاء input من نوع نص للحصول على إدخال نصي من المستخدم عن طريق إضافة السمة `type` مع القيمة `text` إلى عنصر `input`.

# --hints--

لقد قمت إما بحذف عنصر `input` الخاص بك أو أن بناء الجملة (syntax) غير صحيح. إذا كنت قد أضفت سمات، تأكد من أن قيمها محاطة بعلامات اقتباس.

```js
assert($('input').length);
```

يجب أن يحتوي `form` على عنصر `input` فقط. قم بإزالة أي عناصر HTML أو نصوص بين tags العنصر `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

عنصر `input` الخاص بك يجب أن يحتوي على سمة `type` بقيمة `text`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
assert($('input')[0].hasAttribute('type'));
```

عنصر `input` الخاص بك يجب أن يحتوي على سمة `type` بقيمة `text`. إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن قيم السمة يجب أن تكون محاطة بعلامات اقتباس.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

على الرغم من أنك قمت بتعيين عنصر `input` سمة `type` إلى `text`، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

