---
id: 5efb2c990dc218d6c85f89b2
title: الخطوة 43
challengeType: 0
dashedName: step-43
---

# --description--

على الرغم من أنك أضفت الزر الخاص بك أسفل input النص، فإنها تظهر بجانب بعضها البعض على الصفحة. وذلك لأن عناصر `input` و `button` هي <dfn>inline elements</dfn> والتي لا تظهر على سطور جديدة.

سوف يرسل الزر الذي أضفته النموذج (form) بشكل افتراضي. ومع ذلك، فإن الاعتماد على السلوك الافتراضي قد يسبب أرتباك. أضف السمة `type` مع القيمة `submit` إلى عنصر `button` لتوضيح أنه زر إرسال.

# --hints--

يجب أن يحتوي العنصر `button` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('button'));
```

يجب أن يحتوي العنصر `button` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/button\>/));
```

عنصر `button` الخاص بك لا يحتوي على سمة `type`. تحقق من أن هناك مسافة بعد اسم الـ opening tag.

```js
assert($('button')[0].hasAttribute('type'));
```

عنصر `button` الخاص بك يجب أن يحتوي على سمة `type` بقيمة `submit`. إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن قيم السمة يجب أن تكون محاطة بعلامات اقتباس.

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

على الرغم من أنك قمت بتعيين عنصر `button` سمة `type` إلى `submit`، يوصى بأن تحيط دائما قيمة السمة بعلامات اقتباس.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

