---
id: 5efc4f528d6a74d05e68af74
title: الخطوة 56
challengeType: 0
dashedName: step-56
---

# --description--

هناك طريقة أخرى لربط نص عنصر `input` بالعنصر نفسه. يمكنك دمج النص ضمن عنصر `label` وإضافة سمة `for` بنفس القيمة كسمة `id` لعنصر `input`.

اربط النص `Loving` مع checkbox بواسطة إدخال النص `Loving` فقط في عنصر `label` وإعطائه السمة `for`.

# --hints--

يجب عليك التأكد من أن خانة الاختيار لا تزال موجودة.

```js
assert($('input[type="checkbox"]')[0]);
```

يجب أن يكون لمربع للاختيار الخاص بك سمة `id` مع القيمة `loving`.  توقعت عنصر بقيمة `--fcc-expected--`, لكن وجدت `--fcc-actual--`.

```js
assert.equal($('input[type="checkbox"]')[0].id, 'loving');
```

يجب ألا يكون النص `Loving` موجودا مباشرة علي يمين خانة الاختيار الخاصة بك. يجب تغليفه في عنصر `label`. تأكد من وجود مسافة بين العنصرين.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  !checkboxInputElem.nextSibling.nodeValue
    .replace(/\s+/g, ' ')
    .match(/ Loving/i)
);
```

سوف تحتاج إلى إضافة عنصر `label` جديد ليتم فيه دمج النص `Loving`. تأكد من أن لديه علامة فتح وإغلاق.

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

عنصر `label` الجديد يجب أن يكون مباشرة علي يمين خانة الاختيار الخاصة بك. تأكد من وجود مسافة بين العنصرين.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

عنصر `label` الجديد لا يحتوي على سمة `for`. تحقق من أن هناك مسافة بعد اسم الـ opening tag.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

عنصر `label` الخاص بك يجب أن يحتوي على سمة `for` بقيمة `loving`.  توقعت عنصر بقيمة `--fcc-expected--`, لكن وجدت `--fcc-actual--`.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert.equal(labelElem.getAttribute('for'), 'loving');
```

يجب أن يكون النص `Loving` مدمج داخل عنصر `label` الجديد. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

