---
id: 5f05a1d8e233dff4a68508d8
title: الخطوة 47
challengeType: 0
dashedName: step-47
---

# --description--

إنشاء زر راديو آخر أسفل زر الراديو الأول. أدخله داخل عنصر `label` مع `Outdoor` كنص للعنصر `label`. اعطي زر بسمة radio سمة `id` بقيمة `outdoor`.

# --hints--

يجب أن يحتوي زر الراديو `input` الأصلي فقط, بداخله نص `Indoor`, على `id` بقيمة `indoor`. يمكنك استخدام زر إعادة تعيين الخطوة للحصول على HTML الأصلي إذا لزم الأمر.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

لا يجب عليك إجراء أي تغييرات على زر الراديو `Indoor`. يمكنك استخدام زر إعادة تعيين الخطوة للحصول على HTML الأصلي إذا لزم الأمر.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

يجب عليك أن تضيف عنصر `input` واحد جديد في عنصر `label` الجديد. تحقق من أن علامة `label` تحتوي على علامة أفتاح (opening tag) وإغلاق (closing tag) على حد سواء.

```js
assert(document.querySelectorAll('label input').length === 2);
```

لا يجب عليك إضافة أي سمات لعلامة أفتاح (opening tag) في `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

يجب أن يحتوي `input` على سمة `id` بقيمة `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

يجب أن يحتوي عنصرك `input` على سمة `type` بقيمة `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

لا يجب عليك إضافة أي عناصر جديدة غير `input` داخل `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

يجب ألا يكون هناك نص في يسار من `input` الجديد.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

يجب أن يطابق نص عنصر label في زر radio الجديد `Outdoor`. ويشمل ذلك التشكيل (capitalization).

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

يجب أن تكون عناصر زر radio الجديد و label أسفل مباشرة/بعد عناصر زر radio و label من `Indoor`. وينبغي ألا تكون هناك علامات أخرى بينهما.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

