---
id: 5f07c98cdb9413cbd4b16750
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

وقد حان الوقت لإضافة قسم (section) جديد. أضف عنصر `section` الثاني أدني عنصر `section` الحالي.

# --hints--

يجب أن يحتوي العنصر `section` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

يجب عليك إضافة opening tag واحد فقط للـ `section`. الرجاء إزالة أي زيادات.

```js
assert(document.querySelectorAll('section').length === 2);
```

يجب أن يحتوي العنصر `section` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

يجب عليك إضافة closing tag واحد فقط للـ `section`. الرجاء إزالة أي زيادات.

```js
assert(code.match(/<\/section>/g).length === 2);
```

عنصر `section` الثاني يجب أن لا يكون مدمج في عنصر `section` الأول.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

يجب أن يكون العنصران `section` بين العلامتين الافتتاحية والإغلاقية لعنصر `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

