---
id: 5f07be6ef7412fbad0c5626b
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

قبل إضافة أي محتوى جديد، يجب عليك استخدام عنصر `section` لفصل محتوى صور القط عن المحتوى المستقبلي.

خذ عنصر `h2` و التعليق و عنصر `p` و عنصر الرابط (`a`) وأدخلهم في عنصر `section`.

# --hints--

يجب أن يحتوي العنصر `section` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
assert(document.querySelector('section'));
```

يجب أن يحتوي العنصر `section` على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

يجب أن يكون عنصر `section` بأكمله بين علامات فتح وإغلاق عنصر `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

عنصر `h2` الموجود و التعليق و عنصر `p` و عنصر الرابط (`a`) ينبغي ان يكونوا بين العلامتين الافتتاحية والإغلاقية لعنصر `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

يجب ألّا يكون عنصر `h1` داخل عنصر `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

