---
id: 5f0d4d04b435f13ab6550053
title: الخطوة 53
challengeType: 0
dashedName: step-53
---

# --description--

قم بإضافة عنصر `legend` مع النص `What's your cat's personality?` داخل عنصر `fieldset` الثاني.

# --hints--

لقد قمت إما بحذف العنصر `fieldset` الثاني أو انه يفتقر وجود علامة فتح أو علامة إغلاق.

```js
assert(
  document.querySelectorAll('fieldset').length === 2 &&
    code.match(/<\/fieldset>/g).length === 2
);
```

يجب أن يحتوي العنصر `legend` على opening tag. Opening tags تكتب هكذا: `<elementName>`.

```js
const secondFieldset = $('fieldset')[1];
assert(
  secondFieldset &&
    [...secondFieldset.children].filter((child) => child.nodeName === 'LEGEND')
      .length
);
```

يجب أن يحتوي عنصر `legend` الخاص بك على closing tag. Closing tags لها `/` مباشرة بعد رمز `<`.

```js
assert(code.match(/<\/legend\>/g).length === 2);
```

يجب أن يحتوي عنصر `legend` على النص `What's your cat's personality?`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const secondFieldset = $('fieldset')[1];
assert(
  secondFieldset &&
    [...secondFieldset.children].filter((child) => {
      const extraSpacesRemoved = child.innerText.replace(/\s+/g, ' ');
      return (
        child.nodeName === 'LEGEND' &&
        extraSpacesRemoved.match(/What's your cat's personality\??$/i)
      );
    }).length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
--fcc-editable-region--
          <fieldset>

          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

