---
id: 62dabe2ef403a12d5d295273
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

الآن بعد تغيرك نص `cat photos` داخل عنصر `p` إلى رابط، لا تحتاج إلى الرابط الثاني أسفل عنصر `p`. حذف عنصر الرابط (anchor) بِرُمَّته أسفل عنصر `p`.

# --hints--

يجب أن يحتوي كودك فقط على عنصر رابط واحد (`a`).

```js
assert(document.querySelectorAll('a').length === 1);
```

يجب أن يظل داخل عنصر `p` على عنصر الرابط (`a`) داخله.

```js
assert($('p > a').length);
```

يجب أن يكون سمة الرابط `href` بقيمة `https://freecatphotoapp.com`. إما أنك حذفت قيمة `href` أو لديك خطأ إملائي.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

يجب أن يكون نص الرابط `cat photos`. إما أنك حذفت النص أو لديك خطأ إملائي.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

يجب أن يكون محتوى عنصر `p` المرئي في المتصفح `Click here to view more cat photos.` تحقق مرتين من النص أو المسافات أو علامات الترقيم لكل من `p` وعنصر الرابط المدمج.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

لا ينبغي أن يكون هناك `cat photos` نص تحت عنصر `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more <a href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
