---
id: 7cf9b03d81a65668421804c3
title: الخطوة 39
challengeType: 0
dashedName: step-39
---

# --description--

من أجل الوصول إلى بيانات النموذج عن طريق الموقع المحدد في سمة `action` ، يجب أن تعطي حقل النص سمة `name` وأن تعيينه قيمة لتمثيل البيانات التي تم إرسالها. على سبيل المثال، يمكنك استخدام الجملة التالية لحقل نص عنوان البريد الإلكتروني: `<input type="text" name="email">`.

أضف السمة `name` مع القيمة `catphotourl` إلى حقل النص الخاص بك.

# --hints--

لقد قمت إما بحذف عنصر `input` الخاص بك أو أن بناء الجملة (syntax) غير صحيح. يجب أن تكون جميع قيم السمات محاطة بعلامات اقتباس.

```js
assert($('input').length);
```

يجب أن يحتوي `form` على عنصر `input` فقط. إزالة أي عناصر HTML إضافية أو نص داخل عنصر `form`.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

عنصر `input` الخاص بك لا يحتوي على سمة `name`. تحقق من وجود مسافة بعد اسم الـ opening tag و/أو ان هناك مسافات قبل جميع أسماء السمات.

```js
assert($('input')[0].hasAttribute('name'));
```

عنصر `input` الخاص بك يجب أن يحتوي على سمة `name` بقيمة `catphotourl`.  إما أنك حذفت القيمة أو لديك خطأ إملائي. تذكر أن قيم السمة يجب أن تكون محاطة بعلامات اقتباس.

```js
assert(
  $('input')[0]
    .getAttribute('name')
    .match(/^catphotourl$/i)
);
```

على الرغم من أنك قمت بتعيين عنصر `input` سمة `name` إلى `catphotourl`، يوصى بأن تحيط دائما بقيمة السمة بعلامات اقتباس.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

