---
id: 66a2be6afc9ed259e2657eaf
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

step 9 instructions

# --hints--

Test 1

```js

```

# --seed--

## --seed-contents--

```js
function sayHello(retrievedName)
{
  return "Hi" + retrievedName(); 
}

function sayName()
{
  return "Quincy" 
}

function test(title, callback) {
  try {
    callback();
    console.log(`✓ ${title}`);
  } catch (error) {
    console.error(`✕ ${title}`);
    console.error(error);
  }
}

function assert(actual) {
  return {
    deepEqual(expected) {
      for (let i = 0; i < actual.length; i++) {
        if (actual[i] !== expected[i]) {
          throw new Error(
            `Element at index ${i} is different: ${actual[i]} is not equal to ${expected[i]}`
          );
        }
      }
    },
  };
}
--fcc-editable-region--

--fcc-editable-region--
```
