---
id: 66a2be6afc9ed259e2657eac
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

step 6 instructions

# --hints--


`multiplyBy2` is a defined function. 

```js
assert.isFunction(multiplyBy2); 
```

`multiplyBy2` returns a number.

```js
assert.equal(typeof(multiplyBy2(3)),"number");
```

`multiplyBy2` returns a number multiplied by 2.

```js
assert.equal(multiplyBy2(6),12);
```


# --seed--

## --seed-contents--

```js
function test(title, callback) {
  try {
    callback();
    console.log(`✓ ${title}`);
  } catch (error) {
    console.error(`✕ ${title}`);
    console.error(error);
  }
}

function assert(actual) {
  return {
    deepEqual(expected) {
      for (let i = 0; i < actual.length; i++) {
        if (actual[i] !== expected[i]) {
          throw new Error(
            `Element at index ${i} is different: ${actual[i]} is not equal to ${expected[i]}`
          );
        }
      }
    },
  };
}
--fcc-editable-region--

--fcc-editable-region--
```
