---
id: 66a2be6afc9ed259e2657ea9
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

For this lesson, you are going to make it possible to actually pass a function as a variable
by adding an parameter to our `sayHello` function called `retrievedName`.  

# --hints--

You should have a parameter in `sayHello` called `retrievedName`. 

```js
const parameters = helpers.getFunctionParams(sayHello); 
assert.equal(parameters[0].name,"retrievedName"); 
```



# --seed--

## --seed-contents--

```js

--fcc-editable-region--
function sayHello()
{
  return "Hi" 
}
--fcc-editable-region--
function sayName()
{
  return "Quincy" 
}

```
