---
id: 66a2be6afc9ed259e2657ea8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now that we have created a function that returns `"Hi"`, let's create a second function to pass into 
`sayHello` function that returns your name. Name this function `sayName`. 

# --hints--

You have defined the `sayName` function. 

```js
assert.isFunction(sayName); 
```

You need to have the `sayName` function return a string. 

```js
assert.isEqual(typeOf(sayName()),"string"); 
```

# --seed--


## --seed-contents--

```js
function sayHello()
{
  return "Hi" 
}

--fcc-editable-region--

--fcc-editable-region--

```
