---
id: 66a2be6afc9ed259e2657eaa
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

For this lesson, we are going to finally going to use a callback function. Despite being only a variable inside 
of the function definition , it is completely callable. 

We are going to concatenate the value of our `retrievedName` callback to our current return statement in 
the `sayHello` function. 

# --hints--

You should add the result of `retrievedName` to your return statement. 

```js
assert.match(code,/return\s`"Hi"\s+\+\s+retrievedName\(\);?/); 
```

# --seed--

## --seed-contents--

```js

--fcc-editable-region--
function sayHello(retrievedName)
{
  return "Hi"; 
}
--fcc-editable-region--

function sayName()
{
  return "Quincy" 
}
```
