---
id: 60eebd07ea685b0e777b5583
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

مرحبا بكم في مشروع نموذج التسجيل! ابدأ بإضافة تعريف `!DOCTYPE html` في الجزء العلوي من المستند حتى يعرف المتصفح نوع المستند الذي يقرأه.

# --hints--

يجب أن يحتوي الكود الخاص بك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` مع `>` بعد الـ type.

```js
assert(code.match(/html\s*>/gi));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
