---
id: 60f02e7361b68405e27b62a5
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

ضمن `body` ، قم بتوفير سياق عنوان للمحتوى، عن طريق إضافة `h1` بالنص `Registration Form`.

# --hints--

يجب عليك إضافة `h1` داخل `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

يجب عليك إعطاء `h1` النص `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
