---
id: 60f027c87bc98f050395c139
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

بعد ذلك، قم بإضافة علامات فتح وإغلاق `head` و `body` داخل عنصر `html`.

# --hints--

يجب أن يكون لديك الـ opening tag للـ `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك الـ opening tag للـ `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

عناصر `head` و `body` يجب أن يكونا أشقاء (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

عنصر `head` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

عنصر `body` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
