---
id: 60f0286404aefb0562a4fdf9
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

أضف عنصر `title` وعنصر `meta` جديد داخل عنصر `head`. أعطِ مشروعك عنوانًا `Registration Form`، وأعط `charset` سمة بقيمة `UTF-8` لعنصرك `meta`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

عنصر `title` يجب أن يكون داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > title'));
```

يجب أن يحتوي مشروعك على عنوان `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

تذكر أن الـ casing، اي حالة الحرف سواء كبير او صغير (capital or small) والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
