---
id: 60f1a9cbd23023082e149fee
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

هذا أفضل. الآن، اجعل الخلفية سهلة على العينين، بتغيير `background-color` للـ `body` إلى `#1b1b32`. ثم لمشاهدة النص، غيّر `color` إلى `#f5f6f7`.

# --hints--

يجب عليك إضافة `background-color` داخل منتقي العنصر `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

يجب عليك إعطاء `background-color` قيمة `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

يجب عليك إضافة `color` داخل منتقي العنصر `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

يجب أن تعطي `color` القيمة `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
