---
id: 60f1922fcbd2410527b3bd89
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--
يكون `vh` وحدة ارتفاع نافذة العرض (viewport)، وتتناسب مع 1٪ من `height` نافذة العرض (viewport).

لقد حان الوقت للبدء المشروع مع بعض CSS. ابدأ بإعطاء `body` خاصية `width` بقيمة `100%` و `height` بقيمة `100vh`.

# --hints--

يجب عليك استخدام منتقي العنصر `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

يجب عليك إعطاء `body` خاصية `width` بقيمة `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

يجب عليك إعطاء `body` خاصية `height` بقيمة `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
