---
id: 60f5c3e399ff1a05629964e4
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

كما اقترح العنوان، فإنك تقوم بإنشاء نموذج (form). لذا، بعد عنصر `p` ، أدخل `form` مع سمة `action` تستهدف `https://register-demo.freecodecamp.org`.

# --hints--

يجب عليك إضافة عنصر `form` مجاور لعنصر `p`.

```js
assert.exists(document.querySelector('p + form'));
```

يجب أن تعطي عنصر `form` سمة `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

يجب أن تعطي `action` قيمة `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
