---
id: 60f805f813eaf2049bc2ceea
title: الخطوة 17
challengeType: 0
dashedName: step-17
---

# --description--

قم بدمج عنصر `input` داخل كل `label`. تأكد من إضافة كل `input` بعد `label` النص، وإدرج مسافة بعد الـ colon (:).

# --hints--

يجب عليك إضافة أربعة عناصر `input` إلى عنصر `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

يجب عليك دمج عناصر `input` داخل عناصر `label`.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

يجب عليك إضافة أول `input` بعد نص الـ `label` الآتي `Enter Your First Name:`,، وقم بإدراج مسافه بعد الـ colon (:).

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

يجب عليك إضافة ثاني `input` بعد نص الـ `label` الآتي `Enter Your Last Name:`,، وقم بإدراج مسافه بعد الـ colon (:).

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

يجب عليك إضافة ثالث `input` بعد نص الـ `label` الآتي `Enter Your Email:`,، وقم بإدراج مسافه بعد الـ colon (:).

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

يجب عليك إضافة رابع `input` بعد نص الـ `label` الآتي `Create a New Password:`، وقم بإدراج مسافه بعد الـ colon (:).

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
