---
id: 60f803d5241e6a0433a523a1
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--
وحدة `rem` تعني مقدار الأصلي من `em`، وهي ترطبت بحجم خط (font size) عنصر `html`.

حيث إن عناصر `label` على مستوى السطر (inline) بشكل افتراضي، وهم جميعا معرضون جنبا إلى جنب على نفس السطر مما يجعل من الصعب قراءة نصهم. لجعلها تظهر على أسطر منفصلة، أضف `display: block` إلى عنصر `label`، أضف `margin` بالقيمة `0.5rem 0`، لفصلهما عن بعضهما البعض.

# --hints--

يجب عليك استخدام منتقي `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

يجب عليك إضافة خاصية `display` بالقيمة `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

يجب عليك إضافة خاصية `margin` بالقيمة `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
