---
id: 60f84ec41116b209c280ba91
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

مع `type="password"` يمكنك استخدام السمة `pattern` لتعريف الـ regular expression الذي يجب أن يطابق كلمة المرور لاعتبارها صحيحة.

قم بإضافة سمة `pattern` إلى عنصر `input` لكلمة المرور لطلب مطابقة الإدخال: `[a-z0-5]{8,}`

ما ورد أعلاه هو regular expression يتطابق مع ثمانية أحرف صغيرة (lowercase) أو أكثر أو الأرقام من `0` إلى `5`. ثم قم بإزالة السمة `minlength` وقم بالتجربة.

# --hints--

يجب أن تعطي عنصر `input` لكلمة المرور السمة `pattern`.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

يجب أن تعطي سمة `pattern` القيمة `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

يجب عليك إزالة السمة `minlength` من عنصر `input` لكلمة المرور.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
