---
id: 60f81616cff80508badf9ad5
title: الخطوة 21
challengeType: 0
dashedName: step-21
---

# --description--

في هذه المرحلة، يجب أن تكون قادراً على إرسال النموذج (submit the form). ومع ذلك ، قد تلاحظ عدم حدوث الكثير.

لجعل النموذج أكثر تفاعلية، أضف السمة `required` إلى عناصر `input` في أول `fieldset`.

الآن، إذا حاولت إرسال النموذج دون ملء الحقول المطلوبة، سترى رسالة خطأ.

# --hints--

يجب أن تعطي أول عنصر `input` سمة `required`.

```js
assert.equal(document.querySelector('input')?.required, true);
```

يجب أن تعطي ثاني `input` سمة `required`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.required, true);
```

يجب أن تعطي ثالث `input` سمة `required`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.required, true);
```

يجب أن تعطي رابع `input` سمة `required`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.required, true);
```

لا يجب أن تعطي `input` الـ `submit` سمة `required`.

```js
assert.equal(document.querySelector('input[type="submit"]')?.required, false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
