---
id: 60f80e0081e0f2052ae5b505
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

إن تحديد سمة `type` لعنصر النموذج (form) أمر مهم للمتصفح لمعرفة نوع البيانات التي ينبغي أن يتوقعها. إذا لم يتم تحديد `type` ، فسيعتبره المتصفح `text` بشكل افتراضي.

قم بإعطاء أول عنصرين `input` سمة `type` بالقيمة `text`، والثالث سمة `type` بقيمة `email`، والرابع سمة `type` بقيمة `password`.

نوع `email` يسمح فقط بالبريد الإلكتروني مع `@` و `.` في النطاق (domain). نوع `password` يحجب الإدخال، ويحذر إذا كان الموقع لا يستخدم HTTPS.

# --hints--

يجب أن تعطي عنصر `input` الاول سمة `type` بقيمة `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

يجب أن تعطي عنصر `input` الثاني سمة `type` بقيمة `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

يجب أن تعطي عنصر `input` الثالث سمة `type` بقيمة `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

يجب أن تعطي عنصر `input` الرابع سمة `type` بقيمة `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
