---
id: 60fac4095512d3066053d73c
title: الخطوة 33
challengeType: 0
dashedName: step-33
---

# --description--

إن إضافة القائمة المنسدلة إلى النموذج أمر سهل مع عنصر `select`. عنصر `select` هو حاوية لمجموعة من عناصر `option`، ويعمل عنصر `option` كتسمية (label) لكل خيار من خيارات القائمة المنسدلة. وكلا العنصرين يتطلبان علامات إغلاق (closing tags).

Start by adding a `select` element below the two `label` elements. Then nest 5 `option` elements within the `select` element.

# --hints--

يجب عليك إضافة عنصر `select` إلى `fieldset` الثالث.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > select'));
```

يجب عليك تضمين خمس عناصر `option` ضمن عنصر `select`.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > select > option')?.length, 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" /></label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
