---
id: 60fad1cafcde010995e15306
title: الخطوة 41
challengeType: 0
dashedName: step-41
---

# --description--

مع عمليات إرسال النموذج (form submissions)، من المفيد ومن الممارسات الجيدة تزويد كل عنصر قابل للإرسال بسمة `name`. تُستخدم هذه السمة لتحديد العنصر في إرسال النموذج.

قم بإعطاء كل عنصر قابل للإرسال سمة `name` فريدة من اختيارك، باستثناء مدخلات `radio`.

# --hints--

You should give the `input` expecting a first name a `name` attribute. _لو كنت مكانك كنت ساختار `first-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

يجب أن تعطي `input` توقع اسم العائلة (last name) السمة `name`. _لو كنت مكانك كنت ساختار `last-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

يجب عليك إعطاء `email` السمة `name`. _لو كنت مكانك كنت ساختار `email`_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

يجب عليك إعطاء `password` السمة `name`. _لو كنت مكانك كنت ساختار `password`_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

يجب عليك إعطاء `checkbox` السمة `name`. _لو كنت مكانك كنت ساختار `terms`_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

يجب عليك إعطاء `file` السمة `name`. _لو كنت مكانك كنت ساختار `file`_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

يجب عليك إعطاء `number` السمة `name`. _لو كنت مكانك كنت ساختار `age`_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

يجب أن تعطي عنصر `select` السمة `name`. _لو كنت مكانك كنت ساختار `referrer`_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

يجب أن تعطي عنصر `textarea` السمة `name`. _لو كنت مكانك كنت ساختار `bio`_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

لا يجب عليك إعطاء أي من عناصر `option` السمة `name`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

لا يجب عليك إعطاء أي من عناصر `label` السمة `name`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

لا يجب عليك إعطاء أي من عناصر `fieldset` السمة `name`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
