---
id: 62cc5b1779e4d313466f73c5
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أدخل عنصر `link` ذاتي الأغلاق داخل عنصر `head`. اعطيها سمة `rel` بقيمة `stylesheet` و `href` بقيمة `styles.css`.

# --hints--

يجب أن يحتوي كودك على عنصر `link`.

```js
assert.exists(document.querySelector('link'));
```

لا يجب عليك تغيير وسوم `head` الموجودة. تيقن من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

يجب أن يحتوي كودك على عنصر `link` مغلق ذاتيا (self-closing).

```js
assert(document.querySelectorAll('link').length === 1);
```

يجب أن يحتوي عنصرك `link` داخل عنصرك `head`.

```js
assert.exists(document.querySelector('head > link'));
```

يجب أن يحتوي عنصرك `link` على سمة `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

يجب أن يحتوي عنصرك `link` على سمة `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
