---
id: 62ff8b9dab5ac88e4d3d43a3
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

أتباعاً أفضل ممارسات تسهيل المنال (accessibility), أربط عناصر `input` و `label` معاً باستعمال سمة `for`.

استخدم `first-name`, و `last-name`, و `email`, و `new-password` كقيم لسمات `id`.

# --hints--

يجب أن يحتوي أول عنصر `input` علي `id` بقيمة `first-name`.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

يجب أن يحتوي ثاني عنصر `input` علي `id` بقيمة `last-name`.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

يجب أن يحتوي ثالث عنصر `input` علي `id` بقيمة `email`.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

يجب أن يحتوي رابع عنصر `input` علي `id` بقيمة `new-password`.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

يجب أن يحتوي عنصرك الأول `label` على سمة `for` بقيمة `first-name`.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

يجب أن يحتوي عنصرك الثاني `label` على سمة `for` بقيمة `last-name`.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

يجب أن يحتوي عنصرك الثالث `label` على سمة `for` بقيمة `email`.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

يجب أن يحتوي عنصرك الرابع `label` على سمة `for` بقيمة `new-password`.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
