---
id: 62ff8e998d3e7eae14d6ae3b
title: الخطوة 29
challengeType: 0
dashedName: step-29
---

# --description--

اتبع أفضل ممارسات تسهيل المنال (accessibility) عن طريق ربط عناصر `input` و `label` في `fieldset` الثاني.

استخدم `personal-account`, و `business-account`, و `terms-and-conditions` كقيم لسمات `id`.

# --hints--

يجب أن يحتوي أول عنصر `input` علي `id` بقيمة `personal-account`.

```js
assert(document.querySelectorAll('fieldset:nth-of-type(2) input')?.[0]?.matches('#personal-account'))
```

يجب أن يحتوي الثاني عنصر `input` علي `id` بقيمة `business-account`.

```js
assert(document.querySelectorAll('fieldset:nth-of-type(2) input')?.[1]?.matches('#business-account'))
```

يجب أن يحتوي ثالث عنصر `input` علي `id` بقيمة `terms-and-conditions`.

```js
assert(document.querySelectorAll('fieldset:nth-of-type(2) input')?.[2]?.matches('#terms-and-conditions'))
```

يجب أن يحتوي عنصرك الأول `label` على سمة `for` بقيمة `personal-account`.

```js
assert(document.querySelectorAll('fieldset:nth-of-type(2) label')?.[0]?.matches('label[for="personal-account"]'))
```

يجب أن يحتوي عنصرك الثاتي `label` على سمة `for` بقيمة `business-account`.

```js
assert(document.querySelectorAll('fieldset:nth-of-type(2) label')?.[1]?.matches('label[for="business-account"]'))
```

يجب أن يحتوي عنصرك الثالث `label` على سمة `for` بقيمة `terms-and-conditions`.

```js
assert(document.querySelectorAll('fieldset:nth-of-type(2) label')?.[2]?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label><input type="checkbox" required /> I accept the terms and conditions</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
