---
id: 63541ef4f96cd82e8e6c788a
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

تحدد سمة `method` كيفية إرسال بيانات النموذج إلى عنوان URL المحدد في سمة `action`. يمكن إرسال بيانات الاستمارة عبر طلب `GET` كوسائط عنوان URL (بواسطة `method="get"`) أو عبر طلب `POST` كبيانات في هيئة الطلب (بواسطة `method="post"`).

عيّن خاصية `method` لإرسال بيانات نموذجك عبر طلب `POST`.

# --hints--

يجب عليك إلا تضيف عنصر `form` جديد.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

يجب أن يحتوي عنصرك `form` علي سمة `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

يجب تعيين سمتك `method` بقيمة `post`.

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
