---
id: 66a2b162e659291efc5c6085
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, we are going to learn how to use callbacks by building out a testing suite for a small 
application. 

Callbacks are among the most powerful tools in Javascript. Callbacks are used to pass a function into another 
function. This is because all functions in Javascript are `first class functions`. The term first class refers 
to the ability to treat functions as values. The passed in function is then executed inside of the other 
function. The output then influences the other function. 


Let's start building out an example. Begin by creating a function called `sayHello` that returns "Hi". 

# --hints--

You have defined the `sayHello` function. 

```js
assert.isFunction(sayHello); 
```

You need to make `sayHello` return `"Hi"` with no extra characters. 

```js
assert.isEqual(sayHello(),"Hi"); 
```


# --seed--


## --seed-contents--

```js


```
