---
id: 61695c4aad56f95497c19583
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

请记住，`title` 元素为搜索引擎提供了有关页面的额外信息。 It also displays the content of that `title` element in two more ways:

* in the title bar when the page is open
* in the browser tab for the page when you hover on it. Even if that tab is not active, once you hover on the tab, the `title` text is displayed.

在 `head` 元素中，嵌套一个带有文本 `Colored Markers` 的 `title` 元素。

# --hints--

你应该有一个 `title` 开始标签。

```js
assert(code.match(/<title\s*>/i));
```

你应该有一个 `title` 结束标签。

```js
assert(code.match(/<\/title\s*>/i));
```

你的项目应该有标题 `Colored Markers`。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

请记住，标题文本的大小写和拼写很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
