---
id: 61695d1fbc003856628bf561
title: 步骤 4
challengeType: 0
dashedName: step-4
---

# --description--

要告诉浏览器如何对页面上的字符进行编码，请将 `charset` 设置为 `utf-8`。 `utf-8` 是一个通用字符集，几乎包括所有人类语言中的每个字符。

在 `head` 元素内，嵌套一个 `meta` 元素，其中 `charset` 属性值为 `utf-8`。 请记住，`meta` 元素是自闭合的，不需要结束标签。

# --hints--

你应该有一个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

你的 `meta` 元素应该是一个自闭合元素。

```js
assert(code.match(/<\/meta>/i) === null);
```

你的 `meta` 元素应该有一个 `charset` 属性值为 `utf-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
