---
id: 616965351e74d4689eb6de30
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

你可以在一个网页上有多个自闭合的 `meta` 元素。 将一个 `meta` 元素和其他区别开来的是属性。 你应该为你想要指定的每个属性添加一个新的 meta 元素。

在 `head` 中添加另一个自闭合的 `meta` 元素。 给它一个 `name` 属性，值为 `viewport`，和一个 `content` 属性，值为 `width=device-width, initial-scale=1.0`，这样你的页面在各种设备上看起来是一样的。

# --hints--

你应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

你的新 `meta` 元素应该是一个自闭合元素。

```js
assert(code.match(/<\/meta>/i) === null);
```

你的新 `meta` 元素应将 `name` 属性设置为 `viewport`，并将 `content` 属性设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
