---
id: 616971b3cd990186b66c99a1
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

现在你的外部 CSS 文件已设置好，你可以开始设置页面样式了。

提醒一下，以下是定位段落元素并将其向右对齐的方法：

```css
p {
  text-align: right;
}
```

创建一个定位 `h1` 元素的新 CSS 规则，并将其 `text-align` 属性设置为 `center`。

# --hints--

你应该使用 `h1` 选择器。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

你的 `h1` CSS 规则应该将 `text-align` 属性设置为 `center`。

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
