---
id: 61696ef7ac756c829f9e4048
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在这个项目中，你将使用外部 CSS 文件来设置页面样式。 我们已经为你创建了一个 `styles.css` 文件。 但在你使用它之前，你需要将它链接到页面上。

在 `head` 元素中嵌套一个 `link` 元素。 给它一个 `rel` 属性，设置为 `stylesheet`，和一个 `href` 属性，设置为 `styles.css`。

# --hints--

你应该有一个自闭合的 `link` 元素。

```js
assert(document.querySelectorAll('link').length === 1);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert.exists(document.querySelector('head > link'));
```

您的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
