---
id: 61696ef7ac756c829f9e4048
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

在这个项目中，你将使用外部 CSS 文件来设置页面样式。 我们已经为你创建了一个 `styles.css` 文件。 但在你使用它之前，你需要将它链接到页面上。

在 `head` 元素中嵌套一个 `link` 元素。 Give it a `rel` attribute set to `"stylesheet"` and an `href` attribute set to `"styles.css"`.

# --hints--

You should have one `link` element.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Your `link` element should be a void element, it should not have a closing tag `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Your `link` element should have a `rel` attribute with the value `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Your `link` element should be inside the `head` element.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
