---
id: 61695ab9f6ffe951c16d03dd
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素中嵌套一个 `head` 元素。 在 `head` 元素之后，添加一个 `body` 元素。

# --hints--

你应该有一个打开的 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

你应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

你应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

你应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

你的 `head` 和 `body` 元素应该是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

你的 `head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

你的 `body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
