---
id: 616965351e74d4689eb6de30
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

最后，使用 `viewport` `<meta>` 标签确保你的页面在所有设备上看起来都一样。

在 `head` 中嵌套一个自闭合的 `meta` 元素。 给它一个 `name` 属性值为 `viewport`，和一个 `content` 属性值为 `width=device-width, initial-scale=1.0`。

# --hints--

你应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

你的新 `meta` 元素应该是一个自闭合元素。

```js
assert(code.match(/<\/meta>/i) === null);
```

你的新 `meta` 元素应将 `name` 属性设置为 `viewport`，并将 `content` 属性设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
