---
id: 616d525007b8c5d8b3308b1c
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

如果你的笔在页面上居中，看起来会更好。 一个简单的方法是使用 `margin` <dfn>简写属性</dfn>。

在上一个项目中，你使用 `margin-top` 和 `margin-left` 等属性分别设置元素的边距区域。 `margin` 简写属性可以方便地同时设置多个边距区域。

要将笔在页面上居中，请将其 `margin` 属性设置为 `auto`。 这会将 `margin-top`、`margin-right`、`margin-bottom` 和 `margin-left` 全部设置为 `auto`。

# --hints--

你的 `.marker` CSS 规则应该有一个 `margin` 属性设置为 `auto`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
