---
id: 616d55bd160a95e22453a081
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

现在你已经有了一个以颜色为中心的标记，是时候添加其他标记了。

在 `container` `div` 中，再添加两个 `div` 元素，并给每个元素一个 `marker` 类。

# --hints--

你的第一个新 `div` 元素应该有一个开始标签。

```js
assert([...code.matchAll(/<div.*?>/gi)][2]);
```

你的第一个新 `div` 元素应该有一个结束标签。

```js
assert([...code.matchAll(/<\/div\s*>/gi)][2]);
```

你的第二个新 `div` 元素应该有一个开始标签。

```js
assert([...code.matchAll(/<div.*?>/gi)][3]);
```

你的第二个新 `div` 元素应该有一个结束标签。

```js
assert([...code.matchAll(/<\/div\s*>/gi)][3]);
```

你的新 `div` 元素应该在带有 `container` 类的 `div` 中。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.localName === 'div') && containerChildren.length === 3);
```

你的新 `div` 元素都应该有一个 `marker` 类。

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert(containerChildren.every(child => child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    <div class="container">
      <div class="marker">
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
