---
id: 616d3a67ccf800ad94ec89ae
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

现在，你将添加一些元素，最终将它们设置为彩色笔。

首先，在 `body` 元素中，添加一个 `div` 元素，并将其 `class` 属性设置为 `container`。 确保 `div` 元素位于 `h1` 元素下方。

# --hints--

你应该有一个 `div` 开始标签。

```js
assert(code.match(/<div.*>/i));
```

你应该有一个 `div` 结束标签。

```js
assert(code.match(/<\/div\s*>/i));
```

你的 `div` 元素应该在你的 `body` 元素中。

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

你的 `div` 元素应该在 `h1` 元素之后。

```js
assert.exists(document.querySelector('h1 + div'));
```

你的 `div` 元素应该有一个 `class` 属性值为 `container`。

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
