---
id: 616d4a84b756d9c4b8255093
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

是时候为彩笔添加一些颜色了。 请记住，为元素添加颜色的一种方法是使用 <dfn>color 关键字</dfn>，例如 `black`、`cyan` 或 `yellow`。

提醒一下，以下是定位类 `freecodecamp` 的方法：

```css
.freecodecamp {

}
```

创建一个以 `marker` 类为目标的新 CSS 规则，并将其 `background-color` 属性设置为 `red`。

**注意：** 你将不会在添加 CSS 后看到任何更改。

# --hints--

你应该创建一个类选择器来定位 `marker` 类。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

你的 `.marker` CSS 规则应该将 `background-color` 属性设置为 `red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
