---
id: 616d50b93ba424d6282c99cf
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

请注意，你的笔似乎没有任何颜色。 实际应用了背景颜色，但由于彩笔 `div` 元素为空，因此默认情况下它没有任何高度。

在你的 `.marker` CSS 规则中，将 `width` 属性设置为 `200px`，并将 `height` 属性设置为 `25px`。

# --hints--

你的 `.marker` CSS 规则应该将 `width` 属性设置为 `200px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

你的 `.marker` CSS 规则应该将 `height` 属性设置为 `25px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
