---
id: 61764c602bee6974e7790f35
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

在学校里，你可能已经知道红色、黄色和蓝色是原色，并学会了如何通过混合这些颜色来创造新的颜色。 然而，这是一个过时的模型。

目前，有两种主要的颜色模型：用于电子设备的 <dfn>叠加</dfn> RGB（红、绿、蓝）模型，以及 <dfn>消减</dfn> CMYK（青色、品红色、黄色，黑色）用于打印的模型。 在这个项目中，你将使用 RGB 模型。

首先，将 `class` `one` 添加到第一个标记 `div` 元素。

```html
<div class="exampleClass1 exampleClass2">
  ...content
</div>
```

# --hints--

你应该将 class​​ `one` 添加到第一个标记 `div` 元素。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

你的第一个标记 `div` 应该有 `marker` 和 `one` class。

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
