---
id: 61764f7e5240eb7ccc7f6a0a
title: 步骤 18
challengeType: 0
dashedName: step-18
---

# --description--

然后，创建一个定位 `one` class 的新 CSS 规则，并将其 `background-color` 属性设置为 `red`。

# --hints--

您应该使用类选择器来定位 `one` class。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one'));
```

你的 `.one` CSS 规则应该将 `background-color` 属性设置为 `red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
