---
id: 617b65579ce424bf5f02ca73
title: 步骤 21
challengeType: 0
dashedName: step-21
---

# --description--

目前，有两种主要的颜色模型：用于电子设备的 <dfn>叠加</dfn> RGB（红、绿、蓝）模型，以及用于打印的 <dfn>消减</dfn> CMYK（青色、品红色、黄色，黑色）模型。

在这个项目中，将使用 RGB 模型。 这意味着颜色从黑色开始，色值随着红色、绿色和蓝色数值变化而变化。 一个简单的方法是使用 CSS `rgb` 函数。

创建一个定位 `container` 类的新 CSS 规则，并使用 `rgb(0, 0, 0)` 将其 `background-color` 设置为黑色。

# --hints--

你应该使用类选择器来定位类名为 `container`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

你的 `.container` CSS 容器应当包含设一个属性 `background-color` ，它的值被设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
