---
id: 617b8e0d93a8d10d9a90e720
title: 步骤 23
challengeType: 0
dashedName: step-23
---

# --description--

请注意，笔的 `background-color` 仍然是红色的。 这是因为你将 `rgb` 函数的红色值设置为最大值 `255`，即 100% 红色，并将绿色和蓝色值都设置为 `0`。

现在使用 `rgb` 函数来设置其他颜色。

在 `.two` CSS 规则中，使用 `rgb` 函数将 `background-color` 设置为绿色是最大值，而其他颜色的值是 `0`。 在 `.three` CSS 规则中，使用 `rgb` 函数将 `background-color` 设置为蓝色是最大值，而其他颜色的值是 `0`。

# --hints--

你的 `.two` CSS 规则不应使用 `green` 颜色关键字来设置 `background-color` 属性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

你的 `.three` CSS 规则不应使用 `blue` 颜色关键字来设置 `background-color` 属性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

你的 `.three` CSS 规则应该将 `background-color` 属性设置为 `rgb(0, 0, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
