---
id: 617bd6ec666b1da2587e4e37
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

还有三种复色：黄绿色（绿色+黄色）、天蓝色（蓝色+青色）、玫瑰色（红色+品红色）。

要创建黄绿色，更新 `.one` CSS 规则中的 `rgb` 函数，使红色的值是 `127`，绿色是最大值。

要创建天蓝色，更新 `.two` CSS 规则中的 `rgb` 函数，使绿色的值是 `127`，蓝色是最大值。

对于有时被称为亮粉红色的玫瑰色，更新 `.three` CSS 规则中的 `rgb` 函数，使蓝色的值为 `127`，红色为最大值。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(127, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 127, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

你的 `.three` CSS 规则应该将 `background-color` 属性设置为 `rgb(255, 0, 127)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
