---
id: 617bc4824e233180553a8069
title: 步骤 33
challengeType: 0
dashedName: step-33
---

# --description--

要创建复色蓝紫色，需要组合品红色和蓝色。 更新 `.three` CSS 规则中的 `rgb` 函数，使蓝色是最大值，红色的值是 `127`。

# --hints--

你的 `.three` CSS 规则应该将 `background-color` 属性设置为 `rgb(127, 0, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(127, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

--fcc-editable-region--
.three {
  background-color: rgb(255, 0, 255);
}
--fcc-editable-region--

```
