---
id: 617b8b38f32bf2080a140675
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

函数是一个可以输入并执行特定动作的代码。 CSS 的 `rgb` 函数接收红色、绿色和蓝色的数值或 <dfn>arguments（参数）</dfn>，并生成一个颜色。

```css
rgb(red, green, blue);
```

每个红色、绿色和蓝色值都是一个从 `0` 到 `255` 的数字。 `0` 表示该颜色的 0%，并且是黑色。 `255` 意味着该颜色的 100%。

在 `.one` 的CSS规则中，修改 `rgb` 函数中的 `red`的值。 对于 `rgb` 函数，将红色的值设置为 `255`，将绿色的值设置为 `0`，将蓝色的值设置为 `0`。

# --hints--

你的 `.one` CSS 规则不应该使用 `red` 关键字来设置 `background-color`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

你的 `.one` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(255, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
