---
id: 617bc3386dc7d07d6469bf20
title: 步骤 32
challengeType: 0
dashedName: step-32
---

# --description--

注意，要创建橙色，你需要增加红色的强度，并降低绿色 `rgb` 值的强度。 这是因为橙色是红色和黄色的结合，再色轮上处于两种颜色之间。

要创建复色亮绿色，需要组合青色和绿色。 更新 `.two` CSS 规则中的 `rgb` 函数，使绿色是最大值，蓝色的值是 `127`。

# --hints--

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 255, 127)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
