---
id: 617b97abd9f3f61d1590b815
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

在补充 RGB 颜色模型中，<dfn>原色</dfn> 是在组合后成为纯白色的颜色。 但要做到这一点，每种颜色都需要处于最高强度。

在组合颜色之前，将绿色标记设置回纯绿色。 对于 `.two` CSS 规则中的 `rgb` 函数，将绿色设置回最大值 `255`。

# --hints--

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 255, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
