---
id: 618a132676346ac9f7fd59dd
title: 步骤 39
challengeType: 0
dashedName: step-39
---

# --description--

注意到你的眼睛是如何自然地被中心的红色所吸引的吗？ 在设计网站时，你可以用这种效果来吸引人们对重要标题、按钮或链接的注意。

除了补色之外，还有其他几种重要的颜色组合，但稍后你会了解到这些。

现在，使用 `.two` CSS 规则中的 `rgb` 函数将 `background-color` 设置为黑色。

# --hints--

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(0, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
