---
id: 618a1275e873dcc803c2d1aa
title: 步骤 38
challengeType: 0
dashedName: step-38
---

# --description--

下一步，在 `.one` CSS 规则中，使用 `rgb` 函数将 `background-color` 设置为黑色。 下一步，在 `.two` CSS 规则中，使用 `rgb` 函数将 `background-color` 设置为红色。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

你的 `.two` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(255, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
  background-color: rgb(0, 255, 255);
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
