---
id: 618a0927005553b74bfa05ae
title: 步骤 36
challengeType: 0
dashedName: step-36
---

# --description--

色轮是一个圆圈，其中相似的颜色或<dfn>色调</dfn>彼此靠近，而不同的颜色相距较远。 例如，纯红色介于玫瑰色和橙色之间。

色轮上相互对立的两种颜色称为<dfn>补色</dfn>。 如果将两种互补色组合在一起，它们会产生灰色。 但当它们并排放置时，这些颜色会产生强烈的视觉对比，显得更亮。

在 `.one` CSS 规则的 `rgb` 函数中，将红色值设置为 `255` 的最大值以生成纯红色。 在 `.two` CSS 规则的 `rgb` 函数中，将绿色和蓝色的值设置为 `255` 的最大值以生成青色。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(255, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 255, 255)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
