---
id: 618a00ed1ca871a2b3aca0cb
title: 步骤 35
challengeType: 0
dashedName: step-35
---

# --description--

现在你已经学习了色轮上的所有原色、间色和复色，更容易理解其他颜色理论概念和它们如何影响设计。

首先，在 CSS 规则 `.one`、`.two` 和 `.three` 中，调整 `rgb` 函数中的值，将每个元素的 `background-color` 设置为纯黑色。 请记住 `rgb` 函数使用加成色模型。颜色起始为黑色，随红色、绿色和蓝色的值增加而变化。

# --hints--

你的 `.one` CSS 规则应该有一个 `background-color` 属性，设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

你的 `.two` CSS 规则应该有一个 `background-color` 属性设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

你的 `.three` CSS 规则应该将 `background-color` 属性设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
