---
id: 619b7c3c83de403126b69c1e
title: 步骤 47
challengeType: 0
dashedName: step-47
---

# --description--

你可能已经熟悉从 0 到 9 的十进制或以 10 为底的值。 十六进制或以 16 为底的值，从 0 到 9，然后是 A - F：

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

对于十六进制颜色，`00` 是该颜色的 0%，`FF` 是 100%。 所以 `#00FF00` 转换为 0% 红色、100% 绿色和 0% 蓝色，与 `rgb(0, 255, 0)` 相同。

通过将十六进制颜色的绿色值设置为 `7F` 来降低绿色强度。

# --hints--

你的 `.green` CSS 规则应该将 `background-color` 属性设置为 `#007F00`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
