---
id: 619b72a0db211f1c29afb906
title: 步骤 42
challengeType: 0
dashedName: step-42
---

# --description--

更新 `.one` CSS 规则以定位新的 `red` 类。

# --hints--

你的代码不应再有 `.one` 类选择器。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.one'));
```

你应该使用类选择器来定位类 `red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red'));
```

你的 `.red` CSS 规则应该将 `background-color` 属性设置为 `rgb(0, 0, 0)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
