---
id: 61a493ead935c148d2b76312
title: 步骤 50
challengeType: 0
dashedName: step-50
---

# --description--

`linear-gradient` 函数非常灵活——这是你将在本教程中使用的基本语法：

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` 是表明过渡的线的方向。 `color1` 和 `color2` 是颜色参数，它们是将在过渡本身中使用的颜色。 这些可以是任何类型的颜色，包括颜色关键字、十六进制、`rgb` 或 `hsl`。

现在，你将沿 90 度线将红到绿渐变应用到第一支笔。

首先，在 `.red` CSS 规则中，将 `background` 属性设置为 `linear-gradient()`，并将它的值设置为 `90deg`，作为 `gradientDirection`。

# --hints--

你的 `.red` CSS 规则应该有一个 `background` 属性，其值为 `linear-gradient(90deg)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
