---
id: 61a4ae5f29eb5584187201c3
title: 步骤 61
challengeType: 0
dashedName: step-61
---

# --description--

对于这支笔，你将使用十六进制颜色代码生成渐变色。

使用 `linear-gradient` 函数并将 `gradientDirection` 设置为 `180deg`。 对于第一个颜色参数，使用十六进制颜色代码，其值 `55` 表示红色，`68` 表示绿色，`0D` 表示蓝色。

# --hints--

你的 `.green` CSS 规则应该有一个 `background` 属性，其值为 `linear-gradient(180deg, #55680D)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.green\{.*?background:linear-gradient\(180deg,#55680D\)/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}

--fcc-editable-region--
.green {
  background: #007F00;
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
