---
id: 61a489b8579e87364b2d2cdb
title: 步骤 49
challengeType: 0
dashedName: step-49
---

# --description--

你已经了解了几种在 CSS 中设置纯色的方法，但你也可以在元素上使用颜色过渡或<dfn>渐变色</dfn>。

渐变是一种颜色过渡到另一种颜色。 CSS `linear-gradient` 函数可让你控制沿线的过渡方向以及使用的颜色。

要记住的一件事是 `linear-gradient` 函数实际上创建了一个 `image` 元素，并且通常与可以接受图像作为一个值的 `background` 属性配对。

在 `.red` CSS 规则中，将 `background-color` 属性更改为 `background`。

# --hints--

你的 `.red` CSS 规则应该有一个 `background` 属性，其值为 `rgb(255, 0, 0)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
