---
id: 61a49d15bdbb5e57cc6fd280
title: 步骤 54
challengeType: 0
dashedName: step-54
---

# --description--

颜色中间点（color-stop）允许你微调颜色沿渐变线的位置。 它们是像 `px` 这样的长度单位，或者是在 `linear-gradient` 函数中定位颜色的百分比。

例如，在这个红-黑渐变中，从红色到黑色的过渡发生在渐变线的 90% 处，因此红色占据了大部分可用空间：

```css
linear-gradient(90deg, red 90%, black);
```

在 `linear-gradient` 函数中，在第一个红色参数之后添加 `75%` 颜色中间点。 不要将颜色中间点添加到其他颜色参数。

# --hints--

你的 `.red` CSS 规则应该有一个 `background` 属性设置为 `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`。

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
