---
id: 61a5ca57f50ded36d33eef96
title: 步骤 66
challengeType: 0
dashedName: step-66
---

# --description--

现在你会将渐变应用到蓝色笔上，这次使用 `hsl` 函数作为颜色参数。

在 `.blue` CSS 规则中，将 `background-color` 属性更改为 `background`。

# --hints--

你的 `.blue` CSS 规则应该有一个 `background` 属性，其值为 `hsl(240, 100%, 50%)`。

```js
assert.match(__helpers.removeWhiteSpace(code), /\.blue\{.*?background:hsl\(240,100%,50%\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

--fcc-editable-region--
.blue {
  background-color: hsl(240, 100%, 50%);
}
--fcc-editable-region--

```
