---
id: 61a5bfe091060f1d6a629dd0
title: 步骤 64
challengeType: 0
dashedName: step-64
---

# --description--

即使没有颜色中间点，你可能已经注意到绿色笔的颜色在与红色笔相同的点位过渡。 第一种颜色位于渐变线的开头 (0%)，第二种颜色位于中间 (50%)，最后一种颜色位于渐变线的结尾 (100%)。

`linear-gradient` 函数会自动为你计算这些值，并在默认情况下沿渐变线均匀分布颜色。

在 `.red` CSS 规则中，从 `linear-gradient` 函数中移除三个色标，以稍微清理你的代码。

# --hints--

你的 `.red` CSS 规则应该有一个 `background` 属性设置为 `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`。

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}
--fcc-editable-region--

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
