---
id: 61b0936d9e7fc020b43b1bb8
title: 步骤 78
challengeType: 0
dashedName: step-78
---

# --description--

看起来你的笔筒不见了，但别担心——它还在那里。 发生的情况是，你的新笔盖 `div` 占据了笔的整个宽度，并将笔筒推到了下一行。

这是因为 `div` 元素的 `display` 属性默认是 `block`。 因此，当两个 `block` 元素彼此相邻时，它们会像实际的块一样堆叠。 例如，你的彩笔元素都是堆叠在一起的。

要将两个 `div` 元素放置在同一行上，需要将它们的 `display` 属性设置为 `inline-block`。

创建一个新规则以同时定位 `cap` 和 `sleeve` 类，并将 `display` 设置为 `inline-block`。

# --hints--

你应该使用类选择器来定位 `cap` 和 `sleeve` 类。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

你的 CSS 规则应该将 `display` 属性设置为 `inline-block`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
