---
id: 61b093329e7fc020b43b1bb5
title: 步骤 75
challengeType: 0
dashedName: step-75
---

# --description--

你已经熟悉使用 `rgb` 函数来设置颜色。 要将 alpha 通道添加到 `rgb` 颜色，请改用 `rgba` 函数。

`rgba` 函数的工作方式与 `rgb` 函数类似，但在 `0` 到 `1.0` 之间多取一个数字作为 alpha 通道：

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

你也可以使用带有 `hsl` 和 `hex` 颜色的 alpha 通道。 你将很快看到如何实现这一点。

在 `.sleeve` 规则中，使用 `rgba` 函数将 `background-color` 属性设置为具有 50% 不透明度的纯白色。

# --hints--

你的 `.sleeve` CSS 规则应将 `background-color` 属性设置为 `rgba(255, 255, 255, 0.5)`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor === 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
