---
id: 61b093219e7fc020b43b1bb4
title: 步骤 74
challengeType: 0
dashedName: step-74
---

# --description--

另一种设置元素不透明度的方法是使用 <dfn>alpha 通道</dfn>。 与 `opacity` 属性类似，alpha 通道控制颜色的透明或不透明程度。

你已经用命名的颜色和 `opacity` 属性设置了笔筒的不透明度，但你可以给其他 CSS 颜色属性添加一个 alpha 通道。

在 `.sleeve` 规则中，移除 `opacity` 属性和值。

# --hints--

你的 `.sleeve` CSS 规则不应具有 `opacity` 属性和值。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
  opacity: 0.5;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
