---
id: 61b093179e7fc020b43b1bb3
title: 步骤 73
challengeType: 0
dashedName: step-73
---

# --description--

<dfn>透明度</dfn>描述事物的不透明性或不透明性。 例如，一堵实心墙是不透明的，没有光线可以通过。 但是水杯要透明得多，你可以透过玻璃看到另一边。

使用 CSS `opacity` 属性，你可以控制元素的不透明或透明程度。 使用值 `0` 或 0%，元素将完全透明，而在 `1.0` 或 100% 时，元素将像默认情况下一样完全不透明。

在 `.sleeve` CSS 规则中，将 `opacity` 属性设置为 `0.5`。

# --hints--

你的 `.sleeve` CSS 规则应该将 `opacity` 属性设置为 `0.5`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
