---
id: 61b0a44a6b865738ba49b9fb
title: 步骤 86
challengeType: 0
dashedName: step-86
---

# --description--

你要做的最后一件事是为每支笔添加轻微的阴影，使它们看起来更加真实。

`box-shadow` 属性允许你在元素周围应用一个或多个阴影。 这是基本语法：

```css
box-shadow: offsetX offsetY color;
```

下面是 `offsetX` 和 `offsetY` 值的作用：

* `offsetX` 和 `offsetY` 都接受数字值，单位是 `px` 和其他 CSS 单位
* 正 `offsetX` 值将阴影向右移动，负值将它向左移动
* 正 `offsetY` 值将阴影向下移动，负值将它向上移动
* 如果你想要 `offsetX` 和 `offsetY` 其中一个值为零，或两个都为零（`0`），那么你不需要添加单位。 每种浏览器都知道零意味着没有变化。

阴影的高度和宽度由应用阴影的元素的高度和宽度决定。 你也可以使用可选的 `spreadRadius` 值来扩大阴影的范围。 之后会介绍更多。

首先为红色笔添加一个简单的阴影。

在 `.red` CSS 规则中，添加 `box-shadow` 属性，其中 `offsetX` 的值为 `5px`，`offsetY` 为 `5px`，`color` 为 `red`。

# --hints--

你的 `.red` CSS 规则应该有一个 `box-shadow` 简写属性，其值为 `5px 5px red`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
