---
id: 60eebd07ea685b0e777b5583
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Bem-vindo ao projeto do formulário de registro! Comece adicionando a declaração `!DOCTYPE html` no topo do documento para que o navegador saiba que tipo de documento está lendo.

# --hints--

O código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento para `html`.

```js
assert(code.match(/html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/html\s*>/gi));
```

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
