---
id: 60f0286404aefb0562a4fdf9
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Adicione os elementos `title` e `meta` a `head`. Dê ao projeto o título de `Registration Form` e dê ao atributo `charset` um valor de `UTF-8` para o elemento `meta`.

# --hints--

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

O projeto deve ter o título `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
