---
id: 60f027099a15b00485563dd2
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione as tags de abertura e fechamento de `html` abaixo do `DOCTYPE` para que você tenha um lugar para começar a colocar algum código.

# --hints--

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

O elemento `html` deve ter uma tag de abertura.

```js
assert(code.match(/<html\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

As tags `html` devem estar na ordem correta.

```js
assert(code.match(/<html\s*>\s*<\/html\s*>/));
```

Você deve ter apenas um elemento `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
