---
id: 60f0286404aefb0562a4fdf9
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Adicione um elemento `title` e um elemento `meta` dentro do elemento `head`. Dê ao projeto o título de `Registration Form` e adicione um atributo `charset` com um valor de `utf-8` ao elemento `meta`.

# --hints--

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

O projeto deve ter o título `Registration Form`. Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Você deve criar um elemento `meta` ao redor do elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Você deve dar ao elemento `meta` um atributo `charset` com o valor `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

O elemento `meta` deve ser um elemento nulo, sem uma tag final `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
