---
id: 60f027099a15b00485563dd2
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Abaixo do `DOCTYPE`, adicione um elemento `html` com um atributo `lang` definido como `en` para que você tenha um lugar para começar a colocar o código.

# --hints--

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

O elemento `html` deve ter uma tag de abertura e um atributo `lang` com o valor de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

As tags `html` devem estar na ordem correta.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Você deve ter apenas um elemento `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
