---
id: 60f02e7361b68405e27b62a5
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro do `body`, crie um título para dar contexto ao conteúdo adicionando um `h1` com o texto `Registration Form`.

# --hints--

Você deve adicionar o `h1` dentro do `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Você deve atribuir ao elemento `h1` o texto `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
