---
id: 60f030d388cb74067cf291c3
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Abaixo do título, use o texto a seguir dentro de um elemento de parágrafo para incentivar os usuários a se registrar:

```md
Please fill out this form with the required information
```

# --hints--

Você deve adicionar um elemento `p` ao `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Você deve adicionar um elemento `p` abaixo do `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Você deve dar ao elemento `p` o texto `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
