---
id: 60f1922fcbd2410527b3bd89
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--
A unidade `vh` significa viewport height (altura da viewport) e é relativa a 1% da `height` (altura) da viewport.

É hora de melhorar o projeto com um pouco de CSS. Comece fornecendo ao `body` uma `width` de `100%`, e uma `height` de `100vh`.

# --hints--

Você deve usar o elemento `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Você deve dar ao `body` uma `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Você deve dar ao `body` uma `height` de `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
