---
id: 60f1a9cbd23023082e149fee
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Assim é melhor. Agora, para deixar o fundo mais confortável de se ver, altere a `background-color` do `body` para `#1b1b32`. Em seguida, para ver o texto, altere a `color` para `#f5f6f7`.

# --hints--

Você deve adicionar `background-color` dentro do seletor do elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Você deve definir o valor da propriedade `background-color` como `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Você deve adicionar `color` dentro do elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Você deve dar à `color` um valor de `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
