---
id: 60f80e0081e0f2052ae5b505
title: Passo 19
challengeType: 0
dashedName: step-19
---

# --description--

Especificar o atributo `type` de um elemento de formulário é importante para o navegador saber que tipo de dados ele deve esperar. Se `type` não for especificado, o navegador o definirá por padrão como `text`.

Dê aos dois primeiros elementos `input` um atributo `type` de `text`, ao terceiro um atributo `type` de `email` e ao quarto um atributo `type` de `password`.

O tipo `email` só permite e-mails com `@` e um `.` no domínio. O tipo `password` oculta a entrada e avisa se o site não usa HTTPS.

# --hints--

Você deve dar ao primeiro elemento `input` um atributo `type` de `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Você deve dar ao segundo elemento `input` um atributo `type` de `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Você deve dar ao terceiro elemento `input` um atributo `type` de `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Você deve dar ao quarto elemento `input` um atributo `type` de `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
