---
id: 60f8604682407e0d017bbf7f
title: Passo 26
challengeType: 0
dashedName: step-26
---

# --description--

Dentro de cada elemento `label` e imediatamente após o elemento `input`, adicione um espaço e adicione respectivamente o seguinte texto:

```md
Personal 
Business 
```

# --hints--

Você deve dar ao primeiro elemento `label` o texto `Personal`.

```js
assert.include(document.querySelector('fieldset:nth-child(2) > label')?.innerText, 'Personal');
```

Você deve dar ao segundo elemento `label` o texto `Business`.

```js
assert.include(document.querySelector('fieldset:nth-child(2) > label:nth-child(2)')?.innerText, 'Business');
```

Você deve dar ao primeiro elemento `label` espaço no início.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) > label')?.innerText?.[0], ' ');
```

Você deve dar ao segundo elemento `label` espaço no início.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) > label:nth-child(2)')?.innerText?.[0], ' ');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
--fcc-editable-region--
        <label><input type="radio" /></label>
        <label><input type="radio" /></label>
--fcc-editable-region--
      </fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
