---
id: 60f805f813eaf2049bc2ceea
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

Adicione um elemento `input` dentro de cada elemento `label`. Não se esqueça de adicionar cada elemento `input` após o texto da `label` e de incluir um espaço após os dois-pontos.

# --hints--

Você deve adicionar quatro elementos `input` ao elemento `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

Você deve aninhar os elementos `input` dentro dos elementos `label`.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

Você deve adicionar o primeiro `input` após o texto `Enter Your First Name:` do elemento `label` e incluir um espaço depois dos dois-pontos.

```js
assert.equal(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), 'Enter Your First Name: <input>');
```

Você deve adicionar o segundo `input` após o texto `Enter Your Last Name:` do elemento `label` e incluir um espaço depois dos dois-pontos.

```js
assert.equal(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), 'Enter Your Last Name: <input>');
```

Você deve adicionar o terceiro `input` após o texto `Enter Your Email:` do elemento `label` e incluir um espaço depois dos dois-pontos.

```js
assert.equal(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), 'Enter Your Email: <input>');
```

Você deve adicionar o quarto `input` após o texto `Create a New Password:` do elemento `label` e incluir um espaço depois dos dois-pontos.

```js
assert.equal(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), 'Create a New Password: <input>');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
