---
id: 60f8604682407e0d017bbf7f
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

Para os termos e condições, adicione uma `input` com um `type` de `checkbox` para o terceiro elemento `label`. Além disso, como não queremos que os usuários se cadastrem, sem ter lido os termos e condições, faça com que seja `required`.

# --hints--

Você deve adicionar um elemento `input` dentro do terceiro elemento `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(2) label:nth-child(3) input'));
```

Você deve adicionar um atributo `type` de valor `checkbox` ao elemento `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(3) input')?.type, 'checkbox');
```

Você deve adicionar um atributo `required` ao elemento `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(3) input')?.required, true);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name: <input type="text" required /></label>
        <label>Enter Your Last Name: <input type="text" required /></label>
        <label>Enter Your Email: <input type="email" required /></label>
        <label>Create a New Password: <input type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label><input type="radio" /></label>
        <label><input type="radio" /></label>
        <label></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}

label {
    display: block;
    margin: 0.5rem 0;
}

```
