---
id: 60f84ec41116b209c280ba91
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

Com `type="password"` você pode usar o atributo `pattern` para definir uma expressão regular que a senha deve corresponder para ser considerada válida.

Adicione um atributo `pattern` ao elemento `input` de senha para exigir que a entrada corresponda aos valores: `[a-z0-5]{8,}`

O comando acima é uma expressão regular que corresponde a oito ou mais letras minúsculas ou os dígitos de `0` a `5`. Agora, remova o atributo `minlength` e experimente-o.

# --hints--

Você deve dar ao elemento `input` um atributo `pattern`.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

Você deve dar ao atributo `pattern` um valor de `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

Você deve remover o atributo `minlength` do primeiro elemento `input`.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
