---
id: 60f803d5241e6a0433a523a1
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--
A unidade `rem` significa root `em` e é relativa ao tamanho da fonte do elemento `html`.

Como os elementos `label` estão em linha por padrão, serão exibidos lado a lado na mesma linha, tornando o texto difícil de ler. Para fazer com que apareçam em linhas separadas, adicione `display: block` ao elemento `label` e adicione uma `margin` de `0.5rem 0`, para separá-los um do outro.

# --hints--

Você deve usar o seletor `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Você deve definir a propriedade `display` como `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Você deve definir a propriedade `margin` como `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
