---
id: 60f83e7bfc09900959f41e20
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Certos valores de atributo `type` vêm com formulário de validação integrado. Por exemplo, `type="email"` requer que o valor seja um endereço de e-mail válido.

Adicione uma validação personalizada ao elemento `input` de senha, adicionando um atributo `minlength` com um valor de `8`. Isso evita que entradas com menos de 8 caracteres sejam enviadas.

# --hints--

Você deve dar ao elemento `input` um atributo `minlength`.

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

Você deve dar ao atributo `minlength` um valor de `8`.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
