---
id: 60fad1cafcde010995e15306
title: Passo 36
challengeType: 0
dashedName: step-36
---

# --description--

Com envios de formulário, é útil e boa prática fornecer a cada elemento enviado um atributo `name`. Este atributo é usado para identificar o elemento no envio do formulário.

Vá em frente e dê a cada elemento que possa ser enviado um atributo `name` único de sua escolha. _Exceto para as duas entradas `radio`._

# --hints--

Você deve dar ao elemento `input` que pede pelo primeiro nome um atributo `name`. _Dica: eu teria escolhido `first-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

Você deve dar ao elemento `input` que pede pelo sobrenome um atributo `name`. _Dica: eu teria escolhido `last-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

Você deve dar ao elemento `email` um atributo `name`. _Dica: eu teria escolhido `email`_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

Você deve dar ao elemento `password` um atributo `name`. _Dica: eu teria escolhido `password`_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

Você deve dar ao elemento `checkbox` um atributo `name`. _Dica: eu teria escolhido `terms`_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

Você deve dar ao elemento `file` um atributo `name`. _Dica: eu teria escolhido `file`_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

Você deve dar ao elemento `number` um atributo `name`. _Dica: eu teria escolhido `age`_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

Você deve dar ao elemento `select` um atributo `name`. _Dica: eu teria escolhido `referrer`_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

Você deve dar ao elemento `textarea` um atributo `name`. _Dica: eu teria escolhido `bio`_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

Você deve não dar atributos `name` para nenhum elemento `option`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

Você deve não dar atributos `name` para nenhum elemento `label`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

Você deve não dar atributos `name` para nenhum elemento `fieldset`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" type="text/css" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input type="text" required /></label>
        <label>Enter Your Last Name: <input type="text" required /></label>
        <label>Enter Your Email: <input type="email" required /></label>
        <label>Create a New Password: <input type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label>
          <input type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
              </label>
      </fieldset>
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" />
              </label>
        <label>How did you hear about us?
          <select>
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label>Provide a bio:
                  <textarea rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
              </label>
      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}

label {
    display: block;
    margin: 0.5rem 0;
}

```
