---
id: 60fac4095512d3066053d73c
title: Passo 29
challengeType: 0
dashedName: step-29
---

# --description--

Adicionar uma lista suspensa ao formulário é fácil com o elemento `select`. O elemento `select` é um repositório para um grupo de elementos `option` e o elemento `option` funciona como um rótulo para cada opção de seleção. Os dois elementos requerem tags de fechamento.

Comece adicionando um elemento `select` abaixo dos dois elementos `label`. Você deve aninhar 5 elementos `option` dentro do elemento `select`.

# --hints--

Você deve adicionar um elemento `select` ao terceiro `fieldset`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > select'));
```

Você deve aninhar cinco elementos `option` dentro do elemento `select`.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > select > option')?.length, 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Registration Form</title>
      <link rel="stylesheet" type="text/css" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name: <input type="text" required /></label>
        <label>Enter Your Last Name: <input type="text" required /></label>
        <label>Enter Your Email: <input type="email" required /></label>
        <label>Create a New Password: <input type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label>
          <input type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
              </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" />
              </label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
    color: #f5f6f7;
}

label {
    display: block;
    margin: 0.5rem 0;
}

```
