---
id: 63541ef4f96cd82e8e6c788a
title: Etapa 12
challengeType: 0
dashedName: step-12
---

# --description--

O atributo `method` especifica como enviar dados de formulário para o URL especificado no atributo `action`. Os dados do formulário podem ser enviados através de uma solicitação `GET` como parâmetros de URL (com `method="get"`) ou através de uma solicitação `POST` como dados no corpo da solicitação (com `method="post"`).

Defina o atributo `method` para que envie os dados do formulário através de uma solicitação `POST`.

# --hints--

Você não deve adicionar um novo elemento `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

O elemento `form` deve ter um atributo `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

O atributo `method` deve ser definido como `post`.

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
