---
id: 60a3e3396c7b40068ad69977
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Utilice el relleno para ajustar el espaciado dentro de un elemento.

En `.frame`, use la abreviatura `padding` para aumentar el espacio entre los elementos `.frame` y `.canvas` por `50px`. La taquigrafía aumentará el espacio en la parte superior, inferior, izquierda y derecha del borde y el lienzo del elemento dentro.

# --hints--

Debe establecer la propiedad `padding` en `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

El elemento `.frame` debe tener un valor `padding` de `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
