---
id: 60a3e3396c7b40068ad6998b
title: Paso 33
challengeType: 0
dashedName: step-33
---

# --description--

Es útil que sus márgenes empujen en una dirección.

En este caso, el margen inferior del elemento `.one` empuja `.two` 20 píxeles hacia abajo.

En el selector `.two`, use la abreviatura `margin` para establecer el margen superior en `0`, el margen horizontal en `auto` y el margen inferior en `20px`. Esto eliminará su margen superior, lo centrará horizontalmente y establecerá su margen inferior en 20 píxeles.

# --hints--

Debe establecer la propiedad `margin` en `0 auto 20px`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

El elemento `.two` debe tener un valor `margin` de `0 auto 20px`.

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
