---
id: 60a3e3396c7b40068ad69978
title: Paso 15
challengeType: 0
dashedName: step-15
---

# --description--

Use márgenes para ajustar el espaciado fuera de un elemento.

Usando la propiedad `margin`, proporcione al elemento `.frame` un margen vertical de `20px` y un margen horizontal de `auto`. Esto moverá el marco 20 píxeles hacia abajo y lo centrará horizontalmente en la página.

# --hints--

Debe establecer la propiedad `margin` en `20px auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Su elemento `.frame` debe tener un valor `margin` de `20px auto`.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
