---
id: 60a3e3396c7b40068ad69970
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Antes de que pueda comenzar a aplicar estilo al `div` que agregó, debe vincular su CSS a su HTML.

Agregue un elemento `link` para vincular el archivo `styles.css`. Establezca el `href` en `styles.css` y recuerde establecer el atributo `rel` en `stylesheet`.

# --hints--

El código debe tener un elemento `link`.

```js
assert(/<link/.test(code))
```

El elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

El elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
