---
id: 60a3e3396c7b40068ad69979
title: Paso 16
challengeType: 0
dashedName: step-16
---

# --description--

Agrega un nuevo elemento `div` dentro de tu elemento `.canvas`.

Asigne al nuevo atributo `div` el atributo `class` con un valor de `one`. Este será tu primer rectángulo.

# --hints--

Debe crear un nuevo elemento `div`.

```js
assert(document.querySelectorAll('div').length === 3);
```

Debe anidar el nuevo elemento `div` dentro de su elemento `.canvas`.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

El nuevo `div` debe tener un atributo `class` con un valor `one`.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
