---
id: 60a3e3396c7b40068ad6996c
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

En el modelo de cuadro CSS, cada elemento HTML se trata como un cuadro con cuatro áreas.

Imagine que recibe una caja de su minorista en línea favorito: el contenido es el elemento en la caja, o en nuestro caso, un encabezado, párrafo o elemento de imagen.

Cambie el atributo `src` en el atributo `<img>` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` a `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

El elemento `img` debe tener un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

