---
id: 60a3e3396c7b40068ad6996a
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

A estas alturas, debería estar familiarizado con los elementos básicos que debería tener una página HTML.

Configure su código con una declaración de `DOCTYPE`, un elemento `html` con el idioma establecido en inglés, un elemento `head` y un elemento `body`.

# --hints--

Tu código debe tener una declaración `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Debes tener una etiqueta `<html>` de apertura con un atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Tu código debe tener un elemento `head` dentro del elemento `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Tu código debe tener un elemento `body` dentro del elemento `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

Tu elemento `head` debe estar antes de tu elemento `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
