---
id: 60a3e3396c7b40068ad69974
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Cada pintura necesita un marco.

Ajuste el elemento `.canvas` en otro `div`. Asigne a ese `div` la clase `frame`.
# --hints--

Debe agregar un nuevo elemento `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

El elemento `.canvas` debe estar anidado en el nuevo elemento `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

El nuevo `div` debe tener un `class` con el valor `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Su nuevo `div` debe estar dentro de su elemento `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
