---
id: 60a3e3396c7b40068ad69980
title: Paso 23
challengeType: 0
dashedName: step-23
---

# --description--

Agregue otro `div` con un valor `class` de `two` justo debajo de su elemento `one`. Este será tu segundo rectángulo.

# --hints--

No debe cambiar el elemento `.one` existente.

```js
assert(document.querySelectorAll('.one').length === 1);
```

Debes tener un segundo elemento `div` en tu elemento `.canvas`.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

El segundo elemento `div` debe tener un valor `class` de `two`.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

El elemento `.two` debe venir después del elemento `.one`.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
