---
id: 60a3e3396c7b40068ad6996d
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El contenido está rodeado por un espacio llamado relleno, similar a cómo el plástico de burbujas separa un elemento de la caja que lo rodea.

Piense en el borde como la caja de cartón en la que se envió su artículo.

Cambie el atributo `src` a `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

El elemento `img` debe tener un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
