---
id: 60a3e3396c7b40068ad69975
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Escribe una nueva regla usando el selector de clase `.frame`.

Usa la declaración abreviada `border` para darle al elemento `.frame` un borde negro sólido con un ancho de `50px`.

# --hints--

Su código debe tener un selector `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Debe establecer la propiedad `border` en `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

El elemento `.frame` debe tener un `50px solid black` `border`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
