---
id: 65ef190c6b51e9b5a5f7ed29
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Give the `square_root_bisection` method the following parameters:

- `square_target`: The number for which you want to find the square root.
- `tolerance` (optional): The acceptable difference between the square of the approximate root value and the actual target value (default is `1e-7`).
- `max_iterations` (optional): The maximum number of iterations to perform (default is `100`). If the method doesn't converge within this limit, you'll assume the solution is not found.
  
# --hints--

Your function should have these parameters: `source_target`, `tolerance = 1e-7`, and `max_iterations = 100`. The order matters.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const square_root_bisection = __helpers.python.getDef("\n" + transformedCode, "square_root_bisection");
        assert.exists(square_root_bisection);
        const { function_parameters } = square_root_bisection;
        assert.equal(function_parameters.replace(/ /g, ''), "square_target,tolerance=1e-7,max_iterations=100");
    }
})
```

# --seed--

## --seed-contents--

```py

--fcc-editable-region--
def square_root_bisection():
    pass
--fcc-editable-region--

```

