---
id: 688003f589b875af60273dd9
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Still inside the `__init__` method, assign `None` to `self.next`. This sets the `next` pointer to `None`, meaning that the new node does not yet point to any other node in the list.

Later, `next` can be updated to point to another node to form the chain of a linked list.

# --hints--

Inside the `__init__` method, you should initialize the `next` attribute to `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
