---
id: 68857e0a49baf4751f737a2a
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Now, after the `is_empty` method, create a method called `add` inside the `LinkedList` class. The new method should have two parameters, `self` and `element`.

Inside your new method, create a variable named `node` and assign it a new instance of `Node` by using `self.Node(element)`. This will create a new node to be added to the linked list.

# --hints--

You should create an `add` method inside the `LinkedList` class.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Your `add` method should have two parameters, `self` and `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

You should assign `self.Node(element)` to `node` in the `add` method.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
