---
id: 68858a1b4e867acd87f6f7e8
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

If the list is empty, then the new node becomes the first element (the `head`). Inside the `add` method, create an `if` statement that checks if the list is empty using `self.is_empty()` and assign `node` to `self.head`.

# --hints--

You should check if the list is empty by using the `is_empty()` method.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

You should assign `node` to `self.head` inside the `if` statement.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
