---
id: 688649c8673f842f49228a5e
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

If the list is not empty, the method should find the last node and link the new node to its `next` pointer.

To do this, create an `else` block after the `if` block. Inside the `else` block, create a `current_node` variable and assign `self.head` to it.

# --hints--

You should assign `self.head` to `current_node` inside the `else` block.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
