---
id: 68867c3898b09115d8c5140f
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

The `while` loop you've just created enables you to go through each node in the list until the last node is reached. Once that happens, you can add the node to the list.

In the `else` block, right after the `while` block, assign `node` to `current_node.next`.

# --hints--

You should assign `node` to `current_node.next` in the `else` block.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
