---
id: 68868f502659139bed8e5c90
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

After adding a new node to the list, you need to update its length. At the end of the `add` method, increment `self.length` by `1`.

# --hints--

You should increment `self.length` by `1` at the end of the `add` method.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").has_stmt("self.length += 1")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
