---
id: 6888f9827637322d4f309192
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Inside the `remove` method, create a variable named `current_node` and assign it the value of `self.head`.

This starts the traversal of the linked list from the first node. `current_node` will be used to keep track of the node currently being examined as you look for the node containing the specified element to remove.

# --hints--

You should assign `self.head` to `current_node` inside the `remove` method.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
