---
id: 6888fe6c745b6852ed240229
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Create a `while` loop that checks when `current_node` is not `None` and when `current_node.element` is not equal to `element`. Inside the `while` loop, assign `current_node` to `previous_node`.

# --hints--

Your `while` loop should check for when `current_node` is not `None` and when `current_node.element` is not equal to `element`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

You should assign `current_node` to `previous_node` inside the `while` loop.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--
```
