---
id: 6852ef95ccc27660275692cd
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Your `Email` class needs to store information about each email when it's created.

Inside your `Email` class, remove the existing `pass` keyword and create the `__init__` method that takes `self`, `sender`, `receiver`, `subject`, and `body` as parameters.

# --hints--

You should define an `__init__` method inside the `Email` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Your `__init__` method should have parameters: `self`, `sender`, `receiver`, `subject`, and `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
