---
id: 6852ea0c3cadb54e897acdac
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you are going to build an *Email Simulator* that simulates sending, receiving, and managing emails between different users. You'll learn about classes, objects, and how to organize code in an object-oriented way.

Begin by creating a class named `Email` using the `class` keyword.

# --hints--

You should create a class named `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
