---
id: 6852f0c3fcedb964ab959921
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Now let's test that the `read` attribute was added correctly to your `Email` class. Since you already have an email object from the previous steps, print the `read` attribute to see that it's now `False` by default.

# --hints--

You should print `email_obj.read` to test the read attribute.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
