---
id: 6859400f8f7f872e50e1238c
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

The `User` class needs a method to send emails to other users. When sending an email, you'll create a new `Email` object and add it to the receiver's inbox.

Create a method called `send_email` in your `User` class. This method should take parameters for `receiver`, `subject`, and `body`. 

Inside the method, create a new `Email` object with the following parameter values and assign it to a variable named `email`:

- `sender`: `self` (the current user)
- `receiver`: `receiver` (the user receiving the email)
- `subject`: `subject` (the subject of the email)
- `body`: `body` (the body of the email)

# --hints--

You should define a `send_email` method in the `User` class.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

The `send_email` method should have parameters: `self`, `receiver`, `subject`, and `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

The method should create an `Email` object named `email` with the correct parameters.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
