---
id: 688754f3494b17e0b352d67d
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Now it's time to test our complete email system! Let's create some users and see the email functionality in action.

Create two User objects: `alice` with name `"Alice"` and `bob` with name `"Bob"`. This will demonstrate how users can be created and interact with each other.

# --hints--

You should create a User object named `alice` with name `"Alice"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

You should create a User object named `bob` with name `"Bob"`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
