---
id: 6887550fe36f64e27074475a
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Use f-strings to add the sender, reveiver and subject details like this:

```py
From: [sender name]
To: [receiver name]
Subject: [subject]
```

# --hints--

You should print the sender's name using an f-string with `self.sender`. Add `'From: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender}')")`))
    }

})
```

You should print the receiver's name using an f-string with `self.receiver`. Add `'To: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver}')")`))
    }

})
```

You should print the subject using an f-string with `self.subject`. Add `'Subject: '` before it.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Subject: {self.subject}')")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        
--fcc-editable-region--

email_obj = Email(
    sender='quinn@example.com',
    receiver='ash@example.com',
    subject='Meeting Reminder',
    body='Hi Ash, just reminding you about our meeting at 3 PM today.'
)

```
