---
id: 688754d6d59578df1313cb3c
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Now it's time to connect the User and Inbox classes so that emails can actually be delivered! 

Users need to have proper Inbox objects instead of simple lists, and the `send_email` method should deliver emails to the receiver's inbox.

Update the `User` class `__init__` in a way that each user gets an actual Inbox object instead of an empty list created earlier. 

# --hints--

You should update the User class to initialize `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
