---
id: 688754f3494b17e0b352d67d
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now you'll use the `print` function to display the email information. 

In your `display_full_email` method, after calling `mark_as_read()`, print `\n--- Email ---` to start the email display.

# --hints--

You should print the header `'--- Email ---'` at the beginning with a newline character at the start.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call('print("\\\\n--- Email ---")')`))
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        
--fcc-editable-region--

email_obj = Email(
    sender='quinn@example.com',
    receiver='ash@example.com',
    subject='Meeting Reminder',
    body='Hi Ash, just reminding you about our meeting at 3 PM today.'
)

```
