---
id: 6887553a1d2084e5efd42c72
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Print the body of the email with the label `Body:` followed by the actual body text:

```py
Body: [email body]
```

# --hints--

You should print `'Body:'` followed by the email body using `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Body: {self.body}')")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender}')
        print(f'To: {self.receiver}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
       
--fcc-editable-region--

email_obj = Email(
    sender='quinn@example.com',
    receiver='ash@example.com',
    subject='Meeting Reminder',
    body='Hi Ash, just reminding you about our meeting at 3 PM today.'
)

```
