---
id: 68875554c6c32de788afc64e
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Users need a way to read their emails properly.

Add a `display_full_email` method to the `Email` class that takes only `self` as a parameter. Inside this method, call `mark_as_read()` method to mark the email as read when someone views it.

This method would be used to display the email's full content in a nicely formatted way.

# --hints--

You should define a `display_full_email` method in the `Email` class.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

The `display_full_email` method should take only `self` as a parameter.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

The method should call `self.mark_as_read()` to mark the email as read.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
