---
id: 688798fccea91f522d1f245e
title: Step 33
challengeType: 20
dashedName: step-33
---

# --description--

When the inbox is not empty, you'll try to access the email at the given index and display it.

But first, you'll need to handle cases where the inbox is empty or the user provides an invalid email number. Python's `try` statement is perfect for this - it lets you try something that might fail and handle the error gracefully.

Complete the `read_email` method by adding a `try` block that displays the full email message at the given index. After that, add an `except` block that catches an `IndexError` and prints the message `'Invalid email number.\n'`.


# --hints--

You should call `display_full_email()` on the email within the `try` block.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_trys()[0].find_body().is_equivalent("self.emails[index].display_full_email()")`))
    }
})
```

You should add an `except IndexError` block to handle the case when the index is out of range and print the message `'Invalid email number.\n'`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_trys()[0].find_except("IndexError").find_body().is_equivalent("print('Invalid email number.\\\\n')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        
--fcc-editable-region--

```
