---
id: 68b96d6a290d52d5e57a7485
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Remove the `email_obj` and the following print statements.

# --hints--

You should not have the `email_obj` creation line in your code.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`));
    }
})
```

You should not have `print(email_obj.sender)` in your code.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.sender)")`));
    }
})
```

You should not have `print(email_obj.subject)` in your code.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.subject)")`));
    }
})
```

You should not have `print(email_obj.read)` in your code.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

You should not have `email_obj.mark_as_read()` in your code.

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
email_obj.mark_as_read()
print(email_obj.read)
--fcc-editable-region--
```
