---
id: 68b9366fae8b13c1278761b1
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Your inbox needs a way to receive new emails. When someone sends an email to a user, it should be added to their inbox.

Add a method called `receive_email` to your `Inbox` class that takes `self` and an email object `email` as parameters. Within the method body, add the `email` to the `emails` list using the `append` method.

# --hints--

You should define a method named `receive_email` in the `Inbox` class that takes `self` and `email` as parameters.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

You should append the `email` to `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
