---
id: 68be6af1f43cc44541debe41
title: Step 35
challengeType: 20
dashedName: step-35
---

# --description--

You should handle the case where the user tries to delete an email using an invalid index. Within the `delete_email` method:

- Use try-except to handle invalid indices. 
- In the try block, use the `del` keyword to delete the email at the mentioned index and print a confirmation message `Email deleted.\n`.
- In the except block, catch `IndexError` and print the message `Invalid email number.\n`.

# --hints--


Within the `try` block, you should use `del self.emails[index]` to delete the email. You should also print a confirmation message `Email deleted.\n` when the email is deleted.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_trys()[0].find_body().is_equivalent("del self.emails[index]\\nprint('Email deleted.\\\\n')")`))
    }
})
```

You should catch `IndexError` for invalid email numbers and print the message `Invalid email number.\n`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_trys()[0].find_except("IndexError").find_body().is_equivalent("print('Invalid email number.\\\\n')")`))
    }
})

```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        try:
            self.emails[index].display_full_email()
        except IndexError:
            print('Invalid email number.\n')
            
--fcc-editable-region--
    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return

--fcc-editable-region--

```
