---
id: 65f056a405239e1dc4cc2854
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер призначте `y` до атрибута `y` об’єкта `Vector`.

# --hints--

Призначте `y` до `self.y`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").find_variable("self.y").is_equivalent("self.y = y")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
--fcc-editable-region--
```
