---
id: 65f03d9f92eac9183a4d3281
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

На евклідовій площині вектор можна визначити двома координатами (`x` та `y`). Відстань між початком осей та точкою `(x, y)` буде його довжиною або нормою. А напрямок вектора вказуватиме на `(x, y)`.

<img alt="2-вимірний вектор з координатами (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;" />

У межах класу `Vector` створіть метод `__init__` і надайте йому три параметри: `self`, `x` та `y`.

# --hints--

Визначте метод `__init__` в межах класу `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Метод `__init__` має приймати три параметри: `self`, `x` та `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
