---
id: 65f06a8e5a57673d700c79c3
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Поза класом `Vector` створіть екземпляр класу `Vector`, передаючи цілі числа `2` та `3` як значення `x` та `y`, і призначте його до змінної під назвою `v1`.

# --hints--

Оголосіть змінну `v1` та призначте до неї `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
