---
id: 65f0694fb2296f3caadf8347
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Довжина вектора $\mathbf{a}$ (або норма) зазвичай позначається як $\\| \mathbf{a} \\|$. Її можна обчислити як квадратний корінь із суми квадратів його компонентів: \\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2} \\]

Обчисліть векторну норму та поверніть результат із методу `norm`.

# --hints--

Метод `norm` має використовувати надану формулу, щоб обчислити та повернути число, яке представляє норму поточного екземпляра вектора. Не наближайте значення.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        pass
--fcc-editable-region--
```
