---
id: 65f01969115f933073b6be03
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Вектор — це об’єкт, який має довжину (або величину) і напрямок та не може бути виражений одним числом. Вектори зазвичай використовують у фізиці для представлення сили, швидкості, прискорення та багато іншого.

У цьому проєкті ви побудуєте векторний простір — множину, між елементами якої визначено ряд операцій. Незабаром ви дізнаєтесь всі деталі.

А поки розпочніть проєкт з оголошення порожнього класу під назвою `Vector`.

# --hints--

Оголосіть порожній клас під назвою `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
