---
id: 65f07c9b1ffb814d856dcffc
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Перевірте, що `norm` працює належним чином, надрукувавши значення, яке повертає метод.

# --hints--

Викличте `norm()` на `v1` та надрукуйте результат.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
