---
id: 65f3fe07cc763212efe91285
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Вектор може мати `n` вимірів (компонентів). Ось представлення 3-вимірного вектора:

<img alt="3-вимірний вектор з координатами (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;" />

Наразі ви створили 2-вимірний вектор. Вам потрібно представляти вектори з різною кількістю вимірів, не переписуючи необхідний код для кожного конкретного випадку. Для цього ви будете використовувати успадкування.

Для початку перейменуйте клас `Vector` на `R2Vector`, щоб вказати, що цей клас матиме справу з 2-вимірними векторами. Не забудьте також змінити створення екземпляру `v1`.

# --hints--

Перейменуйте клас `Vector` на `R2Vector`. Не забудьте також змінити створення екземпляру `v1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
