---
id: 65f3fa097f26b510db6c710b
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Чудово, `norm` працює як очікувалося. Тепер, якщо ви спробуєте надрукувати `v1`, то отримаєте рядкове представлення об’єкта за замовчуванням (щось на зразок `<__main__.Vector object at 0x11eb778>`).

В межах класу `Vector` оголосіть порожній метод `__str__`, щоб реалізувати читабельне представлення рядка. Не забудьте надати йому параметр `self`.

Зверніть увагу на те, щоб не друкувати `v1`, доки `__str__` не поверне рядок, інакше ви отримаєте `TypeError`, оскільки `__str__` має завжди повертати рядок.

# --hints--

Визначте метод `__str__` в межах класу `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Метод `__str__` повинен мати параметр `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
