---
id: 65f3fa60a93b84110b3f1708
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Вам потрібно повернути значуще представлення рядка. Наприклад, у випадку `v1` потрібно повернути рядок, що містить кортеж зі значеннями компонентів вектора: `(2, 3)`.

Поверніть рядок з методу `__str__`, що представляє вектор у вигляді кортежу, який містить компоненти вектора за порядком. Потім вийдіть за межі класу `Vector` та надрукуйте `v1`, щоб перевірити результат.

# --hints--

Метод `__str__` має повернути рядок, що представляє вектор у вигляді кортежу, який містить компоненти вектора за порядком.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Надрукуйте `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
