---
id: 65f40051d6b09a139f253e8e
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Успадкування дозволяє визначити клас з наявного. Новий клас (який називається дочірнім) успадковує всі методи та властивості наявного класу (який називається батьківським).

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass

Oak().sprout() # Output: Making new leaves!
```

У наведеному вище прикладі дочірній клас `Oak` успадковує від `Tree` та успадковує метод `sprout` від батьківського класу `Tree`.

Створіть новий клас під назвою `R3Vector` та дотримуйтеся прикладу вище, щоб він успадковував від класу `R2Vector`.

# --hints--

Визначте новий клас `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

Новий клас `R3Vector` має успадковувати від `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
