---
id: 65f4520e363e2642f8112e33
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Як бачите, щось не так. Методи `norm` та `__str__`, успадковані від `R2Vector`, не можуть адаптуватися до 3-вимірного вектора. Їхня реалізація має бути більш гнучкою.

Кожен об’єкт у Python має спеціальний атрибут під назвою `__dict__`, який є словником, у якому зберігаються атрибути об’єкта.

Видаліть наявні виклики `print`. Потім надрукуйте атрибут `__dict__` векторів `v1` та `v2`, щоб побачити, як вони виглядають.

# --hints--

Код не повинен містити `print(v2)` та `print(v2.norm())`.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

Надрукуйте атрибут `__dict__` векторів `v1` та `v2`.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
