---
id: 65f407ea37ad6e181b90462e
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Призначте кожен параметр `i` до `self.i`, як ви робили раніше. Хоча в цьому випадку повторяться декілька рядків, цього можна уникнути за допомогою функції `super()`. `super()` дозволяє неявно посилатися на батьківський клас: `super().__init__(x, y)` викликає метод `__init__` батьківського класу.

Додайте виклик `super().__init__(x, y)` до методу `__init__`.

# --hints--

Метод `__init__` повинен мати виклик `super().__init__(x, y)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return (self.x**2 + self.y**2)**0.5

    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
