---
id: 5dc1798ff86c76b9248c6eb3
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Die `h1`- bis `h6`-Überschriftenelemente werden verwendet, um die Wichtigkeit der Inhalte unter ihnen zu signalisieren. Je kleiner die Zahl, desto höher die Wichtigkeit, so dass `h2`-Elemente weniger von Bedeutung sind als `h1`-Elemente.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Benutze nur ein `h1`-Element pro Seite und platziere die Überschriften mit geringer Bedeutung unterhalb der Überschriften mit hoher Bedeutung.

Füge unterhalb des `h1`-Elements ein `h2`-Element mit dem folgenden Text hinzu:

`Cat Photos`

# --hints--

Dein `h1`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Dein `h1`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<` Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

Du solltest nur ein `h1`-Element haben. Entferne die zusätzlichen.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Der Text deines `h1`-Elements sollte 'CatPhotoApp' lauten. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Dein `h2`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Dein `h2`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<` Zeichen.

```js
assert(code.match(/<\/h2\>/));
```

Der Text deines `h2`-Elements sollte 'Cat Photos' sein. Platziere nur den Text `Cat Photos` zwischen das öffnende und schließende `h2`-Tag.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Dein `h2`-Element sollte unter dem `h1`-Element stehen. Das `h1`-Element ist von größerer Bedeutung und muss über dem `h2`-Element liegen.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>

--fcc-editable-region--
  </body>
</html>
```

