---
id: 5dc17dc8f86c76b9248c6eb5
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Commenting allows you to leave messages without affecting the browser display. It also allows you to make code inactive. Ein Kommentar in HTML beginnt mit `<!--`, enthält eine beliebige Anzahl an Textzeilen und endet mit `-->`. Zum Beispiel enthält der Kommentar `<!-- TODO: Remove h1 -->` den Text `TODO: Remove h1`.

Add a comment above the `p` element with this text:

`TODO: Add link to cat photos`

# --hints--

Dein Kommentar sollte mit `<!--` beginnen. You are missing one or more of the characters that define the start of a comment.

```js
assert(code.match(/<!--/));
```

Dein Kommentar sollte mit `-->` enden.  You are missing one or more of the characters that define the end of a comment.

```js
assert(code.match(/-->/));
```

Your code should not have extra opening/closing comment characters. Dir wird ein zusätziches `<!--` oder `-->` im Browser angezeigt.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Dein Kommentar sollte den Text `TODO: Add link to cat photos` enthalten.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

Your comment should be above the `p` element. Sie sind in der falschen Reihenfolge.

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

