---
id: 5dc174fcf86c76b9248c6eb2
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML Elemente haben öffnende Tags wie `<h1>` und schließende Tags wie `</h1>`.

Der Text für ein Element gehört zwischen seinen öffnenden und schließenden Tags.

Finde das `h1`-Element und ändere seinen Text zu:

`CatPhotoApp`

# --hints--

Der Text `CatPhotoApp` sollte im Code vorhanden sein. Möglicherweise möchtest du deine Rechtschreibung überprüfen.

```js
assert(code.match(/catphotoapp/i));
```

Dein `h1`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Dein `h1`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

Du hast mehrere `h1`-Elemente. Entferne das zusätzliche `h1`-Element.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Der Text deines `h1`-Elements sollte `CatPhotoApp` sein. Du hast entweder den Text weggelassen, einen Tippfehler gemacht oder es befindet sich nicht zwischen den öffnenden und schließenden Tags des `h1`-Elements.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Du scheinst eine Browsererweiterung zu verwenden, die die Seite verändert. Stelle sicher, dass du alle Browsererweiterungen deaktivierst.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

