---
id: 5dc174fcf86c76b9248c6eb2
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

HTML-Elemente haben einen offenen und schließenden Tag mit Inhalt dazwischen.

Hier ist die grundlegende Syntax:

```html
<openingTag>content</closingTag>
```

Das erste Element, über das du etwas lernen wirst, ist das Element `h1`. Das Element `h1` ist eine Überschrift und wird für die Hauptüberschrift einer Webseite verwendet.

```html
<h1>This is a main heading</h1>
```

Ändere den Text des unten angegebenen `h1` Elements von `Hello World` zu `CatPhotoApp` und sieh dir die Veränderung im Browser-Fenster an.

Wenn du fertig bist, klicke auf "Check Your Code", um zu sehen, ob es korrekt ist.

# --hints--

Der Text `CatPhotoApp` sollte im Code vorhanden sein. Möglicherweise möchtest du deine Rechtschreibung überprüfen.

```js
assert(code.match(/catphotoapp/i));
```

Dein `h1`-Element sollte ein öffnendes Tag haben. Öffnende Tags besitzen die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Dein `h1`-Element sollte ein schließendes Tag haben. Abschließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/h1\>/));
```

Du hast mehr als ein `h1`-Element. Entferne das zusätzliche `h1`-Element.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Der Text deines `h1`-Elements sollte `CatPhotoApp` lauten. Du hast entweder den Text weggelassen, einen Tippfehler gemacht oder es befindet sich nicht zwischen den öffnenden und schließenden Tags des `h1`-Elements.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Du scheinst eine Browsererweiterung zu verwenden, die die Seite verändert. Stelle sicher, dass du alle Browsererweiterungen deaktivierst.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

