---
id: 5dc17dc8f86c76b9248c6eb5
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Kommentare erlauben dir, Nachrichten zu hinterlassen, ohne die Anzeige des Browsers zu beeinflussen. Sie erlauben dir auch, Code inaktiv zu machen. Ein Kommentar in HTML beginnt mit `<!--`, enthält eine beliebige Anzahl an Textzeilen und endet mit `-->`.

Here is an example of a comment with the `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Füge einen Kommentar über das `p`-Element mit diesem Text hinzu:

`TODO: Add link to cat photos`

# --hints--

Dein Kommentar sollte mit `<!--` beginnen. Es fehlen ein oder mehrere Zeichen, die den Beginn eines Kommentars definieren.

```js
assert(code.match(/<!--/));
```

Dein Kommentar sollte mit `-->` enden.  Es fehlen ein oder mehrere Zeichen, die das Ende eines Kommentars definieren.

```js
assert(code.match(/-->/));
```

Dein Code sollte keine zusätzlichen öffnenden/abschließenden Kommentarzeichen enthalten. Dir wird ein zusätziches `<!--` oder `-->` im Browser angezeigt.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.match(/<!--/g).length < 2 && noSpaces.match(/-->/g).length < 2);
```

Dein Kommentar sollte den Text `TODO: Add link to cat photos` enthalten.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i));
```

Dein Kommentar sollte über dem `p`-Element stehen. Sie sind in der falschen Reihenfolge.

```js
assert(
  code
    .replace(/\s/g, '')
    .match(
      /<!--todo:addlinktocatphotos--><p>seemorecatphotosinourgallery\.?<\/p>/i
    )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--

    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

