---
id: 5dc24165f86c76b9248c6ebc
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Alle `img`-Elemente sollten ein `alt`-Attribut haben. Der Text des `alt`-Attributs wird für Screenreader verwendet, um die Zugänglichkeit zu verbessern und wird angezeigt, wenn das Bild nicht geladen wird.

Here is an example of an `img` element with an `alt` attribute:

```html
<img src="cat.jpg" alt="A cat">
```

Füge in dem `img`-Element ein `alt`-Attribut mit diesem Text hinzu:

`A cute orange cat lying on its back`

# --hints--

Dein Code sollte ein `img`-Element haben. Du hast das `img`-Element aus einem vorherigen Schritt entfernt.

```js
assert(document.querySelector('img'));
```

Dein `img`-Element hat kein `alt`-Attribut. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

Der `alt`-Attributwert deines `img`-Elements ist auf etwas anderes als 'Eine süße orangefarbene Katze liegt auf ihrem Rücken' gesetzt. Stelle sicher, dass der Wert des `alt`-Attributs zwischen Anführungszeichen steht.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

