---
id: 5dc24073f86c76b9248c6ebb
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>attributes</dfn> sind spezielle Wörter, die innerhalb des öffnenden Tags eines Elements verwendet werden, um das Verhalten des Elements zu steuern. Das `src`-Attribut in einem `img`-Element gibt die URL des Bildes an (wo sich das Bild befindet).

Hier ist ein Beispiel für ein `img`-Element mit einem `src`-Attribut, das auf das freeCodeCamp-Logo verweist:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Inside the existing `img` element, add a `src` attribute with this URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Dein Code sollte ein `img`-Element enthalten. Möglicherweise hast du das `img`-Element entfernt oder die `src`-Attributwerte nicht mit Anführungszeichen umgeben.

```js
assert(document.querySelector('img'));
```

Your `img` element should have a `src` attribute. Du hast entweder das Attribut weggelassen oder einen Tippfehler gemacht. Stelle sicher, dass zwischen dem Namen des Elements und dem Namen des Attributs ein Leerzeichen vorhanden ist.

```js
assert(document.querySelector('img').src);
```

Das `src`-Attribut deines `img`-Elements sollte auf '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`' gesetzt sein. Du hast entweder die URL weggelassen oder einen Tippfehler gemacht. Die Groß- und Kleinschreibung der URL ist wichtig.

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Obwohl du die `src` des `img`-Elements auf die richtige URL gesetzt hast, wird empfohlen, die Werte eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

Your `img` element should be closed either with `>` or `/>`.

```js
assert(/<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

