---
id: 5dc2385ff86c76b9248c6eb7
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 hat einige Elemente, die unterschiedliche Inhaltsbereiche identifizieren. Diese Elemente machen dein HTML leichter lesbar und helfen bei der Suchmaschinenoptimierung (SEO) und der Barrierefreiheit.

The `main` element is used to represent the main content of the body of an HTML document. Content inside the `main` element should be unique to the document and should not be repeated in other parts of the document.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifiziere den Hauptteil dieser Seite, indem du ein `<main>` öffnendes Tag vor dem `h1`-Element und ein `</main>` schließendes Tag nach dem `p`-Element hinzufügst.

# --hints--

Dein `main`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('main'));
```

Dein `main`-Element sollte ein schließendes Tag haben. Abschließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/main\>/));
```

Das öffnende Tag deines `main`-Elements sollte unterhalb des öffnenden Tags des `body`-Elements stehen. Sie sind in falscher Reihenfolge.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Das öffnende Tag deines `main`-Elements sollte über dem `h1`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

Das schließende Tag deines `main`-Elements sollte unter dem `p`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/see more cat photos in our gallery/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

