---
id: 5dc24614f86c76b9248c6ebd
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

You can link to another page with the anchor (`a`) element.

Here is an example linking to `https://www.freecodecamp.org`:

```html
<a href="https://www.freecodecamp.org"></a>
```

Füge ein Ankerelement nach dem Absatz, der auf `https://freecatphotoapp.com` verlinkt, hinzu. An dieser Stelle wird der Link nicht in der Vorschau erscheinen.

# --hints--

Dein (`a`)-Ankerelement sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Dein (`a`)-Ankerelement sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/a\>/));
```

Dein (`a`)-Ankerelement sollte unter dem `p`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('a, p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('A'));
```

Dein (`a`)-Element hat kein `href`-Attribut. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert(document.querySelector('a').hasAttribute('href'));
```

Dein (`a`)-Ankerelement sollte auf `https://freecatphotoapp.com` verlinken. Du hast entweder die URL weggelassen oder einen Tippfehler gemacht.

```js
assert(
  document.querySelector('a').getAttribute('href') ===
    'https://freecatphotoapp.com'
);
```

Obwohl du das `href`-Attribut des ('a') Ankerelements auf den richtigen Link gesetzt hast, wird empfohlen, die Werte eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(
  !/\<a\s+href\s*=\s*https:\/\/www.freecodecamp.org\/cat-photos/.test(code)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

