---
id: 5dc23f9bf86c76b9248c6eba
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Du kannst Bilder zu deiner Webseite hinzufügen, indem du das `img`-Element verwendest. `img`-Elemente haben ein öffnendes Tag, aber kein schließendes Tag. Ein Tag eines Elements ohne ein schließendes Tag wird als <dfn>self-closing tag</dfn> bezeichnet.

Füge ein `img`-Element unter dem `p`-Element hinzu. An diesem Punkt wird kein Bild im Browser angezeigt.

# --hints--

Dein `img`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Dein `img`-Element sollte kein schließendes Tag haben. Abschließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(!code.match(/<\/img\>/));
```

Du solltest nur ein `img`-Element haben. Entferne alles Zusätzliche.

```js
assert(document.querySelectorAll('img').length === 1);
```

Dein `img`-Element sollte unter dem `p`-Element stehen. Sie sind in falscher Reihenfolge.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

