---
id: 5dfa22d1b521be39a3de7be0
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Im vorherigen Schritt hast du die Wörter `link to cat pictures` in einen Link verwandelt, indem du die Wörter zwischen einem öffnenden und schließenden (`a`)-Ankertag platziert hast. Dasselbe gilt für Wörter innerhalb eines Elements, z. B. eines `p`-Elements.

Verwandle im Text deines `p`-Elements die Wörter `cat photos` in einen Link zu `https://freecatphotoapp.com`, indem du öffnende und schließende (`a`)-Ankertags um diese Wörter hinzufügst.


# --hints--

Du solltest ein neues Anker-Element (`a`) innerhalb des `p`-Elements einbetten.

```js
assert($('p > a').length);
```

Der `href`-Wert des Links sollte `https://freecatphotoapp.com` sein. Du hast entweder den `href`-Wert weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

Der Linktext sollte `cat photos` sein. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Nachdem das (`a`)-Ankerelement eingebettet wurde, sollte der einzige sichtbare `p`-Elementinhalt im Browser `See more cat photos in our gallery.` sein. Überprüfe den Text, die Abstände oder die Satzzeichen des `p`s und des eingebetteten Ankerelements.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

