---
id: 5dfa22d1b521be39a3de7be0
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Du kannst jeden beliebigen Text in einen Link verwandeln, zum Beispiel den Text innerhalb eines `p`-Elements.

``` html
<p>I think <a href="https://www.freecodecamp.org">freeCodeCamp</a> is great.</p>
```

Verwandle im Text deines `p`-Elements die Wörter `cat photos` in einen Link, indem du öffnende und abschließende (`a`)-Ankertags um diese Wörter herum hinzufügst. Setze dann das `href`-Attribut auf `https://freecatphotoapp.com`


# --hints--

Du solltest ein neues Anker-Element (`a`) innerhalb des `p`-Elements einbetten.

```js
const nestedAnchor = document.querySelector(`p > a`);
assert.isNotNull(nestedAnchor)
```

Der `href`-Wert des Links sollte `https://freecatphotoapp.com` sein. Du hast entweder den `href`-Wert weggelassen oder einen Tippfehler gemacht.

```js
assert(
  document.querySelector('p > a').href === 'https://freecatphotoapp.com/'
);
```

Der Linktext sollte `cat photos` sein. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = document.querySelector(`p > a`);
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Nachdem das (`a`)-Ankerelement eingebettet wurde, sollte der einzige sichtbare `p`-Elementinhalt im Browser `See more cat photos in our gallery.` sein. Überprüfe den Text, die Abstände oder die Satzzeichen des `p`s und des eingebetteten Ankerelements.

```js
assert.match(code, /<p>see more <a[^>]*>cat photos<\/a> in our gallery\.?<\/p>/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

