---
id: 5dfb6250eacea3f48c6300b2
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

After the unordered list, add a new image with a `src` attribute value set to:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`

Und setze dessen Wert für das `alt`-Attribut auf:

`A slice of lasagna on a plate.`

# --hints--

Es sollte ein `img`-Element direkt nach dem schließenden `</ul>`-Tag geben.

```js
assert($('section')[1].lastElementChild.nodeName === 'IMG');
```

Das neue Bild besitzt kein `alt`-Attribut. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert($('section')[1].lastElementChild.hasAttribute('alt'));
```

Das neue Bild sollte `A slice of lasagna on a plate.` als `alt`-Wert enthalten. Stelle sicher, dass der Wert des `alt`-Attributs in Anführungszeichen gesetzt wird.

```js
assert(
  $('section')[1]
    .lastElementChild.getAttribute('alt')
    .replace(/\s+/g, ' ')
    .match(/^A slice of lasagna on a plate\.?$/i)
);
```

The new image does not have a `src` attribute. Überprüfe, ob nach dem Namen des öffnenden Tags ein Leerzeichen steht und/oder ob vor allen Attributnamen Leerzeichen stehen.

```js
assert($('section')[1].lastElementChild.hasAttribute('src'));
```

The new image should have a `src` value of `https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg`. Stelle sicher, dass der Wert des `src`-Attributs in Anführungzeichen gesetzt wird.

```js
assert(
  $('section')[1].lastElementChild.getAttribute('src') ===
    'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

Obwohl du den `src`-Wert des neuen Bilds auf die richtige URL gesetzt hast, wird es empfohlen, den Wert immer in Anführungszeichen zu setzen.

```js
assert(!/\<img\s+.+\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/lasagna\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
--fcc-editable-region--
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

