---
id: 5dfb655eeacea3f48c6300b3
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Das `figure`-Element repräsentiert eigenständige Inhalte und ermöglicht dir, ein Bild mit einer Beschriftung zu verbinden.

Bette das Bild, das du gerade hinzugefügt hast, in ein `figure`-Element ein.

# --hints--

Dein `figure`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

Dein `figure`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen.

```js
assert(code.match(/<\/figure\>/));
```

Ein `figure`-Element sollte genau über dem schließenden Tag des zweiten `section`-Elements stehen.

```js
assert($('section')[1].lastElementChild.nodeName === 'FIGURE');
```

Das Lasagne-`img`-Element sollte innerhalb des `figure`-Elements eingebettet sein.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

