---
id: 5ef9b03c81a63668521804db
title: Schritt 41
challengeType: 0
dashedName: step-41
---

# --description--

Um zu verhindern, dass ein Benutzer dein Formular absendet, wenn erforderliche Informationen fehlen, musst du das `required`-Attribut zu einem `input`-Element hinzufügen. Es ist nicht nötig, dem `required`-Attribut einen Wert zu geben. Füge stattdessen einfach das Wort `required` zum `input`-Element hinzu und stelle sicher, dass zwischen ihm und anderen Attributen Leerzeichen vorhanden sind.

# --hints--

Du hast entweder dein `input`-Element gelöscht oder es hat eine ungültige Syntax. Alle Attributwerte sollten zwischen Anführungszeichen stehen.

```js
assert($('input').length);
```

Deine `form` sollte nur das `input`-Element enthalten. Remove any additional HTML elements or text within the `form` element.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Dein `input`-Element sollte das `required`-Attribut besitzen. Denke daran, dass du nur das Wort `required` in das Tag des `input`-Elements hinzufügst.

```js
assert($('input')[0].hasAttribute('required'));
```

Dem `required`-Attribut sollte kein Wert gegeben werden.

```js
assert($('input')[0].getAttribute('required') === '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

