---
id: 5ef9b03c81a63668521804d4
title: Schritt 32
challengeType: 0
dashedName: step-32
---

# --description--

Das `strong`-Element wird verwendet, um anzuzeigen, dass mancher Text von großer Bedeutung oder dringend ist.

In the `figcaption` you just added, indicate that `hate` is of strong importance by wrapping it in a `strong` element.

# --hints--

Dein `strong`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('strong'));
```

Dein `strong`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<` Zeichen.

```js
assert(code.match(/<\/strong\>/));
```

Dein `strong`-Element sollte das Wort `hate` im Text `Cats hate other cats.` umschließen Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.toLowerCase() === 'hate'
);
```

Der Text von `figcaption` sollte `Cats hate other cats.` sein Prüfe auf Tippfehler und darauf, ob die erforderlichen Leerzeichen um das öffnende und schließende Tag des `strong`-Elements vorhanden sind.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
--fcc-editable-region--

          <figcaption>Cats hate other cats.</figcaption>  

--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

