---
id: 5ef9b03c81a63668521804d7
title: Schritt 36
challengeType: 0
dashedName: step-36
---

# --description--

The `action` attribute indicates where form data should be sent.

Here is an example of a `form` element with an `action` attribute:

```html
<form action="/submit-url"></form>
```

In the example, `action="/submit-url"` tells the browser that the form data should be sent to the path `/submit-url`.

Füge ein `action`-Attribut mit dem Wert `https://freecatphotoapp.com/submit-cat-photo` zum `form`-Element hinzu.

# --hints--

Dein `form`-Element sollte ein öffnendes und abschließendes Tag in der richtigen Reihenfolge haben. Es kann sein, dass eine oder beide der erforderlichen Tags fehlen oder in der falschen Reihenfolge sind.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Dein `form`-Element, das im letzten `section`-Element eingebettet wurde, sollte unter deinem `h2`-Element liegen. Du hast das `h2`-Element und das `form`-Element in der falschen Reihenfolge.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Dein `form`-Element sollte keinen Inhalt haben. Entferne jedes HTML-Element oder jeglichen Text zwischen den Tags des `form`-Elements.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

Dein `form`-Element hat kein `action`-Attribut. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags und/oder Leerzeichen vor allen Attributnamen vorhanden sind.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Dein `form`-Element sollte ein `action`-Attribut mit dem Wert `https://freecatphotoapp.com/submit-cat-photo` haben.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Obwohl du das `action`-Attribut auf die richtige URL gesetzt hast, wird empfohlen, den Wert eines Attributs immer zwischen Anführungszeichen anzugeben.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

