---
id: 5ef9b03c81a63668521804d2
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Der Code für eine geordnete Liste (`ol`) ist einer ungeordneten Liste ähnlich, nur dass die Listenelemente in einer geordneten Liste nummeriert werden, wenn sie angezeigt werden.

Füge nach dem letzten `h3`-Element des zweiten `section`-Elements eine geordnete Liste mit diesen drei Elementen hinzu:

`flea treatment` `thunder` `other cats`

# --hints--

Dein `ol`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben diese Syntax: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

Dein `ol`-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/ol>/));
```

The `ol` element should be above the second `section` element's closing tag. Sie sind in falscher Reihenfolge.

```js
assert($('main > section')[1].lastElementChild.nodeName === 'OL');
```

Die drei `li`-Elemente sollten in dem `ol`-Element eingebettet sein.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Du solltest drei `li`-Elemente mit dem Text `flea treatment`, `thunder` und `other cats` in beliebiger Reihenfolge haben.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Du solltest nur ein `ol`-Element haben.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

