---
id: 5ef9b03c81a63668521804e9
title: Schritt 64
challengeType: 0
dashedName: step-64
---

# --description--

Verwandle den `freeCodeCamp.org`-Text in einen Link, indem du ihn in ein (`a`)-Ankerelement einschließt. Das `href`-Attribut sollte auf `https://www.freecodecamp.org` gesetzt sein.

# --hints--

Dein (`a`)-Ankerelement sollte innerhalb des `p`-Elements im `footer` eingebettet sein. Stelle sicher, dass du ein öffnendes Tag und ein schließendes Tag für die (`a`)-Ankerelemente hinzufügst.

```js
assert($('footer > p > a').length);
```

Dein (`a`)-Element sollte ein schließendes Tag haben. Schließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
const aElemClosingTags = code.match(/<\/a\>/g);
assert(aElemClosingTags && aElemClosingTags.length === 3);
```

Dein (`a`)-Ankerelement sollte ein `href`-Attribut mit dem Wert `https://www.freecodecamp.org` haben. Möglicherweise hast du das Attribut/den Wert weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = $('footer > p > a')[0];
assert(nestedAnchor.getAttribute('href').toLowerCase() === 'https://www.freecodecamp.org');
```

Der Linktext sollte `freeCodeCamp.org` sein. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const nestedAnchor = $('footer > p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'freecodecamp.org'
);
```

Nach der Einbettung des (`a`)-Ankerelements sollte der einzige Inhalt des `p`-Elements, der im Browser sichtbar ist, `No Copyright - freeCodeCamp.org` sein. Überprüfe den Text, die Abstände und die Zeichensetzung des `p` und des eingebetteten Ankerelements.

```js
const pText = $('footer > p')[0].innerText.toLowerCase().replace(/\s+/g, ' ');
assert(pText.match(/^no copyright - freecodecamp.org$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor" checked> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving" checked> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic">Energetic</label>
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
    <footer>
--fcc-editable-region--
      <p>No Copyright - freeCodeCamp.org</p>
--fcc-editable-region--
    </footer>
  </body>
</html>
```

