---
id: 5ef9b03c81a63668521804dc
title: Schritt 44
challengeType: 0
dashedName: step-44
---

# --description--

Du kannst Optionsschaltflächen für Fragen verwenden, bei denen du nur eine Antwort aus mehreren Optionen möchtest.

Here is an example of a radio button with the option of `cat`:

```html
<input type="radio"> cat
```

Beachte, dass `input`-Elemente <dfn>self-closing</dfn> sind.

Füge vor der Texteingabe einen Radio-Button hinzu, der diese Option enthält:

`Indoor`

# --hints--

Du solltest für deine Optionsschaltfläche ein `input`-Element erstellen. Überprüfe die Syntax.

```js
assert($('form > input').length >= 2);
```

Dein `input`-Element sollte ein öffnendes Tag haben, aber kein schließendes Tag.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

Du solltest nur ein `input`-Element für deinen Radio-Button hinzugefügt haben. Entferne alles Zusätzliche.

```js
assert($('form > input').length === 2);
```

Dein neues `input`-Element sollte über dem vorhandenen `input` liegen, mit dem `type`-Attribut gesetzt auf `text`. Sie sind in der falschen Reihenfolge.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Dein neues `input`-Element hat kein `type`-Attribut. Überprüfe, ob es ein Leerzeichen nach dem Namen des öffnenden Tags gibt.

```js
assert($('input')[0].hasAttribute('type'));
```

Dein neues `input`-Element sollte nur ein `type`-Attribut haben. Entferne alles Zusätzliche.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Dein neues `input`-Element sollte ein `type`-Attribut mit dem Wert `radio` haben. Du hast entweder den Wert weggelassen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen stehen sollten.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Obwohl du das `type`-Attribut des neuen `input`-Elements auf `radio` gesetzt hast, wird empfohlen, den Wert eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Der `Indoor`-Text des `radio`-Buttons sollte danach und nicht davor stehen.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

Der Text `Indoor` sollte sich direkt rechts neben deinem `radio`-Button befinden. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

