---
id: 5ef9b03c81a63668521804d1
title: Schritt 26
challengeType: 0
dashedName: step-26
---

# --description--

Füge nach dem `figure`-Element ein weiteres `h3`-Element mit dem folgenden Text hinzu:

`Top 3 things cats hate:`

# --hints--

Es sollte ein `h3`-Element direkt über dem schließenden Tag des zweiten `section`-Elements geben. Stelle sicher, dass es ein öffnendes und abschließendes Tag enthält.

```js
assert(
  document.querySelectorAll('main > section')[1].lastElementChild.nodeName ===
    'H3' && code.match(/<\/h3\>/g).length === 2
);
```

Das neue `h3`-Element sollte den Text `Top 3 things cats hate:` enthalten. Stelle sicher, dass der Doppelpunkt am Ende des Textes eingefügt wurde.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'top 3 things cats hate:'
);
```

There should be a `figure` above the new `h3` element. Möglicherweise hast du versehentlich das `figure`-Element gelöscht.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  .lastElementChild;
assert(
  secondSectionLastElemNode.nodeName === 'H3' &&
    secondSectionLastElemNode.previousElementSibling.nodeName === 'FIGURE'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

