---
id: 5ef9b03c81a63668521804d0
title: Schritt 25
challengeType: 0
dashedName: step-25
---

# --description--

To place emphasis on a specific word or phrase, you can use the `em` element.

Hebe das Wort `love` im `figcaption`-Element hervor, indem du es in ein `em`-Element, das betont wird, umbrichst.

# --hints--

Dein `em`-Betonungselement sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Dein `em`-Betonungselement sollte ein schließendes Tag haben. Abschließende Tags haben einen `/` direkt nach dem `<`-Zeichen haben.

```js
assert(code.match(/<\/em\>/));
```

Du hast entweder das `figcaption` Element gelöscht oder es fehlt ein öffnendes oder schließendes Tag.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Dein `em`-Betonungselement sollte den Text `love` umgeben. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Der `figcaption`-Text sollte `Cats love lasagna` lauten. Prüfe, ob Tippfehler und die notwendigen Leerzeichen um das öffnende und abschließende Tag des `em`-Elements vorhanden sind.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```

