---
id: 5ef9b03c81a63668521804e2
title: Schritt 54
challengeType: 0
dashedName: step-54
---

# --description--

In Formularen werden häufig Kontrollkästchen für Fragen verwendet, die mehr als eine Antwort haben können. The `input` element with a `type` attribute set to `checkbox` creates a checkbox.

Füge unter das `legend`-Element, dass du eben hinzugefügt hast, einen `input` mit seinem `type`-Attribut gesetzt auf `checkbox` und gib ihm eine Option von:

`Loving`

# --hints--

Das `input`-Element für dein Kontrollkästchen sollte ein öffnendes Tag, aber kein schließendes Tag haben.

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

Du solltest nur ein Eingabe-Element für dein Kontrollkästchen hinzugefügt haben. Entferne alle zusätzlichen.

```js
assert($('fieldset > input').length === 1);
```

Dein neues `input`-Element sollte sich unter dem `legend`-Element mit dem Text `What's your cat's personality?` befinden. Sie sind in falscher Reihenfolge.

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

Dein neues `input`-Element hat kein `type`-Attribut. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags vorhanden ist.

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

Dein neues `input`-Element sollte ein `type`-Attribut mit dem Wert `checkbox` haben. Du hast entweder den Wert weggelassen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen stehen sollten.

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

Obwohl du das `type`-Attribut des neuen `input`-Elements auf `checkbox` gesetzt hast, wird empfohlen, den Wert eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Der Text `Loving` sollte sich direkt rechts neben deinem Kontrollkästchen befinden. Stelle sicher, dass zwischen dem Element und dem Text ein Leerzeichen ist. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

