---
id: 5ef9b03c81a63668521804d6
title: Schritt 35
challengeType: 0
dashedName: step-35
---

# --description--

Nun wirst du ein Webformular hinzufügen, um Informationen von Benutzern zu sammeln.

The `form` element is used to get information from a user like their name, email, and other details.

Füge nach der `Cat Form`-Überschrift ein `form`-Element hinzu.

# --hints--

Dein `form`-Element sollte ein öffnendes und ein abschließendes Tag enthalten. Es kann sein, dass eine oder beide der erforderlichen Tags fehlen oder in der falschen Reihenfolge liegen.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Deine `form`-Element-Tags sind in der falschen Reihenfolge.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Das `form`-Element, das im letzten `section`-Element eingebettet ist, sollte unter dem `h2`-Element stehen. Das `h2`-Element und das `form`-Element stehen in der falschen Reihenfolge.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Das `form`-Element sollte keinen Inhalt haben. Entferne jedes HTML-Element oder jeglichen Text zwischen den Tags des `form`-Elements.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

