---
id: 5efb23e70dc218d6c85f89b1
title: Schritt 38
challengeType: 0
dashedName: step-38
---

# --description--

Es gibt viele Arten von Eingängen, die du erstellen kannst, indem du das `type`-Attribut verwendest. Du kannst ganz einfach ein Passwortfeld, einen Reset-Knopf, oder eine Steuerung erstellen, damit Benutzer eine Datei von ihrem Computer auswählen können.

Erstelle ein Textfeld, um die Texteingabe eines Benutzers zu erhalten, indem du das `type`-Attribut mit dem Wert `text` zu dem `input`-Element hinzufügst.

# --hints--

Du hast entweder dein `input`-Element gelöscht oder es hat eine ungültige Syntax. Wenn du Attribute hinzugefügt hast, stelle sicher, dass ihre Werte zwischen Anführungszeichen stehen.

```js
assert($('input').length);
```

Dein `form`-Element sollte nur das `input`-Element enthalten. Entferne jedes zusätzliche HTML-Element oder jeglichen Text, zwischen den Tags des `form`-Elements.

```js
assert(
  $('form')[0].children.length === 1 &&
    $('form')[0].innerText.trim().length === 0
);
```

Dein `input`-Element hat kein `type`-Attribut mit dem Wert `text`. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags und/oder Leerzeichen vor allen Attributnamen vorhanden sind.

```js
assert($('input')[0].hasAttribute('type'));
```

Dein `input`-Element sollte ein `type`-Attribut mit dem Wert `text` haben. Du hast entweder den Wert weggelasssen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen stehen sollten.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .replace(/\s+/g, ' ')
    .match(/^text$/i)
);
```

Obwohl du das `type`-Attribut des `input`-Elements auf `text` gesetzt hast, wird empfohlen, den Wert eines Attributs zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

