---
id: 5efb2c990dc218d6c85f89b2
title: Schritt 43
challengeType: 0
dashedName: step-43
---

# --description--

Obwohl du deinen Knopf unter der Texteingabe hinzugefügt hast, erscheinen sie auf der Seite nebeneinander. That's because both `input` and `button` elements are <dfn>inline elements</dfn>, which don't appear on new lines.

The button you added will submit the form by default. Allerdings kann es zu Verwechslungen kommen, wenn man sich auf das Standardverhalten verlässt. Füge das `type`-Attribut mit dem Wert `submit` zum `button` hinzu, um klarzustellen, dass es sich um einen Bestätigungsbutton handelt.

# --hints--

Dein `button`-Element sollte ein öffnendes Tag haben. Öffnende Tags haben folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Dein `button`-Element sollte ein schließendes Tag haben. Öffnende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/button\>/));
```

Dein `button`-Element hat kein `type`-Attribut. Überprüfe, ob es ein Leerzeichen nach dem Namen des öffnenden Tags gibt.

```js
assert($('button')[0].hasAttribute('type'));
```

Dein `button`-Element sollte ein `type`-Attribut mit dem Wert `submit` haben. Du hast entweder den Wert weggelassen oder einen Tippfehler gemacht. Denke daran, dass Attributwerte zwischen Anführungszeichen geschrieben werden sollten.

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

Obwohl du das `type`-Attribut des `button`-Elements auf `submit` gesetzt hast, wird empfohlen, den Wert eines Attributs immer zwischen Anführungszeichen zu schreiben.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

