---
id: 5efc4f528d6a74d05e68af74
title: Schritt 56
challengeType: 0
dashedName: step-56
---

# --description--

Es gibt eine andere Möglichkeit, den Text eines `input`-Elements mit dem Element selbst zu verknüpfen. Du kannst den Text innerhalb eines `label`-Elements verschachteln und ein `for`-Attribut mit dem gleichen Wert wie das `id`-Attribut des `input`-Elements hinzufügen.

Here is an example of using the `for` attribute to associate a `label` with an `input` element:

```html
<label for="breakfast"> Breakfast </label>
<input id="breakfast" type="radio" name="meal" value="breakfast">
```

Verknüpfe den Text `Loving` mit der Checkbox, indem du nur den Text `Loving` in einem `label`-Element verschachtelst und ihm ein entsprechendes `for`-Attribut gibst.

# --hints--

Du solltest sicher stellen, dass das Kontrollkästchen noch vorhanden ist.

```js
assert($('input[type="checkbox"]')[0]);
```

Dein Kontrollkästchen sollte noch immer ein `id`-Attribut mit dem Wert `loving` haben.  Es ist `--fcc-expected--` erwartet worden, aber es wurde `--fcc-actual--` gefunden.

```js
assert.equal($('input[type="checkbox"]')[0].id, 'loving');
```

Der Text `Loving` sollte in einem `label`-Element eingebettet werden.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  /Loving/i.test(checkboxInputElem.nextElementSibling.innerText.replace(/\s+/g, ' '))
);
```

Du wirst ein neues `label`-Element hinzufügen müssen, in dem du den Text `Loving` einbetten kannst. Stelle sicher, dass es sowohl ein öffnendes als auch abschließendes Tag enthält.

```js
assert(
  document.querySelectorAll('label').length === 3 &&
    code.match(/<\/label\>/g).length === 3
);
```

Das neue `label`-Element sollte sich direkt rechts neben deinem Kontrollkästchen befinden. Stelle sicher, dass zwischen den zwei Elementen ein Leerzeichen vorhanden ist.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(checkboxInputElem.nextElementSibling.nodeName === 'LABEL');
```

Das neue `label`-Element hat kein `for`-Attribut. Überprüfe, ob ein Leerzeichen nach dem Namen des öffnenden Tags vorhanden ist.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert(labelElem.hasAttribute('for'));
```

Dein neues `label`-Element sollte ein `for`-Attribut mit dem Wert `loving` haben.  Es ist `--fcc-expected--` erwartet worden, aber es wurde `--fcc-actual--` gefunden.

```js
const labelElem = $('input[type="checkbox"]')[0].nextElementSibling;
assert.equal(labelElem.getAttribute('for'), 'loving');
```

Der Text `Loving` sollte innerhalb des neuen `label`-Elements eingebettet werden. Du hast entweder den Text weggelassen oder einen Tippfehler gemacht.

```js
const labelElem = document.querySelector('label[for="loving"]');
assert(labelElem.textContent.replace(/\s/g, '').match(/Loving/i));
```

Es sollte ein Leerzeichen zwischen deiner Checkbox und deinem neuen `label`-Element vorhanden sein.

```js
assert.match(code, />\s+<label\s+for\s*=\s*('|")loving/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> Loving
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

