---
id: 5f07c98cdb9413cbd4b16750
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Es ist an der Zeit, einen neuen Abschnitt hinzuzufügen. Füge ein zweites `section`-Element unter dem vorhandenen `section`-Element hinzu.

# --hints--

Dein `section`-Element sollte ein öffnendes Tag haben. Öffnende Tags besitzen die folgende Syntax: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Du solltest nur ein öffnendes `section`-Tag haben. Entferne bitte alle zusätzlichen.

```js
assert(document.querySelectorAll('section').length === 2);
```

Dein `section`-Element sollte ein abschließendes Tag enthalten. Abschließende Tag haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Du solltest nur ein abschließendes `section`-Tag hinzufügen. Entferne bitte alle zusätzlichen.

```js
assert(code.match(/<\/section>/g).length === 2);
```

Das zweite `section`-Element sollte nicht im ersten `section`-Element eingebettet sein.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Beide `section`-Elemente sollten sich zwischen dem öffnenden und abschließenden Tag deines `main`-Elements befinden.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

