---
id: 5f07fb1579dc934717801375
title: Schritt 33
challengeType: 0
dashedName: step-33
---

# --description--

Es ist an der Zeit, einen neuen Abschnitt hinzuzufügen. Füge ein drittes `section`-Element unter dem zweiten Element hinzu.

# --hints--

Dein `section`-Element sollte ein öffnendes Tag haben. Öffnende Tags besitzen diese Syntax: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

Du solltest nur ein öffnendes `section`-Tag haben. Entferne bitte alle zusätzlichen.

```js
assert(document.querySelectorAll('section').length === 3);
```

Dein `section`-Element sollte ein schließendes Tag haben. Abschließende Tags haben ein `/` genau nach dem `<`-Zeichen.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

Du solltest nur ein abschließendes `section`-Tag hinzufügen. Entferne bitte alle zusätzlichen.

```js
assert(code.match(/<\/section>/g).length === 3);
```

Alle `section`-Elemente sollten zwischen dem öffnenden und abschließenden Tag des `main`-Elements stehen.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

Das letzte `section`-Element sollte keinen Inhalt enthalten. Entferne innerhalb des `section`-Elements jedes HTML-Element oder Text.

```js
assert($('main > section')[2].children.length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

