---
id: 5f07be6ef7412fbad0c5626b
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Bevor du neuen Inhalt hinzufügst, solltest du von einem `section`-Element Gebrauch machen, um die Katzenfotos vom zukünftigen Inhalt zu trennen.

Take your `h2`, comment, `p`, and anchor (`a`) elements and nest them in a `section` element.

# --hints--

Dein `section`-Element sollte ein öffnendes Tag enthalten. Öffnende Tags haben die folgende Syntax: `<elementName>`.

```js
assert(document.querySelector('section'));
```

Dein `section`-Element sollte ein abschließendes Tag enthalten. Abschließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/section\s*>/i));
```

Das ganze `section`-Element sollte sich zwischen dem öffnenden und abschließenden Tag des `main`-Elements befinden.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

The existing `h2`, comment, `p` element, and anchor (`a`) element should be between the opening and closing tags of the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

Das `h1`-Element sollte nicht im `section`-Element verschachtelt werden.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

