---
id: 5f0d48e7b435f13ab6550051
title: Schritt 51
challengeType: 0
dashedName: step-51
---

# --description--

Das `legend`-Element fungiert in dem `fieldset`-Element als Untertitel für den Inhalt. Dies gibt den Benutzern Hinweise darüber, was sie in diesen Teil des Formulars eintragen sollen.

Füge ein `legend`-Element mit dem Text `Is your cat an indoor or outdoor cat?` über beide Radio-Buttons hinzu.

# --hints--

Dein `legend`-Element sollte ein öffnendes Tag enthalten. Opening tags have the following syntax: `<elementName>`.

```js
assert(document.querySelector('legend'));
```

Dein `legend`-Element sollte ein abschließendes Tag haben. Abschließende Tags haben ein `/` direkt nach dem `<`-Zeichen.

```js
assert(code.match(/<\/legend\>/));
```

Dein `legend`-Element sollte das erste Element direkt unter dem öffnenden Tag des `fieldset`-Elements und vor dem öffnenden `label`-Tag des ersten Radio-Buttons sein. Es ist nicht in der richtigen Position.

```js
const fieldsetElem = document.querySelector('fieldset');
const fieldsetElemChildren = fieldsetElem.children;
assert(
  fieldsetElem.firstElementChild.nodeName === 'LEGEND' &&
    fieldsetElemChildren[1].nodeName === 'LABEL' &&
    fieldsetElemChildren[1].children[0].nodeName === 'INPUT' &&
    fieldsetElemChildren[1].children[0].id === 'indoor'
);
```

Der Text des `legend`-Elements sollte `Is your cat an indoor or outdoor cat?` lauten. Du hast entweder den Text weggelassen, einen Tippfehler gemacht oder er befindet sich nicht zwischen den öffnenden und abschließenden Tags des `legend`-Elements.

```js
const extraSpacesRemoved = document
  .querySelector('legend')
  .innerText.replace(/\s+/g, ' ');
assert(extraSpacesRemoved.match(/Is your cat an indoor or outdoor cat\??$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>

            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

