---
id: 62dabe2ef403a12d5d295273
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Now that you turned the text `cat photos` inside the `p` element into a link, you don't need the second link below the `p` element. Lösche das gesamte Ankerelement unterhalb des `p`-Elements.

# --hints--

Dein Code sollte nur ein (`a`)-Ankerelement enthalten.

```js
assert(document.querySelectorAll('a').length === 1);
```

Das `p`-Element sollte immer noch das (`a`)-Ankerelement enthalten.

```js
assert($('p > a').length);
```

The link's `href` value should be `https://freecatphotoapp.com`. You have either omitted the `href` value or have a typo.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

The link's text should be `cat photos`. You have either omitted the text or have a typo.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

The `p` element content visible in the browser should be `See more cat photos in our gallery.` Double check the text, spacing, or punctuation of both the `p` and nested anchor element.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

Unter dem `p`-Element sollte kein `cat photos`-Text vorhanden sein.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
