---
id: 62dabe2ef403a12d5d295273
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Nun, da du den Text `cat photos` innerhalb des `p`-Elements in einen Link verwandelt hast, benötigst du den zweiten Link unterhalb des `p`-Elements nicht mehr. Lösche das gesamte Ankerelement unterhalb des `p`-Elements.

# --hints--

Dein Code sollte nur ein (`p`)-Paragraphen-Element enthalten.

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

Dein Code sollte nur ein (`a`)-Ankerelement enthalten.

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

Verändere keine Elemente, die du im vorherigen Schritt erstellt hast.

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

Unter dem `p`-Element sollte kein `cat photos`-Text vorhanden sein.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
