---
id: 5d822fd413a79914d39e98da
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Your first building looks pretty good now. Nest three new `div` elements in the `.background-buildings` container and give them the classes of `bb2`, `bb3`, and `bb4` in that order. These will be three more buildings for the background.

# --hints--

You should create a `div` with a class of `bb2`.

```js
assert.exists(document.querySelector('div.bb2'));
```

You should create a `div` with a class of `bb3`.

```js
assert.exists(document.querySelector('div.bb3'));
```

You should create a `div` with a class of `bb4`.

```js
assert.exists(document.querySelector('div.bb4'));
```

You should create 3 new `div` elements.

```js
assert.equal(document.querySelectorAll('div')?.length, 9);
```

You should place these `div` elements within the `.background-buildings` element.

```js
assert.equal(document.querySelector('div.background-buildings')?.children?.length, 4);
```

You should place the elements in the correct order.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1','bb2') && __t('div.bb2','bb3') && __t('div.bb3','bb4'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}
    
```

