---
id: 660f0a55847d6cc485f29ba5
title: Step 26
challengeType: 1
dashedName: step-26
---

# --description--

A <dfn>method</dfn> in JavaScript,  arrays are used to store multiple values in a single variable. One of the methods you can use with arrays is `.push()`. This method allows you to add a value to the end of an array. For example:".

```js
array.push("freeCodeCamp");
```

Use `.push()` to add the string `"freeCodeCamp"` to the end of your `rows` array. Add this code before your `console.log` so you can see the change you made to your array.

# --hints--

You should use `.push()` in your code.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

You should use the `.push()` method of your `rows` array.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

You should pass the string `"freeCodeCamp"` to your `.push()` method. 

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
