---
id: 660f4377a359972c521d3f4b
title: Step 78
challengeType: 1
dashedName: step-78
---

# --description--

You'll see the string printed in the console, because `true` is in fact true.

Change the condition of your `if` statement to the boolean `false`.

# --hints--

Your `if` condition should have `false` as the condition.

```js
assert.match(code, /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
