---
id: 660f1cedf3676fe26122ebf6
title: Step 38
challengeType: 1
dashedName: step-38
---

# --description--

You should see the numbers zero through seven printed in your console, one per line. This will serve as the foundation for generating your pyramid.

Replace your log statement with a statement to push `i` to your `rows` array.

# --hints--

You should not have a `console.log` call.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

You should call `.push()` on your `rows` array.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

You should push `i` to your `rows` array.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Your `.push()` should happen in your `for` loop. 

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
