---
id: 6610c16c4fa0df12c0e30675
title: Step 49
challengeType: 1
dashedName: step-49
---

# --description--

You are calling your `padRow` function, but not doing anything with that function call. All functions in JavaScript <dfn>return</dfn> a value, meaning they provide the defined result of calling them for you to use elsewhere.

To see the result of calling your `padRow` function, declare a `call` variable and assign your existing `padRow` call to that variable.

# --hints--

You should declare a `call` variable.

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+call/);
```

You should use `const` to declare your `call` variable.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

You should assign `call` the result of your `padRow` call.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
