---
id: 6610c8cfe4cf4d278e35c156
title: Step 63
challengeType: 1
dashedName: step-63
---

# --description--

As expected, your function now returns `undefined` again. Your `call` variable is not necessary any more, so remove the `call` declaration and the `console.log` for the `call` variable.

# --hints--

You should not have a `call` declaration.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s*call/);
```

You should not log your `call` variable.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
