---
id: 6610c1d97b1671140f95cfbb
title: Step 50
challengeType: 1
dashedName: step-50
---

# --description--

Now add a log statement to print the value of your `call` variable.

# --hints--

You should add a `console.log` call.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

You should log your `call` variable. Don't forget the semicolon.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
