---
id: 6610bbed59bc2a0194d85533
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Notice how the value inside your `rows` array has been changed directly? This is called <dfn>mutation</dfn>. As you learn more about arrays, you will learn when to mutate an array, and when you should not.

Before moving on, this is a great opportunity to learn a common array use. Currently, your code accesses the last element in the array with `rows[2]`. But you may not know how many elements are in an array when you want the last one.

You can make use of the `.length` property of an array - this returns the number of elements in the array. For example, your `rows` array has 3 elements, so `rows.length` would be `3`.

Since you know the last element is at index `2`, you can compare that with the length of the array and see it is one less. The last element of an array will always be accessible at the index `length - 1`. You can use that in your bracket notation.

Update your `rows[2]` to access the `rows.length - 1` index instead of the `2` index. You should not see anything change in your console.

# --hints--

You should use the `.length` property of your `rows` array.

```js
assert.match(code, /rows\.length/);
```

You should use `rows.length` in your bracket notation.

```js
assert.match(code, /rows\[\s*rows\.length/);
```

You should subtract `1` from the length in your bracket notation.

```js
assert.match(code, /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
