---
id: 664599653fcd6e97104f9261
title: Step 76
challengeType: 1
dashedName: step-76
---

# --description--

In programming, it is common to have multiple conditions that need to be checked.

In the code, an extra `if` statement with the condition of `5 < 10` has been provided for you.

But instead of using multiple `if` statements, you can use `else if` statements to check multiple conditions in a single block of code.

Here is the syntax for an `else if` statement:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

If the first condition is `false`, JavaScript will check the next condition in the chain. If the second condition is `false`, JavaScript will check the third condition, and so on.

Update the second `if` statement to use an `else if` statement instead. 

Check the console to see the results.

# --hints--

You should update the second `if` statement to an `else if (5 < 10)` statement.

```js
assert.match(code, /else\s+if\s*\(\s*5\s*<\s*10\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = "";

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
