---
id: 66458f0a05df478aa627629e
title: Step 75
challengeType: 1
dashedName: step-75
---

# --description--

The text is gone again! Empty strings evaluate to `false`, making them a <dfn>falsy value</dfn>. You will learn more about truthy and falsy values in future projects.

When the condition in the `if` statement is `false`, the code inside the `if` block will not run. In these situations, you can use an `else` statement to run code when the `if` condition is `false`. 

Here is an example of an `if` statement with an `else` block:

```js

if (condition) {
  // this code will run if condition is true
} else {
  // this code will run 
  // if the condition is false
}
```

Add an `else` block to the `if` block. Inside the `else` block, log the string `"This is the else block"` to the console.

Check the console to see the results.

# --hints--

You should have an `else` block.

```js
assert.match(code, /else\s*\{/);
```

Your `else` block should log the string `"This is the else block"` to the console.

```js
assert.match(code, /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = "";

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
