---
id: 664ed97c55a99e5ffab759dc
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Now it is time to review the concepts you have learned so far. This review step is to ensure you have a good grasp on the basics so you can apply them to the pyramid generator project later on. 

In this step you will need to do the following:

- Use the `let` keyword to declare a variable called `profession` and assign it the string `"teacher"`.
- Use the `let` keyword to declare a variable called `age`. **Do not** assign it a value.
- Log to the console the value of the `age` variable. You should see `undefined` in the console.
- Log to the console the value of the `profession` variable. You should see `"teacher"` in the console.

# --hints--

You should declare a variable called `profession`.

```js
assert.match(code, /profession/);
```

You should use `let` to declare the `profession` variable.

```js
assert.match(code, /let\s+profession/);
```

You should assign the string `"teacher"` to the `profession` variable.

```js
assert.equal(profession, "teacher");
```

You should have a variable called `age`

```js
assert.match(code, /age/);
```

You should use `let` to declare the `age` variable.

```js
assert.match(code, /let\s+age/);
```

You should log `age` to the console.

```js
assert.match(code, /console\.log\(age\);/);
```

You should log `profession` to the console.

```js
assert.match(code, /console\.log\(profession\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--

--fcc-editable-region--
```
