---
id: 664ed9344ba3495e1f67f8ea
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Now that you have finished with the review, you can remove all of the code you just wrote in the previous step. 

# --hints--

You should not have a variable called profession.

```js
assert.notMatch(code, /profession/);
```

You should not have a variable called age.

```js
assert.notMatch(code, /age/);
```

You should not have a console.log statement for age.

```js
assert.notMatch(code, /console\.log\(\s*age\s*\);/);
```

You should not have a console.log statement for profession.

```js
assert.notMatch(code, /console\.log\(\s*profession\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let profession = 'teacher';
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
