---
id: 676f03523846a9359b200d16
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Your `call` variable has an `undefined` value, even though you defined it! This is because your `padRow` function does not currently return a value. By default, functions return `undefined` as their value.

In order to return something else, you need to use the `return` keyword. Here is an example of a function that returns the string `"Functions are cool!"`:

```js
function demo() {
  return "Functions are cool!";
}
```

Use the `return` keyword to have your function return the string `"Hello!"`.

# --hints--


Your `padRow` function should use the `return` keyword.

```js
assert.match(padRow.toString(), /return/);
```

Your `padRow` function should return the string `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```
# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);
```
