---
id: 676f03aa08e64b3797f82794
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

With that quick review complete, you should remove your `addTwoNumbers` function, `sum` variable, and log statement.

# --hints--

You should not have a `addTwoNumbers` function in your code.

```js
assert.notMatch(code, /addTwoNumbers/);
```

You should not have a `sum` variable in your code.

```js
assert.notMatch(code, /sum/);
```

You should not log the `sum` variable.

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);


function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);

```
