---
id: 676f0342743d7134ef2fe380
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Programs often can get unwieldy when code with similar logic is repeated. That is why programmers
use functions. 

A <dfn>function</dfn> is a block of code that can be reused throughout your application. Functions are declared with the following syntax:

```js
function name(parameter) {

}
```

The `function` keyword tells JavaScript that the `name` variable is going to be a function. `parameter` is a variable that represents a value that is passed into the function when it is used. A function may have as many, or as few, <dfn>parameters</dfn> as you'd like. The space between the curly braces is the <dfn>function body</dfn> and holds
the operations for the function.

Declare a `padRow` function. Do not create any parameter variables yet. The function body should be empty. Remember that you need to use camel case for your naming convention.

# --hints--


You should use the `function` keyword.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

You should declare a `padRow` function.

```js
assert.isFunction(padRow);
```

Your `padRow()` function should not have any parameters.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Your `padRow()` function should have an empty body.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```


# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
