---
id: 676f03c25d4ae33856a0b7e8
title: Step 32
challengeType: 0
dashedName: step-32
---

# --description--

Now your `call` variable has the value `"Testing"`. But your function is no longer using the `name` parameter.

Remove the `name` parameter from your function declaration, then remove your `"CamperChan"` string from the `padRow` call.

Also, remove both `console.log` from the `padRow` function.

# --hints--


Your `padRow` function should not have a `name` parameter.

```js
assert.notMatch(padRow.toString(), /name/);
```

You should not pass `"CamperChan"` to your `padRow` call.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

You should still call your `padRow` function.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

You should not have a `console.log` before your `return` keyword.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

You should not have a `console.log` after your `return` keyword.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

```
