---
id: 676f035eef6b1a35eb437386
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

When you have a value that is explicitly written in your code, like the `"Hello!"` string in your function, it is considered to be <dfn>hard-coded</dfn>. Hard-coding a value inside a function might not make it as reusable as you'd like.

Instead, you can define <dfn>parameters</dfn> for the function. Parameters are special variables that are given a value when you call the function, and can be used in your function to dynamically change the result of the function's code.

To add a parameter to your function, you need to add a variable name inside the parentheses. For example, this `demo` function has a `name` parameter:

```js
function demo(name) {

}
```

`name` sounds like a useful parameter, so go ahead and add it to your `padRow` function.

# --hints--


Your `padRow` function should have a `name` parameter.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);
```
