---
id: 676f034bc0adac3545249dad
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

You are calling your `padRow` function, but not doing anything with that function call. All functions in JavaScript <dfn>return</dfn> a value, meaning they provide the defined result of calling them for you to use elsewhere.

To see the result of calling your `padRow` function, declare a `call` variable and assign your existing `padRow` call to that variable.


# --hints--


You should declare a `call` variable.

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+call/);
```

You should use `const` to declare your `call` variable.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

You should assign `call` the result of your `padRow` call.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--
```
