---
id: 676f0366c1ee033641f409d5
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

If you open your console again, you'll see that your `padRow` function is returning `undefined`, even though you defined a return value! This is because parameters need to be given a value when you **call** the function.

When you pass a value to a function call, that value is referred to as an <dfn>argument</dfn>. Here is an example of calling a `demo` function and passing `"Naomi"` as the argument for the `name` parameter.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Pass your own name as the argument for the `name` parameter in your `padRow` call. Remember that your name is a string, so you'll need to use quotes.


# --hints--

You should pass a string to your `padRow()` call.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```


# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);
```
