---
id: 676f0346b39d42351a30ff5f
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

In order to use a function, you need to call it. A <dfn>function call</dfn> tells your application to run the code from the function wherever you choose to call it. The syntax for a function call is the function name followed by parentheses. For example, this code defines and calls a `test` function.


```js
function test() {

}
test();
```

Call your `padRow` function.

# --hints--


You should call the `padRow` function.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Your `padRow` function body should be empty.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
```
