---
id: 676f03d41665c638e3bee5af
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

In order to know how to format a row, your `padRow` function will need to know which row number you are on, and how many rows in total are being generated.

The best way to do this is by creating function parameters for them. Give your `padRow` function a `rowNumber` and `rowCount` parameter. Multiple parameters are separated by a comma:

```js
function name(first, second) {

}
```

# --hints--

Your `padRow` function should have a `rowNumber` parameter.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

You should add a comma after your `rowNumber` parameter.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Your `padRow` function should have a `rowCount` parameter.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--


```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--

```

