---
id: 676f03c57ca6d03882095b17
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

Because your function was no longer using the parameter, changing the argument did not affect it.

Go ahead and remove the `test` declaration and `return` statement from your `padRow` function, so the function is empty again.

# --hints--


Your `padRow` function should not have a `test` variable.

```js
assert.notMatch(padRow.toString(), /test/);
```

Your `padRow` function should not return a value.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

```
