---
id: 676f03625f17903616f2efc5
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

A function does not have to return a hard-coded value. It can return the value stored in a variable. Parameters are special variables for a function, so they can also be returned.

Change your `padRow` function to `return` the `name` parameter directly.


# --hints--


Your `padRow` function should return the value of the `name` parameter.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);

```
