---
id: 676f034f104662357388254f
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Now add a log statement to print the value of your `call` variable.

# --hints--


You should add a `console.log` call.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

You should log your `call` variable. Don't forget the semicolon.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--
```
