---
id: 676f03ca04d03538ae92ddd3
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

As expected, your function now returns `undefined` again. Your `call` variable is not necessary any more, so remove the `call` declaration and the `console.log` for the `call` variable.


# --hints--


You should not have a `call` declaration.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s*call/);
```

You should not log your `call` variable.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
console.log(count + 1);

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--
```
