---
id: 612e8279827a28352ce83a72
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Sasa nakili seti ya vipengele saba vya `.key`, na ubandike seti mbili zaidi kwenye div ya `.keys`.

# --hints--

Unapaswa kuwa na vipengele 21 vya `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Unapaswa kuwa na vipengele 15 vya `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Vipengele vyote vya `.key` vinapaswa kuwa ndani ya kipengele chako cha `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Vipengele vyote vya `.black--key` vinapaswa kuwa ndani ya kipengele chako cha `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
