---
id: 612eb4893b63c75bb9251ddf
title: Hatua ya 22
challengeType: 0
dashedName: step-22
---

# --description--

Anza kupamba nembo kwa kuunda cha kichaguzi `.logo`. Weka `width` kuwa `200px`, `position` ya `absolute` na `top` kuwa `23px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.logo`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo'));
```

Kichaguzi chako cha `.logo` kinafaa kuwa na sifa ya `width` iliyowekwa kuwa `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo')?.width === '200px');
```

Kichaguzi chako cha `.logo` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo')?.position === 'absolute');
```

Kichaguzi chako cha `.logo` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `23px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.logo')?.top === '23px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <img class="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg" alt="freeCodeCamp Logo" />
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

.key.black--key::after {
  background-color: #1d1e22;
  content: "";
  position: absolute;
  left: -18px;
  width: 32px;
  height: 100px;
}

--fcc-editable-region--

--fcc-editable-region--
```
