---
id: 612e83ec2eca1e370f830511
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

Ongeza kipengele cha `link` ndani ya kipengele chako cha `head`. Kwa kipengele hicho cha `link`, weka sifa ya `rel` kuwa `stylesheet` na `href` kuwa `./styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/)
```

Hupaswi kubadilisha tagi zako za `head` zilizopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Unapaswa kuwa na kipengele kimoja cha kujifunga cha `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `rel` yenye thamani `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
