---
id: 612ea4c4993aba52ab4aa32e
title: Hatua ya 18
challengeType: 0
dashedName: step-18
---

# --description--

Sasa ni wakati wa kutumia vichaguzi vya pseudo ulivyotayarisha hapo awali. Ili kuunda vitufe vyeusi, ongeza kichaguzi kipya cha `.key.black--key::after`. Hii italenga vipengee vilivyo na darasa la `key black--key`, na kuchagua kipengele cha pseudo baada ya vipengele hivi katika HTML.

Katika kichaguzi kipya, weka `background-color` kuwa `#1d1e22`. Pia weka kipengele cha `content` kuwa `""`. Hii itafanya vipengee vya pseudo kuwa tupu.

Sifa ya `content` inatumika kuweka au kubatilisha maudhui ya kipengele. Kwa chaguo-msingi, vipengele vya pseudo vilivyoundwa na vichaguzi vya pseudo vya `::before` na `::after` ni tupu, na vipengele havitatolewa kwa ukurasa. Kuweka kipengele cha `content` kuwa mfuatano tupu `""` kutahakikisha kipengele kinatolewa kwa ukurasa huku kikiwa tupu.

Ikiwa ungependa kufanya majaribio, jaribu kuondoa `background-color` na kuweka thamani tofauti za kipengele cha `content`, kama vile `"♥"`. Kumbuka kutendua mabadiliko haya unapomaliza ili majaribio yapite.

# --hints--

Unapaswa kuwa na kichaguzi cha `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Kichaguzi chako cha `.key.black--key::after` kinafaa kuwa na sifa ya `background-color` iliyowekwa kuwa `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

Kichaguzi chako cha `.key.black--key::after` kinafaa kuwa na sifa ya `content` iliyowekwa kuwa `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
