---
id: 612e7d1c29fb872d6384379c
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Weka `div` ya pili ndani ya `div` yako iliyopo, na uweke `class` kuwa `keys`.

# --hints--

Unapaswa kuunda kipengele cha pili cha `div`.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

Kipengele chako kipya cha `div` kinafaa kuwa ndani ya kipengee cha `div` kilichopo.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

Kipengee chako kipya cha `div` kinapaswa kuwa na `class` iliyowekwa kuwa `keys`.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
