---
id: 612e813b3ba67633222cbe54
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Kumbuka kwamba sifa ya `class` inaweza kuwa na thamani nyingi. Ili kutenganisha vitufe vyako vyeupe na vitufe vyako vyeusi, utaongeza thamani ya pili ya `class` ya `black--key`. Ongeza hii kwenye vipengele vyako vya `.key` vya pili, tatu, tano, sita na saba.

# --hints--

Kipengele chako cha pili cha `.key` kinapaswa kuwa na `class` iliyowekwa kuwa `black--key`.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

`.key` yako ya tatu inapaswa kuwa na `black--key` katika `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

`.key` yako ya tano inapaswa kuwa na `black--key` katika `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

`.key` yako ya sita inapaswa kuwa na `black--key` katika `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

`.key` yako ya saba inapaswa kuwa na `black--key` katika `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

Unapaswa kuwa na vipengele vitano vya `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

Unapaswa kuwa na vipengele saba vya `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
