---
id: 612e96fc87fe8e44f69f7ec5
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Wakati wa kupamba vitufe (keys). Chini ya sheria ya `#piano`, lenga kipengele cha `.keys` kwa kutumia kichaguzi cha `class`. Ipe sheria mpya sifa ya `background-color` ya `#040404`, sifa ya `width` ya `949px` na sifa ya `height` ya `180px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.keys`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

Kichaguzi chako cha `.keys` kinafaa kuwa na sifa ya `background-color` iliyowekwa kuwa `#040404`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

Kichaguzi chako cha `.keys` kinafaa kuwa na sifa ya `width` iliyowekwa kuwa `949px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

Kichaguzi chako cha `.keys` kinafaa kuwa na sifa ya `height` iliyowekwa kuwa `180px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
