---
id: 612e804c54d5e7308d7ebe56
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Ndani ya kipengele chako cha `.keys`, ongeza vipengele saba vya `div`. Zipe zote class ya `key`.

# --hints--

Unapaswa kuunda vipengee saba vipya vya `div`.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

Vipengele vyako saba vipya vya `div` vinapaswa kuwa ndani ya kipengele chako `.keys`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

Vipengele vyako saba vipya vya `div` vinapaswa kuwa na `class` iliyowekwa kuwa `key`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
