---
id: 5f332a88dc25a0fd25c7687a
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O nome da cafeteria é `CAMPER CAFE`. Você deve criar um elemento `h1` dentro do elemento `main`. Coloque o nome da cafeteria em letras maiúsculas para que ele se destaque.

# --hints--

Você deve acrescentar uma tag de abertura `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Você deve acrescentar uma tag de fechamento `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

O elemento preexistente `main` deve permanecer o mesmo.

```js
assert($('main').length === 1);
```

O elemento `h1` deve estar dentro do elemento `main`.

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

O elemento `h1` deve conter o texto `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

