---
id: 5f33071498eb2472b87ddee4
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS é a linguagem usada para estilizar um documento HTML. Ela descreve como elementos HTML devem ser exibidos na tela.

Como você aprendeu nos últimos passos do CatPhotoApp, existe uma estrutura básica para começar a construir sua página da web. Cada documento HTML deve ter uma declaração `DOCTYPE` e um elemento `html`. O `DOCTYPE` diz ao navegador em que versão do HTML o documento está. O elemento `html` representa o elemento raiz que contém todos os outros elementos.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Adicione a tag `<!DOCTYPE html>` e um elemento `html` com um atributo `lang` de `en`.

# --hints--

Você deve ter uma declaração `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Você deve ter uma tag de abertura `<html lang="en">` com o idioma definido para inglês.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Você deve acrescentar uma tag de fechamento `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
