---
id: 5f331e55dfab7a896e53c3a1
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

O `title` é um dos vários elementos que fornecem informações extras que não são visíveis na página da web, mas ele é útil para mecanismos de busca ou até mesmo em como a página é exibida.

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8` to tell the browser how to encode characters for the page.

# --hints--

Você deve criar uma tag `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

A tag `meta` deve ter um atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Seu atributo `charset` deve ter um valor de `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
