---
id: 5f33294a6af5e9188dbdb8f3
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

O nome da cafeteria é `CAMPER CAFE`. Você deve criar um elemento `h1` dentro do elemento `body`. Coloque o nome da cafeteria em letras maiúsculas para que ele se destaque.

# --hints--

Você deve acrescentar uma tag de abertura `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Você deve acrescentar uma tag de fechamento `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Você deve manter o elemento `body` preexistente.

```js
assert($('body').length === 1);
```

O elemento `h1` deve estar dentro do elemento `body`.

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

O elemento `h1` deve conter o texto `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
<html>
```
