---
id: 5f331e55dfab7a896e53c3a1
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

O `title` é um dos vários elementos que fornecem informações extras que não são visíveis na página da web, mas ele é útil para mecanismos de busca ou até mesmo em como a página é exibida.

Dentro do elemento `head`, adicione um elemento `meta` com um atributo chamado `charset` definido para o valor `utf-8` para dizer ao navegador como codificar caracteres para a página. Note que os elementos `meta` fecham em si mesmos.

# --hints--

Você deve criar uma tag `meta`.

```js
assert(code.match(/<meta.*?\/?>/is));
```

A tag `meta` deve ter um atributo `charset`.

```js
assert(code.match(/<meta charset=/i));
```

Seu atributo `charset` deve ter um valor de `utf-8`.

```js
assert(code.match(/charset=('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
