---
id: 5f3313e74582ad9d063e3a38
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione um elemento `head` dentro do elemento `html` para que você possa adicionar um elemento `title`. O texto do elemento `title` deve ser `Cafe Menu`.

# --hints--

Você deve acrescentar uma tag de abertura`<head>`.

```js
assert(code.match(/<head>/i));
```

Você deve acrescentar uma tag de encerramento `</head>`.

```js
assert(code.match(/<head>/i));
```

Você deve acrescentar uma tag de abertura`<title>`.

```js
assert(code.match(/<title>/i));
```

Você deve acrescentar uma tag de encerramento `</title>`.

```js
assert(code.match(/<\/title>/i));
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

O elemento `title` deve conter o texto `Cafe Menu`. Não se esqueça de verificar a ortografia.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
