---
id: 5f33294a6af5e9188dbdb8f3
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

É hora de adicionar o conteúdo do menu. Adicione um elemento `main` dentro do elemento `body` existente. Mais pra frente ele conterá informações sobre preços dos cafés e sobremesas oferecidas pela cafeteria.

# --hints--

Você deve acrescentar uma tag de abertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Você deve acrescentar uma tag de fechamento `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Você deve manter o elemento preexistente `body`. Verifique se você não excluiu a tag de fechamento.

```js
assert($('body').length === 1);
```

A tag `main` deve estar dentro da tag `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
