---
id: 5f33071498eb2472b87ddee4
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Como você aprendeu nos últimos passos do CatPhotoApp, existe uma estrutura básica para começar a construir sua página da web.

Adicione a tag `<!DOCTYPE html>` e um elemento `html` com um atributo `lang` de `en`.

# --hints--

Você deve criar uma declaração `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Você deve ter uma tag de abertura `<html>` com um atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Você deve acrescentar uma tag de encerramento `<html>`. Lembre-se de que as tags de fechamento têm um `/` logo depois do colchete de abertura `<`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
