---
id: 5f33310c1851c6c4da013250
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

É hora de adicionar o conteúdo do menu. Adicione um elemento `main` abaixo do elemento existente `header`. Mais pra frente ele conterá informações sobre preços dos cafés e sobremesas oferecidas pela cafeteria.

# --hints--

Você deve acrescentar uma tag de abertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Você deve acrescentar uma tag de fechamento `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Você deve manter o elemento preexistente `header`. Verifique se você não excluiu a tag de fechamento.

```js
assert($('header').length === 1);
```

O elemento `main` deve estar depois do elemento `header`.

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
