---
id: 5f332b23c2045fb843337579
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Para que os visitantes saibam que o café foi fundado em 2020, adicione um elemento `p` abaixo do elemento `h1` com o texto `Est. 2020`.

# --hints--

Você deve acrescentar uma tag de abertura`<p>`.

```js
assert(code.match(/<p>/i));
```

Você deve acrescentar uma tag de encerramento `</p>`.

```js
assert(code.match(/<\/p>/i));
```

Você deve manter o elemento `h1` preexistente. Verifique se você não excluiu a tag de fechamento.

```js
assert($('h1').length === 1);
```

O elemento `p` deve estar abaixo do elemento `h1`.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

O elemento `p` deve conter o texto `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

