---
id: 5f332b23c2045fb843337579
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Já que o elemento `p` adicionado na etapa anterior fornece informações complementares sobre a cafeteria, aninhe os elementos `h1` e `p` em um elemento `header`.

# --hints--

Você deve acrescentar uma tag de abertura `<header>`.

```js
assert(code.match(/<header>/i));
```

Você deve acrescentar uma tag de fechamento `</header>`.

```js
assert(code.match(/<\/header>/i));
```

O elemento `h1` deve estar dentro do elemento `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

O elemento `p` deve estar dentro do elemento `header`.

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
<html>
```

