---
id: 5f3477ae9675db8bb7655b30
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

No passo anterior, você usou um <dfn>seletor de tipo</dfn> para estilizar o elemento `h1`. Centralize os elementos `h2` e `p` adicionando um novo seletor de tipo para cada um dos elementos `style` existentes.

# --hints--

O elemento preexistente `h1` deve permanecer o mesmo.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Você não deve adicionar uma nova tag `style`. Adicione as novas regras de CSS à tag `style` existente.

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

Você deve usar um novo seletor `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Você deve usar um novo seletor `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

O elemento `h1` deve ter o `text-align` para `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

O elemento `h2` deve ter o `text-align` de `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

O elemento `p` deve ter o `text-align` para `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

