---
id: 5f3477cb2e27333b1ab2b955
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Agora você precisa vincular o arquivo `styles.css` para que os estilos sejam aplicados novamente. Adicione um elemento `link` de fechamento automático ao elemento `head`. Dê a ele o valor do atributo `rel` de `stylesheet` e o valor do atributo `href` de `styles.css`.

# --hints--

O código deve ter um elemento `link`.

```js
const link = document.querySelector('link');
assert(link);
```

Você não deve mudar o elemento `head` existente. Verifique se você não excluiu a tag de fechamento.

```js
assert($('head').length === 1);
```

Você deve ter um elemento `link` de autofechamento.

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
const link = document.querySelector('head > link');
assert(link);
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
