---
id: 5f3477cb2e27333b1ab2b955
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

Agora você precisa vincular o arquivo `styles.css` para que os estilos sejam aplicados novamente. Adicione um elemento `link` de fechamento automático ao elemento `head`. Dê a ele um atributo `rel` com o valor `stylesheet`, um atributo `type` com o valor `text/css` e um atributo `href` com o valor `styles.css`.

# --hints--

O código deve ter um elemento `link`.

```js
// link is removed -> if exists, replaced with style
const link = document.querySelector('head > style');
assert(link);
```

Você não deve mudar o elemento `head` existente. Verifique se você não excluiu a tag de fechamento.

```js
assert($('head').length === 1);
```

O elemento `link` deve ser de fechamento automático.

```js
assert(code.match(/<link[\w\W\s]+\/?>/i));
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
assert(code.match(/rel=('|")stylesheet\1/i));
```

O elemento `link` deve ter o atributo `type` com o valor `text/css`.

```js
assert(code.match(/type=('|")text\/css\1/i));
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert(code.match(/href=('|")styles.css\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
    <main>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
<html>
```

```css
h1, h2, p {
  text-align: center;
}
```
