---
id: 5f3477cb303c5cb61b43aa9b
title: Passo 18
challengeType: 0
dashedName: step-18
---

# --description--

O texto está centralizado novamente, então o link para o arquivo CSS está funcionando. Adicione outro estilo ao arquivo que altere a propriedade `background-color` para `brown` para o elemento `body`.

# --hints--

Você deve usar o seletor `body`.

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

Você deve definir o valor da propriedade `background-color` para `brown`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

O elemento `body` deve ter um fundo `brown` (marrom).

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

