---
id: 5f3477cbcb6ba47918c1da92
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

Para que a estilização da página pareça semelhante no celular, em um desktop ou em um laptop, você precisa adicionar um elemento `meta` com um atributo especial `content`.

Adicione o seguinte dentro do elemento `head`:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

O código deve ter dois elementos `meta`.

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

O elemento `meta` deve ter um atributo `name` com o valor `viewport`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

O elemento `meta` deve ter um atributo `content` com o valor `width=device-width, initial-scale=1.0`.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
