---
id: 5f3477ae8466a9a3d2cc953c
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Agora que você tem o CSS no arquivo `styles.css`, vá em frente e remova o elemento `style` e todo seu conteúdo. Uma vez removido, o texto centralizado voltará para a esquerda.

# --hints--

Você deve excluir as tags `style` do seu código.

```js
assert(!code.match(/style/i));
```

Você não deve ter nenhum seletor CSS no seu arquivo HTML.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

