---
id: 5f344fbc22624a2976425065
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Crie um elemento `h2` no elemento `section` e dê a ele o texto `Coffee`.

# --hints--

Você deve acrescentar uma tag de abertura `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

Você deve acrescentar uma tag de fechamento `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

Você deve manter o elemento `section` preexistente. Verifique se você não excluiu a tag de fechamento.

```js
assert($('section').length === 1);
```

O elemento `h2` deve estar dentro do elemento `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

O elemento `h2` deve conter o texto `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
