---
id: 5f3477ae34c1239cafe128be
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora você tem três tipos de seletores exatamente com o mesmo estilo. Você pode adicionar o mesmo grupo de estilos a muitos elementos, criando uma lista de seletores. Cada seletor é separado por vírgulas assim:

```css
selector1, selector2 {
  property: value;
}
```

Exclua os três seletores de tipos existentes e substitua-os por uma lista de seletores que centraliza o texto para os elementos `h1`, `h2` e `p`.

# --hints--

Você deve usar um seletor de tipo único para todos os três elementos, `h1, h2, p`. Não se esqueça de lançar os elementos na sequência solicitada.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Você deve ter apenas um seletor no seu elemento `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

