---
id: 5f35e5c44359872a137bd98f
title: Passo 28
challengeType: 0
dashedName: step-28
---

# --description--

Como o principal produto para venda na cafeteria é o café, você pode usar uma imagem de grãos de café para o fundo da página.

Exclua o comentário e o conteúdo dentro do seletor de tipos `body`. Em seguida, adicione uma propriedade `background-image` e defina seu valor como `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Você deve remover o comentário com a propriedade `background-color`.

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

O seletor `body` não deve conter nenhum comentário.

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

Defina a propriedade `background-image` com o valor `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

O elemento `body` deve ter como imagem de fundo os grãos de café.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

