---
id: 5f356ed63c7807a4f1e6d054
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Agora o objetivo é fazer a `div` não ocupar toda a largura da página. A propriedade CSS `width` é perfeita para isso.

Você pode usar o seletor de `id` para direcionar um elemento específico com um atributo `id`. Um <dfn>seletor de id</dfn> é definido colocando o símbolo hash `#` diretamente na frente do valor de `id` do elemento. Por exemplo, se um elemento tiver o `id` `cat`, você direcionaria esse elemento assim:

```css
#cat {
  width: 250px;
}
```

Use o seletor `#menu` para dar ao elemento uma largura de `300px`.

# --hints--

Você deve ter um seletor `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

Você deve definir a propriedade `width` como `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

A `div` deve ter uma largura de 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

