---
id: 5f356ed6199b0cdef1d2be8f
title: Passo 26
challengeType: 0
dashedName: step-26
---

# --description--

Até agora, você tem usado seletores de tipo e de id para elementos de estilo. No entanto, é mais comum usar um seletor diferente para estilizar os elementos.

Um <dfn>seletor de classe</dfn> é definido por um nome com um ponto diretamente na frente dele, como este:

```css
.class-name {
  styles
}
```

Altere o seletor `#menu` existente para um seletor de classe substituindo `#menu` por uma classe chamada `.menu`.

# --hints--

Você deve usar um seletor de classe `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Você não deve ter um seletor `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

