---
id: 5f356ed6199b0cdef1d2be8f
title: Passo 27
challengeType: 0
dashedName: step-27
---

# --description--

Até agora, você tem usado seletores de tipo para elementos de estilo. Um <dfn>seletor de classe</dfn> é definido por um nome com um ponto diretamente à sua frente, assim:

```css
.class-name {
  styles
}
```

Altere o seletor `div` existente para um seletor de classe substituindo `div` por uma classe chamada `menu`.

# --hints--

Você deve usar um seletor de classe `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Você não deve ter um seletor `div`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('div');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet" type="text/css" />
  </head>
  <body>
    <div>
      <header>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
      </header>
      <main>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
<html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
div {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

