---
id: 5f3c866de7a5b784048f94b1
title: Passo 36
challengeType: 0
dashedName: step-36
---

# --description--

Agora está mais próximo do que você quer, mas seria legal se o sabor e o preço estivessem na mesma linha. Elementos `p` são elementos <dfn>em nível de bloco</dfn>, então eles ocupam toda a largura do seu elemento pai.

Para fazer com que fiquem na mesma linha, você precisa aplicar algum tipo de estilo aos elementos `p`, de forma que eles comecem a se comportar mais como elementos <dfn>inline</dfn>. Adicione um atributo `class` com o valor `item` para o primeiro elemento `article` sob o título `Coffee`.

# --hints--

Você deve aplicar a classe `item` ao elemento `article`.

```js
assert(code.match(/<article\s*class=('|")item\1\s*>/i))
```

Você deve ter apenas uma classe `item`.

```js
assert($('.item').length === 1);
```

O elemento `article` deve ter a classe `item`.

```js
assert($('article')[0].className === 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
