---
id: 5f3c866daec9a49519871816
title: Passo 30
challengeType: 0
dashedName: step-30
---

# --description--

Elementos `article` comumente contêm vários elementos que possuem informações relacionadas. Neste caso, ele vai conter um sabor de café e um preço por esse sabor. Aninhe dois elementos `p` dentro do elemento `article`. O primeiro texto deve ser `French Vanilla` e o segundo texto deve ser `3.00`.

# --hints--

Você deve mudar o elemento `article` pre-existente.

```js
assert($('article').length === 1);
```

O elemento `article` deve ter dois elementos `p`.

```js
assert($('article').children('p').length === 2);
```

O primeiro elemento `p` deve ter o texto `French Vanilla`.

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

O segundo elemento `p` deve ter o texto `3.00`.

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
