---
id: 5f716ad029ee4053c7027a7a
title: Passo 49
challengeType: 0
dashedName: step-49
---

# --description--

Coloque dois elementos `p` dentro do elemento `article`. O texto do primeiro elemento deve ser `Donut` e o texto do segundo deve ser `1.50`. Coloque os dois na mesma linha, garantindo que não haja espaço entre eles.

# --hints--

Você deve mudar o elemento `article` existente.

```js
assert($('article').length === 6);
```

O novo elemento `article` deve ter dois elementos `p`.

```js
assert($('article').last().children('p').length === 2);
```

O primeiro elemento `p` deve ter o texto `Donut`.

```js
assert($('article').last().children('p')[0].innerText.match(/Donut/i));
```

O segundo elemento `p` deve ter o texto `1.50`.

```js
assert($('article').last().children('p')[1].innerText.match(/1\.50/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

