---
id: 178d9d193fae0e36c3d0a932
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

The second clue involves a keypad with mysterious symbols. You need to search through a grid of symbols to find a specific target pattern.

Create a function called `patternSearch` that accepts two parameters: `grid` (an array of arrays of strings) and `target` (a string). Use a `for...of` loop to iterate over each row in the grid, and within that loop use another `for...of` loop to iterate over each symbol in the row.

If the `target` is found in any row, return `true` immediately. If the loop finishes without finding the target, return `false`.

For example:

```js
const grid = [
  ["alpha", "beta", "gamma"],
  ["delta", "epsilon", "zeta"],
];
patternSearch(grid, "epsilon"); // true
patternSearch(grid, "theta"); // false
```

# --hints--

You should create a function named `patternSearch`.

```js
assert.isFunction(patternSearch);
```

Your `patternSearch` function should have `grid` and `target` parameters.

```js
const regex = __helpers.functionRegex('patternSearch', ['grid', 'target']);
assert.match(__helpers.removeJSComments(patternSearch.toString()), regex);
```

`patternSearch` should return `true` when the target exists in the grid.

```js
const grid = [
  ["alpha", "beta", "gamma"],
  ["delta", "epsilon", "zeta"],
];
assert.isTrue(patternSearch(grid, "epsilon"));
```

`patternSearch` should return `false` when the target does not exist in the grid.

```js
const grid = [
  ["alpha", "beta", "gamma"],
  ["delta", "epsilon", "zeta"],
];
assert.isFalse(patternSearch(grid, "theta"));
```

Your `patternSearch` function should use `for...of` loops.

```js
assert.match(
  __helpers.removeJSComments(patternSearch.toString()),
  /for\s*\(\s*(?:const|let|var)\s+\w+\s+of\s+/
);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
