---
id: 1e34744901a1eec0f5af186a
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

The escape room has a bonus challenge: a master control grid where you must locate a specific sequence of symbols. Unlike the previous pattern search, this grid is three levels deep — an array of arrays of arrays.

Create a function called `masterGridSearch` that accepts `grid` (a 3D array of strings) and `target` (a string). Use a manual flag variable called `found` initialized to `false`. Then, use three nested `for` loops to iterate through all three levels of the grid.

When you find the `target`, set `found = true` and use it to break out of each loop level. The function should return `true` if found and `false` otherwise.

For example:

```js
const grid3D = [
  [["a", "b"], ["c", "d"]],
  [["e", "f"], ["g", "h"]],
];
masterGridSearch(grid3D, "g"); // true
masterGridSearch(grid3D, "z"); // false
```

# --hints--

You should create a function named `masterGridSearch`.

```js
assert.isFunction(masterGridSearch);
```

Your `masterGridSearch` function should have `grid` and `target` parameters.

```js
const regex = __helpers.functionRegex('masterGridSearch', ['grid', 'target']);
assert.match(__helpers.removeJSComments(masterGridSearch.toString()), regex);
```

`masterGridSearch` should return `true` when the target is found in a 3D grid.

```js
const grid3D = [
  [["a", "b"], ["c", "d"]],
  [["e", "f"], ["g", "h"]],
];
assert.isTrue(masterGridSearch(grid3D, "g"));
```

`masterGridSearch` should return `false` when the target is not found in a 3D grid.

```js
const grid3D = [
  [["a", "b"], ["c", "d"]],
  [["e", "f"], ["g", "h"]],
];
assert.isFalse(masterGridSearch(grid3D, "z"));
```

Your `masterGridSearch` function should use a manual flag variable.

```js
assert.match(
  __helpers.removeJSComments(masterGridSearch.toString()),
  /(?:const|let|var)\s+found\s*=\s*false/
);
```

Your `masterGridSearch` function should use the flag to break out of the nested loops.

```js
assert.match(
  __helpers.removeJSComments(masterGridSearch.toString()),
  /if\s*\(\s*found\s*\)\s*break/
);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
  outer: for (const row of grid) {
    for (const symbol of row) {
      if (symbol === target) {
        progressState.patternSearch = true;
        console.log("Clue 2 solved: keypad unlocked");
        return true;
      }
    }
  }
  return false;
}

function inputValidator(code) {
  if (code.length < 4) return false;
  let i = 0;
  while (i < code.length) {
    if (!/\d/.test(code[i])) {
      return false;
    }
    i++;
  }
  progressState.inputValidator = true;
  console.log("Clue 3 solved: code validator unlocked");
  return true;
}

--fcc-editable-region--

--fcc-editable-region--
```
