---
id: 29dd9d1be7f4cdcf2015a322
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

The third clue involves a security code validator on the door's keypad. You need to check whether a given code meets specific criteria using a `while` loop.

Create a function called `inputValidator` that accepts a `code` parameter (a string). The function should use a `while` loop to iterate through each character of the code and verify that every character is a digit (`0`-`9`).

The loop should use an index variable starting at `0` and continue as long as the index is less than the length of `code`. If a non-digit character is found, the loop should short-circuit and the function should return `false` immediately. If all characters are digits, the function should return `true`.

# --hints--

You should create a function named `inputValidator`.

```js
assert.isFunction(inputValidator);
```

Your `inputValidator` function should have a `code` parameter.

```js
const regex = __helpers.functionRegex('inputValidator', ['code']);
assert.match(__helpers.removeJSComments(inputValidator.toString()), regex);
```

`inputValidator("12345")` should return `true`.

```js
assert.isTrue(inputValidator("12345"));
```

`inputValidator("123a5")` should return `false`.

```js
assert.isFalse(inputValidator("123a5"));
```

`inputValidator("")` should return `true`.

```js
assert.isTrue(inputValidator(""));
```

`inputValidator("abc")` should return `false`.

```js
assert.isFalse(inputValidator("abc"));
```

Your `inputValidator` function should use a `while` loop.

```js
assert.match(
  __helpers.removeJSComments(inputValidator.toString()),
  /while\s*\(/);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
  outer: for (const row of grid) {
    for (const symbol of row) {
      if (symbol === target) {
        progressState.patternSearch = true;
        console.log("Clue 2 solved: keypad unlocked");
        return true;
      }
    }
  }
  return false;
}

--fcc-editable-region--

--fcc-editable-region--
```
