---
id: 49bd4e328738648e76bf2116
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now, update your `inputValidator` function to integrate with the escape room progress. When the code passes validation (all digits), the function should set `progressState.inputValidator = true`, log `"Clue 3 solved: code validator unlocked"` to the console, and return `true`.

Additionally, add a second validation rule: the code must be at least `4` characters long. If the code is shorter than `4` characters, return `false` immediately without modifying `progressState`.

# --hints--

`inputValidator("12345")` should return `true`.

```js
progressState.inputValidator = false;
assert.isTrue(inputValidator("12345"));
```

`inputValidator("123a5")` should return `false`.

```js
progressState.inputValidator = false;
assert.isFalse(inputValidator("123a5"));
```

`inputValidator("123")` (shorter than 4 characters) should return `false`.

```js
progressState.inputValidator = false;
assert.isFalse(inputValidator("123"));
```

`inputValidator("0000")` should return `true`.

```js
progressState.inputValidator = false;
assert.isTrue(inputValidator("0000"));
```

When the code passes validation, `progressState.inputValidator` should be set to `true`.

```js
progressState.inputValidator = false;
inputValidator("12345");
assert.isTrue(progressState.inputValidator);
```

When the code passes validation, it should print `"Clue 3 solved: code validator unlocked"`.

```js
const log = [];
const originalLog = console.log;
console.log = (...args) => log.push(args.join(' '));
progressState.inputValidator = false;
inputValidator("12345");
console.log = originalLog;
assert.include(log.join(' '), 'Clue 3 solved: code validator unlocked');
```

When the code fails validation, `progressState.inputValidator` should remain `false`.

```js
progressState.inputValidator = false;
inputValidator("abc");
assert.isFalse(progressState.inputValidator);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
  outer: for (const row of grid) {
    for (const symbol of row) {
      if (symbol === target) {
        progressState.patternSearch = true;
        console.log("Clue 2 solved: keypad unlocked");
        return true;
      }
    }
  }
  return false;
}

function inputValidator(code) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
