---
id: 7a69bf16d9213280b093e0fb
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now it's time to bring everything together. Create a function called `tryUnlockDoor` that checks whether all three clues have been solved.

The function should accept the `progressState` object as a parameter and return `true` only if `progressState.countdown`, `progressState.patternSearch`, and `progressState.inputValidator` are all `true`. If any clue is still unsolved, the function should return `false`.

Use a `for...in` loop to iterate over the properties of `progressState`. If any property is `false`, return `false` immediately (short-circuit). If the loop completes without finding a `false` value, return `true`.

# --hints--

You should create a function named `tryUnlockDoor`.

```js
assert.isFunction(tryUnlockDoor);
```

Your `tryUnlockDoor` function should have a `progress` parameter.

```js
const regex = __helpers.functionRegex('tryUnlockDoor', ['progress']);
assert.match(__helpers.removeJSComments(tryUnlockDoor.toString()), regex);
```

`tryUnlockDoor` should return `true` when all clues are solved.

```js
const state = { countdown: true, patternSearch: true, inputValidator: true };
assert.isTrue(tryUnlockDoor(state));
```

`tryUnlockDoor` should return `false` when any clue is unsolved.

```js
const state = { countdown: true, patternSearch: true, inputValidator: false };
assert.isFalse(tryUnlockDoor(state));
```

`tryUnlockDoor` should return `false` when all clues are unsolved.

```js
const state = { countdown: false, patternSearch: false, inputValidator: false };
assert.isFalse(tryUnlockDoor(state));
```

Your `tryUnlockDoor` function should use a `for...in` loop.

```js
assert.match(
  __helpers.removeJSComments(tryUnlockDoor.toString()),
  /for\s*\(\s*(?:const|let|var)\s+\w+\s+in\s+/
);
```

Your `tryUnlockDoor` function should short-circuit when a `false` value is found.

```js
assert.match(
  __helpers.removeJSComments(tryUnlockDoor.toString()),
  /return\s+false/
);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
  outer: for (const row of grid) {
    for (const symbol of row) {
      if (symbol === target) {
        progressState.patternSearch = true;
        console.log("Clue 2 solved: keypad unlocked");
        return true;
      }
    }
  }
  return false;
}

function inputValidator(code) {
  if (code.length < 4) return false;
  let i = 0;
  while (i < code.length) {
    if (!/\d/.test(code[i])) {
      return false;
    }
    i++;
  }
  progressState.inputValidator = true;
  console.log("Clue 3 solved: code validator unlocked");
  return true;
}

function masterGridSearch(grid, target) {
  let found = false;
  for (let i = 0; i < grid.length; i++) {
    for (let j = 0; j < grid[i].length; j++) {
      for (let k = 0; k < grid[i][j].length; k++) {
        if (grid[i][j][k] === target) {
          found = true;
        }
        if (found) break;
      }
      if (found) break;
    }
    if (found) break;
  }
  return found;
}

--fcc-editable-region--

--fcc-editable-region--
```
