---
id: 7d46ab547a36b668cb3ca12a
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Your first clue involves a countdown mechanism on the door's control panel. The panel displays a timer that counts down from a given number to `1`. When the timer reaches `1`, the first lock disengages.

Create a function called `countdown` that accepts a `start` parameter. Inside the function, use a `for` loop to iterate from `start` down to `1`. During each iteration, add the current number to an array called `steps`. Once the loop finishes, return the `steps` array.

For example, `countdown(5)` should return `[5, 4, 3, 2, 1]`.

# --hints--

You should create a function named `countdown`.

```js
assert.isFunction(countdown);
```

Your `countdown` function should have a `start` parameter.

```js
const regex = __helpers.functionRegex('countdown', ['start']);
assert.match(__helpers.removeJSComments(countdown.toString()), regex);
```

`countdown(5)` should return `[5, 4, 3, 2, 1]`.

```js
assert.deepEqual(countdown(5), [5, 4, 3, 2, 1]);
```

`countdown(3)` should return `[3, 2, 1]`.

```js
assert.deepEqual(countdown(3), [3, 2, 1]);
```

Your `countdown` function should use a `for` loop.

```js
assert.match(__helpers.removeJSComments(countdown.toString()), /for\s*\(/);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

--fcc-editable-region--

--fcc-editable-region--
```
