---
id: 87b4c57502b3aae30fe06f4d
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Now, update your `patternSearch` function to short-circuit when it finds the target. When the target symbol is found, the function should immediately set `progressState.patternSearch = true`, log `"Clue 2 solved: keypad unlocked"` to the console, and return `true`.

Additionally, you need to break out of the inner loop without using a `break` statement — instead, use a labeled statement on the outer loop. Label the outer loop `outer:` and use `break outer;` to exit both loops when the target is found.

# --hints--

`patternSearch` should return `true` when the target is found and short-circuit immediately.

```js
const grid = [
  ["a", "b", "c"],
  ["d", "e", "f"],
  ["g", "h", "i"],
];
assert.isTrue(patternSearch(grid, "e"));
```

`patternSearch` should return `false` when the target is not found.

```js
const grid = [
  ["a", "b", "c"],
  ["d", "e", "f"],
];
assert.isFalse(patternSearch(grid, "z"));
```

When the target is found, `progressState.patternSearch` should be set to `true`.

```js
progressState.patternSearch = false;
const grid = [["a", "b"], ["c", "d"]];
patternSearch(grid, "c");
assert.isTrue(progressState.patternSearch);
```

When the target is found, it should print `"Clue 2 solved: keypad unlocked"`.

```js
const log = [];
const originalLog = console.log;
console.log = (...args) => log.push(args.join(' '));
progressState.patternSearch = false;
const grid = [["a", "b"], ["c", "d"]];
patternSearch(grid, "c");
console.log = originalLog;
assert.include(log.join(' '), 'Clue 2 solved: keypad unlocked');
```

Your `patternSearch` function should use a labeled loop (`outer:`) to break out of the nested loop.

```js
assert.match(
  __helpers.removeJSComments(patternSearch.toString()),
  /outer\s*:\s*for/
);
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
