---
id: a3e6f67e80752513c0625585
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

It's time for the final escape! Create a function called `escapeRoom` that orchestrates the entire escape room experience.

The function should:

1. Accept `grid` (a 2D array of strings), `target` (a string), and `code` (a string) as parameters.
2. Call `countdown(10)` to solve clue 1.
3. Call `patternSearch(grid, target)` to solve clue 2.
4. Call `inputValidator(code)` to solve clue 3.
5. Call `tryUnlockDoor(progressState)` to check if the door can be unlocked.
6. If `tryUnlockDoor` returns `true`, log `"Congratulations! You escaped the room!"` and return `true`.
7. If `tryUnlockDoor` returns `false`, log `"The door remains locked. Solve all clues to escape."` and return `false`.

Finally, call your `escapeRoom` function at the bottom of the file with the following test data:

```js
const symbolGrid = [
  ["alpha", "beta", "gamma"],
  ["delta", "epsilon", "zeta"],
];

escapeRoom(symbolGrid, "epsilon", "4321");
```

# --hints--

You should create a function named `escapeRoom`.

```js
assert.isFunction(escapeRoom);
```

Your `escapeRoom` function should call `countdown(10)`.

```js
assert.match(
  __helpers.removeJSComments(escapeRoom.toString()),
  /countdown\s*\(\s*10\s*\)/
);
```

Your `escapeRoom` function should call `patternSearch(grid, target)`.

```js
assert.match(
  __helpers.removeJSComments(escapeRoom.toString()),
  /patternSearch\s*\(\s*grid\s*,\s*target\s*\)/
);
```

Your `escapeRoom` function should call `inputValidator(code)`.

```js
assert.match(
  __helpers.removeJSComments(escapeRoom.toString()),
  /inputValidator\s*\(\s*code\s*\)/
);
```

Your `escapeRoom` function should call `tryUnlockDoor(progressState)`.

```js
assert.match(
  __helpers.removeJSComments(escapeRoom.toString()),
  /tryUnlockDoor\s*\(\s*progressState\s*\)/
);
```

`escapeRoom` should print `"Congratulations! You escaped the room!"` when all clues are solved.

```js
const log = [];
const originalLog = console.log;
console.log = (...args) => log.push(args.join(' '));
progressState.countdown = false;
progressState.patternSearch = false;
progressState.inputValidator = false;
const symbolGrid = [
  ["alpha", "beta", "gamma"],
  ["delta", "epsilon", "zeta"],
];
escapeRoom(symbolGrid, "epsilon", "4321");
console.log = originalLog;
assert.include(log.join(' '), 'Congratulations! You escaped the room!');
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
  outer: for (const row of grid) {
    for (const symbol of row) {
      if (symbol === target) {
        progressState.patternSearch = true;
        console.log("Clue 2 solved: keypad unlocked");
        return true;
      }
    }
  }
  return false;
}

function inputValidator(code) {
  if (code.length < 4) return false;
  let i = 0;
  while (i < code.length) {
    if (!/\d/.test(code[i])) {
      return false;
    }
    i++;
  }
  progressState.inputValidator = true;
  console.log("Clue 3 solved: code validator unlocked");
  return true;
}

function masterGridSearch(grid, target) {
  let found = false;
  for (let i = 0; i < grid.length; i++) {
    for (let j = 0; j < grid[i].length; j++) {
      for (let k = 0; k < grid[i][j].length; k++) {
        if (grid[i][j][k] === target) {
          found = true;
        }
        if (found) break;
      }
      if (found) break;
    }
    if (found) break;
  }
  return found;
}

function tryUnlockDoor(progress) {
  for (const key in progress) {
    if (!progress[key]) {
      return false;
    }
  }
  return true;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
  if (start <= 0) return false;
  for (let i = start; i >= 1; i--) {
    if (i === 1) {
      progressState.countdown = true;
      console.log("Clue 1 solved: countdown unlocked");
      return true;
    }
  }
}

function patternSearch(grid, target) {
  outer: for (const row of grid) {
    for (const symbol of row) {
      if (symbol === target) {
        progressState.patternSearch = true;
        console.log("Clue 2 solved: keypad unlocked");
        return true;
      }
    }
  }
  return false;
}

function inputValidator(code) {
  if (code.length < 4) return false;
  let i = 0;
  while (i < code.length) {
    if (!/\d/.test(code[i])) {
      return false;
    }
    i++;
  }
  progressState.inputValidator = true;
  console.log("Clue 3 solved: code validator unlocked");
  return true;
}

function masterGridSearch(grid, target) {
  let found = false;
  for (let i = 0; i < grid.length; i++) {
    for (let j = 0; j < grid[i].length; j++) {
      for (let k = 0; k < grid[i][j].length; k++) {
        if (grid[i][j][k] === target) {
          found = true;
        }
        if (found) break;
      }
      if (found) break;
    }
    if (found) break;
  }
  return found;
}

function tryUnlockDoor(progress) {
  for (const key in progress) {
    if (!progress[key]) {
      return false;
    }
  }
  return true;
}

function escapeRoom(grid, target, code) {
  countdown(10);
  patternSearch(grid, target);
  inputValidator(code);

  if (tryUnlockDoor(progressState)) {
    console.log("Congratulations! You escaped the room!");
    return true;
  } else {
    console.log("The door remains locked. Solve all clues to escape.");
    return false;
  }
}

const symbolGrid = [
  ["alpha", "beta", "gamma"],
  ["delta", "epsilon", "zeta"],
];

escapeRoom(symbolGrid, "epsilon", "4321");
```
