---
id: c7199cf412a0b40c35102419
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now, modify your `countdown` function so it short-circuits when it reaches `1`. Instead of collecting all steps into an array, the function should return `true` immediately once the loop counter equals `1`, indicating the countdown clue has been solved.

When the function returns `true`, you should also set `progressState.countdown = true` and log the message `"Clue 1 solved: countdown unlocked"` to the console.

If `start` is `0` or a negative number, the function should return `false` without modifying `progressState`.

# --hints--

`countdown(5)` should return `true`.

```js
const prevState = { ...progressState };
const result = countdown(5);
assert.isTrue(result);
```

`countdown(1)` should return `true`.

```js
const prevState = { ...progressState };
const result = countdown(1);
assert.isTrue(result);
```

`countdown(0)` should return `false`.

```js
const prevState = { ...progressState };
const result = countdown(0);
assert.isFalse(result);
```

`countdown(-3)` should return `false`.

```js
const prevState = { ...progressState };
const result = countdown(-3);
assert.isFalse(result);
```

When `countdown` returns `true`, `progressState.countdown` should be set to `true`.

```js
progressState.countdown = false;
countdown(5);
assert.isTrue(progressState.countdown);
```

When `countdown` returns `true`, it should print `"Clue 1 solved: countdown unlocked"`.

```js
const log = [];
const originalLog = console.log;
console.log = (...args) => log.push(args.join(' '));
progressState.countdown = false;
countdown(5);
console.log = originalLog;
assert.include(log.join(' '), 'Clue 1 solved: countdown unlocked');
```

# --seed--

## --seed-contents--

```js
const progressState = {
  countdown: false,
  patternSearch: false,
  inputValidator: false,
};

function countdown(start) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
