---
id: fad43da5927d9e5a3079ee87
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

Welcome to the Escape Room Code Runner workshop! You have been locked inside a mysterious room and must solve a series of loop-driven puzzles to unlock the door and escape.

In this workshop, you will build three clue-solving modules — `countdown`, `patternSearch`, and `inputValidator` — each using a different loop construct. You will also maintain a shared `progressState` object that tracks which clues have been solved.

To begin, create a `progressState` object with three boolean properties: `countdown`, `patternSearch`, and `inputValidator`, all initialized to `false`.

# --hints--

You should create a variable named `progressState`.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+progressState\s*=/);
```

`progressState` should be an object.

```js
assert.isObject(progressState);
```

`progressState` should have a `countdown` property set to `false`.

```js
assert.strictEqual(progressState.countdown, false);
```

`progressState` should have a `patternSearch` property set to `false`.

```js
assert.strictEqual(progressState.patternSearch, false);
```

`progressState` should have an `inputValidator` property set to `false`.

```js
assert.strictEqual(progressState.inputValidator, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
