---
id: 6a036fa04efd843ffc4bf573
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

To manually trigger an error, we use the `throw` keyword. 

Inside your `try` block, use `throw` to create a new `Error` with the message `"Invalid age"` if the `age` parameter is less than 0.

# --hints--

You should throw a new `Error` with the message `"Invalid age"`.

```js
try {
  checkAge(-1);
} catch (e) {
  assert.equal(e.message, "Invalid age");
}
```

# --seed--
## --seed-contents--

```js
--fcc-editable-region--
function checkAge(age) {
  try {

  } catch (error) {
    console.log(error.message);
  }
}
--fcc-editable-region--
```
