---
id: 6a036fa04efd843ffc4bf572
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Now that you have your function, it's time to add the error-handling structure. 

A `try` block contains the code that might throw an error, while the `catch` block contains the code to execute if an error occurs.

Inside your `checkAge` function, add an empty `try` block followed by a `catch` block that takes `error` as a parameter.

# --hints--

Your `checkAge` function should contain a `try...catch` statement.

```js
const checkAgeStr = checkAge.toString();
assert.match(checkAgeStr, /try\s*\{[\s\S]*\}\s*catch\s*\(\s*error\s*\)\s*\{[\s\S]*\}/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function checkAge() {

}
--fcc-editable-region--
```
