---
id: 6a036fa04efd843ffc4bf571
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

Welcome to the try/catch workshop! In JavaScript, runtime errors can completely crash your application. To handle these gracefully, we use `try...catch` blocks.

To get started, let's create a scenario. Declare an empty function named `checkAge`. We will use this function throughout this workshop to practice throwing and catching custom errors.

# --hints--

You should declare a function named `checkAge`.

```js
assert.isFunction(checkAge);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
