---
id: 6a036fa04efd843ffc4bf574
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

The `finally` block is the last part of error handling. It will execute regardless of whether an error was thrown or caught. This is often used for "cleanup" tasks.

Add a `finally` block after your `catch` block that logs the string `"Validation attempt complete"` to the console.

# --hints--

You should have a `finally` block at the end of your `try...catch` statement.

```js
const checkAgeStr = checkAge.toString();
assert.match(checkAgeStr, /finally\s*\{[\s\S]*\}/);
```

You should log `"Validation attempt complete"` in the `finally` block.

```js
let logs = [];
const oldLog = console.log;
console.log = (msg) => logs.push(msg);
checkAge(25);
console.log = oldLog;
assert.include(logs, "Validation attempt complete");
```

# --seed--
## --seed-contents--

```js
--fcc-editable-region--
function checkAge(age) {
  try {
    if (age < 0) {
      throw new Error("Invalid age");
    }
  } catch (error) {
    console.log(error.message);
  } finally {

  }
}
--fcc-editable-region--
```
