---
id: 657732654845d8e2fb1217e6
title: Schritt 62
challengeType: 20
dashedName: step-62
---

# --description--

`all()` is a built-in Python function that returns `True` if all the elements inside a given iterable evaluate to `True`. Otherwise, it returns `False`.

You can combine the `all()` function with the list comprehension syntax to make your code more concise.

Modify your `if` condition by passing the list comprehension you created in the previous step to the `all()` function.

# --hints--

You should pass your list comprehension to the `all()` function.

```js
({ test: () => assert.match(code, /^(\s*)if\s+all\s*\(\s*\[\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\]\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        if [constraint <= len(re.findall(pattern, password))
        for constraint, pattern in constraints]:
--fcc-editable-region--        
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
