---
id: 6577320da0d4c2e594d418e2
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Pass `right_part` as the argument to the `merge_sort()` call you added in the last step.

# --hints--

You should pass `right_part` as the argument to your last `merge_sort()` call.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*right_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort()
--fcc-editable-region--
```
