---
id: 68af274e60e009597dd861e4
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

In the previous steps, you created a basic movie object. Now, you'll add some validation to make sure that movies cannot be instantiated with invalid attributes.

At the beginning of the `__init__` method, create an `if` statement that checks if the `title` is either empty or contains whitespace only (you can use the `strip` method to build the condition).

If so, raise a `ValueError` with the message `Title cannot be empty`.

# --hints--

You should raise a `ValueError` with the message `Title cannot be empty` when `title` is either empty or contains whitespace only.

```js
({
  test: () => runPython(`
  _msg = 'Title cannot be empty'
  _expected_error = f"Expected to raise 'ValueError: {_msg}'"
  try:
    Movie(' ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error

  try:
    Movie('   ', 1999, 'The Wachowskis', 136)
  except ValueError as e:
    assert str(e) == _msg
  else:
    assert False, _expected_error
  `)
})
```

You should not raise a `ValueError` when `title` is a valid string.

```js
({
    test: () => runPython(`
    try:
      Movie('The Matrix', 1999, 'The Wachowskis', 136)
    except ValueError:
      assert False, 'Expected not to raise a ValueError with valid title'
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):
        
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
