---
id: 68af274e60e009597dd861e4
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

In the previous steps, you created a basic movie object. Now, you'll add some validation to make sure that movies cannot be instantiated with invalid attributes.

At the beginning of the `__init__` method, create an `if` statement that checks if the `title` is empty (you can use the `strip` method to build the condition).

If so, raise a `ValueError` with the message `Title cannot be empty`.

# --hints--

T

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Movie").has_function("__str__")`)
})
```

T

```js
({
    test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
    print(movie1)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "The Matrix (1999) - 136 min, The Wachowskis" in output
    `)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
--fcc-editable-region--
    def __init__(self, title, year, director, duration):

        self.title = title
        self.year = year
        self.director = director
        self.duration = duration
--fcc-editable-region--
    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
print(movie1)
```
