---
id: 68af274560e009597dd861e3
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Now, create an instance of `Movie` passing in arguments of your choice and assign it to a variable named `movie1`. Then, print `movie1` to the console.

# --hints--

You should have a variable named `movie1`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("movie1")`)
})
```

You should assign an instance of `Movie` to the `movie1` variable.

```js
({
    test: () => runPython(`
    assert isinstance(movie1, Movie)
    `)
})
```

You should print `movie1` to the console.

```js
({
    test: () => runPython(`assert _Node(_code).has_call("print(movie1)")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'
--fcc-editable-region--

--fcc-editable-region--
```
