---
id: 68b704862dd6559e890c2a86
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

After the `if` statement, create a variable named `result`. This will be the string to return from the method.

So assign it an f-string with the format `Media Catalogue (N items):`. Replace `N` with the number of items in the `items` list and make the string end with two newline characters.

# --hints--

You should have a variable named `result` in your `__str__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_variable("result")`)
})
```

You should assign `f'Media Catalogue ({len(self.items)} items):\n\n'` to the `result` variable.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("MediaCatalogue").find_function("__str__").has_stmt("result = f'Media Catalogue ({len(self.items)} items):\\\\n\\\\n'")`)
})
```

# --seed--

## --seed-contents--

```py
class Movie:
    def __init__(self, title, year, director, duration):
        if not title.strip():
            raise ValueError('Title cannot be empty')
        if year < 1895:
            raise ValueError('Year must be 1895 or later')
        if not director.strip():
            raise ValueError('Director cannot be empty')
        if duration <= 0:
            raise ValueError('Duration must be positive')
        self.title = title
        self.year = year
        self.director = director
        self.duration = duration

    def __str__(self):
        return f'{self.title} ({self.year}) - {self.duration} min, {self.director}'

class MediaCatalogue:
    def __init__(self):
        self.items = []

    def add(self, media_item):
        self.items.append(media_item)
--fcc-editable-region--
    def __str__(self):
        if not self.items:
            return 'Media Catalogue (empty)'
        
--fcc-editable-region--
try:
    movie1 = Movie('The Matrix', 1999, 'The Wachowskis', 136)
except ValueError as e:
    print(f'Validation Error: {e}')
```
