---
id: 5a23c84252665b21eecc803c
title: Sudoku
challengeType: 5
forumTopicId: 302329
dashedName: sudoku
---

# --description--

Write a function to solve a partially filled-in normal 9x9 [Sudoku](https://en.wikipedia.org/wiki/Sudoku) grid and return the result. The blank fields are represented by 0s. [Algorithmics of Sudoku](https://en.wikipedia.org/wiki/Algorithmics_of_sudoku) may help implement this.

# --hints--

`solveSudoku` should be a function.

```js
assert(typeof solveSudoku == 'function');
```

`solveSudoku([[8, 1, 9, -1, -1, 5, -1, -1, -1],[-1, -1, 2, -1, -1, -1, 7, 5, -1],[-1, 3, 7, 1, -1, 4, -1, 6, -1],[4, -1, -1, 5, 9, -1, 1, -1, -1],[7, -1, -1, 3, -1, 8, -1, -1, 2],[-1, -1, 3, -1, 6, 2, -1, -1, 7],[-1, 5, -1, 7, -1, 9, 2, 1, -1],[-1, 6, 4, -1, -1, -1, 9, -1, -1],[-1, -1, -1, 2, -1, -1, 4, 3, 8]])` should return an array.

```js
assert(
  Array.isArray(
    solveSudoku([
      [8, 1, 9, -1, -1, 5, -1, -1, -1],
      [-1, -1, 2, -1, -1, -1, 7, 5, -1],
      [-1, 3, 7, 1, -1, 4, -1, 6, -1],
      [4, -1, -1, 5, 9, -1, 1, -1, -1],
      [7, -1, -1, 3, -1, 8, -1, -1, 2],
      [-1, -1, 3, -1, 6, 2, -1, -1, 7],
      [-1, 5, -1, 7, -1, 9, 2, 1, -1],
      [-1, 6, 4, -1, -1, -1, 9, -1, -1],
      [-1, -1, -1, 2, -1, -1, 4, 3, 8]
    ])
  )
);
```

`solveSudoku([[8, 1, 9, -1, -1, 5, -1, -1, -1],[-1, -1, 2, -1, -1, -1, 7, 5, -1],[-1, 3, 7, 1, -1, 4, -1, 6, -1],[4, -1, -1, 5, 9, -1, 1, -1, -1],[7, -1, -1, 3, -1, 8, -1, -1, 2],[-1, -1, 3, -1, 6, 2, -1, -1, 7],[-1, 5, -1, 7, -1, 9, 2, 1, -1],[-1, 6, 4, -1, -1, -1, 9, -1, -1],[-1, -1, -1, 2, -1, -1, 4, 3, 8]])` should return `[[8, 1, 9, 6, 7, 5, 3, 2, 4],[6, 4, 2, 9, 8, 3, 7, 5, 1],[5, 3, 7, 1, 2, 4, 8, 6, 9],[4, 2, 6, 5, 9, 7, 1, 8, 3],[7, 9, 5, 3, 1, 8, 6, 4, 2],[1, 8, 3, 4, 6, 2, 5, 9, 7],[3, 5, 8, 7, 4, 9, 2, 1, 6],[2, 6, 4, 8, 3, 1, 9, 7, 5],[9, 7, 1, 2, 5, 6, 4, 3, 8]]`.

```js
assert.deepEqual(
  solveSudoku([
    [8, 1, 9, -1, -1, 5, -1, -1, -1],
    [-1, -1, 2, -1, -1, -1, 7, 5, -1],
    [-1, 3, 7, 1, -1, 4, -1, 6, -1],
    [4, -1, -1, 5, 9, -1, 1, -1, -1],
    [7, -1, -1, 3, -1, 8, -1, -1, 2],
    [-1, -1, 3, -1, 6, 2, -1, -1, 7],
    [-1, 5, -1, 7, -1, 9, 2, 1, -1],
    [-1, 6, 4, -1, -1, -1, 9, -1, -1],
    [-1, -1, -1, 2, -1, -1, 4, 3, 8]
  ]),
  [
    [8, 1, 9, 6, 7, 5, 3, 2, 4],
    [6, 4, 2, 9, 8, 3, 7, 5, 1],
    [5, 3, 7, 1, 2, 4, 8, 6, 9],
    [4, 2, 6, 5, 9, 7, 1, 8, 3],
    [7, 9, 5, 3, 1, 8, 6, 4, 2],
    [1, 8, 3, 4, 6, 2, 5, 9, 7],
    [3, 5, 8, 7, 4, 9, 2, 1, 6],
    [2, 6, 4, 8, 3, 1, 9, 7, 5],
    [9, 7, 1, 2, 5, 6, 4, 3, 8]
  ]
);
```

`solveSudoku([[5, 3, -1, -1, 2, 4, 7, -1, -1],[-1, -1, 2, -1, -1, -1, 8, -1, -1],[1, -1, -1, 7, -1, 3, 9, -1, 2],[-1, -1, 8, -1, 7, 2, -1, 4, 9],[-1, 2, -1, 9, 8, -1, -1, 7, -1],[7, 9, -1, -1, -1, -1, -1, 8, -1],[-1, -1, -1, -1, 3, -1, 5, -1, 6],[9, 6, -1, -1, 1, -1, 3, -1, -1],[-1, 5, -1, 6, 9, -1, -1, 1, -1]])` should return `[[5, 3, 9, 8, 2, 4, 7, 6, 1],[6, 7, 2, 1, 5, 9, 8, 3, 4],[1, 8, 4, 7, 6, 3, 9, 5, 2],[3, 1, 8, 5, 7, 2, 6, 4, 9],[4, 2, 5, 9, 8, 6, 1, 7, 3],[7, 9, 6, 3, 4, 1, 2, 8, 5],[8, 4, 1, 2, 3, 7, 5, 9, 6],[9, 6, 7, 4, 1, 5, 3, 2, 8],[2, 5, 3, 6, 9, 8, 4, 1, 7]]`.

```js
assert.deepEqual(
  solveSudoku([
    [5, 3, -1, -1, 2, 4, 7, -1, -1],
    [-1, -1, 2, -1, -1, -1, 8, -1, -1],
    [1, -1, -1, 7, -1, 3, 9, -1, 2],
    [-1, -1, 8, -1, 7, 2, -1, 4, 9],
    [-1, 2, -1, 9, 8, -1, -1, 7, -1],
    [7, 9, -1, -1, -1, -1, -1, 8, -1],
    [-1, -1, -1, -1, 3, -1, 5, -1, 6],
    [9, 6, -1, -1, 1, -1, 3, -1, -1],
    [-1, 5, -1, 6, 9, -1, -1, 1, -1]
  ]),
  [
    [5, 3, 9, 8, 2, 4, 7, 6, 1],
    [6, 7, 2, 1, 5, 9, 8, 3, 4],
    [1, 8, 4, 7, 6, 3, 9, 5, 2],
    [3, 1, 8, 5, 7, 2, 6, 4, 9],
    [4, 2, 5, 9, 8, 6, 1, 7, 3],
    [7, 9, 6, 3, 4, 1, 2, 8, 5],
    [8, 4, 1, 2, 3, 7, 5, 9, 6],
    [9, 6, 7, 4, 1, 5, 3, 2, 8],
    [2, 5, 3, 6, 9, 8, 4, 1, 7]
  ]
);
```

`solveSudoku([[-1, -1, 3, -1, 2, -1, 6, -1, -1],[9, -1, -1, 3, -1, 5, -1, -1, 1],[-1, -1, 1, 8, -1, 6, 4, -1, -1],[-1, -1, 8, 1, -1, 2, 9, -1, -1],[7, -1, -1, -1, -1, -1, -1, -1, 8],[-1, -1, 6, 7, -1, 8, 2, -1, -1],[-1, -1, 2, 6, -1, 9, 5, -1, -1],[8, -1, -1, 2, -1, 3, -1, -1, 9],[-1, -1, 5, -1, 1, -1, 3, -1, -1]])` should return `[[4, 8, 3, 9, 2, 1, 6, 5, 7],[9, 6, 7, 3, 4, 5, 8, 2, 1],[2, 5, 1, 8, 7, 6, 4, 9, 3],[5, 4, 8, 1, 3, 2, 9, 7, 6],[7, 2, 9, 5, 6, 4, 1, 3, 8],[1, 3, 6, 7, 9, 8, 2, 4, 5],[3, 7, 2, 6, 8, 9, 5, 1, 4],[8, 1, 4, 2, 5, 3, 7, 6, 9],[6, 9, 5, 4, 1, 7, 3, 8, 2]]`.

```js
assert.deepEqual(
  solveSudoku([
    [-1, -1, 3, -1, 2, -1, 6, -1, -1],
    [9, -1, -1, 3, -1, 5, -1, -1, 1],
    [-1, -1, 1, 8, -1, 6, 4, -1, -1],
    [-1, -1, 8, 1, -1, 2, 9, -1, -1],
    [7, -1, -1, -1, -1, -1, -1, -1, 8],
    [-1, -1, 6, 7, -1, 8, 2, -1, -1],
    [-1, -1, 2, 6, -1, 9, 5, -1, -1],
    [8, -1, -1, 2, -1, 3, -1, -1, 9],
    [-1, -1, 5, -1, 1, -1, 3, -1, -1]
  ]),
  [
    [4, 8, 3, 9, 2, 1, 6, 5, 7],
    [9, 6, 7, 3, 4, 5, 8, 2, 1],
    [2, 5, 1, 8, 7, 6, 4, 9, 3],
    [5, 4, 8, 1, 3, 2, 9, 7, 6],
    [7, 2, 9, 5, 6, 4, 1, 3, 8],
    [1, 3, 6, 7, 9, 8, 2, 4, 5],
    [3, 7, 2, 6, 8, 9, 5, 1, 4],
    [8, 1, 4, 2, 5, 3, 7, 6, 9],
    [6, 9, 5, 4, 1, 7, 3, 8, 2]
  ]
);
```

# --seed--

## --seed-contents--

```js
function solveSudoku(puzzle) {

}
```

# --solutions--

```js
const solveSudoku = puzzle => {
    const isValid = (puzzle, row, col, k) => {
        for (let i = 0; i < 9; i++)
            if (puzzle[row][i] === k 
            || puzzle[i][col] === k 
            || puzzle[3 * Math.floor(row / 3) + Math.floor(i / 3)][3 * Math.floor(col / 3) + i % 3] === k) 
                return false;
        return true;
    }
    for (let i = 0; i < 9; i++) 
        for (let j = 0; j < 9; j++) 
            if (puzzle[i][j] === -1) {
                for (let k = 1; k <= 9; k++) 
                    if (isValid(puzzle, i, j, k)) {
                        puzzle[i][j] = k;
                        if (solveSudoku(puzzle)) 
                            return puzzle;
                        else 
                            puzzle[i][j] = -1;
                    
                    }    
                return false;
            }  
    return puzzle;
}
```
