---
id: 68e52994bc7ea2e3dec530a6
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

Now let's test the complete discount system! Add the `if __name__ == "__main__":` pattern and move `product = Product("Wireless Mouse", 50.0)` inside it.

After that, inside the `if __name__ == "__main__":` block, create a product, set the user tier to `"Premium"`, create a list of strategies containing `PercentageDiscount(10)`, `FixedAmountDiscount(5)`, and `PremiumUserDiscount()`.

Then, create a `DiscountEngine` instance, calculate the best price and store it in a variable named `best_price`.

# --hints--

You should use the `if __name__ == "__main__":` pattern.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent('__name__ == "__main__"')`)
})
```

You should move `product = Product("Wireless Mouse", 50.0)` inside the `if __name__ == "__main__":` block.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('product = Product("Wireless Mouse", 50.0)')`)
})
```

Within the `if __name__ == "__main__":` block, you should set `user_tier` to `"Premium"`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('user_tier = "Premium"')`)
})
```

You should create a list of strategies containing `PercentageDiscount(10)`, `FixedAmountDiscount(5)`, and `PremiumUserDiscount()`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('strategies = [PercentageDiscount(10), FixedAmountDiscount(5), PremiumUserDiscount()]')`)
})
```

You should create a `DiscountEngine` instance with the `strategies` list.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('engine = DiscountEngine(strategies)')`)
})
```

You should call `engine.calculate_best_price(product, user_tier)` and store the result in `best_price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('best_price = engine.calculate_best_price(product, user_tier)')
`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
from typing import List

class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price

    def __str__(self):
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent):
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount):
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.80

class DiscountEngine:
    def __init__(self, strategies: List[DiscountStrategy]):
        self.strategies = strategies

    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]

        for strategy in self.strategies:
            if strategy.is_applicable(product, user_tier):
                discounted = strategy.apply_discount(product)
                prices.append(discounted)

        return min(prices)
--fcc-editable-region--
product = Product("Wireless Mouse", 50.0)
--fcc-editable-region--
```
