---
id: 68e529226b0defdff26c4a8b
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now you'll start building the discount system. Python provides a way to define abstract classes that serve as blueprints for other classes. This is done using the `ABC` (Abstract Base Class) module.

At the top of your code, add the following import statement:

```py
from abc import ABC, abstractmethod
```

This imports the `ABC` class and the `abstractmethod` decorator, which you'll use to define methods that must be implemented by any class that inherits from your abstract class.

# --hints--

You should have an import statement at the top of your code.

```js
({
  test: () => runPython(`assert _Node(_code).has_import("abc")`)
})
```

You should import `ABC` from the `abc` module.

```js
({
  test: () => runPython(`assert _Node(_code).has_import_from("abc", "ABC")`)
})
```

You should import `abstractmethod` from the `abc` module.

```js
({
  test: () => runPython(`assert _Node(_code).has_import_from("abc", "abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price

    def __str__(self):
        return f'{self.name} - ${self.price}'

product = Product("Wireless Mouse", 50.0)
print(product)
```
