---
id: 68e529226b0defdff26c4a8a
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

To convert the default object representation into a readable string,create a `__str__` method in your `Product` class that returns a string with the format `name - $price` (where `name` and `price` should be replaced with the values of the corresponding attributes).

# --hints--

Your `Product` class should have a `__str__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Your `__str__` method should return a string in the format `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

When you print a `Product` with name `"Wireless Mouse"` and price `50.0`, it should display `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product("Wireless Mouse", 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--

--fcc-editable-region--
product = Product("Wireless Mouse", 50.0)
print(product)
```
