---
id: 68e5293bd00d2fe134f58984
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Now it's time to create a concrete implementation of your abstract class. Create a class named `PercentageDiscount` that inherits from `DiscountStrategy`.

Give it an `__init__` method that takes `self` and `percent` as parameters. Give a type hint of `int` to `percent` and set the return type of the method to `None`. Inside the `__init__` method, store `percent` as an instance attribute.

# --hints--

You should create a class named `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Your `PercentageDiscount` class should inherit from `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Your `PercentageDiscount` class should have an `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Your `__init__` method should have parameters `self` and `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Your `__init__` method should have a return type hint of `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

You should assign `percent` to `self.percent` inside your `__init__` method.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
