---
id: 68e524c94e2e4ab7b4bfc7e2
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you are going to build a discount calculator that can apply different discount strategies to products. The system will determine the best price for a customer based on multiple discount rules.

Start by defining a class named `Product`. Give your class an `__init__` method with the following parameters: `self`, `name`, and `price`. Inside the `__init__` method, assign `name` and `price` parameters to instance attributes with the same name.

# --hints--

You should create a class named `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Your `Product` class should have an `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Your `__init__` method should have the following parameters: `self`, `name`, and `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

You should assign `name` to `self.name` inside your `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

You should assign `price` to `self.price` inside your `__init__` method.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
