---
id: 68ef53ef276acfa40c83cb12
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Add a return type hint to the `is_applicable` method so that it clearly indicates it returns a `bool`. In Python, you do this by adding `-> bool` after the parameter list in the method definition.

Here is an example of a method with both parameter and return type hints:

```py
def example_method(self, value: int) -> bool:
  pass
```

# --hints--

Your `is_applicable` method should have a return type hint of `-> bool`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_returns("bool")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price

    def __str__(self):
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str):
        pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
