---
id: 68f119ad905b29ff59d858c8
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

Just like the previous step, in your `__str__` method, add a return type hint of `str` since this method returns a string representation of the object.

# --hints--

Your `__str__` method should have a return type hint of `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
