---
id: 66faa60000000000000000ab
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
A camera that scrolls is meaningless without a world to scroll across. Render 16 colored rectangles spaced along a 4000-pixel-wide world so the camera has something to chase.

Inside `MainScene.create`:

```js
const colors = [0x4dccbd, 0xff6b6b, 0xffd166, 0x6c8aff, 0xff9ce8];
for (let i = 0; i < 16; i++) {
  const x = 120 + i * 240;
  const y = 80 + (i % 4) * 120;
  const w = 60 + (i % 3) * 30;
  const h = 60 + ((i + 1) % 3) * 30;
  this.add.rectangle(x, y, w, h, colors[i % colors.length]).setDepth(0);
}
```

At this point only the leftmost rectangles are visible — the camera hasn't been told to follow anything yet.

# --hints--
You should declare a `colors` array with five hex values.

```js
assert.match(code, /const\s+colors\s*=\s*\[\s*0x4dccbd\s*,\s*0xff6b6b\s*,\s*0xffd166\s*,\s*0x6c8aff\s*,\s*0xff9ce8\s*\]/);
```

You should iterate `i` from `0` to `16` and call `this.add.rectangle` inside the loop.

```js
assert.match(code, /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*16\s*;[\s\S]*this\.add[\s\S]*\.rectangle\(/);
```

You should compute `x` as `120 + i * 240` so rectangles spread across the 4000px world.

```js
assert.match(code, /const\s+x\s*=\s*120\s*\+\s*i\s*\*\s*240/);
```

You should pick the colour with `colors[i % colors.length]`.

```js
assert.match(code, /colors\[\s*i\s*%\s*colors\.length\s*\]/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Camera Playground</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
--fcc-editable-region--
--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
