---
id: 66faa60000000000000000aa
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the camera playground. Across 25 steps you'll build a wide scrolling world, fly the camera through every knob Phaser exposes (`startFollow`, `setLerp`, `setDeadzone`, `shake`, `flash`, `fade`, `setZoom`), add a minimap, and ship a `prefers-reduced-motion`-aware accessibility toggle.

Declare `MainScene` extending `Phaser.Scene` with a constructor that calls `super("MainScene")` and empty `create` and `update(time, delta)` methods.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {}

  update(time, delta) {}
}
```

The `config` object below is already filled in — a 480×320 canvas with arcade physics and zero gravity.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["\']MainScene["\']\s*\)/);
```

You should declare a `create` method.

```js
assert.match(code, /\bcreate\s*\(\s*\)\s*\{/);
```

You should declare an `update` method that takes `time` and `delta`.

```js
assert.match(code, /\bupdate\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Camera Playground</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
