---
id: 66faa60000000000000000b3
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--
Lerp smooths the camera's reaction. A **deadzone** controls *whether* the camera reacts at all for small movements. Add a `Dead` button that toggles a 200×200 deadzone on and off. The actual `setDeadzone` call lives in an `applyDeadzone` closure so future code (the visualizer in step 11) can keep its visibility in sync without duplicating logic.

After the smooth toggle:

```js
this.applyDeadzone = () => {
  const w = this.deadzoneOn ? 200 : 0;
  const h = this.deadzoneOn ? 200 : 0;
  this.cameras.main.setDeadzone(w, h);
  if (this.deadzoneViz) this.deadzoneViz.setVisible(this.deadzoneOn);
};
this.deadzoneOn = false;
const deadBtn = this.makeButton(58, 4, "Dead: off", () => {
  this.deadzoneOn = !this.deadzoneOn;
  deadBtn.setLabel(this.deadzoneOn ? "Dead: on" : "Dead: off");
  this.applyDeadzone();
});
```

# --hints--
You should define `this.applyDeadzone = () => { ... }` that calls `this.cameras.main.setDeadzone(w, h)`.

```js
assert.match(code, /this\.applyDeadzone\s*=\s*\(\s*\)\s*=>\s*\{[\s\S]*this\.cameras\.main\.setDeadzone\(\s*w\s*,\s*h\s*\)/);
```

`applyDeadzone` should sync the visualizer with `if (this.deadzoneViz) this.deadzoneViz.setVisible(this.deadzoneOn)`.

```js
assert.match(code, /if\s*\(\s*this\.deadzoneViz\s*\)\s*this\.deadzoneViz\.setVisible\(\s*this\.deadzoneOn\s*\)/);
```

Add the `Dead: off` toggle via `this.makeButton(58, 4, "Dead: off", ...)`.

```js
assert.match(code, /this\.makeButton\(\s*58\s*,\s*4\s*,\s*["\']Dead: off["\']/);
```

The handler should toggle `this.deadzoneOn` and call `this.applyDeadzone()`.

```js
assert.match(code, /this\.deadzoneOn\s*=\s*!\s*this\.deadzoneOn[\s\S]*this\.applyDeadzone\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Camera Playground</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    const colors = [0x4dccbd, 0xff6b6b, 0xffd166, 0x6c8aff, 0xff9ce8];
    for (let i = 0; i < 16; i++) {
      const x = 120 + i * 240;
      const y = 80 + (i % 4) * 120;
      const w = 60 + (i % 3) * 30;
      const h = 60 + ((i + 1) % 3) * 30;
      this.add
        .rectangle(x, y, w, h, colors[i % colors.length])
        .setDepth(0);
    }
    this.player = this.add
      .rectangle(80, 160, 16, 16, 0xffffff)
      .setDepth(20);
    this.physics.add.existing(this.player);
    this.cameras.main.startFollow(this.player);
    this.physics.world.setBounds(0, 0, 4000, 600);
    this.player.body.setCollideWorldBounds(true);
    this.cameras.main.setBounds(0, 0, 4000, 600);
    this.uiLayer = this.add.layer();
    this.uiLayer.add(
      this.add
        .text(8, 300, "Arrows to move - try the buttons", {
          fontSize: "9px",
          color: "#cccccc"
        })
        .setScrollFactor(0)
        .setDepth(110)
    );
    this.makeButton = (x, y, label, onClick) => {
      const w = 50;
      const h = 16;
      const bg = this.add
        .rectangle(x, y, w, h, 0x222244, 1)
        .setOrigin(0)
        .setScrollFactor(0)
        .setDepth(100);
      const txt = this.add
        .text(x + w / 2, y + h / 2, label, {
          fontSize: "9px",
          color: "#ffffff"
        })
        .setOrigin(0.5)
        .setScrollFactor(0)
        .setDepth(101);
      bg.setInteractive().on("pointerdown", onClick);
      this.uiLayer.add([bg, txt]);
      return { bg, txt, setLabel: (s) => txt.setText(s) };
    };
    this.lerpOn = false;
    const smoothBtn = this.makeButton(4, 4, "Smooth: off", () => {
      this.lerpOn = !this.lerpOn;
      smoothBtn.setLabel(this.lerpOn ? "Smooth: on" : "Smooth: off");
      const k = this.lerpOn ? 0.1 : 1;
      this.cameras.main.setLerp(k, k);
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    if (!this.cursors) {
      this.cursors = this.input.keyboard.createCursorKeys();
    }
    const speed = 200;
    const body = this.player.body;
    body.setVelocity(0);
    if (this.cursors.left.isDown) body.setVelocityX(-speed);
    else if (this.cursors.right.isDown) body.setVelocityX(speed);
    if (this.cursors.up.isDown) body.setVelocityY(-speed);
    else if (this.cursors.down.isDown) body.setVelocityY(speed);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
