---
id: 66faa60000000000000000c0
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--
Toggling reduce-motion on every reload is annoying. Persist it to `localStorage` and read the saved value at scene create-time.

Replace the entire reduce-motion block — both the initial-value line and the button definition. The new version reads the persisted value, sets the initial label accordingly, and writes back on every toggle.

```js
this.reduceMotion = localStorage.getItem("cam:reduce") === "true";
this.reduceBtn = this.makeButton(
  328,
  4,
  this.reduceMotion ? "Motion: off" : "Motion: on",
  () => {
    this.reduceMotion = !this.reduceMotion;
    localStorage.setItem("cam:reduce", String(this.reduceMotion));
    this.reduceBtn.setLabel(
      this.reduceMotion ? "Motion: off" : "Motion: on"
    );
  }
);
```

Reload the preview and the toggle remembers its state.

# --hints--
Initialise `this.reduceMotion` from `localStorage.getItem("cam:reduce") === "true"`.

```js
assert.match(code, /this\.reduceMotion\s*=\s*localStorage\.getItem\(\s*["\']cam:reduce["\']\s*\)\s*===\s*["\']true["\']/);
```

The initial label (third arg of `makeButton`) should depend on the persisted value: `this.reduceMotion ? "Motion: off" : "Motion: on"`.

```js
assert.match(code, /makeButton\(\s*328\s*,\s*4\s*,\s*this\.reduceMotion\s*\?\s*["\']Motion: off["\']\s*:\s*["\']Motion: on["\']/);
```

The handler should call `localStorage.setItem("cam:reduce", String(this.reduceMotion))`.

```js
assert.match(code, /localStorage\.setItem\(\s*["\']cam:reduce["\']\s*,\s*String\(\s*this\.reduceMotion\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Camera Playground</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    const colors = [0x4dccbd, 0xff6b6b, 0xffd166, 0x6c8aff, 0xff9ce8];
    for (let i = 0; i < 16; i++) {
      const x = 120 + i * 240;
      const y = 80 + (i % 4) * 120;
      const w = 60 + (i % 3) * 30;
      const h = 60 + ((i + 1) % 3) * 30;
      this.add
        .rectangle(x, y, w, h, colors[i % colors.length])
        .setDepth(0);
    }
    this.player = this.add
      .rectangle(80, 160, 16, 16, 0xffffff)
      .setDepth(20);
    this.physics.add.existing(this.player);
    this.cameras.main.startFollow(this.player);
    this.physics.world.setBounds(0, 0, 4000, 600);
    this.player.body.setCollideWorldBounds(true);
    this.cameras.main.setBounds(0, 0, 4000, 600);
    this.uiLayer = this.add.layer();
    this.uiLayer.add(
      this.add
        .text(8, 300, "Arrows to move - try the buttons", {
          fontSize: "9px",
          color: "#cccccc"
        })
        .setScrollFactor(0)
        .setDepth(110)
    );
    this.makeButton = (x, y, label, onClick) => {
      const w = 50;
      const h = 16;
      const bg = this.add
        .rectangle(x, y, w, h, 0x222244, 1)
        .setOrigin(0)
        .setScrollFactor(0)
        .setDepth(100);
      const txt = this.add
        .text(x + w / 2, y + h / 2, label, {
          fontSize: "9px",
          color: "#ffffff"
        })
        .setOrigin(0.5)
        .setScrollFactor(0)
        .setDepth(101);
      bg.setInteractive().on("pointerdown", onClick);
      this.uiLayer.add([bg, txt]);
      return { bg, txt, setLabel: (s) => txt.setText(s) };
    };
    this.lerpOn = false;
    const smoothBtn = this.makeButton(4, 4, "Smooth: off", () => {
      this.lerpOn = !this.lerpOn;
      smoothBtn.setLabel(this.lerpOn ? "Smooth: on" : "Smooth: off");
      const k = this.lerpOn ? 0.1 : 1;
      this.cameras.main.setLerp(k, k);
    });
    this.applyDeadzone = () => {
      const w = this.deadzoneOn ? 200 : 0;
      const h = this.deadzoneOn ? 200 : 0;
      this.cameras.main.setDeadzone(w, h);
      if (this.deadzoneViz) this.deadzoneViz.setVisible(this.deadzoneOn);
    };
    this.deadzoneOn = false;
    const deadBtn = this.makeButton(58, 4, "Dead: off", () => {
      this.deadzoneOn = !this.deadzoneOn;
      deadBtn.setLabel(this.deadzoneOn ? "Dead: on" : "Dead: off");
      this.applyDeadzone();
    });
    this.deadzoneViz = this.add
      .rectangle(240, 160, 200, 200, 0x000000, 0)
      .setStrokeStyle(2, 0xff66cc)
      .setScrollFactor(0)
      .setDepth(60)
      .setVisible(false);
    this.uiLayer.add(this.deadzoneViz);
    this.applyDeadzone();
    this.bgIndex = 0;
    const bgColors = [0x101020, 0x1f3030, 0x2a1a3a, 0x000000];
    this.makeButton(274, 4, "BG", () => {
      this.bgIndex = (this.bgIndex + 1) % bgColors.length;
      this.cameras.main.setBackgroundColor(bgColors[this.bgIndex]);
    });
--fcc-editable-region--
    this.reduceMotion = false;
    this.reduceBtn = this.makeButton(
      328,
      4,
      "Motion: on",
      () => {
        this.reduceMotion = !this.reduceMotion;
        this.reduceBtn.setLabel(
          this.reduceMotion ? "Motion: off" : "Motion: on"
        );
      }
    );
--fcc-editable-region--
    this.shakeIntensity = 0.012;
    this.makeButton(112, 4, "Shake", () => {
      if (!this.reduceMotion) {
        this.cameras.main.shake(300, this.shakeIntensity);
      }
    });
    this.shakeLevel = 1;
    const shakeLevels = [0.005, 0.012, 0.025];
    const shakeNames = ["low", "med", "high"];
    const shakeLbl = this.makeButton(174, 24, "Shake: med", () => {
      this.shakeLevel = (this.shakeLevel + 1) % 3;
      this.shakeIntensity = shakeLevels[this.shakeLevel];
      shakeLbl.setLabel("Shake: " + shakeNames[this.shakeLevel]);
    });
    this.makeButton(166, 4, "Flash", () => {
      if (!this.reduceMotion) {
        this.cameras.main.flash(150, 255, 255, 255);
      }
    });
    this.makeButton(220, 4, "Fade", () => {
      if (this.reduceMotion) return;
      this.cameras.main.fade(600, 0, 0, 0);
      this.cameras.main.once("camerafadeoutcomplete", () => {
        this.cameras.main.fadeIn(400, 0, 0, 0);
      });
    });
    this.zoom = 1;
    this.applyZoom = (target) => {
      this.zoom = target;
      this.tweens.add({
        targets: this.cameras.main,
        zoom: target,
        duration: 400,
        ease: "Sine.InOut"
      });
      this.zoomLabel.setLabel("Zoom: " + target.toFixed(1));
    };
    this.minimap = this.cameras.add(340, 256, 130, 60);
    this.minimap.setBounds(0, 0, 4000, 600);
    this.minimap.setBackgroundColor(0x111133);
    this.minimap.startFollow(this.player);
    this.minimap.setZoom(0.06);
    this.minimap.ignore([this.uiLayer, this.player]);
    this.zoomLabel = this.makeButton(38, 24, "Zoom: 1.0", () => {});
    this.makeButton(4, 24, "-", () => this.applyZoom(this.zoom - 0.2));
    this.makeButton(92, 24, "+", () => this.applyZoom(this.zoom + 0.2));
  }

  update(time, delta) {
    if (!this.cursors) {
      this.cursors = this.input.keyboard.createCursorKeys();
    }
    const speed = 200;
    const body = this.player.body;
    body.setVelocity(0);
    if (this.cursors.left.isDown) body.setVelocityX(-speed);
    else if (this.cursors.right.isDown) body.setVelocityX(speed);
    if (this.cursors.up.isDown) body.setVelocityY(-speed);
    else if (this.cursors.down.isDown) body.setVelocityY(speed);
    if (!this.marker) {
      this.marker = this.add
        .rectangle(this.player.x, this.player.y, 80, 80, 0xffff66, 1)
        .setDepth(70);
      this.cameras.main.ignore(this.marker);
    }
    this.marker.setPosition(this.player.x, this.player.y);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
