---
id: 66faa60000000000000000ad
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
The player needs input. Bind cursor keys lazily inside `update` (creating them on the first frame, reusing them after) and translate the held arrow keys into velocity on the player's body. Lazy-init keeps the keybinding alongside the input that consumes it — easier to reason about than splitting it across `create` and `update`.

Inside `MainScene.update`:

```js
if (!this.cursors) {
  this.cursors = this.input.keyboard.createCursorKeys();
}
const speed = 200;
const body = this.player.body;
body.setVelocity(0);
if (this.cursors.left.isDown) body.setVelocityX(-speed);
else if (this.cursors.right.isDown) body.setVelocityX(speed);
if (this.cursors.up.isDown) body.setVelocityY(-speed);
else if (this.cursors.down.isDown) body.setVelocityY(speed);
```

The player should now glide around at 200 px/s. The camera still sits at `(0, 0)` though — it doesn't follow yet.

# --hints--
You should lazily create `this.cursors` with `this.input.keyboard.createCursorKeys()`.

```js
assert.match(code, /if\s*\(\s*!\s*this\.cursors\s*\)\s*\{\s*this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)\s*;\s*\}/);
```

You should declare `const speed = 200`.

```js
assert.match(code, /const\s+speed\s*=\s*200/);
```

You should zero the body's velocity with `body.setVelocity(0)`.

```js
assert.match(code, /body\.setVelocity\(\s*0\s*\)/);
```

You should map `cursors.left.isDown` to `body.setVelocityX(-speed)`.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*body\.setVelocityX\(\s*-\s*speed\s*\)/);
```

You should map `cursors.up.isDown` to `body.setVelocityY(-speed)`.

```js
assert.match(code, /this\.cursors\.up\.isDown[\s\S]*body\.setVelocityY\(\s*-\s*speed\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Camera Playground</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    const colors = [0x4dccbd, 0xff6b6b, 0xffd166, 0x6c8aff, 0xff9ce8];
    for (let i = 0; i < 16; i++) {
      const x = 120 + i * 240;
      const y = 80 + (i % 4) * 120;
      const w = 60 + (i % 3) * 30;
      const h = 60 + ((i + 1) % 3) * 30;
      this.add
        .rectangle(x, y, w, h, colors[i % colors.length])
        .setDepth(0);
    }
    this.player = this.add
      .rectangle(80, 160, 16, 16, 0xffffff)
      .setDepth(20);
    this.physics.add.existing(this.player);
  }

  update(time, delta) {
--fcc-editable-region--
--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
