---
id: 66faa60000000000000000ac
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
The camera needs a target to follow. Add a 16×16 white rectangle and register it with arcade physics — that gives it a `body` for velocity and bounds-clamping later.

After the rectangle loop in `create`:

```js
this.player = this.add
  .rectangle(80, 160, 16, 16, 0xffffff)
  .setDepth(20);
this.physics.add.existing(this.player);
```

# --hints--
You should create a 16×16 white rectangle at `(80, 160)` and store it on `this.player`.

```js
assert.match(code, /this\.player\s*=\s*this\.add\s*\.rectangle\(\s*80\s*,\s*160\s*,\s*16\s*,\s*16\s*,\s*0xffffff\s*\)/);
```

You should give the player `setDepth(20)`.

```js
assert.match(code, /\.setDepth\(\s*20\s*\)/);
```

You should register the player with arcade physics via `this.physics.add.existing(this.player)`.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.player\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Camera Playground</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  create() {
    const colors = [0x4dccbd, 0xff6b6b, 0xffd166, 0x6c8aff, 0xff9ce8];
    for (let i = 0; i < 16; i++) {
      const x = 120 + i * 240;
      const y = 80 + (i % 4) * 120;
      const w = 60 + (i % 3) * 30;
      const h = 60 + ((i + 1) % 3) * 30;
      this.add
        .rectangle(x, y, w, h, colors[i % colors.length])
        .setDepth(0);
    }
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
