---
id: 65688efcc78c9495e73acfc9
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Add two parameters to the `__init__` method, order matters:

- `self`: This is a reference to the instance of the class. It is a convention to name this parameter self.
- `board`: The board parameter is expected to be passed when creating an instance of the `Board` class.

# --hints--

You should add the parameter `self` and `board` to the method.

```js
assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*board\s*\):/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__():

--fcc-editable-region--
```
