---
id: 690ddb5e6ee94eb73ed172b0
title: Debug a Donation Form
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

A local charity has built a donation form website, but there are several issues that need to be fixed. The form isn't accessible and has some HTML syntax errors.

Your job is to fix all of the errors so the form works correctly and is accessible to all users. Complete the items in the user stories below and click "Run the Tests" to see if you fixed all the errors.

**User Stories:**

1. The `input` elements are void elements and should not have closing tags. Remove all `</input>` closing tags from the form.
2. Add `<label>` elements for each form input field so users know what each field is for. The label text should match what's currently next to each input.
3. Use the `for` attribute on each `<label>` element and the corresponding `id` attribute on each `<input>` element to explicitly associate them together:

    - Set the `for` attribute of the "Full Name:" label to `full-name` and the corresponding input's `id` to `full-name`.
    - Set the `for` attribute of the "Email Address:" label to `email` and the corresponding input's `id` to `email`.
    - Set the `for` attribute of the "Donation Amount ($):" label to `amount` and the corresponding input's `id` to `amount`.
    - Set the `for` attribute of the "Subscribe" label to `newsletter` and the corresponding input's `id` to `newsletter`.

4. Add the `required` attribute to the text, email, and number input fields (but not the checkbox or submit button) to ensure users fill in the required information.

# --hints--

You should not have any `</input>` closing tags in your code.

```js
assert.notMatch(code, /<\/input>/);
```

You should have exactly four `<label>` elements in your form.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Your first label should have the text `Full Name:` and a `for` attribute set to `full-name`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0].textContent, /Full\s*Name\s*:/i);
assert.strictEqual(labels[0].getAttribute("for"), "full-name");
```

Your first input should have an `id` attribute set to `full-name` and a `required` attribute.

```js
const firstInput = document.querySelector("input[name='name']");
assert.strictEqual(firstInput.getAttribute("id"), "full-name");
assert.isTrue(firstInput.hasAttribute("required"));
```

Your second label should have the text `Email Address:` and a `for` attribute set to `email`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1].textContent, /Email\s*Address\s*:/i);
assert.strictEqual(labels[1].getAttribute("for"), "email");
```

Your second input should have an `id` attribute set to `email` and a `required` attribute.

```js
const secondInput = document.querySelector("input[name='email']");
assert.strictEqual(secondInput.getAttribute("id"), "email");
assert.isTrue(secondInput.hasAttribute("required"));
```

Your third label should have the text `Donation Amount ($):` and a `for` attribute set to `amount`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2].textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
assert.strictEqual(labels[2].getAttribute("for"), "amount");
```

Your third input should have an `id` attribute set to `amount` and a `required` attribute.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.strictEqual(thirdInput.getAttribute("id"), "amount");
assert.isTrue(thirdInput.hasAttribute("required"));
```

Your fourth label should have the text `Subscribe` and a `for` attribute set to `newsletter`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3].textContent, /Subscribe/i);
assert.strictEqual(labels[3].getAttribute("for"), "newsletter");
```

Your checkbox input should have an `id` attribute set to `newsletter` and should not have a `required` attribute.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.strictEqual(checkboxInput.getAttribute("id"), "newsletter");
assert.isFalse(checkboxInput.hasAttribute("required"));
```

Your submit button should not have a `required` attribute.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="text" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
