---
id: 6720b1f2680ec706ed534acf
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Start the project by linking the `styles.css` file.

# --hints--

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
