---
id: 6728fda085006772a247e779
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Give `.card` a `max-width` of `400px`, a `padding` of `40px` on all sides, and a `margin` of `0` for top and bottom and `auto` for left and right.

# --hints--

`.card` should have a `max-width` of `400px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('max-width'), "400px");
```

`.card` should have a `padding` of `40px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('max-width'), "40px");
assert(false, "consider other possible ways to write that padding")
```

`.card` should have a `margin` of `0 auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('margin'), "0 auto");
assert(false, "other possible ways to write that margin"); 
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
  	</div>
  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

--fcc-editable-region--

.card {
  background-color: white;
}

--fcc-editable-region--

```
