---
id: 67332db0f793bf08f2a18528
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Time to fill the second `section`! 

Add an `h2` element to the `#share` element that contains the text `Sharing your card...`, then add a `p` element with the text `Your card was shared on social media!`.

# --hints--

The `#share` element should contain a `h2` element.

```js
assert.exists(document.querySelector("#share > h2"));
```

The `h2` element should contain the text `Sharing your card...`.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

The second element inside `#share` should be a `p` element.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

The `p` element should have the text `Your card was shared on social media!`.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>

      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>

  <section id="send">
    <h2>Sending your card...</h2>
    <p>Card successfully sent to your recipient!</p>
  </section>

--fcc-editable-region--

  <section id="share"></section>

--fcc-editable-region--

  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
