---
id: 673333130a2c40204162f0c7
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

The `.message` element needs some styling. Give it:

- a `font-size` of `1.2em`,
- a `color` of `gray`,
- a `margin-bottom` of `20px`.

# --hints--

You should have a `.message` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".message"));
```

The `.message` selector should have a `font-size` property set to `1.2em`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".message")?.getPropertyValue("font-size"), "1.2em");
```

The `.message` selector should have a `color` property set to `gray`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".message")?.getPropertyValue("color"), "gray");
```

The `.message` selector should have a `margin-bottom` property set to `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".message")?.getPropertyValue("margin-bottom"), "20px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>

      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>

  <section id="send">
    <h2>Sending your card...</h2>
    <p>Card successfully sent to your recipient!</p>
  </section>

  <section id="share">
    <h2>Sharing your card...</h2>
    <p>Your card was shared on social media!</p>
  </section>

  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

--fcc-editable-region--



--fcc-editable-region--

```
