---
id: 6733352ce2a5642827b4f7f0
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Add a `.card-links` selector, and set the `margin-top` property to `20px`.

Then, to align the two links, add a `display` property set to `flex`, and a `justify-content` set to `space-around`.

# --hints--

You should have a `.card-links` selector.

```js
assert(false)
```

The `.card-links` selector should have a `margin-top` property set to `20px`.

```js
assert(false)
```

The `.card-links` selector should have a `display` property set to `flex`.

```js
assert(false)
```

The `.card-links` selector should have a `justify-content` property set to `space-around`.

```js
assert(false)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>

      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>

      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>

  <section id="send">
    <h2>Sending your card...</h2>
    <p>Card successfully sent to your recipient!</p>
  </section>

  <section id="share">
    <h2>Sharing your card...</h2>
    <p>Your card was shared on social media!</p>
  </section>

  </body>
</html>

```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  backgroung-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "&#129395; ";
}

h1::after {
  content: " &#129395;";
}

.message {
  font-size: 1.2em;
  color: gray;
  margin-bottom: 20px;
}

--fcc-editable-region--



--fcc-editable-region--

```
