---
id: bad87fee1348bd9aede08718
title: Використовування властивостей RGB до Кольорових Елементів
challengeType: 0
videoUrl: 'https://scrimba.com/c/cRkp2fr'
forumTopicId: 18369
dashedName: use-rgb-values-to-color-elements
---

# --description--

Інша можливість представлення кольорів у CSS - це використання конструкції `RGB`.

Наприклад, представлення `RGB` для чорного кольору виглядає наступним чином:

```css
rgb(0, 0, 0)
```

Представлення `RGB` для білого кольору виглядає так:

```css
rgb(255, 255, 255)
```

Замість використання шести шістнадцяткових цифр, як у випадку з шістнадцятковим кодом, у разі `RGB` яскравість кожного кольору задається числом від 0 до 255.

Якщо підрахувати, то два значення для одного кольору дорівнює 16 разів по 16, що приносить у підсумку 256 значень. Таким чином, `RGB` починаючи відлік з нуля, має таку ж саму кількість можливих значень, як і шістнадцятковий код.

Ось приклад того, як можна змінити `body` тла на помаранчевий, використовуючи його RGB код.

```css
body {
  background-color: rgb(255, 165, 0);
}
```

# --instructions--

Давайте змінимо шістнадцятковий код у `body` властивості кольору тла на значення RGB для чорного кольору: `rgb(0, 0, 0)`

# --hints--

Ваш елемент `body` повинен мати чорне тло.

```js
assert($('body').css('background-color') === 'rgb(0, 0, 0)');
```

Слід використовувати `rgb` для надання чорного кольору тлу елемента `body`.

```js
assert(code.match(/rgb\s*\(\s*0\s*,\s*0\s*,\s*0\s*\)/gi));
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: #F00;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: rgb(0, 0, 0);
  }
</style>
```
