---
id: 587d7db6367417b2b2512b98
title: Збіги невказаного символа
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

До цього часу ви створювали набір символів, для яких хочете найти збіги. Однак ви можете створювати й набори символів, для яких не потрібно знаходити збіги. Такі набори символів називають <dfn>набором заперечних символів</dfn>.

Щоб створити набір заперечних символів, поставте символ карет (`^`) після відкриваючої дужки та перед символами, для яких не хочете шукати збіги.

Наприклад, `/[^aeiou]/gi` знайде збіги для всіх символів, що не є голосними. Зверніть увагу, що для символів типу `.`, `!`, `[`, `@`, `/` та пробілів також будуть знайдені збіги: набір заперечних голосних символів голосних виключає лише символи голосних.

# --instructions--

Створіть регулярний вираз, що знаходить збіги для всіх символів, які не є числами чи голосними. Не забудьте використати відповідні прапорці у регулярному виразі.

# --hints--

Ваш регулярний вираз `myRegex` повинен знайти збіги для 9 елементів.

```js
assert(result.length == 9);
```

Ваш регулярний вираз `myRegex` має використати глобальний прапорець.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Ваш регулярний вираз `myRegex` має використати прапорець без урахування регістру.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
