---
id: 612e96fc87fe8e44f69f7ec5
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

حان الوقت لتصميم المفاتيح (style the keys). أسفل قاعدة `#piano`، استهدف عنصر `.keys` باستخدام منتقي `class`. اعطي القاعدة الجديدة خاصية `background-color` بقيمة `#040404` و خاصية `width` بقيمة `949px` و خاصية `height` بقيمة `180px`.

# --hints--

يجب أن يكون لديك منتقي `.keys`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

يجب أن يحتوي منتقي `.keys` على خاصية `background-color` بقيمة `#040404`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

يجب أن يكون لمنتقي `.keys` الخاص بك الخاصية `width` بقيمة `949px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

يجب أن يكون لمنتقي `.keys` الخاص بك الخاصية `height` بقيمة `180px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
