---
id: 612e7d1c29fb872d6384379c
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

قم بدمج `div` ثاني داخل `div` الموجود، وقم بتعيين `class` ليكون `keys`.

# --hints--

يجب عليك إضافة عنصر `div` ثاني.

```js
const divDiv = document.querySelectorAll('div');
assert(divDiv?.length === 2);
```

يجب أن يكون عنصر `div` الجديد الخاص بك داخل عنصر `div` الحالي.

```js
const div = document.querySelector('div');
assert(div?.children?.length === 1);
assert(div?.children?.[0]?.localName === 'div');
```

عنصر `div` الجديد الخاص بك يجب أن يحتوي على `class` يسمي `keys`.

```js
const div = document.querySelector('div');
assert(div?.children?.[0]?.className === 'keys');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    <div id="piano">
  --fcc-editable-region--

  --fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
