---
id: 612e804c54d5e7308d7ebe56
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

داخل عنصر `.keys` قم باضافة سبع عناصر `div`. اعطهم كلهم class الـ `key`.

# --hints--

يجب عليك إضافة سبع عناصر `div` جديدة.

```js
const divDivDiv = document.querySelectorAll('div');
assert(divDivDiv?.length === 9);
```

العناصر `div` السبعة الجديدة يجب ان تكون داخل عنصر `.keys`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].length === 7);
assert([...keys?.children].every(child => child?.tagName === 'DIV'));
```

عناصر `div` السبع الجديدة يجب ان تحتوي جميعها علي `class` معين الي `key`.

```js
const keys = document.querySelector('.keys');
assert([...keys?.children].every(child => child?.classList?.contains('key')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys"></div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
