---
id: 612e813b3ba67633222cbe54
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

تذكر ان سمة `class` ممكن ان تحتوي علي عدة قيم. لفصل المفاتيح البيضاء عن المفاتيح السوداء، ستضيف قيمة `class` بقيمة `black--key`. أضف هذا إلى عناصر `.key` الثانية والثالثة والخامسة والسادسة والسابعة الخاصة بك.

# --hints--

يجب أن يحتوي عنصر `.key` الثاني على `class` من `black--key`.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

يجب أن يكون لـ `.key` الثالث الخاص بك `black--key` في `class`.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

يجب أن يكون لـ `.key` الخامس الخاص بك `black--key` في `class`.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

يجب أن يكون لـ `.key` السادس الخاص بك `black--key` في `class`.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

يجب أن يكون لـ `.key` السابع الخاص بك `black--key` في `class`.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

يجب أن يكون لديك 5 عناصر `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

يجب أن يكون لديك 7 عناصر `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
