---
id: 612e77aba7ca691f598feb02
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف وسمين `meta` واحدة لتحسين صفحتك للأجهزة المحمولة، وواحدة لتحديد `charset` مقبول للصفحة.

# --hints--

يجب عليك إضافة عنصرين `meta` إلى الصفحة الخاصة بك.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

عنصر `meta` واحد يجب أن يحتوي على `name` تم تعيينه إلى `viewport`، و `content` تم تعيينه إلى `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

عنصر `meta` الآخر يجب أن يحتوي على السمة `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
