---
id: 612e78af05201622d4bab8aa
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

حان الوقت لبدء العمل على البيانو. قم بإنشاء عنصر `div` داخل عنصر `body` باستخدام `id` بقيمة `piano`.

# --hints--

يجب عليك إضافة عنصر `div` جديد.

```js
const div = document.querySelector('div');
assert.exists(div);
```

يجب أن يكون `div` الجديد ضمن عنصر `body` الخاص بك.

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

`div` الخاص بك يجب أن يكون له `id` بقيمة `piano`.

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
