---
id: 612e83ec2eca1e370f830511
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

قم بإضافة عنصر `link` ضمن عنصر `head`. لعنصر `link` هذا، قم بتعيين سمة `rel` إلى `stylesheet` و `href` إلى `./styles.css`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `link`.

```js
assert.match(code, /<link/)
```

لا يجب عليك تغيير وسوم `head` الموجودة. تأكد من أنك لم تقم بحذف علامة الإغلاق (closing tag).

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

الكود الخاص بك يجب أن يحتوي على عنصر `link` مغلق ذاتيا.

```js
assert(document.querySelectorAll('link').length === 1);
```

عنصر `link` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert.exists(document.querySelector('head > link'));
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
