---
id: 660ef55dd468079679ee0092
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

You are now ready to declare your next variable. Remove both `console.log` statements, and the `character` reassignment.

Leave the `character` declaration alone.

# --hints--

You should not have any `console.log()` statements in your code.

```js
assert.notMatch(code, /console/);
```

You should not reassign the value of `character`.

```js
assert.isAtMost(code.match(/character/).length, 1);
```

Your `character` variable should have the value `"Hello"`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);
--fcc-editable-region--
```
