---
id: 660ef0f7c4b8e68ccd1f0786
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

JavaScript has eight primitive data types. `String` is one of the primitive data types. A <dfn>string</dfn> in JavaScript represents a series of characters, and is denoted with single (`'`) or double (`"`) quotes.

Note that strings are <dfn>immutable</dfn>, meaning they cannot be altered once created.

Change your `"Hello"` string to use single quotes.

# --hints--

You should not have double quotes in your code.

```js
assert.notMatch(code, /"/);
```

You should use single quotes for your `"Hello"` string.

```js
assert.match(code, /'Hello'/);
```

You should still use `let` in your code.

```js
assert.match(code, /let/);
```

You should still declare a `character` variable.

```js
assert.match(code, /let\s+character/);
```

Your `character` variable should still have the string `"Hello"` for its value.

```js
assert.equal(character, "Hello");
```

Your statement should still end with a semi-colon.

```js
assert.match(code, /let\s+character\s+=\s+'Hello';/)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
