---
id: 660f033cf051ebb50ea3bf48
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

It is time for another variable type. An <dfn>array</dfn> is a special data type that can hold a series of values. Arrays are denoted using square brackets (`[]`). Here is an example of a variable with the value of an empty array:

```js
let array = [];
```

Declare a `rows` variable and assign it an empty array.

# --hints--

You should have a `rows` variable.

```js
assert.match(code, /rows/);
```

You should use `let` to declare your `rows` variable.

```js
assert.match(code, /let\s+rows/);
```

You should assign an empty array to your `rows` variable.

```js
assert.deepEqual(rows, []);
```

Your new statement should end with a semi-colon.

```js
assert.match(code, /let\s+rows\s*=\s*\[\s*\];/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
