---
id: 660f23b53db70af0f2620e78
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

Now all of your numbers are appearing on the same line. This will not work for creating a pyramid.

You will need to add a new line to each row. But you cannot use the return key to add a line break between quotes - JavaScript will throw an error when trying to parse that. Instead, you need to use the special <dfn>escape sequence</dfn> `\n`. This is replaced by a new line when the string is logged. Here is an example:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Use a second addition operator to concatenate a new line between the existing `result` value and the added `row` value.

# --hints--

You should use the `\n` escape sequence. Remember that it needs to be a string, so it is wrapped in quotes.

```js
assert.match(code, /('|")\\n\1/);
```

You should concatenate the `\n` escape sequence to your `result` variable.

```js
assert.match(code, /result\s*\+\s*('|")\\n\1/);
```

You should concatenate your `row` variable to your `\n` escape sequence.

```js
assert.match(code, /result\s*\+\s*('|")\\n\1\s*\+\s*row/);
```

You should assign the entire concatenation back to your `result` variable. Don't forget your semi-colon.

```js
assert.match(code, /result\s*=\s*result\s*\+\s*('|")\\n\1\s*\+\s*row;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i)
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
