---
id: 660f280dda5040f707c76b4a
title: Step 35
challengeType: 0
dashedName: step-35
---

# --description--

Now you have a bunch of `#` characters, but still no pyramid. Thankfully, your `i` variable represents the number of the "row" your loop is currently on, and you can use that to create a pyramid-like structure.

To do this, you will use the `.repeat()` method that strings have. The `.repeat()` method takes a number, which indicates how many times you would like to repeat the given string.

Use the `.repeat()` method on your `character`, and give it `i` for the number.

# --hints--

You should use the `.repeat()` method.

```js
assert.match(code, /\.repeat\(/);
```

You should use the `.repeat()` method on your `character` variable.

```js
assert.match(code, /character\.repeat\(/);
```

You should pass `i` to your `.repeat()` method.

```js
assert.match(code, /character\.repeat\(\s*i\s*\)/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
