---
id: 660f0c34aad72dc712b97624
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

The other method you will need to know for this project is the `.pop()` method. This method "pops" the last value from the array and <dfn>returns</dfn> it.

When a method returns a value, you can think of it as giving the value back to you so you can use it elsewhere.

Declare a `popped` variable, and assign it the result of `rows.pop()`. Then, log your `popped` variable.

# --hints--

You should declare a `popped` variable.

```js
assert.match(code, /popped/);
```

You should use `let` to declare your `popped` variable.

```js
assert.match(code, /let\s+popped/);
```

You should call the `.pop()` method.

```js
assert.match(code, /\.pop\(\s*\)/);
```

You should call the `.pop()` method on your `rows` array. Don't forget your semi-colon.

```js
assert.match(code, /rows\.pop\(\s*\)/);
```

You should log your `popped` variable. Don't forget your semi-colon.

```js
assert.match(code, /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
