---
id: 660f165270622fd4ec0da3f7
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

The `let` keyword allows a variable to be reassigned. This means you could change `character` later to be a completely different value.

For this project, you will not want to change these variable values. So instead, you should use `const` to declare them. `const` tells JavaScript that they are constants, and it will throw an error if you try to reassign them.

Replace your `let` keywords with `const`.

# --hints--

You should use `const` to declare your `character` variable.

```js
assert.match(code, /const\s+character/);
```

You should use `const` to declare your `count` variable.

```js
assert.match(code, /const\s+count/);
```

You should use `const` to declare your `rows` variable.

```js
assert.match(code, /const\s+rows/);
```

You should not use `let` in your code.

```js
assert.notMatch(code, /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
