---
id: 660f09a2694b59c3a10ee304
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Notice how the value inside your `rows` array has been changed directly? This is called <dfn>mutation</dfn>. As you learn more about arrays, you will learn when to mutate an array, and when you should not.

For now, remove your first console log and your `rows[2]` assignment. Leave the second `rows` log statement for later.

# --hints--

You should remove your `console.log(rows[0])` statement.

```js
assert.notMatch(code, /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

You should remove your `rows[2]` reassignment.

```js
assert.notMatch(code, /rows\[2\]/);
```

You should not remove your `console.log(rows)` statement.

```js
assert.match(code, /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
