---
id: 68eaa668d4d9fbe529a0e16b
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Now, add a `label` element with the text `Any specific concerns or topics you'd like to discuss?` and a `for` attribute set to `notes`.

Below this `label`, add a `textarea` element. Set its `id` to `notes`, `name` to `notes`, `rows` to `4`, and `cols` to `50`.

# --hints--

Your last `fieldset` should contain a `label` element.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

The `label` element should have the text `Any specific concerns or topics you'd like to discuss?`

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

The `label` element's `for` attribute should be set to `"notes"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

Your last `fieldset` should contain a `textarea` element.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

Your `textarea` element should be below the `label` element.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

The `textarea` element should have an `id` attribute set to `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

The `textarea` element should have a `name` attribute set to `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

The `textarea` element should have a `rows` attribute set to `4`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

The `textarea` element should have a `cols` attribute set to `50`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
