---
id: 68eb323ff5a6dc1cc19fa300
title: Step 32
challengeType: 0
dashedName: step-32
---

# --description--

Now, you should create the inner circle that will appear when the radio button is checked. You will use the `::before` pseudo-element for this.

A pseudo-element like `::before` lets you insert extra content before the actual element. This is often used for decorative purposes.

Targeting `.contact-method-radio-btn::before` pseudo-element, set a `display` of `block`, `content` of `" "`, `width` of `10px`, `height` of `10px`, and `border-radius` of `50%`.

# --hints--

You should have a `.contact-method-radio-btn::before` selector.

```js
assert(new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before'));
```

Set the `display` property to `block`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.display, 'block');
```

Set the `content` property to `" "`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.content, '" "');
```

Set the `width` property to `10px`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.width, '10px');
```

Set the `height` property to `10px`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.height, '10px');
```

Set the `border-radius` property to `50%`.

```js
const style = new __helpers.CSSHelp(document).getStyle('.contact-method-radio-btn::before');
assert.strictEqual(style?.borderRadius, '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```

```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

.container {
    background-color: #ffffff1a;
    width: 80%;
    max-width: 600px;
    border-radius: 10px;
    margin: 20px auto;
    padding: 10px 20px;
    box-shadow: 0 5px 15px black;  
}

.center {
  text-align: center;
}

.description {
  font-size: 1.2rem;
}

fieldset {
  border: 1px solid gray;
  border-radius: 5px;
  margin: 20px 0;
  padding: 20px;
}

fieldset legend {
  font-size: 1.3rem;
  font-weight: 600;
}

label {
  font-size: 1.2rem;
}

label:not(.contact-method) {
  display: block;
  margin: 10px 0;
}

input:not(.contact-method-radio-btn),
textarea {
  background-color: #ffffff1a;
  width: 95%;
  border: 1px solid gray;
  border-radius: 5px;
  padding: 10px;
}

input,
input::placeholder,
textarea {
  color: whitesmoke;
}


.contact-method-radio-btn {
  appearance: none;
  width: 20px;
  height: 20px;
  border-radius: 50%;
  border: 2px solid gray;
  vertical-align: bottom;
}

--fcc-editable-region--


--fcc-editable-region--
```

