---
id: 5d5a813321b9e3db6c106a46
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript 是一种强大的语言，它允许你构建交互式网站。 首先，使用 `DOCTYPE`、`html`、`head` 和 `body` 创建你的标准 HTML 样板。 包括一个 `title` 元素和一个 `link` 用于样式表，以及 `meta` 标记用于 `charset`。 Then, create a `div` element with `id` set to `game` within your `body`. 给你的 `title` 元素添加文本 `RPG - Dragon Repeller`。

# --hints--

代码应包含 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

应该在 `DOCTYPE` 声明之后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

应该将文档类型定义为 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

应该在类型之后使用 `>` 关闭 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 元素应该有一个开始标签。 不要忘记 `lang` 属性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 声明应位于 HTML 的开头。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

应该有一个 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

应该有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

应该有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

应该有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

代码应该有一个 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 元素应该有一个 `charset` 属性，其值为 `UTF-8`。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

代码应该有一个 `link` 元素。

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

代码应该有一个 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` 元素应该有一个 `id` 属性，其值为 `game`。

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

`div` 元素应该在 `body` 元素内。

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
