---
id: 62a1166ed9a56d439c0770e7
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `#stats` 元素中创建三个 `span` 元素。 给每个元素一个类 `stat`。 然后给第一个元素设置文本 `XP: 0`，给第二个元素设置文本 `Health: 100`，给第三个元素设置文本 `Gold: 50`。

# --hints--

`#stats` 元素中应该有三个 `span` 元素。

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

应该给新的三个 `span` 元素一个 `stat` 类。

```js
assert.exists(document.querySelectorAll('#stats > .stat')?.length, 3);
```

第一个 `.stat` 元素应具有提供的文本 `XP: 0`。

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

第二个 `.stat` 元素应具有提供的文本 `Health: 100`。

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

第三个 `.stat` 元素应具有提供的文本 `Gold: 50`。

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
--fcc-editable-region--
<body>
    <div id="game">
        <div id="stats">

        </div>
        <div id="controls"></div>
        <div id="monsterStats"></div>
        <div id="text"></div>
    </div>
</body>
--fcc-editable-region--
</html>
```
