---
id: 62a255dae245b52317da824a
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

你的开发者控制台是最强大的工具之一。 根据你的浏览器不同，可能需要通过按 `F12` 或 `Ctrl+Shift+I` 来打开它。 On Mac, you can press `Option + ⌘ + C` and select "Console". 也可以点击预览窗口上方的“控制台”按钮来查看我们的内置控制台。

开发者控制台将包含你的代码产生的错误，但你也可以使用它来查看代码中变量的值，这有助于调试。

Add a `console.log("Hello World");` line between your `script` tags. Then click the "Console" button to open the console. You should see the text `"Hello World"`.


# --hints--

应该在代码中添加 `console.log("Hello World");` 行。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

`console.log("Hello World");` 行应该在 `script` 标记之间。

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
