---
id: 62a2401b9842721796b72850
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

为你的 `#text` 元素提供以下文本：

```markup
Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. 
You are in the town square. Where do you want to go? Use the buttons above.
```

# --hints--

`#text` 元素应该有上面引用的文本。

```js
const text = document.querySelector('#text');
assert.equal(text.innerText, "Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
--fcc-editable-region--
        <div id="text">

        </div>
--fcc-editable-region--
    </div>
</body>
</html>
```
