---
id: 62a2509ba163e020bb9d84ea
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

现在可以开始编写 JavaScript。 首先创建一个 `script` 元素。 此元素用于将 JavaScript 加载到 HTML 文件中。

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

应该有一个 `script` 元素。

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

`script` 元素应该有一个开始标记。

```js
assert.match(code, /<script\s*>/i);
```

`script` 元素应该有一个结束标记。

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
