---
id: 62a23d1c5f1c93161f3582ae
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

与 `#stats` 元素类似，`#monsterStats` 元素需要两个 `span` 元素。 给它们类 `stat`，并给第一个元素文本 `Monster Name:`，给第二个元素文本 `Health:`。 在每个文本之后，添加一个 `strong` 元素和一个空的嵌套 `span` 元素。

# --hints--

Your `#monsterStats` element should have two `span` elements.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

新的 `span` 元素都应该具有 `stat` 的 `class` 值。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

第一个 `span` 元素应该包含文本 `Monster Name:`。 请确保间隔正确。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].innerText, 'Monster Name: ');
```

第二个 `span` 元素应包含文本 `Health:`。 请确保间隔正确。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[1].textContent, 'Health: ');
```

第一个 `span` 元素应该有一个 `strong` 元素和一个空的嵌套 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

第二个 `span` 元素应该有一个 `strong` 元素和一个空的嵌套 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

`strong` 和 `span` 元素应该在文本之后。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name: <strong>/);
assert.match(spans[1].innerHTML, /Health: <strong>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
--fcc-editable-region--
        <div id="monsterStats">

        </div>
--fcc-editable-region--
        <div id="text"></div>
    </div>
</body>
</html>
```
