---
id: 62a23cb9bc467a147516b500
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

在 `#controls` 元素中，创建三个 `button` 元素。 第一个应该将 `id` 设置为 `button1`，以及文本设置为 `Go to store`。 第二个应该将 `id` 设置为 `button2`，以及文本设置为 `Go to cave`。 第三个应该将 `id` 设置为 `button3`，以及文本设置为 `Fight dragon`。

# --hints--

应该将三个 `button` 元素添加到 `#controls` 元素中。

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

第一个按钮应该将 `id` 设置为 `button1`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

第一个按钮应该有文本 `Go to store`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

第二个按钮应该将 `id` 设置为 `button2`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

第二个按钮应该有文本 `Go to cave`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

第三个按钮应该将 `id` 设置为 `button3`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

第三个按钮应该有文本 `Fight dragon`。

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
--fcc-editable-region--
        <div id="controls">

        </div>
--fcc-editable-region--
        <div id="monsterStats"></div>
        <div id="text"></div>
    </div>
</body>
</html>
```
