---
id: 62a3b79d520a7f3d0e25afd6
title: 步骤 29
challengeType: 0
dashedName: step-29
---

# --description--

JavaScript 使用<dfn>文档对象模型</dfn>（又称 DOM）与 HTML 交互。 DOM 是一个代表 HTML 的对象树。 You can access the HTML using the `document` object, which represents your entire HTML document.

在你的 HTML 中找到特定元素的一个方法是使用 `querySelector()` 函数。 `querySelector()` 方法需要一个 CSS 选择器作为参数，并返回与该选择器匹配的第一个元素。 例如，要在你的 HTML 中找到 `<h1>` 元素，你应该写下：

```js
let h1 = document.querySelector("h1");
```

Note that `h1` is a string and matches the CSS selector you would use.

Create a `button1` variable and use `querySelector()` to assign it your element with the `id` of `button1`. Remember that CSS `id` selectors are prefixed with a `#`.

# --hints--

You should use `let` to declare a `button1` variable.

```js
assert.match(code, /let button1/);
```

You should use `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

You should use the `#button1` selector.

```js
assert.match(code, /querySelector\(('|")#button1\1\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
