---
id: 62a3a75d8466a12e009eff76
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

初始化一个值为 `100` 的变量 `health` 和一个值为 `50` 的变量 `gold`。

# --hints--

You should use `let` to declare a variable called `health`.

```js
assert.match(code, /let\s+health/);
```

你应该初始化一个值为 `100` 的变量 `health`。

```js
assert.match(code, /let\s+health\s*=\s*100/);
```

You should use `let` to declare a variable called `gold`.

```js
assert.match(code, /let\s+gold/);
```

你应该初始化一个值为 `50` 的变量 `gold`。

```js
assert.match(code, /let\s+gold\s*=\s*50/);
```

`health` 变量的值为 `100`。

```js
assert.equal(health, 100);
```

`gold` 变量的值应为 `50`。

```js
assert.equal(gold, 50);
```

`xp` 变量的值仍然为 `0`。

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
--fcc-editable-region--
```
