---
id: 62a3b506dbaead396f58a701
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

The player's inventory in your game will be able to hold multiple items. You will need to use a data type that can do this. An <dfn>array</dfn> can be used to hold multiple values. For example:

```js
let order = ["first", "second", "third"];
```

This is an array which holds three values. Notice how the values are separated by commas.

Change your `inventory` variable to be an array with the strings `"stick"`, `"dagger"`, and `"sword"`.

# --hints--

Your `inventory` variable should be an array.

```js
assert.isArray(inventory);
```

Your `inventory` variable should have three values.

```js
assert.lengthOf(inventory, 3);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

Your `inventory` variable should include the string `"dagger"`.

```js
assert.include(inventory, "dagger");
```

Your `inventory` variable should include the string `"sword"`.

```js
assert.include(inventory, "sword");
```

Your `inventory` variable should have the values in the correct order.

```js
assert.deepEqual(inventory, ["stick", "dagger", "sword"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory = "stick";
--fcc-editable-region--
```
