---
id: 62a3a0a3c0a4b32915d26a6e
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

删除 `console.log("Hello World");` ，然后开始写你自己的项目代码。

在 JavaScript 中，使用<dfn>变量</dfn>来保存一个值。 若要使用变量，你必须先<dfn>声明</dfn>。 例如，要声明一个名为 `camperbot` 的变量，你应该写：

```js
let camperbot;
```

The `let` keyword tells JavaScript you are declaring a variable. 声明一个名为 `xp` 的变量。

# --hints--

你的代码中不应该有 `console.log("Hello World");`。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

You should use the `let` keyword to declare your variable.

```js
assert.match(code, /let/);
```

你应该声明一个名为 `xp` 的变量。

```js
assert.match(code, /let\s+xp/);
```

你不应将一个值分配给你的变量。

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

别忘了代码末尾的分号。

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
