---
id: 62a3b79d520a7f3d0e25afd6
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript 使用<dfn>文档对象模型</dfn>（又称 DOM）与 HTML 交互。 DOM 是一个代表 HTML 的对象树。 你可以使用 `document` 对象来访问 HTML，其代表你的整个HTML文档。

在你的 HTML 中找到特定元素的一个方法是使用 `querySelector()` 函数。 `querySelector()` 方法需要一个 CSS 选择器作为参数，并返回与该选择器匹配的第一个元素。 例如，要在你的 HTML 中找到 `<h1>` 元素，你应该写下：

```js
let h1 = document.querySelector("h1");
```

请注意， `h1` 是一个字符串，它匹配你将要使用的 CSS 选择器。

创建一个 `button1` 变量。并使用 `querySelector()` 来给它赋值 `id` 为 `button1` 的元素。 请记住，CSS 的 `id` 选择器前缀为 `#`。

# --hints--

你应该使用 `let` 关键词来声明 `button1` 变量。

```js
assert.match(code, /let\s+button1/);
```

你应该使用 `document.querySelector()` 。

```js
assert.match(code, /document\.querySelector/);
```

You should use the `"#button1"` selector.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
