---
id: 62a3b41c9494f937560640ab
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

声明另外两个变量，名为 `monsterHealth` 和 `inventory`，但不要初始化。

# --hints--

你应该使用 `let` 关键词声明你的 `monsterHealth` 变量。

```js
assert.match(code, /let\s+monsterHealth/i);
```

You should use camelCase to name your `monsterHealth` variable.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 不应该有值。

```js
assert.isUndefined(monsterHealth);
```

你不应该为你的 `monsterHealth` 变量分配一个值。 别忘了分号。

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

你应该使用 `let` 关键词声明你的 `inventory` 变量。

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` 不应该有值。

```js
assert.isUndefined(inventory);
```

你不应该为你的 `inventory` 变量分配一个值。 别忘了分号。

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
