---
id: 66b5928b86e907fdfbf34e56
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Now it is time to set the bot's name.

Create a variable called `botName` and assign it the string value of `"teacherBot"`.

# --hints--

You should have a variable called `botName`.

```js
assert.isNotNull(botName);
```

Your `botName` variable should be a string.

```js
assert.isString(botName);
```

Your `botName` variable should have the value of `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
