---
id: 66b6e62423e8a031d6c1f03d
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Now it is time to get the length of the `topic` string.

Start by outputting the message `Here is an example of using the length property on the word [topic].` to the console. 

Remember to replace `[topic]` with the `topic` variable and use proper template literal syntax like you did in the previous steps.

Then, add a second `console.log` statement that outputs the length of the `topic` string to the console.

# --hints--

You should have a `console` statement.

```js
assert.lengthOf(code.match(/console.log\(.*\)/), 6);
```

You should use template literal syntax to output the message `Here is an example of using the length property on the word [topic].` to the `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+$\{topic\}\.`\)/);
```

You should output the length of the `topic` variable to the `console`.

```js
assert.match(code, /console\.log\($\{topic\}\.length\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

const strLength = subject.length;
console.log(strLength);

--fcc-editable-region--

--fcc-editable-region--
```
