---
id: 66b6efddeca35833cd6f0b03
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

To access a character from a string, you can use `bracket notation` and the index number of the character you want to access.

Here is an example of accessing the first letter in the name `Jessica`:

```js
// returns 'J'
firstName[0];
```  

Index numbers start at `0`, so the first letter in a string will always be at index `0`.

Add another console statement. Inside the statement, output the first letter of the `subject` variable using bracket notation and the correct index number.

# --hints--

You should have a `console` statement.

```js
assert.lengthOf(code.match(/console.log\(.*\)/), 9);
```

Your `console` statement should output the first letter of the `subject` variable using bracket notation index `0`.

```js
assert.match(code, /console.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

const strLength = subject.length;
console.log(strLength);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
