---
id: 66b6e39031393e30f2c48d0f
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

To get the length of a string, you can use the <dfn>length</dfn> property. This property returns the number of characters in a string.

Here is an example:

```js
const greeting = "Hello, world!";

// Output: 13
greeting.length
```

Start by creating a variable called `strLength`.

Use the `length` property to get the length of the `subject` string and assign it to the `strLength` variable.

Finally, log the `strLength` variable to the console.

# --hints--

You should have a `strLength` variable.

```js
assert.isNotNull(strLength);
```

Your `strLength` variable should be a number.

```js
assert.isNumber(strLength);
```

You should use the `length` property to get the length of the `subject` string.

```js
assert.match(code, /subject\.length/);
```

You should assign the length of the `subject` string to the `strLength` variable.

```js
assert.equal(strLength, subject.length);
```

You should have a `console` statement.

```js
assert.lengthOf(code.match(/console.log\(.*\)/), 5);
```

You should log the `strLength` variable to the console.

```js
assert.match(code, /console\.log\(\s*strLength\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
