---
id: 66b7142e588c4a407a51cdb6
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

The last console statement outputs `-1` because the substring `"learning"` is not found in the `"Learning is fun.` sentence.

The `indexOf` method is case-sensitive, so the substring `"learning"` is not the same as the substring `"Learning"`.

Now that you understand how some common string methods work, you can complete the workshop by logging one last message to the `console`.

Add a `console` statement that outputs the message `"I hope you enjoyed learning today."` to the `console`.

And with that final message, you have completed the workshop!

# --hints--

You should have a `console` statement.

```js
assert.lengthOf(code.match(/console\.log\(.*\)/), 18);
```

Your `console` statement should output the message `"I hope you enjoyed learning today."`.

```js
assert.match(code, /console\.log\(\s*(\('")I\s+hope\s+you\s+enjoyed\s+learning\s+today.\('"))\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

const strLength = subject.length;
console.log(strLength);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

console.log(sentence.indexOf("fun"));
console.log(sentence.indexOf("learning"));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

const strLength = subject.length;
console.log(strLength);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

console.log(sentence.indexOf("fun"));
console.log(sentence.indexOf("learning"));

console.log("I hope you enjoyed learning today.");
```
