---
id: 66b70b8a611cbf3bcc5c6c5f
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

In the previous lecture videos you learned how to work with the `indexOf` method like this:

```js
const sentence = "I love to learn.";

// returns index 2
sentence.indexOf("love");

// returns -1
sentence.indexOf("hate");
```

If the substring is found, the `indexOf` method will return the index of the first occurrence of the substring in the string.

If the substring is not found in the string, the `indexOf` method will return `-1`.

Add a new `console` statement that outputs the result of using the `indexOf` method to find the position of the substring `"Learning"` in the `learningIsFunSentence` variable.

# --hints--

You should have a `console` statement.

```js
assert.lengthOf(code.match(/console\.log\(.*\)/), 15);
```

Your `console` should output the index position of the substring `"Learning"` from the `learningIsFunSentence` variable.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*\('")Learning\('")\s*\)\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

const strLength = subject.length;
console.log(strLength);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
