---
id: 61537a8054753e2f1f2a1574
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل عنصر `head` قم بإضافة وسم `meta` مع `name` تم تعيينه إلى `viewport` و `content` تم تعيينه إلى `width=device-width, initial-scale=1`.

أضف أيضًا وسم `meta` مع `charset` بقيمة `UTF-8`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصرين `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

عنصر `meta` واحد يجب أن يحتوي على `name` تم تعيينه إلى `viewport`، و `content` تم تعيينه إلى `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي عنصر `meta` الآخر على السمة `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
