---
id: 61537bb9b1a29430ac15ad38
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل عنصر `head`، قم بإضافة عنصر `title` مع نص بقيمة `Photo Gallery`، و عنصر `link` لإضافة ملف `styles.css` إلى الصفحة.

# --hints--

يجب أن يحتوي عنصر `link` الخاص بك على خاصية `href` مع القيمة `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لمشروعك العنوان `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

تذكر أن الـ casing، اي حالة الحرف سواء كبير او صغير (capital or small) والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
