---
id: 61537c9eecea6a335db6da79
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أسفل عنصرك `.header`، أنشئ عنصر `div` جديد مع `class` يسمى `gallery`. سيكون `div` حاوية لصور المَعْرِض.

في عنصر `.gallery`، أنشئ تسع عناصر `img`.

# --hints--

يجب عليك إنشاء عنصر `div` في عنصرك `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

يجب أن يحتوي عنصرك `div` على `class` يسمى `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

يجب أن يأتي عنصرك `div` بعد عنصرك `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

يجب أن يحتوي عنصر `.gallery` على تسع عناصر `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
