---
id: 61537d86bdc3dd343688fceb
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

بعد ذلك، اعطي كل `img` سمة `src` وفقا لترتيبهم في المستند. يجب أن يحتوي عنصر `img` الأول على `src` بقيمة `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. يجب أن تكون البقية مثلها، باستثناء استبدل `1` بالرقم ترتيب `img` في المستند.

# --hints--

جميع العناصر التسع من `img` الخاصة بك يجب أن تحتوي على سمة `src`.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

يجب أن يحتوي العنصر `img` الأول على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

يجب أن يحتوي العنصر `img` الثاني على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

يجب أن يحتوي العنصر `img` الثالث على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

يجب أن يحتوي العنصر `img` الرابع على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

يجب أن يحتوي العنصر `img` الخامس على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

يجب أن يحتوي العنصر `img` السادس على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

يجب أن يحتوي العنصر `img` السابع على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

يجب أن يحتوي العنصر `img` الثامن على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

يجب أن يحتوي العنصر `img` التاسع على `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` قيمة سمة `src`.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
