---
id: 61537485c4f2a624f18d7794
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

ابدأ بنموذج HTML المعتاد الخاص بك. قم بإضافة تعريف `DOCTYPE`، وعنصر `html` و عنصر `head` وعنصر `body`.

أضف خاصية `lang` إلى علامة فتح (opening tag) بشكل `<html>` بقيمة `en`.

# --hints--

يجب أن يحتوي الكود الخاص بك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

عليك أن تضيف مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

ينبغي تعيين `html` نوعا للمستند.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

ينبغي إغلاق إعلان `DOCTYPE` بواسطة <code dir="ltr">&gt;</code> بعد تعيين النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` على علامة فتح (opening tag).

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصرك `html` على علامة أغلاق (closing tag).

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يحتوي عنصر `<html>` على خاصية `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

يجب أن يكون إعلانك `DOCTYPE` في بداية HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك علامة فتح (opening tag) الآتية للعنصر `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك علامة إغلاق (closing tag) الآتية لعنصر `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك العلامة فتح (opening tag) الآتية للعنصر `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك العلامة إغلاق (closing tag) الآتية للعنصر `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

عناصر `head` و `body` يجب أن يكونا متوازيان (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
