---
id: 6153908a366afb4d57185c8d
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

لدى Flexbox محور رئيس وعارض. المحور الرئيس يتم تعريفه بواسطة خاصية `flex-direction` التي لها أربع قيم ممكنة:

- `row` (الافتراضي): محور أفقي مع عناصر flex من اليسار إلى اليمين
- `row-reverse`: محور أفقي مع عناصر flex من اليمين إلى اليسار
- `column`: المحور العمودي مع عناصر flex من الأعلى إلى الأسفل
- `column-reverse`: المحور العمودي مع عناصر flex من الأسفل إلى الأعلى

**تلميح**: المحاور والاتجاهات ستكون مختلفة تبعاً لاتجاه النص. القيم المعروضة هي لاتجاه النص من اليسار إلى اليمين.

حاول أن ترى القيم المختلفة كيف تؤثر على التخطيط.

عند الانتهاء من ذلك، عيّن `flex-direction` بقيمة `row` إلى عنصر `.gallery`.

# --hints--

يجب أن يحتوي منتقي `.gallery` على خاصية `flex-direction` بقيمة `row`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection === 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
