---
id: 615f171d05def3218035dc85
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

صورك كبيرة جداً. أنشئ منتقي `.gallery img` لاستهدافهم. اعطيهم جميعا `width` بقيمة `100%` و `max-width` بقيمة `350px` حتى يتقلص حسب الحاجة، لكنهم لا يصبحون أكبر مما ينبغي.

قم أيضا بتعيين خاصية `height` إلى `300px` للحفاظ على حجم موحد للصور.

# --hints--

يجب أن يكون لديك منتقي `.gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

يجب أن يحتوي منتقي `.gallery img` على خاصية `width` بقيمة `100%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

يجب أن يحتوي منتقي `.gallery img` على خاصية `max-width` بقيمة `350px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

يجب أن يحتوي منتقي `.gallery img` على خاصية `height` بقيمة `300px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
