---
id: 612e89562043183c86df287c
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

I browser possono applicare dei valori predefiniti per il margine e il padding di elementi specifici. Per assicurarti che il piano abbia l'aspetto desiderato, devi resettare il box model.

Aggiungi un selettore `html` al file CSS e imposta la proprietà `box-sizing` su `border-box`.

# --hints--

Dovresti avere un selettore `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

Il selettore `html` dovrebbe avere la proprietà `box-sizing` impostata su `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
