---
id: 612e89d254fe5d3df7d6693d
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Ora che hai resettato il box model `html`, devi anche applicarlo agli elementi all'interno. Per fare ciò, puoi impostare la proprietà `box-sizing` su `inherit`, che farà in modo che gli elementi selezionati utilizzino lo stesso valore dell'elemento genitore.

Avrai anche bisogno di selezionare gli pseudo-elementi, identificati con parole chiave speciali che seguono un selettore. I due pseudo-elementi che userai sono `::before` e `::after`.

`::before` crea uno pseudo-elemento che è il primo figlio dell'elemento selezionato; `::after` crea uno pseudo-elemento che è l'ultimo figlio dell'elemento selezionato. Questi pseudo-elementi sono spesso utilizzati per creare contenuti estetici, che vedrai più avanti in questo progetto.

Per ora, crea un selettore CSS per selezionare tutti gli elementi con `*` e includi gli pseudo-elementi con `::before` e `::after`. Imposta la proprietà `box-sizing` su `inherit`.

# --hints--

Dovresti avere un selettore `*, ::before, ::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Il selettore `*, ::before, ::after` dovrebbe avere una proprietà `box-sizing` impostata su `inherit`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
