---
id: 612ea4c4993aba52ab4aa32e
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Adesso è tempo di usare gli pseudo-selettori che hai preparato prima. Per creare i tasti neri, aggiungi un nuovo selettore `.key.black--key::after`. Questo selezionerà gli elementi con la classe `key black--key` e lo pseudo-elemento dopo questi elementi nell'HTML.

Nel nuovo selettore, imposta il `background-color` a `#1d1e22`. Inoltre, imposta la proprietà `content` a `""`. Questo renderà gli pseudo-elementi vuoti.

La proprietà `content` viene utilizzata per impostare o sovrascrivere il contenuto dell'elemento. Di default, gli pseudo-elementi creati con gli pseudo-selettori `::before` e `::after` sono vuoti e gli elementi non sono renderizzati nella pagina. Impostare la proprietà `content` su una stringa vuota `""` assicurerà che l'elemento sia renderizzato nella pagina mentre è ancora vuoto.

Se vuoi sperimentare, prova a rimuovere la proprietà `background-color` e impostare valori diversi per la proprietà `content`, come `"♥"`. Ricordati di annullare queste modifiche quando hai finito, in modo da superare i test.

# --hints--

Dovresti avere un selettore `.key.black--key::after`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Il selettore `.key.black--key::after` dovrebbe avere una proprietà `background-color` impostata su `#1d1e22`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

Il selettore `.key.black--key::after` dovrebbe avere una proprietà `content` con il valore `""`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
