---
id: 612e96fc87fe8e44f69f7ec5
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

È ora di dedicarci allo stile dei tasti. Sotto la regola `#piano`, crea un selettore per l'attributo `class` con il valore `.keys`. Assegna alla nuova regola una proprietà `background-color` di `#040404`, una proprietà `width` di `949px` e una proprietà `height` di `180px`.

# --hints--

Dovresti avere un selettore `.keys`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys'));
```

Il selettore `.keys` dovrebbe avere una proprietà `background-color` impostata a `#040404`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.backgroundColor === 'rgb(4, 4, 4)');
```

Il selettore `.keys` dovrebbe avere la proprietà `width` impostata a `949px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.width === '949px');
```

Il selettore `.keys` dovrebbe avere la proprietà `height` impostata a `180px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.keys')?.height === '180px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
