---
id: 612e8279827a28352ce83a72
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Ora copia il set di sette elementi `.key` e incolla altri due set nel div `.keys`.

# --hints--

Dovresti avere 21 elementi `.key`.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 21);
```

Dovresti avere 15 elementi `.black--key`.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 15);
```

Tutti gli elementi `.key` dovrebbero essere all'interno dell'elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.key')?.length === 21);
```

Tutti gli elementi `.black--key` dovrebbero essere all'interno dell'elemento `.keys`.

```js
const keys = document.querySelector('.keys');
assert(keys?.querySelectorAll('.black--key')?.length === 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
