---
id: 612e83ec2eca1e370f830511
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Aggiungi un elemento `link` all'interno dell'elemento `head`. Imposta l'attributo `rel` dell'elemento `link` su `stylesheet` e `href` su `./styles.css`.

# --hints--

Il codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/)
```

Non dovresti cambiare i tag `head` esistenti. Assicurati di non aver eliminato il tag di chiusura.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Dovresti avere un elemento `link` auto-concludente.

```js
assert(document.querySelectorAll('link').length === 1);
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
