---
id: 5d822fd413a79914d39e98d2
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Aninhe uma `div` com uma classe de `bb1` no contêiner dos edifícios em segundo plano. Abra o arquivo `styles.css` e dê a `.bb1` uma `width` de `10%` e uma `height` de `70%`. "bb" é a abreviação para "background building" (construção em segundo plano). Esse será o primeiro prédio.

# --hints--

Você deve criar um elemento `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Você deve dar à nova `div` a classe `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Você deve usar um seletor de classe `.bb1` para estilizar o elemento.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Você deve dar ao elemento `.bb1` uma `width` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Você deve dar ao elemento `.bb1` uma `height` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

