---
id: 5d822fd413a79914d39e98d1
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Dê ao elemento `.background-buildings` uma `width` e uma `height` de `100%` para torná-las a largura e a altura totais do elemento pai, o `body`.

# --hints--

Você deve usar a classe de `background-buildings` para selecionar o elemento correto.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

O elemento `.background-buildings` deve ter uma `width` com o valor de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

O elemento `.background-buildings` deve ter uma `height` com o valor de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--

```

