---
id: 5d822fd413a79914d39e98cd
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

No CSS, você pode selecionar tudo com um asterisco. Adicione uma borda a tudo usando o seletor `*` e dando a ela uma `border` de `1px solid black`. Esse é um truque que ajuda a visualizar onde os elementos estão e seu tamanho. Você removerá isso mais tarde.

# --hints--

Você deve usar o seletor `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Você deve usar a propriedade `border` para estilizar todos os elementos.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Todos os elementos devem ter uma borda `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
