---
id: 5d822fd413a79914d39e98e5
title: Passo 29
challengeType: 0
dashedName: step-29
---

# --description--

Você vai querer que o contêiner `.foreground-buildings` fique diretamente na parte superior do elemento `.background-buildings`. Dê a ele uma `width` e uma `height` de `100%` e defina a `position` como `absolute` e o `top` como `0`. Isso fará com que seja do mesmo tamanho que o corpo e que se mova seu início para o canto superior esquerdo.

# --hints--

Você deve usar o seletor `.foreground-buildings`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.foreground-buildings'));
```

Você deve dar ao elemento `.foreground-buildings` uma `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.width, '100%');
```

Você deve dar ao elemento `.foreground-buildings` uma `height` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.height, '100%');
```

Você deve dar ao elemento `.foreground-buildings` uma `position` de `absolute`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.position, 'absolute');
```

Você deve dar ao elemento `.foreground-buildings` um `top` de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.foreground-buildings')?.top, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings"></div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}
--fcc-editable-region--

--fcc-editable-region--

```

