---
id: 5d822fd413a79914d39e98cc
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro de `head`, aninhe um elemento `meta` com um `charset` de `UTF-8`, um elemento `title` com o título `City Skyline` e um elemento `link` que faça a associação ao arquivo `styles.css`.

# --hints--

Você deve criar um elemento `meta` ao redor do elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Você deve dar à tag `meta` um `charset` de `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

O projeto deve ter o título de `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

Você não deve alterar as tags `head` existentes. Certifique-se de que você não excluiu a tag de fechamento.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Você deve ter um elemento `link` de autofechamento.

```js
assert(document.querySelectorAll('link').length === 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
