---
id: 5d822fd413a79914d39e98cf
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Você pode ver o `body` (ele é a caixa mais interna na página); a caixa em torno dele é o elemento `html`. Faça o `body` preencher toda a viewport dando a ele uma `height` de `100vh`. Remova a `margin` padrão do `body` definindo a `margin` como `0`. Finalmente, defina a propriedade `overflow` como `hidden` para ocultar quaisquer barras de rolagem que aparecem quando algo se estende além da janela de visualização.

# --hints--

Você deve usar o seletor `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

O `body` deve ter uma `height` de `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

O `body` deve ter uma `margin` de `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

O `body` deve ter a propriedade `overflow` definida como `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

