---
id: 5d822fd413a79914d39e98ce
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Também adicione um `box-sizing` de `border-box` para tudo. Isto fará com que a borda que você adicionou não adicione qualquer tamanho aos elementos.

# --hints--

Você deve usar a propriedade `box-sizing`.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

Você deve fazer uso do seletor pré-existente `*`.

```js
// Two selectors create two CSSStyleRule objects
assert.equal(new __helpers.CSSHelp(document).getStyleDeclarations('*').length, 1);
```

Todos os elementos devem ter um `box-sizing` de `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

