---
id: 5d822fd413a79914d39e98d7
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Para usar uma variável, coloque o nome da variável entre parênteses precedidos pela palavra `var` desta forma: `var(--variable-name)`. Não importa o valor que você forneceu à variável, ele será aplicado a qualquer propriedade que você usar.

Adicione a variável `--building-color1` que você criou no passo anterior como o valor da propriedade `background-color` da classe `.bb1a`.

# --hints--

A `background-color` do elemento `.bb1a` deve ser definida.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

Você deve usar `var(--building-color1)` para definir a `background-color` do elemento `.bb1a`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
