---
id: 5d822fd413a79914d39e98d4
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Forneça às partes da construção as propriedades `width` e `height` com estes valores: `70%` e `10%` para `.bb1a`, `80%` e `10%` para `.bb1b`, `90%` e `10%` para `.bb1c` e `100%` e `70%` para `.bb1d`. Lembre-se de que essas porcentagens são relativas ao pai e observe que as alturas somarão 100% - enchendo verticalmente o contêiner.

# --hints--

Você deve usar um seletor de classe para estilizar `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Você deve dar a `.bb1a` uma `width` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Você deve dar a `.bb1a` uma `height` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Você deve usar um seletor de classe para estilizar `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Você deve dar a `.bb1b` uma `width` de `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Você deve dar a `.bb1b` uma `height` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Você deve usar um seletor de classe para estilizar o `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Você deve dar a `.bb1c` uma `width` de `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Você deve dar a `.bb1c` uma `height` de `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Você deve usar um seletor de classe para estilizar o `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Você deve dar a `.bb1d` uma `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Você deve dar a `.bb1d` uma `height` de `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

