---
id: 5d822fd413a79914d39e98d0
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

É difícil ver agora, mas há uma borda no limite da pré-visualização: é o `body`. Crie um elemento `div` no `body` com a classe de `background-buildings`. Este será um recipiente para um grupo de edifícios.

# --hints--

Você deve criar um elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

O elemento `div` deve estar dentro de `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

O elemento `div` deve ter a classe `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>    
  <head>
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" type="text/css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

