---
id: 5d822fd413a79914d39e98cc
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Adicione um elemento `title` ao `head` e dê ao seu projeto um título de `City Skyline`. Da mesma forma, aninhe um elemento `link` de fechamento automático ao elemento `head`. Dê a ele o valor do atributo `rel` de `stylesheet` e o valor do atributo `href` de `styles.css`.

# --hints--

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

O projeto deve ter o título de `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

O código deve ter um elemento `link`.

```js
assert.match(code, /<link/)
```

Você não deve alterar as tags `head` existentes. Certifique-se de que você não excluiu a tag de fechamento.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

O elemento `link` deve ser um elemento de fechamento automático.

```js
assert(code.match(/<link[\w\W\s]+\/>/i));
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

O elemento `link` deve ter um atributo `rel` com o valor `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
