---
id: 69aaafeaa4037197c75587cd
title: Implement a Queue
challengeType: 26
dashedName: implement-a-queue
---

# --description--

In this lab, you will implement a queue data structure using functions. A queue is a First-In-First-Out (FIFO) data structure where elements are added to the back of the queue and removed from the front.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have an `enqueue` function that adds an element to the back of the queue.
1. You should have a `dequeue` function that removes and returns the front element of the queue.
1. You should have a `front` function that returns the front element of the queue without removing it.
1. You should have a `size` function that returns the number of elements in the queue.
1. You should have an `isEmpty` function that returns `true` if the queue is empty and `false` otherwise.

**Note:** Store the elements of the queue in the `collection` array of the object returned by `initQueue`, with the front of the queue at index `0`.

# --hints--

You should have an `enqueue` function.

```js
assert.isFunction(enqueue);
```

`enqueue` function should add an element to the back of the queue.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.deepEqual(queue.collection, ['first', 'second']);
```

You should have a `dequeue` function.

```js
assert.isFunction(dequeue);
```

`dequeue` function should remove and return the front element of the queue.

```js
const queue = initQueue();
queue.collection.push('first', 'second', 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.deepEqual(queue.collection, ['third']);
```

`dequeue` should return `undefined` when called on an empty queue.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` should continue to add elements to the back of the queue after elements are removed from the front.

```js
const queue = initQueue();
queue.collection.push('first', 'second');
queue.collection.shift();
enqueue(queue, 'third');
assert.deepEqual(queue.collection, ['second', 'third']);
```

Different queue instances should not share state.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.deepEqual(queueOne.collection, ['first']);
assert.deepEqual(queueTwo.collection, ['second']);
```

You should have a `front` function.

```js
assert.isFunction(front);
```

`front` function should return the front element of the queue without removing it.

```js
const queue = initQueue();
queue.collection.push('first', 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.deepEqual(queue.collection, ['first', 'second']);
```

`front` should return `undefined` when called on an empty queue.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

You should have a `size` function.

```js
assert.isFunction(size);
```

`size` function should return the correct number of elements in the queue.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
queue.collection.push('first');
assert.strictEqual(size(queue), 1);
queue.collection.push('second', 'third');
assert.strictEqual(size(queue), 3);
queue.collection.shift();
assert.strictEqual(size(queue), 2);
```

You should have an `isEmpty` function.

```js
assert.isFunction(isEmpty);
```

`isEmpty` function should return `true` for an empty queue.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

`isEmpty` function should return `false` for a non-empty queue.

```js
const queue = initQueue();
queue.collection.push('element');
assert.isFalse(isEmpty(queue));
```

`isEmpty` function should return `true` after all elements are removed from the queue.

```js
const queue = initQueue();
queue.collection.push('first', 'second');
queue.collection.shift();
queue.collection.shift();
assert.isTrue(isEmpty(queue));
```

The `enqueue`, `dequeue`, `front`, `size`, and `isEmpty` functions should work together to manage the queue in First-In-First-Out order.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.isFalse(isEmpty(queue));
assert.strictEqual(size(queue), 2);
assert.strictEqual(front(queue), 'first');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 2);
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
assert.isTrue(isEmpty(queue));
assert.strictEqual(size(queue), 0);
assert.isUndefined(front(queue));
assert.isUndefined(dequeue(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
