---
id: 61438b5b66d76a6264430f2a
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Within each of your new `a` elements, add an `i` element and give them the following classes:

- Your first `i` element should have the class `fab fa-facebook-f`
- Your second `i` element should have the class `fab fa-twitter`
- Your third `i` element should have the class `fab fa-instagram`
- Your fourth `i` element should have the class `fab fa-linkedin-in`
- Your fifth `i` element should have the class `fab fa-youtube`

Because these icons are decorative, each `i` element should have `aria-hidden` set to `true`, so screen readers will use only the parent link’s `aria-label`.

Also add an `aria-label` attribute to each `a` element describing its destination:

- First `a`: `aria-label` set to `freeCodeCamp Facebook`
- Second `a`: `aria-label` set to `freeCodeCamp Twitter`
- Third `a`: `aria-label` set to `freeCodeCamp Instagram`
- Fourth `a`: `aria-label` set to `freeCodeCamp LinkedIn`
- Fifth `a`: `aria-label` set to `freeCodeCamp YouTube`

# --hints--

You should have five `i` elements.

```js
assert.lengthOf(document.querySelectorAll('i'), 5);
```

Each `a` element should only have one `i` element.

```js
const aaaaa = document.querySelectorAll('.social-icons a');
assert.lengthOf(aaaaa, 5);
aaaaa?.forEach(
  a => {
    assert.lengthOf(a?.children, 1);
    assert.equal(a?.children?.[0]?.localName, 'i');
  }
);
```

Each `i` element should have a `class` attribute which includes `fab`.

```js
const iiiii = document.querySelectorAll('i');
assert.lengthOf(iiiii, 5);
iiiii?.forEach(
  i => {
    assert.isTrue(i?.classList?.contains('fab'));
  }
);
```

Each `i` element should have `aria-hidden` set to `true`.

```js
const icons = [...document.querySelectorAll('.social-icons i')];
assert.isTrue(icons.every(i => i.getAttribute('aria-hidden') === 'true'));
```

The first `i` element should have a `class` attribute which includes `fa-facebook-f`.

```js
assert.isTrue(document.querySelectorAll('i')?.[0]?.classList?.contains('fa-facebook-f'));
```

The second `i` element should have a `class` attribute which includes `fa-twitter`.

```js
assert.isTrue(document.querySelectorAll('i')?.[1]?.classList?.contains('fa-twitter'));
```

The third `i` element should have a `class` attribute which includes `fa-instagram`.

```js
assert.isTrue(document.querySelectorAll('i')?.[2]?.classList?.contains('fa-instagram'));
```

The fourth `i` element should have a `class` attribute which includes `fa-linkedin-in`.

```js
assert.isTrue(document.querySelectorAll('i')?.[3]?.classList?.contains('fa-linkedin-in'));
```

The fifth `i` element should have a `class` attribute which includes `fa-youtube`.

```js
assert.isTrue(document.querySelectorAll('i')?.[4]?.classList?.contains('fa-youtube'));
```

Your first `a` element should have an `aria-label` set to `freeCodeCamp Facebook`.

```js
assert.equal(document.querySelectorAll('.social-icons a')?.[0]?.getAttribute('aria-label'), 'freeCodeCamp Facebook');
```

Your second `a` element should have an `aria-label` set to `freeCodeCamp Twitter`.

```js
assert.equal(document.querySelectorAll('.social-icons a')?.[1]?.getAttribute('aria-label'), 'freeCodeCamp Twitter');
```

Your third `a` element should have an `aria-label` set to `freeCodeCamp Instagram`.

```js
assert.equal(document.querySelectorAll('.social-icons a')?.[2]?.getAttribute('aria-label'), 'freeCodeCamp Instagram');
```

Your fourth `a` element should have an `aria-label` set to `freeCodeCamp LinkedIn`.

```js
assert.equal(document.querySelectorAll('.social-icons a')?.[3]?.getAttribute('aria-label'), 'freeCodeCamp LinkedIn');
```

Your fifth `a` element should have an `aria-label` set to `freeCodeCamp YouTube`.

```js
assert.equal(document.querySelectorAll('.social-icons a')?.[4]?.getAttribute('aria-label'), 'freeCodeCamp YouTube');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <address class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date"><time datetime="2019-03-07">March 7, 2019</time></p>
        </address>
--fcc-editable-region--
        <nav aria-label="Social media links" class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
            
          </a>
          <a href="https://twitter.com/freecodecamp/">
            
          </a>
          <a href="https://instagram.com/freecodecamp">
            
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            
          </a>
          <a href="https://www.youtube.com/freecodecamp">
            
          </a>
        </nav>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
