---
id: 6165d3b702a5d92ad970b30c
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

After your `img` element, add an `hgroup` element with the `class` set to `hero-title`. An `hgroup` groups a heading with related secondary content, such as a subtitle.

Inside the `hgroup`, add an `h1` element with the text `OUR NEW CURRICULUM`, followed by a `p` element with the `class` set to `hero-subtitle` and the text `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

You should create an `hgroup` element.

```js
assert.exists(document.querySelector('hgroup'));
```

Your `hgroup` element should come after your `img` element.

```js
assert.equal(document.querySelector('hgroup')?.previousElementSibling?.localName, 'img');
```

Your `hgroup` element should have the `class` set to `hero-title`.

```js
assert.equal(document.querySelector('hgroup')?.className, 'hero-title');
```

You should create an `h1` element inside your `hgroup` element.

```js
assert.exists(document.querySelector('hgroup > h1'));
```

Your `h1` element should be the first element in your `hgroup` element.

```js
assert.equal(document.querySelector('hgroup')?.firstElementChild?.localName, 'h1');
```

Your `h1` element should have the text set to `OUR NEW CURRICULUM`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

You should create a new `p` element.

```js
assert.exists(document.querySelector('p'));
```

Your `p` element should come after your `h1` element inside the `hgroup`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Your `p` element should have the `class` set to `hero-subtitle`.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

Your `p` element should have the text set to `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <article>
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </header>
      </article>
    </main>
  </body>
</html>
```

```css

```
