---
id: 6169cd8a558aa8434e0ad7f6
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

The `rel` attribute describes the relationship between the current page and a linked resource. The `author` value identifies the linked page as information about the article's author, while `noreferrer` prevents the browser from sending the current page's address in the `Referer` HTTP header.

Give your `a` element a `rel` attribute containing both `author` and `noreferrer`.

# --hints--

Your `a` element should have a `rel` attribute containing `author` and `noreferrer`.

```js
const relValues = document.querySelector('address a')?.relList;
assert.isTrue(relValues?.contains('author'));
assert.isTrue(relValues?.contains('noreferrer'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <article>
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <hgroup class="hero-title">
            <h1>OUR NEW CURRICULUM</h1>
            <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
            </p>
          </hgroup>
          <address class="author">
          <p class="author-name">
            By
--fcc-editable-region--
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
--fcc-editable-region--
          </p>
          </address>
          <p class="publish-date"><time datetime="2019-03-07">March 7, 2019</time></p>
        </header>
      </article>
    </main>
  </body>
</html>
```

```css

```
