---
id: 5daa813381b9e3db6c126b43
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

PYTHON 

Set the `hello` variable to "world". Then print the value.

# --hints--

The `hello` variable should equal "world".

```js
assert.equal(hello, "world");
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
one = 1
hello = "goodbye"
print(one)
--fcc-editable-region--
```
