---
id: 64a5229b99ff0e8250cd9a72
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

step 2 instructions

# --hints--

Test 1

```js
({ input: ["Beau"], test: () => assert.equal(__pyodide.globals.get("name"), "Beau") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
name = input('What is your name?')
print('Hi ' + name)
--fcc-editable-region--
```
