---
id: 64de7357cc75bd18bdced920
title: Step 33
challengeType: 20
dashedName: step-33
---

# --description--

Now change each occurrence of `target` into `rod2`.

# --hints--

You should change each occurrence of `target` inside `make_allowed_move` into `rod2`.

```js
({ test: () => {
  assert.match(code, /if\s+not\s+rods\s*\[\s*rod2\s*\]\s*:/);
  assert.match(code, /elif\s+rods\s*\[\s*rod1\s*\]\s+and\s+rods\s*\[\s*rod1\s*\]\s*\[\s*-\s*1\s*\]\s*<\s*rods\s*\[\s*rod2\s*\]\s*\[\s*-\s*1\s*\]\s*:/);
  assert.match(code, /if\s+forward(\s*==\s*True)?\s*:\s+print\s*\(\s*f('|")Moving\sdisk\s\{\s*rods\s*\[\s*rod1\s*\]\s*\[\s*-\s*1\s*\]\s*\}\sfrom\s\{\s*rod1\s*\}\sto\s\{\s*rod2\s*\}\2\s*\)\s+rods\s*\[\s*rod2\s*\]\s*\.append\s*\(\s*rods\s*\[\s*rod1\s*\]\s*\.pop\s*\(\s*\)\s*\)/);
  assert.match(code, /else\s*:\s+print\s*\(\s*f('|")Moving\sdisk\s\{\s*rods\s*\[\s*rod2\s*\]\s*\[\s*-\s*1\s*\]\s*\}\sfrom\s\{\s*rod2\s*\}\sto\s\{\s*rod1\s*\}\1\s*\)/);
  assert.match(code, /rods\s*\[\s*rod1\s*\]\s*\.append\s*\(\s*rods\s*\[\s*rod2\s*\]\s*\.pop\s*\(\s*\)\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[target]:
        forward = True
    elif rods[rod1] and rods[rod1][-1] < rods[target][-1]:
        forward = True              
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {target}')
        rods[target].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[target][-1]} from {target} to {rod1}')
        rods[rod1].append(rods[target].pop())
    
    # display our progress
    print(rods)
--fcc-editable-region--

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')            
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
