---
id: 64de73f6c2486518e3064fec
title: Step 34
challengeType: 20
dashedName: step-34
---

# --description--

Now call `make_allowed_move()` and pass in `source` and `target` as the arguments.

# --hints--

You should call `make_allowed_move()` passing `source` and `target` as the arguments.

```js
({ test: () => assert.match(code, /make_allowed_move\s*\(\s*source\s*,\s*target\s*\)(?=\s+elif\s+remainder\s*==\s*2\s*:)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[rod2]:
        forward = True
    elif rods[rod1] and rods[rod1][-1] < rods[rod2][-1]:
        forward = True              
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {rod2}')
        rods[rod2].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[rod2][-1]} from {rod2} to {rod1}')
        rods[rod1].append(rods[rod2].pop())
    
    # display our progress
    print(rods)

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--    
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            
--fcc-editable-region--            
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')            

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
