---
id: 64df346f4c86461419974c1e
title: Step 41
challengeType: 20
dashedName: step-41
---

# --description--

Finally, add an `else` clause that prints the allowed move and call `make_allowed_move`. Try to figure out the correct arguments.

# --hints--

You should add an `else` clause.

```js
const tCode = code.replace(/\r/g, '');
const elifBlock = __helpers.python.getBlock(tCode, "elif remainder == 2");
assert.match(elifBlock.block_body, /else:/);
```

You should have `print(f'Move {i + 1} allowed between {source} and {target}')` within the `else` block.

```js
const tCode = code.replace(/\r/g, '');
const elifBlock = __helpers.python.getBlock(tCode, "elif remainder == 2");
assert.match(elifBlock.block_body, /print\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s{\s*source\s*\}\sand\s\{\s*target\s*\}\1\s*\)/);
```

You should have `make_allowed_move(source, target)` within the `else` block.

```js
const tCode = code.replace(/\r/g, '');
const elifBlock = __helpers.python.getBlock(tCode, "elif remainder == 2");
assert.match(elifBlock.block_body, /make_allowed_move\(\s*source\s*,\s*target\s*\)/);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[rod2]:
        forward = True
    elif rods[rod1] and rods[rod1][-1] < rods[rod2][-1]:
        forward = True              
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {rod2}')
        rods[rod2].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[rod2][-1]} from {rod2} to {rod1}')
        rods[rod1].append(rods[rod2].pop())
    
    # display our progress
    print(rods, '\n')

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            if n % 2 != 0:
                print(f'Move {i + 1} allowed between {source} and {target}')
                make_allowed_move(source, target)
            else:
                print(f'Move {i + 1} allowed between {source} and {auxiliary}')
                make_allowed_move(source, auxiliary)            
--fcc-editable-region--
        elif remainder == 2:
            if n % 2 != 0:
                print(f'Move {i + 1} allowed between {source} and {auxiliary}')
                make_allowed_move(source, auxiliary)
--fcc-editable-region--
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')
            make_allowed_move(auxiliary, target)
           
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
