---
id: 66faa70000000000000000cb
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--
Layer 4: camera shake. Insert `this.cameras.main.shake(80, 0.005)` right after the SFX call.

Replace `onHit` body with:

```js
if (!enemy.active) return;
this.sound.play("sfx-hit", { volume: 0.5 });
this.cameras.main.shake(80, 0.005);
enemy.setFillStyle(0xffffff);
this.time.delayedCall(80, () => {
  if (enemy.active) enemy.setFillStyle(0xff5252);
});
this.tweens.add({
  targets: enemy,
  scaleX: 1.3,
  scaleY: 0.7,
  duration: 60,
  yoyo: true,
  ease: "Power2"
});
```

Shake intensity is fraction-of-screen (`0.005` = 0.5%). Step 25 scales it by hit magnitude.

# --hints--
Call `this.cameras.main.shake(80, 0.005)` once per hit.

```js
assert.match(code, /this\.cameras\.main\.shake\(\s*80\s*,\s*0\.005\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
    this.enemies = this.physics.add.group();
    this.spawnEnemy = () => {
      const x = Phaser.Math.Between(50, 750);
      const y = Phaser.Math.Between(50, 550);
      const enemy = this.add
        .rectangle(x, y, 30, 30, 0xff5252)
        .setOrigin(0.5);
      this.physics.add.existing(enemy);
      enemy.body.setCollideWorldBounds(true);
      enemy.hp = 2;
      this.enemies.add(enemy);
      return enemy;
    };
    for (let i = 0; i < 5; i += 1) this.spawnEnemy();
    this.hitbox = this.add
      .rectangle(0, 0, 40, 60, 0xff0000, 0)
      .setOrigin(0.5);
    this.physics.add.existing(this.hitbox);
    this.hitbox.body.setAllowGravity(false);
    this.hitbox.active = false;
    this.attack = () => {
      if (this.attacking) return;
      this.attacking = true;
      const face = this.lastFacing;
      let dx = 0;
      let dy = 0;
      if (face === "right") dx = 30;
      else if (face === "left") dx = -30;
      else if (face === "up") dy = -30;
      else if (face === "down") dy = 30;
      this.hitbox.setPosition(this.player.x + dx, this.player.y + dy);
      this.hitbox.active = true;
      this.sound.play("sfx-ui-click", { volume: 0.3 });
      this.time.delayedCall(150, () => {
        this.attacking = false;
        this.hitbox.active = false;
      });
    };
    this.input.on("pointerdown", () => this.attack());
    this.input.keyboard.on("keydown-J", () => this.attack());
    this.score = 0;
    this.combo = 0;
    this.bestScore = 0;
    this.fx = this.add.particles(0, 0, "particle", {
      speed: { min: -200, max: 200 },
      lifespan: 500,
      scale: { start: 1, end: 0 },
      gravityY: 200,
      emitting: false
    });
--fcc-editable-region--
    this.onHit = (enemy) => {
      if (!enemy.active) return;
      this.sound.play("sfx-hit", { volume: 0.5 });
      enemy.setFillStyle(0xffffff);
      this.time.delayedCall(80, () => {
        if (enemy.active) enemy.setFillStyle(0xff5252);
      });
      this.tweens.add({
        targets: enemy,
        scaleX: 1.3,
        scaleY: 0.7,
        duration: 60,
        yoyo: true,
        ease: "Power2"
      });
    };
--fcc-editable-region--
    this.physics.add.overlap(this.hitbox, this.enemies, (hb, enemy) => {
      if (!this.hitbox.active) return;
      this.onHit(enemy);
    });
  }

  update(time, delta) {
    if (!this.player.active) return;
    const speed = 150;
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
    if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
    if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
    if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
    this.player.body.setVelocity(vx, vy);
    if (vx > 0) this.lastFacing = "right";
    else if (vx < 0) this.lastFacing = "left";
    else if (vy < 0) this.lastFacing = "up";
    else if (vy > 0) this.lastFacing = "down";
    if (vx !== 0 || vy !== 0) {
      this.player.anims.play("walk-" + this.lastFacing, true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
