---
id: 66faa70000000000000000db
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--
20% of hits become crits - 1.5× score and a bonus particle burst with `sfx-coin`.

Insert at the top of `onHit` (after the combo update):

```js
const isCrit = Math.random() < 0.2;
const impact = (this.heavyMode ? 2 : 1) * (isCrit ? 1.5 : 1);
```

Then in the kill branch:

```js
if (isCrit) this.fx.explode(30, enemy.x, enemy.y);
```

And in the non-fatal branch:

```js
if (isCrit) this.sound.play("sfx-coin", { volume: 0.4 });
// ...
if (isCrit) this.fx.explode(20, enemy.x, enemy.y);
```

We multiply the existing `impact` by `1.5` on crit so score scales but everything else (shake, base particles) keeps the same base impact. The bonus burst is the visual payoff.

# --hints--
Roll `const isCrit = Math.random() < 0.2;`.

```js
assert.match(code, /const\s+isCrit\s*=\s*Math\.random\(\s*\)\s*<\s*0\.2/);
```

Bake crit into impact: `(this.heavyMode ? 2 : 1) * (isCrit ? 1.5 : 1)`.

```js
assert.match(code, /const\s+impact\s*=\s*\(\s*this\.heavyMode\s*\?\s*2\s*:\s*1\s*\)\s*\*\s*\(\s*isCrit\s*\?\s*1\.5\s*:\s*1\s*\)/);
```

Add a bonus 30-particle burst on crit kills.

```js
assert.match(code, /if\s*\(\s*isCrit\s*\)\s*this\.fx\.explode\(\s*30\s*,/);
```

Play `"sfx-coin"` on crit non-fatal hits.

```js
assert.match(code, /if\s*\(\s*isCrit\s*\)\s*this\.sound\.play\(\s*["\']sfx-coin["\']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
    this.enemies = this.physics.add.group();
    this.spawnEnemy = () => {
      const x = Phaser.Math.Between(50, 750);
      const y = Phaser.Math.Between(50, 550);
      const enemy = this.add
        .rectangle(x, y, 30, 30, 0xff5252)
        .setOrigin(0.5);
      this.physics.add.existing(enemy);
      enemy.body.setCollideWorldBounds(true);
      enemy.hp = 2;
      this.enemies.add(enemy);
      return enemy;
    };
    for (let i = 0; i < 3; i += 1) this.spawnEnemy();
    this.spawnDelay = 2000;
    const scheduleSpawn = () => {
      this.time.delayedCall(this.spawnDelay, () => {
        if (!this.scene.isActive("MainScene")) return;
        this.spawnEnemy();
        this.spawnDelay = Math.max(400, this.spawnDelay - 50);
        scheduleSpawn();
      });
    };
    scheduleSpawn();
    this.hitbox = this.add
      .rectangle(0, 0, 40, 60, 0xff0000, 0)
      .setOrigin(0.5);
    this.physics.add.existing(this.hitbox);
    this.hitbox.body.setAllowGravity(false);
    this.hitbox.active = false;
    this.attack = (heavy = false) => {
      if (this.attacking) return;
      this.attacking = true;
      this.heavyMode = heavy;
      const face = this.lastFacing;
      let dx = 0;
      let dy = 0;
      let baseAngle = 0;
      if (face === "right") { dx = 30; baseAngle = 0; }
      else if (face === "left") { dx = -30; baseAngle = 180; }
      else if (face === "up") { dy = -30; baseAngle = -90; }
      else if (face === "down") { dy = 30; baseAngle = 90; }
      this.hitbox.setPosition(this.player.x + dx, this.player.y + dy);
      this.hitbox.active = true;
      if (!this.sword) {
        this.sword = this.add
          .rectangle(0, 0, 32, 4, 0xffeeaa)
          .setOrigin(0, 0.5)
          .setVisible(false);
      }
      if (!this.swingTrail) {
        this.swingTrail = this.add.graphics();
      }
      this.swingTrail.clear();
      this.swingTrail.lineStyle(heavy ? 8 : 4, heavy ? 0xff5555 : 0xffeeaa, 0.8);
      this.swingTrail.beginPath();
      this.swingTrail.arc(
        this.player.x,
        this.player.y,
        heavy ? 48 : 36,
        Phaser.Math.DegToRad(baseAngle - 60),
        Phaser.Math.DegToRad(baseAngle + 60)
      );
      this.swingTrail.strokePath();
      this.swingTrail.alpha = 1;
      this.tweens.add({
        targets: this.swingTrail,
        alpha: 0,
        duration: 200
      });
      this.sword.setVisible(true)
        .setPosition(this.player.x, this.player.y)
        .setScale(heavy ? 1.5 : 1)
        .setAngle(baseAngle - 60);
      this.tweens.add({
        targets: this.sword,
        angle: baseAngle + 60,
        duration: heavy ? 220 : 150,
        ease: "Power2",
        onComplete: () => this.sword.setVisible(false)
      });
      this.sound.play(heavy ? "sfx-jump" : "sfx-ui-click", { volume: 0.4 });
      this.time.delayedCall(heavy ? 220 : 150, () => {
        this.attacking = false;
        this.hitbox.active = false;
        this.heavyMode = false;
      });
    };
    this.input.on("pointerdown", () => this.attack());
    this.input.keyboard.on("keydown-J", () => this.attack());
    this.input.keyboard.on("keydown-K", () => this.attack(true));
    this.score = 0;
    this.combo = 0;
    this.bestScore = 0;
    this.fx = this.add.particles(0, 0, "particle", {
      speed: { min: -200, max: 200 },
      lifespan: 500,
      scale: { start: 1, end: 0 },
      gravityY: 200,
      emitting: false
    });
--fcc-editable-region--
    this.onHit = (enemy) => {
      if (!enemy.active) return;
      this.combo += 1;
      if (this._comboReset) this._comboReset.remove();
      this._comboReset = this.time.delayedCall(2000, () => {
        this.combo = 0;
      });
      const impact = this.heavyMode ? 2 : 1;
      if (this.heavyMode) {
        this.time.timeScale = 0.4;
        this.time.delayedCall(120, () => {
          this.time.timeScale = 1;
        });
      }
      enemy.hp -= this.heavyMode ? 2 : 1;
      if (enemy.hp <= 0) {
        this.score += 5 * impact;
        this.sound.play("sfx-explosion", { volume: 0.5 });
        this.cameras.main.shake(120 * impact, 0.008 * impact);
        this.fx.explode(40 * impact, enemy.x, enemy.y);
        enemy.destroy();
        return;
      }
      this.score += 1 * impact;
      this.sound.play("sfx-hit", { volume: 0.5 });
      this.cameras.main.shake(80 * impact, 0.005 * impact);
      this.physics.world.pause();
      this.time.delayedCall(40, () => this.physics.world.resume());
      this.fx.explode(20 * impact, enemy.x, enemy.y);
      enemy.setFillStyle(0xffffff);
      this.time.delayedCall(80, () => {
        if (enemy.active) enemy.setFillStyle(0xff5252);
      });
      this.tweens.add({
        targets: enemy,
        scaleX: 1.3,
        scaleY: 0.7,
        duration: 60,
        yoyo: true,
        ease: "Power2"
      });
      const dx = enemy.x - this.player.x;
      const dy = enemy.y - this.player.y;
      const len = Math.hypot(dx, dy) || 1;
      const startX = enemy.x;
      const startY = enemy.y;
      this.tweens.chain({
        targets: enemy,
        tweens: [
          { x: startX + (dx / len) * 30, y: startY + (dy / len) * 30, duration: 60, ease: "Power2" },
          { x: startX, y: startY, duration: 120, ease: "Power2" }
        ]
      });
    };
--fcc-editable-region--
    this.physics.add.overlap(this.hitbox, this.enemies, (hb, enemy) => {
      if (!this.hitbox.active) return;
      this.onHit(enemy);
    });
  }

  update(time, delta) {
    if (!this.player.active) return;
    const speed = 150;
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
    if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
    if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
    if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
    this.player.body.setVelocity(vx, vy);
    if (vx > 0) this.lastFacing = "right";
    else if (vx < 0) this.lastFacing = "left";
    else if (vy < 0) this.lastFacing = "up";
    else if (vy > 0) this.lastFacing = "down";
    if (vx !== 0 || vy !== 0) {
      this.player.anims.play("walk-" + this.lastFacing, true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: true });
  }

  create() {
    this.scoreText = this.add
      .text(8, 8, "SCORE 0", { fontSize: "20px", color: "#ffffff" })
      .setScrollFactor(0);
    this.comboText = this.add
      .text(8, 36, "", { fontSize: "16px", color: "#ffeeaa" })
      .setScrollFactor(0);
  }

  update() {
    const main = this.scene.get("MainScene");
    if (!main) return;
    this.scoreText.setText(`SCORE ${main.score ?? 0}`);
    const combo = main.combo ?? 0;
    if (combo > 1) {
      this.comboText.setText(`COMBO x${combo}`);
      this.comboText.setScale(1 + combo * 0.05);
      this.comboText.setTint(
        Phaser.Display.Color.GetColor(255, Math.max(80, 255 - combo * 12), 100)
      );
    } else {
      this.comboText.setText("");
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene, UIScene]
};

new Phaser.Game(config);
```
