---
id: 66faa70000000000000000c2
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
Load the Chapter-2 hero sprite sheet plus the Chapter-7 audio bundle, set up the master juice multiplier, generate a 1×1 particle texture, register the walk/idle animations, and place the player.

In `preload`:

```js
this.load.spritesheet(
  "hero",
  "/curriculum-assets/phaser/spritesheets/hero-walk.png",
  { frameWidth: 16, frameHeight: 16 }
);
[
  "sfx-hit",
  "sfx-explosion",
  "sfx-coin",
  "sfx-jump",
  "sfx-ui-click",
  "sfx-ping"
].forEach((key) => {
  this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
});
this.load.audio("music-ambient", [
  "audio/music-ambient.ogg",
  "audio/music-ambient.mp3"
]);
```

Then in `create`, declare the helpers and the player:

```js
this.juice = 1;
const g = this.add.graphics();
g.fillStyle(0xffffff, 1);
g.fillCircle(4, 4, 4);
g.generateTexture("particle", 8, 8);
g.destroy();
[
  ["idle", 0, 0],
  ["walk-down", 0, 3],
  ["walk-up", 4, 7],
  ["walk-left", 8, 11],
  ["walk-right", 12, 15]
].forEach(([key, start, end]) => {
  this.anims.create({
    key,
    frames: this.anims.generateFrameNumbers("hero", { start, end }),
    frameRate: 8,
    repeat: -1
  });
});
this.player = this.physics.add
  .sprite(400, 300, "hero")
  .setScale(2);
this.player.setCollideWorldBounds(true);
this.lastFacing = "down";
```

`this.juice = 1` is the master juice multiplier - step 33 will flip it to `0.2` when `prefers-reduced-motion` is set. Generating the `particle` texture at runtime avoids shipping a sprite asset.

# --hints--
Inside `preload`, load the `"hero"` spritesheet from `/curriculum-assets/phaser/spritesheets/hero-walk.png` with 16×16 frames.

```js
assert.match(code, /this\.load\.spritesheet\(\s*["\']hero["\'][\s\S]*hero-walk\.png[\s\S]*frameWidth:\s*16\s*,\s*frameHeight:\s*16/);
```

Load each Chapter-7 SFX (`sfx-hit`, `sfx-explosion`, `sfx-coin`, `sfx-jump`, `sfx-ui-click`, `sfx-ping`) plus `music-ambient` with `.ogg` and `.mp3` fallbacks.

```js
assert.match(code, /\[\s*"sfx-hit",[\s\S]*"sfx-explosion",[\s\S]*"sfx-coin",[\s\S]*"sfx-jump",[\s\S]*"sfx-ui-click",[\s\S]*"sfx-ping"\s*\][\s\S]*this\.load\.audio\(\s*key/);
```

Initialise `this.juice = 1` so later steps can scale juice in one place.

```js
assert.match(code, /this\.juice\s*=\s*1\b/);
```

Generate a runtime `"particle"` texture with `Graphics.generateTexture("particle", 8, 8)`.

```js
assert.match(code, /g\.generateTexture\(\s*["\']particle["\']\s*,\s*8\s*,\s*8\s*\)/);
```

Build `this.player` as a physics sprite at `(400, 300, "hero")` scaled `2x` with `setCollideWorldBounds(true)`.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\s*\.sprite\(\s*400\s*,\s*300\s*,\s*["\']hero["\']\s*\)[\s\S]*\.setScale\(\s*2\s*\)[\s\S]*setCollideWorldBounds\(\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
