---
id: 66faa70000000000000000c1
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the arcade brawler. Across 40 steps you'll wire a one-screen fighter where every enemy hit fires the full Chapter-7 juice stack (squash + tint flash + screen shake + hit-pause + particle burst + audio cue), plus a sword swing, combo counter, crits, slow-mo, boss fight, persistent best score, mute toggle, touch controls, and `prefers-reduced-motion` damping.

Start with a `MainScene` extending `Phaser.Scene` with a constructor that calls `super("MainScene")`, an empty `preload`, an empty `create`, and an `update(time, delta)`.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {}

  create() {}

  update(time, delta) {}
}
```

The 800x600 Arcade-Physics `config` is already wired in below.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["\']MainScene["\']\s*\)/);
```

You should declare a `preload` method.

```js
assert.match(code, /\bpreload\s*\(\s*\)\s*\{/);
```

You should declare a `create` method.

```js
assert.match(code, /\bcreate\s*\(\s*\)\s*\{/);
```

You should declare an `update(time, delta)` method.

```js
assert.match(code, /\bupdate\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
