---
id: 66faa70000000000000000c5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
Add an invisible hitbox rectangle that activates only during the attack frames and is positioned in front of the player based on `this.lastFacing`.

Replace the entire attack block from step 4 with the version below. The diff: a new `this.hitbox` rectangle (zero alpha, physics-enabled) at the top of the block, and an extended `this.attack` closure that positions+activates the hitbox per facing.

```js
this.hitbox = this.add
  .rectangle(0, 0, 40, 60, 0xff0000, 0)
  .setOrigin(0.5);
this.physics.add.existing(this.hitbox);
this.hitbox.body.setAllowGravity(false);
this.hitbox.active = false;
this.attack = () => {
  if (this.attacking) return;
  this.attacking = true;
  const face = this.lastFacing;
  let dx = 0;
  let dy = 0;
  if (face === "right") dx = 30;
  else if (face === "left") dx = -30;
  else if (face === "up") dy = -30;
  else if (face === "down") dy = 30;
  this.hitbox.setPosition(this.player.x + dx, this.player.y + dy);
  this.hitbox.active = true;
  this.sound.play("sfx-ui-click", { volume: 0.3 });
  this.time.delayedCall(150, () => {
    this.attacking = false;
    this.hitbox.active = false;
  });
};
this.input.on("pointerdown", () => this.attack());
this.input.keyboard.on("keydown-J", () => this.attack());
```

The hitbox is invisible (alpha `0`); it exists purely to overlap with enemies in step 7. The `if (this.attacking) return;` guard prevents spam from chaining hits while the swing is mid-frame.

# --hints--
Build `this.hitbox` as a `40×60` zero-alpha rectangle and add a physics body via `physics.add.existing`.

```js
assert.match(code, /this\.hitbox\s*=[\s\S]*\.rectangle\(\s*0\s*,\s*0\s*,\s*40\s*,\s*60[\s\S]*this\.physics\.add\.existing\(\s*this\.hitbox\s*\)/);
```

Disable gravity on the hitbox body and start with `active = false`.

```js
assert.match(code, /this\.hitbox\.body\.setAllowGravity\(\s*false\s*\)[\s\S]*this\.hitbox\.active\s*=\s*false/);
```

Inside `attack`, compute `dx`/`dy` from `this.lastFacing` (±30 horizontal/vertical) and call `this.hitbox.setPosition(this.player.x + dx, this.player.y + dy)`.

```js
assert.match(code, /this\.hitbox\.setPosition\(\s*this\.player\.x\s*\+\s*dx\s*,\s*this\.player\.y\s*\+\s*dy\s*\)/);
```

Set `this.hitbox.active = true` for the attack window and back to `false` in the `delayedCall`.

```js
assert.match(code, /this\.hitbox\.active\s*=\s*true[\s\S]*this\.time\.delayedCall\(\s*150[\s\S]*this\.hitbox\.active\s*=\s*false/);
```

Guard against attack spam with `if (this.attacking) return;` at the top of `attack`.

```js
assert.match(code, /if\s*\(\s*this\.attacking\s*\)\s*return\s*;/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
--fcc-editable-region--
    this.attack = () => {
      this.attacking = true;
      this.sound.play("sfx-ui-click", { volume: 0.3 });
      this.time.delayedCall(150, () => {
        this.attacking = false;
      });
    };
    this.input.on("pointerdown", () => this.attack());
    this.input.keyboard.on("keydown-J", () => this.attack());
--fcc-editable-region--
  }

  update(time, delta) {
    if (!this.player.active) return;
    const speed = 150;
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
    if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
    if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
    if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
    this.player.body.setVelocity(vx, vy);
    if (vx > 0) this.lastFacing = "right";
    else if (vx < 0) this.lastFacing = "left";
    else if (vy < 0) this.lastFacing = "up";
    else if (vy > 0) this.lastFacing = "down";
    if (vx !== 0 || vy !== 0) {
      this.player.anims.play("walk-" + this.lastFacing, true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
