---
id: 66faa70000000000000000c6
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Add a `physics.add.group()` of enemies plus a `this.spawnEnemy = () => ...` closure so future steps can spawn on a timer.

Insert under `/* create-world */`:

```js
this.enemies = this.physics.add.group();
this.spawnEnemy = () => {
  const x = Phaser.Math.Between(50, 750);
  const y = Phaser.Math.Between(50, 550);
  const enemy = this.add
    .rectangle(x, y, 30, 30, 0xff5252)
    .setOrigin(0.5);
  this.physics.add.existing(enemy);
  enemy.body.setCollideWorldBounds(true);
  enemy.hp = 2;
  this.enemies.add(enemy);
  return enemy;
};
for (let i = 0; i < 5; i += 1) this.spawnEnemy();
```

Five red 30×30 rectangles spawn at random positions. Each carries `enemy.hp = 2` (step 15 uses it to gate destroy).

# --hints--
Create `this.enemies = this.physics.add.group()`.

```js
assert.match(code, /this\.enemies\s*=\s*this\.physics\.add\.group\(\s*\)/);
```

Define `this.spawnEnemy = () => { ... }` as a closure.

```js
assert.match(code, /this\.spawnEnemy\s*=\s*\(\s*\)\s*=>\s*\{/);
```

Inside `spawnEnemy`, place a 30×30 red (`0xff5252`) rectangle at a random `(x, y)` with `Phaser.Math.Between(50, 750/550)`.

```js
assert.match(code, /Phaser\.Math\.Between\(\s*50\s*,\s*750\s*\)[\s\S]*Phaser\.Math\.Between\(\s*50\s*,\s*550\s*\)[\s\S]*\.rectangle\(\s*x\s*,\s*y\s*,\s*30\s*,\s*30\s*,\s*0xff5252\s*\)/);
```

Add a physics body, enable world-bounds collision, set `enemy.hp = 2`, and add to the group.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*enemy\s*\)[\s\S]*setCollideWorldBounds\(\s*true\s*\)[\s\S]*enemy\.hp\s*=\s*2[\s\S]*this\.enemies\.add\(\s*enemy\s*\)/);
```

Spawn `5` enemies with `for (let i = 0; i < 5; i += 1) this.spawnEnemy();`.

```js
assert.match(code, /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*5\s*;\s*i\s*\+=\s*1\s*\)\s*this\.spawnEnemy\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
--fcc-editable-region--

--fcc-editable-region--
    this.hitbox = this.add
      .rectangle(0, 0, 40, 60, 0xff0000, 0)
      .setOrigin(0.5);
    this.physics.add.existing(this.hitbox);
    this.hitbox.body.setAllowGravity(false);
    this.hitbox.active = false;
    this.attack = () => {
      if (this.attacking) return;
      this.attacking = true;
      const face = this.lastFacing;
      let dx = 0;
      let dy = 0;
      if (face === "right") dx = 30;
      else if (face === "left") dx = -30;
      else if (face === "up") dy = -30;
      else if (face === "down") dy = 30;
      this.hitbox.setPosition(this.player.x + dx, this.player.y + dy);
      this.hitbox.active = true;
      this.sound.play("sfx-ui-click", { volume: 0.3 });
      this.time.delayedCall(150, () => {
        this.attacking = false;
        this.hitbox.active = false;
      });
    };
    this.input.on("pointerdown", () => this.attack());
    this.input.keyboard.on("keydown-J", () => this.attack());
  }

  update(time, delta) {
    if (!this.player.active) return;
    const speed = 150;
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
    if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
    if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
    if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
    this.player.body.setVelocity(vx, vy);
    if (vx > 0) this.lastFacing = "right";
    else if (vx < 0) this.lastFacing = "left";
    else if (vy < 0) this.lastFacing = "up";
    else if (vy > 0) this.lastFacing = "down";
    if (vx !== 0 || vy !== 0) {
      this.player.anims.play("walk-" + this.lastFacing, true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
