---
id: 66faa70000000000000000e4
title: Step 36
challengeType: 0
dashedName: step-36
---

# --description--
The 150ms cooldown still lets perfect-tap players chain attacks. Add a 200ms minimum gap between attacks, even after the cooldown fires.

Replace the entry guard in `attack`:

```js
if (this.attacking) return;
const now = this.time.now;
if (now - (this._lastAttack ?? 0) < 200) return;
this._lastAttack = now;
```

We use the SCENE's `time.now` (millis since boot) so it respects `time.timeScale` from step 26 - if you slow-mo, attacks slow too. This is intentional: pacing should be consistent.

# --hints--
Read `this.time.now` and stash on `this._lastAttack`.

```js
assert.match(code, /const\s+now\s*=\s*this\.time\.now/);
```

Reject attacks within 200ms of the last one with `if (now - (this._lastAttack ?? 0) < 200) return;`.

```js
assert.match(code, /if\s*\(\s*now\s*-\s*\(\s*this\._lastAttack\s*\?\?\s*0\s*\)\s*<\s*200\s*\)\s*return/);
```

Update `this._lastAttack = now;` after passing the cooldown check.

```js
assert.match(code, /this\._lastAttack\s*=\s*now/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
    this.enemies = this.physics.add.group();
    this.spawnEnemy = () => {
      const x = Phaser.Math.Between(50, 750);
      const y = Phaser.Math.Between(50, 550);
      const enemy = this.add
        .rectangle(x, y, 30, 30, 0xff5252)
        .setOrigin(0.5);
      this.physics.add.existing(enemy);
      enemy.body.setCollideWorldBounds(true);
      enemy.hp = 2;
      this.enemies.add(enemy);
      return enemy;
    };
    for (let i = 0; i < 3; i += 1) this.spawnEnemy();
    this.spawnDelay = 2000;
    const scheduleSpawn = () => {
      this.time.delayedCall(this.spawnDelay, () => {
        if (!this.scene.isActive("MainScene")) return;
        this.spawnEnemy();
        this.spawnDelay = Math.max(400, this.spawnDelay - 50);
        scheduleSpawn();
      });
    };
    scheduleSpawn();
    this.spawnBoss = () => {
      const boss = this.add
        .rectangle(400, 80, 80, 80, 0x880000)
        .setStrokeStyle(3, 0xffeeaa)
        .setOrigin(0.5);
      this.physics.add.existing(boss);
      boss.body.setCollideWorldBounds(true);
      boss.hp = 20;
      boss.isBoss = true;
      this.enemies.add(boss);
      return boss;
    };
    this.time.delayedCall(20000, () => this.spawnBoss());
    this.hitbox = this.add
      .rectangle(0, 0, 40, 60, 0xff0000, 0)
      .setOrigin(0.5);
    this.physics.add.existing(this.hitbox);
    this.hitbox.body.setAllowGravity(false);
    this.hitbox.active = false;
--fcc-editable-region--
    this.attack = (heavy = false) => {
      if (this.attacking) return;
      this.attacking = true;
      this.heavyMode = heavy;
      const face = this.lastFacing;
      let dx = 0;
      let dy = 0;
      let baseAngle = 0;
      if (face === "right") { dx = 30; baseAngle = 0; }
      else if (face === "left") { dx = -30; baseAngle = 180; }
      else if (face === "up") { dy = -30; baseAngle = -90; }
      else if (face === "down") { dy = 30; baseAngle = 90; }
      this.hitbox.setPosition(this.player.x + dx, this.player.y + dy);
      this.hitbox.active = true;
      if (!this.sword) {
        this.sword = this.add
          .rectangle(0, 0, 32, 4, 0xffeeaa)
          .setOrigin(0, 0.5)
          .setVisible(false);
      }
      if (!this.swingTrail) {
        this.swingTrail = this.add.graphics();
      }
      this.swingTrail.clear();
      this.swingTrail.lineStyle(heavy ? 8 : 4, heavy ? 0xff5555 : 0xffeeaa, 0.8);
      this.swingTrail.beginPath();
      this.swingTrail.arc(
        this.player.x,
        this.player.y,
        heavy ? 48 : 36,
        Phaser.Math.DegToRad(baseAngle - 60),
        Phaser.Math.DegToRad(baseAngle + 60)
      );
      this.swingTrail.strokePath();
      this.swingTrail.alpha = 1;
      this.tweens.add({
        targets: this.swingTrail,
        alpha: 0,
        duration: 200
      });
      this.sword.setVisible(true)
        .setPosition(this.player.x, this.player.y)
        .setScale(heavy ? 1.5 : 1)
        .setAngle(baseAngle - 60);
      this.tweens.add({
        targets: this.sword,
        angle: baseAngle + 60,
        duration: heavy ? 220 : 150,
        ease: "Power2",
        onComplete: () => this.sword.setVisible(false)
      });
      this.sound.play(heavy ? "sfx-jump" : "sfx-ui-click", { volume: 0.4 });
      this.time.delayedCall(heavy ? 220 : 150, () => {
        this.attacking = false;
        this.hitbox.active = false;
        this.heavyMode = false;
      });
    };
--fcc-editable-region--
    this.input.on("pointerdown", () => this.attack());
    this.input.keyboard.on("keydown-J", () => this.attack());
    this.input.keyboard.on("keydown-K", () => this.attack(true));
    this.score = 0;
    this.combo = 0;
    try {
      this.bestScore = parseInt(localStorage.getItem("brawlerBest") ?? "0", 10) || 0;
    } catch (e) {
      this.bestScore = 0;
    }
    this.fx = this.add.particles(0, 0, "particle", {
      speed: { min: -200, max: 200 },
      lifespan: 500,
      scale: { start: 1, end: 0 },
      gravityY: 200,
      emitting: false
    });
    this.onHit = (enemy) => {
      if (!enemy.active) return;
      if (this._juiceComputed !== true) {
        this.juice = window.matchMedia(
          "(prefers-reduced-motion: reduce)"
        ).matches
          ? 0.2
          : 1;
        this._juiceComputed = true;
      }
      this.combo += 1;
      if (this._comboReset) this._comboReset.remove();
      this._comboReset = this.time.delayedCall(2000, () => {
        this.combo = 0;
      });
      const isCrit = Math.random() < 0.2;
      const impact = (this.heavyMode ? 2 : 1) * (isCrit ? 1.5 : 1) * this.juice;
      if (this.heavyMode && this.juice >= 0.8) {
        this.time.timeScale = 0.4;
        this.time.delayedCall(120, () => {
          this.time.timeScale = 1;
        });
      }
      enemy.hp -= this.heavyMode ? 2 : 1;
      if (enemy.hp <= 0) {
        if (enemy.isBoss) {
          this.score += 100;
          this.cameras.main.shake(800, 0.02);
          for (let i = 0; i < 8; i += 1) {
            this.time.delayedCall(i * 100, () => {
              this.fx.explode(
                50,
                enemy.x + Phaser.Math.Between(-50, 50),
                enemy.y + Phaser.Math.Between(-50, 50)
              );
              this.sound.play("sfx-explosion", { volume: 0.5 });
            });
          }
          enemy.destroy();
          return;
        }
        this.score += 5 * impact;
        this.sound.play("sfx-explosion", { volume: 0.5 });
        this.cameras.main.shake(120 * impact, 0.008 * impact);
        this.fx.explode(40 * impact, enemy.x, enemy.y);
        if (isCrit) this.fx.explode(30, enemy.x, enemy.y);
        enemy.destroy();
        return;
      }
      this.score += 1 * impact;
      if (this.score > this.bestScore) {
        this.bestScore = this.score;
        try {
          localStorage.setItem("brawlerBest", String(this.bestScore));
        } catch (e) {}
      }
      this.sound.play("sfx-hit", { volume: 0.5 });
      if (isCrit) this.sound.play("sfx-coin", { volume: 0.4 });
      this.cameras.main.shake(80 * impact, 0.005 * impact);
      this.physics.world.pause();
      this.time.delayedCall(40, () => this.physics.world.resume());
      this.fx.explode(20 * impact, enemy.x, enemy.y);
      if (isCrit) this.fx.explode(20, enemy.x, enemy.y);
      enemy.setFillStyle(0xffffff);
      this.time.delayedCall(80, () => {
        if (enemy.active) enemy.setFillStyle(0xff5252);
      });
      this.tweens.add({
        targets: enemy,
        scaleX: 1.3,
        scaleY: 0.7,
        duration: 60,
        yoyo: true,
        ease: "Power2"
      });
      const dx = enemy.x - this.player.x;
      const dy = enemy.y - this.player.y;
      const len = Math.hypot(dx, dy) || 1;
      const startX = enemy.x;
      const startY = enemy.y;
      this.tweens.chain({
        targets: enemy,
        tweens: [
          { x: startX + (dx / len) * 30, y: startY + (dy / len) * 30, duration: 60, ease: "Power2" },
          { x: startX, y: startY, duration: 120, ease: "Power2" }
        ]
      });
    };
    this.physics.add.overlap(this.hitbox, this.enemies, (hb, enemy) => {
      if (!this.hitbox.active) return;
      this.onHit(enemy);
    });
    this.playerHp = 3;
    this.invulnerable = false;
    this.physics.add.overlap(this.player, this.enemies, () => {
      if (this.invulnerable || !this.player.active) return;
      this.invulnerable = true;
      this.cameras.main.shake(150, 0.015);
      this.cameras.main.flash(150, 255, 0, 0);
      this.player.setTintFill(0xff0000);
      this.tweens.add({
        targets: this.player,
        alpha: 0.3,
        duration: 100,
        yoyo: true,
        repeat: 3
      });
      this.time.delayedCall(200, () => this.player.clearTint());
      this.sound.play("sfx-hit", { volume: 0.7 });
      this.playerHp -= 1;
      if (this.playerHp <= 0) {
        this.player.active = false;
        this.player.body.setVelocity(0, 0);
        this.time.delayedCall(400, () => {
          this.scene.pause("MainScene");
          this.scene.launch("GameOverScene", {
            score: this.score,
            best: this.bestScore
          });
        });
        return;
      }
      this.time.delayedCall(800, () => {
        this.invulnerable = false;
        this.player.alpha = 1;
      });
    });
  }

  update(time, delta) {
    if (!this.player.active) return;
    const speed = 150;
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
    if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
    if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
    if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
    this.player.body.setVelocity(vx, vy);
    if (vx > 0) this.lastFacing = "right";
    else if (vx < 0) this.lastFacing = "left";
    else if (vy < 0) this.lastFacing = "up";
    else if (vy > 0) this.lastFacing = "down";
    if (vx !== 0 || vy !== 0) {
      this.player.anims.play("walk-" + this.lastFacing, true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

class UIScene extends Phaser.Scene {
  constructor() {
    super({ key: "UIScene", active: true });
  }

  create() {
    this.scoreText = this.add
      .text(8, 8, "SCORE 0", { fontSize: "20px", color: "#ffffff" })
      .setScrollFactor(0);
    this.comboText = this.add
      .text(8, 36, "", { fontSize: "16px", color: "#ffeeaa" })
      .setScrollFactor(0);
    this.bestText = this.add
      .text(792, 8, "BEST 0", { fontSize: "16px", color: "#ffeeaa" })
      .setOrigin(1, 0)
      .setScrollFactor(0);
    const muteLabel = () => (this.sound.mute ? "\u266b OFF" : "\u266b ON");
    this.muteText = this.add
      .text(792, 32, muteLabel(), {
        fontSize: "16px",
        color: "#ffffff",
        backgroundColor: "#2a3045",
        padding: { x: 8, y: 4 }
      })
      .setOrigin(1, 0)
      .setScrollFactor(0)
      .setInteractive();
    const toggleMute = () => {
      this.sound.mute = !this.sound.mute;
      this.muteText.setText(muteLabel());
    };
    this.muteText.on("pointerup", toggleMute);
    this.input.keyboard.on("keydown-M", toggleMute);
    this.attachTouchControls();
  }

  update() {
    const main = this.scene.get("MainScene");
    if (!main) return;
    this.scoreText.setText(`SCORE ${main.score ?? 0}`);
    this.bestText.setText(`BEST ${main.bestScore ?? 0}`);
    const combo = main.combo ?? 0;
    if (combo > 1) {
      this.comboText.setText(`COMBO x${combo}`);
      this.comboText.setScale(1 + combo * 0.05);
      this.comboText.setTint(
        Phaser.Display.Color.GetColor(255, Math.max(80, 255 - combo * 12), 100)
      );
    } else {
      this.comboText.setText("");
    }
  }

  attachTouchControls() {
    const main = this.scene.get("MainScene");
    if (!main) return;
    const atk = this.add
      .rectangle(720, 540, 120, 80, 0x6c8aff, 0.6)
      .setScrollFactor(0)
      .setStrokeStyle(2, 0xffffff)
      .setInteractive();
    this.add
      .text(720, 540, "ATK", { fontSize: "22px", color: "#ffffff" })
      .setOrigin(0.5)
      .setScrollFactor(0);
    const heavy = this.add
      .rectangle(580, 540, 100, 80, 0xff5252, 0.6)
      .setScrollFactor(0)
      .setStrokeStyle(2, 0xffffff)
      .setInteractive();
    this.add
      .text(580, 540, "HEAVY", { fontSize: "16px", color: "#ffffff" })
      .setOrigin(0.5)
      .setScrollFactor(0);
    atk.on("pointerdown", () => main.attack && main.attack());
    heavy.on("pointerdown", () => main.attack && main.attack(true));
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }
  create(data) {
    this.add
      .rectangle(400, 300, 800, 600, 0x000000, 0.7)
      .setOrigin(0.5);
    this.add
      .text(400, 220, "GAME OVER", { fontSize: "48px", color: "#ff5252" })
      .setOrigin(0.5);
    this.add
      .text(400, 290, `SCORE ${data?.score ?? 0}`, { fontSize: "28px", color: "#ffffff" })
      .setOrigin(0.5);
    this.add
      .text(400, 330, `BEST  ${data?.best ?? 0}`, { fontSize: "20px", color: "#ffeeaa" })
      .setOrigin(0.5);
    const btn = this.add
      .text(400, 410, "PLAY AGAIN", { fontSize: "22px", color: "#ffffff", backgroundColor: "#2a3045", padding: { x: 14, y: 8 } })
      .setOrigin(0.5)
      .setInteractive();
    btn.on("pointerup", () => {
      this.scene.stop("GameOverScene");
      this.scene.stop("UIScene");
      this.scene.start("MainScene");
      this.scene.launch("UIScene");
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
   physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene, UIScene, GameOverScene]
};

new Phaser.Game(config);
```
