---
id: 66faa70000000000000000c4
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
Wire the attack action. We'll route everything through a `this.attack = () => ...` closure so later steps can grow it into the full juice pipeline without rewriting the input handlers.

In `create` (under `/* create-methods */`):

```js
this.attack = () => {
  this.attacking = true;
  this.sound.play("sfx-ui-click", { volume: 0.3 });
  this.time.delayedCall(150, () => {
    this.attacking = false;
  });
};
this.input.on("pointerdown", () => this.attack());
this.input.keyboard.on("keydown-J", () => this.attack());
```

Click the canvas or press `J` to hear the swing - no visible hit yet (that comes in step 5).

# --hints--
Define `this.attack = () => { ... }` as an arrow-function closure.

```js
assert.match(code, /this\.attack\s*=\s*\(\s*\)\s*=>\s*\{/);
```

Inside `attack`, set `this.attacking = true` and play `sfx-ui-click`.

```js
assert.match(code, /this\.attacking\s*=\s*true[\s\S]*this\.sound\.play\(\s*["\']sfx-ui-click["\']/);
```

Reset `this.attacking = false` after `150ms` via `this.time.delayedCall`.

```js
assert.match(code, /this\.time\.delayedCall\(\s*150[\s\S]*this\.attacking\s*=\s*false/);
```

Wire pointer click via `this.input.on("pointerdown", () => this.attack())`.

```js
assert.match(code, /this\.input\.on\(\s*["\']pointerdown["\']\s*,\s*\(\s*\)\s*=>\s*this\.attack\(\s*\)\s*\)/);
```

Wire the `J` key via `this.input.keyboard.on("keydown-J", () => this.attack())`.

```js
assert.match(code, /this\.input\.keyboard\.on\(\s*["\']keydown-J["\']\s*,\s*\(\s*\)\s*=>\s*this\.attack\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    if (!this.player.active) return;
    const speed = 150;
    let vx = 0;
    let vy = 0;
    if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
    if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
    if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
    if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
    this.player.body.setVelocity(vx, vy);
    if (vx > 0) this.lastFacing = "right";
    else if (vx < 0) this.lastFacing = "left";
    else if (vy < 0) this.lastFacing = "up";
    else if (vy > 0) this.lastFacing = "down";
    if (vx !== 0 || vy !== 0) {
      this.player.anims.play("walk-" + this.lastFacing, true);
    } else {
      this.player.anims.play("idle", true);
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
