---
id: 66faa70000000000000000c3
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Wire eight-direction movement using BOTH cursor keys AND WASD via the action-map pattern. Track the last non-zero facing so the step-5 hitbox knows which side to swing.

In `create`:

```js
this.cursors = this.input.keyboard.createCursorKeys();
this.keyMap = this.input.keyboard.addKeys("W,A,S,D,J,K,M");
```

And in `update(time, delta)`:

```js
if (!this.player.active) return;
const speed = 150;
let vx = 0;
let vy = 0;
if (this.cursors.left.isDown || this.keyMap.A.isDown) vx -= speed;
if (this.cursors.right.isDown || this.keyMap.D.isDown) vx += speed;
if (this.cursors.up.isDown || this.keyMap.W.isDown) vy -= speed;
if (this.cursors.down.isDown || this.keyMap.S.isDown) vy += speed;
this.player.body.setVelocity(vx, vy);
if (vx > 0) this.lastFacing = "right";
else if (vx < 0) this.lastFacing = "left";
else if (vy < 0) this.lastFacing = "up";
else if (vy > 0) this.lastFacing = "down";
if (vx !== 0 || vy !== 0) {
  this.player.anims.play("walk-" + this.lastFacing, true);
} else {
  this.player.anims.play("idle", true);
}
```

Move with WASD or arrow keys to verify.

# --hints--
Create `this.cursors` with `this.input.keyboard.createCursorKeys()`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)/);
```

Add the `"W,A,S,D,J,K,M"` key map (J/K for attacks, M for mute).

```js
assert.match(code, /this\.keyMap\s*=\s*this\.input\.keyboard\.addKeys\(\s*["\']W,A,S,D,J,K,M["\']\s*\)/);
```

In `update`, accumulate `vx` and `vy` from the cursor or keyMap input at speed `150`.

```js
assert.match(code, /const\s+speed\s*=\s*150[\s\S]*let\s+vx\s*=\s*0[\s\S]*let\s+vy\s*=\s*0/);
```

Apply the movement with `this.player.body.setVelocity(vx, vy)`.

```js
assert.match(code, /this\.player\.body\.setVelocity\(\s*vx\s*,\s*vy\s*\)/);
```

Play `"walk-" + this.lastFacing` when moving and `"idle"` when stationary.

```js
assert.match(code, /this\.player\.anims\.play\(\s*["\']walk-["\']\s*\+\s*this\.lastFacing[\s\S]*this\.player\.anims\.play\(\s*["\']idle["\']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Arcade Brawler</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    [
      "sfx-hit",
      "sfx-explosion",
      "sfx-coin",
      "sfx-jump",
      "sfx-ui-click",
      "sfx-ping"
    ].forEach((key) => {
      this.load.audio(key, [`audio/${key}.ogg`, `audio/${key}.mp3`]);
    });
    this.load.audio("music-ambient", [
      "audio/music-ambient.ogg",
      "audio/music-ambient.mp3"
    ]);
  }

  create() {
    this.juice = 1;
    const g = this.add.graphics();
    g.fillStyle(0xffffff, 1);
    g.fillCircle(4, 4, 4);
    g.generateTexture("particle", 8, 8);
    g.destroy();
    [
      ["idle", 0, 0],
      ["walk-down", 0, 3],
      ["walk-up", 4, 7],
      ["walk-left", 8, 11],
      ["walk-right", 12, 15]
    ].forEach(([key, start, end]) => {
      this.anims.create({
        key,
        frames: this.anims.generateFrameNumbers("hero", { start, end }),
        frameRate: 8,
        repeat: -1
      });
    });
    this.player = this.physics.add
      .sprite(400, 300, "hero")
      .setScale(2);
    this.player.setCollideWorldBounds(true);
    this.lastFacing = "down";
    --fcc-editable-region--
  }

  update(time, delta) {
--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 800,
  height: 600,
  parent: "game-container",
  backgroundColor: "#101020",
  physics: { default: "arcade", arcade: { debug: false } },
  scene: [MainScene]
};

new Phaser.Game(config);
```
