---
id: 660ef5c1904955978a986a5c
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

Use `let` to declare a `count` variable. Assign it the <dfn>number</dfn> `8`. When using a number value, you do not use quotes. For example:

```js
let money = 100;
```

# --hints--

You should have a second `let` keyword in your code.

```js
assert.lengthOf(code.match(/let/g), 2);
```

You should use `let` to declare a `count` variable.

```js
assert.match(code, /let\s+count/);
```

You should assign the number `8` to your `count` variable.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
