---
id: 65648e4c5b316c4ec5e4fddc
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

At the bottom of your function, create a `password` variable and assign it an empty string.

# --hints--

You should have a `password` variable inside your function with the value of an empty string. Pay attention to the indentation.

```js
({ test: () =>
  {
    const generate_psw = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_psw;    
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):

    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
