---
id: 6565cbe0a411c37145c55551
title: Step 67
challengeType: 20
dashedName: step-67
---

# --description--

Finally, delete the `count` variable.

With that the password generator project is complete.

# --hints--

You should delete the `count = 0` line.

```js
({ test: () => assert.isFalse( /count\s*=\s*0/.test(code)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):

    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints
        count = 0
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):            
            break
    
    return password
    

if __name__ == '__main__':
    new_password = generate_password()
    print('Generated password:', new_password)
--fcc-editable-region--
```

# --solutions--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):

    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):            
            break
    
    return password
    

if __name__ == '__main__':
    new_password = generate_password()
    print('Generated password:', new_password)
    
```
