---
id: 6565c54c9497796c2b9f6adb
title: Step 66
challengeType: 20
dashedName: step-66
---

# --description--

The `all()` function returns `True` if all the elements inside a given iterable evaluate to `True`. Otherwise it returns `False`. You can combine the `all()` function with the list comphehension syntax to make your code more concise. Here's an example:

```py
all(expression for element in iterable)
```

HELP, THIS STEP NEEDS TO BE EXPLAINED BETTER

Replace your existing `for` loop and two `if` statements with a single `if` statement. For the `if` condition, call the `all()` function passing a suitable iterable to it.

# --hints--

You should use the `all()` function to check that each `constraint <= len(re.findall(pattern, password))`.

```js
({ test: () => assert.match(e.code.original["main.py"], /^(\s*)if\s+all\s*\(\s*constraint\s*<=\s*len\s*\(\s*re\.findall\s*\(\s*pattern\s*,\s*password\s*\)\s*\)\s*for\s+constraint\s*,\s*pattern\s+in\s+constraints\s*\)\s*:\s*^\1\s{4}break/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1):

    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints
        count = 0
--fcc-editable-region--        
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
        
        if count == 4:
--fcc-editable-region--        
            break
    
    return password
    

if __name__ == '__main__':
    new_password = generate_password()
    print('Generated password:', new_password)

```
