---
id: 6565c234de8cdf673c96bdf3
title: Step 66
challengeType: 20
dashedName: step-66
---

# --description--

It works, but there are still a couple of things you can improve. First of all, calling a function with 5 arguments can create confusion.

Modify your function declaration to accept default arguments. Use `16` for the `length` and `1` for the other constraints.

# --hints--

Your function should accept default parameters.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect
    foo = __locals.get("generate_password")
    sig = str(inspect.signature(foo))
    sig == '(length=16, nums=1, special_chars=1, uppercase=1, lowercase=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string

--fcc-editable-region--
def generate_password(length, nums, special_chars, uppercase, lowercase):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    

new_password = generate_password(8, 1, 1, 1, 1)
print(new_password)

```
