---
id: 69b83e35f19ba26ba1fa517a
title: Build a Traffic Light Sequencer
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

In this lab, you will simulate configurable traffic light cycles and log anomalies.

You'll work with config objects that describe the phases of a traffic light. Each config object has the following properties:

- `fault`: a boolean flag that triggers early termination when `true`.
- `phases`: an array of phase objects.

Each phase object inside `phases` has the following properties:

- `color`: a string representing the light color (`"green"`, `"yellow"`, or `"red"`).
- `duration`: a positive integer representing how long the phase lasts in seconds.

You can refer to the provided `config1`, `config2`, `config3`, and `config4` objects as examples of valid configuration objects.

**Note:** Do not add any extra `console.log()` statements, as they may cause the tests to fail.

**Objective**: Fulfill the user stories below and get all the tests to pass to complete the lab.

# User Stories

1. You should have a function named `runSequence` with two parameters: `config` and `cycles`, where `config` represents a config object and `cycles` represents the maximum number of times the sequence can run before stopping.

1. You should implement `runSequence(config, cycles)` using a `for` or `while` loop to iterate through each phase across the given number of cycles.

1. The `runSequence` function should:
    - Log `No phases found` and immediately return if `config.phases` is empty.
    - Log `Faulted phase!` and stop the simulation early if `config.fault` is set to `true`.
    - Log `Invalid phase detected` if `duration <= 0`.
    - Log `Switching to [color] for [duration] s` for each valid phase. Replace `[color]` and `[duration]` with the corresponding properties of the phase object.
    - For example, `runSequence(exampleConfig, 1)` should log:

    ```md
    Switching to green for 5 s
    Switching to yellow for 2 s
    Switching to red for 4 s
    ```

1. You should have a function named `generateTimeline` with two parameters: `config` and `cycles`.

1. The `generateTimeline` function should:
    - Record the cumulative elapsed time after each phase across the cycles into an array, adding each phase's `duration` to the running total as you iterate.
    - Return the array of cumulative timestamps.
    - For example, `generateTimeline(exampleConfig, 1)` should return the array `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

You should not make changes in the `config1`, `config2`, `config3`, and `config4` objects.

```js
const _config1 = {
    fault: false,
    phases: [
        { color: "green", duration: 5 },
        { color: "yellow", duration: 2 },
        { color: "red", duration: 4 }
    ]
};

const _config2 = {
    fault: false,
    phases: [
        { color: "red", duration: 3 },
        { color: "yellow", duration: -2 },
        { color: "green", duration: 6 }
    ]
};

const _config3 = {
    fault: true,
    phases: [
        { color: "green", duration: 5 },
        { color: "yellow", duration: 2 },
        { color: "red", duration: 6 }
    ]
};

const _config4 = {
    fault: false,
    phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

You should have a function named `runSequence`.

```js
assert.isFunction(runSequence);
```

Your `runSequence` function should accept two parameters, `config` and `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` should log `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` in order.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
    ["Switching to green for 5 s"],
    ["Switching to yellow for 2 s"],
    ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` should log `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s` in order.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
    ["Switching to green for 5 s"],
    ["Switching to yellow for 2 s"],
    ["Switching to red for 4 s"],
    ["Switching to green for 5 s"],
    ["Switching to yellow for 2 s"],
    ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` should log `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s` in order.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
    ["Switching to red for 3 s"],
    ["Invalid phase detected"],
    ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` should log `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` should log `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

You should have a function named `generateTimeline`.

```js
assert.isFunction(generateTimeline);
```

Your `generateTimeline` function should accept two parameters, `config` and `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` should return `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` should return `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` should return `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` should return `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` should return `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
    fault: false,
    phases: [
        { color: "green", duration: 5 },
        { color: "yellow", duration: 2 },
        { color: "red", duration: 4 }
    ]
};

const config2 = {
    fault: false,
    phases: [
        { color: "red", duration: 3 },
        { color: "yellow", duration: -2 },
        { color: "green", duration: 6 }
    ]
};

const config3 = {
    fault: true,
    phases: [
        { color: "green", duration: 5 },
        { color: "yellow", duration: 2 },
        { color: "red", duration: 6 }
    ]
};

const config4 = {
    fault: false,
    phases: []
};
```

# --solutions--

```js
const config1 = {
    fault: false,
    phases: [
        { color: "green", duration: 5 },
        { color: "yellow", duration: 2 },
        { color: "red", duration: 4 }
    ]
};

const config2 = {
    fault: false,
    phases: [
        { color: "red", duration: 3 },
        { color: "yellow", duration: -2 },
        { color: "green", duration: 6 }
    ]
};

const config3 = {
    fault: true,
    phases: [
        { color: "green", duration: 5 },
        { color: "yellow", duration: 2 },
        { color: "red", duration: 6 }
    ]
};

const config4 = {
    fault: false,
    phases: []
};

function runSequence(config, cycles) {
    if (config.phases.length === 0) {
        console.log("No phases found");
        return;
    }
    if (config.fault === true) {
        console.log("Faulted phase!");
        return;
    }
    for (let cycle = 1; cycle <= cycles; cycle++) {
        for (let i = 0; i < config.phases.length; i++) {
            let { color, duration } = config.phases[i];
            if (duration <= 0) {
                console.log("Invalid phase detected");
            }
            else {
                console.log(`Switching to ${color} for ${duration} s`);
            }
        }
    }
}

function generateTimeline(config, cycles) {
    let timeLine = [];
    let currTime = 0;
    for (let cycle = 1; cycle <= cycles; cycle++) {
        for (let i = 0; i < config.phases.length; i++) {
            let { duration } = config.phases[i];
            currTime += duration;
            timeLine.push(currTime);
        }
    }
    return timeLine;
}
```
