---
id: 694ca14c8ef00056ad146e51
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

After the duplicate ID check, push the `astronaut` into the `crew` array and log `Added ${astronaut.name} as ${astronaut.role}` to the console. You can log this in either of the following ways: 

```js
// method 1: with backticks
console.log(`Added ${astronaut.name} as ${astronaut.role}`);
// method 2: concatenation
console.log("Added " + astronaut.name + " as " + astronaut.role);
```

# --hints--

Your function should add the astronaut to the crew when the astronaut's id is unique.

```js
const crew = [];
const astro = { id: 2, name: "testAstronaut", role: "Scout", isEVAEligible: false, priority: 1 };
addCrewMember(crew, astro);
assert.deepEqual(crew[0], astro);
```

Your function should log the new astronaut's name when successfully added ("Added {astronaut.name} as {astronaut.role}).

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\.log\(\s*(?:`Added\s*\${astronaut\.name}\s*as\s*\${astronaut\.role}`|["']Added\s*["']\s*\+\s*astronaut\.name\s*\+\s*["']\s*as\s*["']\s*\+\s*astronaut\.role)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember (crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
--fcc-editable-region--


--fcc-editable-region--
}
```
