---
id: 694cabb9f1afa9446f36ffc0
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now you can add the first astronaut to the roster. Call `addCrewMember` with your `squad` array and `firstAstronaut`.

# --hints--

You should call your `addCrewMember` function with `squad` and `firstAstronaut`.

```js
assert.match(__helpers.removeJSComments(code), /addCrewMember\s*\(\s*squad\s*,\s*firstAstronaut\s*\)/);
```

After calling `addCrewMember`, your `squad` array should contain one astronaut.

```js
assert.strictEqual(squad.length, 1);
```

After calling `addCrewMember`, the only astronaut in the `squad` array should be `firstAstronaut`.

```js
assert.strictEqual(squad[0], firstAstronaut);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
  console.log(`Added ${astronaut.name} as ${astronaut.role}`);
}

--fcc-editable-region--

--fcc-editable-region--
```
