---
id: 694ca14c8ef00056ad146e51
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Push the `astronaut` object into the `crew` array and log `Added ${astronaut.name} as ${astronaut.role}` to the console. You can use either template literals or string concatenation for the log message.

# --hints--

Your `addCrewMember()` function should add the astronaut object to the `crew` array when the astronaut's `id` is unique.

```js
const crew = [];
const astro = { id: 2, name: "testAstronaut", role: "Test", isEVAEligible: false, priority: 1 };
addCrewMember(crew, astro);
assert.deepEqual(crew[0], astro);
```

Your `addCrewMember()` function should log the new astronaut's `name` when successfully added (`Added {astronaut.name} as {astronaut.role}`).

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(addCrewMember.toString()));
assert.match(
  cleaned,
  /console\.log\((?:`Added\$\{astronaut\.name\}as\$\{astronaut\.role\}`|["']Added["']\+astronaut\.name\+["']as["']\+astronaut\.role|["']Added["']\.concat\(astronaut\.name,["']as["']\)\.concat\(astronaut\.role\))\)/
);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember (crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
