---
id: 694c40d86fe6a78d20b6559c
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will leverage JavaScript to build a roster of astronauts for a space mission. 

To begin, create an empty array for your roster named `squad`.

# --hints--

You should create a `squad` variable.

```js
assert.isDefined(squad);
```

Your `squad` variable should be an array.

```js
assert.isArray(squad);
```

Your `squad` array should be empty.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
