---
id: 694c84330cbe604821b9deef
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Mission control requires a method for adding new members to the roster. Create an empty function named `addCrewMember` that accepts two parameters, `crew` and `astronaut`.

# --hints--

You should create a function named `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

Your `addCrewMember` function should have parameters `crew` and `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
