---
id: 694c84106e9023f289c03b84
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Your roster begins with a single commander astronaut. Create an object named `firstAstronaut` with the following properties:

| Key | Value   |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

You should create a `firstAstronaut` variable.

```js
assert.isDefined(firstAstronaut);
```

Your `firstAstronaut` should be an object.

```js
assert.isObject(firstAstronaut);
```

Your `firstAstronaut` should include the correct property values.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
