---
id: 694c9c186c089f7bdff0f3c6
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Generally, it is good practice to validate your inputs inside a function. You should first loop through the input `crew` array to check for a duplicate `id` (an id equal to the input astronaut's). If you find one, use `throw new Error("Duplicate ID: " + astronaut.id);` to indicate the error.

`throw new Error` creates an Error object and immediately stops the current function to signal a problem. If the Error is uncaught (e.g. outside of a `try/catch` block), as it is in our case, the error message will log to the browser console. Here is a quick example:

```js
if (object1.id === object2.id) {
    throw new Error("Duplicate ID: " + id);
}
```

# --hints--

Your function should use a `for` loop to scan existing `crew` members (any valid `for (...)` form is fine).

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*\)/);
```

Your function should throw an error when a duplicate `id` is found (use `throw new Error("Duplicate ID: " + astronaut.id);`).

```js
assert.throws(() => addCrewMember([{ id: 9 }], { id: 9 }), /duplicate\s+id:\s*9/i);
try {
  addCrewMember([{ id: 1 }], { id: 2 });
  assert.ok(true);
} catch (err) {
  assert.fail();
}
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember (crew, astronaut) {
--fcc-editable-region--

--fcc-editable-region--
}
```
