---
id: 694c84330cbe604821b9deef
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Mission control requires a method for adding new members to the roster. Create an empty `addCrewMember()` function that accepts two parameters, `crew` and `astronaut`. Note: for this step, the function should be empty.

# --hints--

You should create an `addCrewMember` function.

```js
assert.isFunction(addCrewMember);
```

Your function should have `crew` and `astronaut` parameters.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
