---
id: 694d8c16f3672c498d4656ac
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

As before, you should first validate the arguments of `swapCrewMembers()`. Ensure that the indices of the astronauts to be swapped are valid. Specifically, if: 

- `fromIndex` is negative

- OR `fromIndex` is greater than or equal to `crew.length`

- OR `toIndex` is negative

- OR `toIndex` is greater than or equal to `crew.length`

then, you should call `throw new Error` with the error message `"Invalid crew indices"`

# --hints--

Your `swapCrewMembers()` function should validate indices and, if invalid, throw an error with the message `"Invalid crew indices"`.

```js
const testCrew = [{ id: 1 }, { id: 2 }];
const errPattern = /Invalid\s+crew\s+indices/i;
assert.throws(() => swapCrewMembers(testCrew, -1, 1), errPattern);
assert.throws(() => swapCrewMembers(testCrew, 0, 5), errPattern);
assert.doesNotThrow(() => swapCrewMembers(testCrew, 0, 1));
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
--fcc-editable-region--

--fcc-editable-region--
}
```
