---
id: 694d98a87c2a7a33ca92cfa9
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Loop through `updatedCrew` and log each astronaut's `name`. Finally, after the loop, return `updatedCrew` to complete the `swapCrewMembers()` function.

# --hints--

Your function should log each astronaut's name in `updatedCrew` using a `for(...)` loop and `console.log(... .name ...)`.

```js
const cleaned = __helpers.removeJSComments(code).replace(/\s+/g, "");
assert.match(cleaned, /for\([^)]*updatedCrew[^)]*\)/, 'A for(...) loop over updatedCrew is required');
assert.match(cleaned, /console\.log[^;]*\.name/, 'console.log should log the .name property');
```

You should return `updatedCrew` after logging their names.

```js
const sample = [
  { name: "A" },
  { name: "B" },
  { name: "C" }
];
const result = swapCrewMembers(sample, 0, 2);
assert.isArray(result);
assert.deepEqual(result.map(a => a.name), ["C", "B", "A"]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) { throw new Error("Invalid crew indices"); }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];
--fcc-editable-region--

--fcc-editable-region--
}
```
