---
id: 694d8ddd106340237efb0c80
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

The `splice()` method is a powerful way to modify arrays, allowing you to add or remove elements from any position in an array, including the middle. `splice()` returns an array of items removed from the array. You can also swap two elements in an array using `splice()` with the following technique:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Here is an example applying the technique after copying an array:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

The technique (at the third line):

- removes the original element at `j` (index 3 in this example, or `55`)

- inserts the original element at `copyArray[1]` in its place (`97` is now at index 3)

- returns an array containing the removed element, which is assigned back to `copyArray[1]`, completing the swap

Use the `splice()` method on the `updatedCrew` array to swap the positions of its member astronauts at `fromIndex` and `toIndex`.

# --hints--

You should use `splice()` to perform the swap.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Calling `swapCrewMembers()` should not mutate the original `crew` array (you should only modify the local `updatedCrew` copy array).

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

Your `splice()` swap should assign with the `updatedCrew` array (this should only be one line, as in the examples).

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\s*\[\s*fromIndex\s*\]\s*=\s*updatedCrew\.splice\(\s*toIndex\s*,\s*1\s*,\s*updatedCrew\s*\[\s*fromIndex\s*\]\s*\)\s*\[\s*0\s*\]|updatedCrew\s*\[\s*toIndex\s*\]\s*=\s*updatedCrew\.splice\(\s*fromIndex\s*,\s*1\s*,\s*updatedCrew\s*\[\s*toIndex\s*\]\s*\)\s*\[\s*0\s*\]/;
assert.match(cleaned, regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) { throw new Error("Invalid crew indices"); }

  const updatedCrew = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
