---
id: 694dfd604c32cfe8e298784c
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

Mission control has requested you to swap two astronauts' positions, specifically `squad` members at indices `2` and `5`. Use `swapCrewMembers()` to swap them and store the result in a new array named `updatedSquad`.

# --hints--

You should store the result of `swapCrewMembers(squad, 2, 5)` in a new array named `updatedSquad`.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
assert.match(cleaned, /(let|const)updatedSquad=swapCrewMembers\(squad,?2,?5\)/);
```

The swap result should place `Felix` at index `2` and `Caroline` at index `5` while keeping length intact.

```js
const result = [...updatedSquad];
assert.strictEqual(result.length, squad.length);
assert.strictEqual(result[2].name, "Felix");
assert.strictEqual(result[5].name, "Caroline");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) { throw new Error("Invalid crew indices"); }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];
  for (let i = 0; i < updatedCrew.length; i++) {
    console.log(updatedCrew[i].name);
  }

  return updatedCrew; 
}

--fcc-editable-region--

--fcc-editable-region--
```
