---
id: 694d8ddbfc4bb6940b921bbc
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

After validating indices, you should make a shallow copy of `crew`. There are a few ways to accomplish this, but you must use the `slice()` method. The copy should be named `updatedCrew`.

Recall that `slice()` can be invoked without arguments to copy an entire array as such:

```js
const copyArray = originalArray.slice();
```

# --hints--

Ensure that the `swapCrewMembers()` method exists without error.

```js
const regex = __helpers.functionRegex('swapCrewMembers', ['crew', 'fromIndex', 'toIndex'], { capture: true });
const match = __helpers.removeJSComments(code).match(regex) || [];
assert.isArray(match, 'swapCrewMembers definition not found');
```

You should copy the input `crew` array with `slice`.

```js
const regex = __helpers.functionRegex('swapCrewMembers', ['crew', 'fromIndex', 'toIndex'], { capture: true });
const match = __helpers.removeJSComments(code).match(regex) || [];
assert.isArray(match, 'swapCrewMembers definition not found');
assert.isString(match[1], 'swapCrewMembers capture missing');
const noComments = __helpers.removeJSComments(code);
assert.match(noComments, /swapCrewMembers\s*\([^)]*\)\s*\{[\s\S]*?(let|const)\s+updatedCrew\s*=\s*crew\.slice\(\)/);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
  console.log(`Added ${astronaut.name} as ${astronaut.role}`);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    throw new Error("Invalid crew indices");
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
