---
id: 694eeb57faae3b2331add3e7
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

Back inside the `getEVAReadyCrew()` function, call `sortByPriorityDescending()` on the local `eligible` array so that it is sorted before being returned.

# --hints--

You should call `sortByPriorityDescending()` on the local `eligible` array inside of `getEVAReadyCrew()`.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(getEVAReadyCrew.toString()));
assert.match(cleaned, /(?:^|[^a-zA-Z])sortByPriorityDescending\(eligible\)/);
```

`getEVAReadyCrew()` should return only EVA-eligible astronauts, sorted by priority descending.

```js
const sample = [
  { name: "Low", isEVAEligible: true, priority: 1 },
  { name: "High", isEVAEligible: true, priority: 5 },
  { name: "Mid", isEVAEligible: true, priority: 3 },
  { name: "NoEVA", isEVAEligible: false, priority: 10 }
];
const result = getEVAReadyCrew(sample);
assert.deepEqual(result.map(a => a.name), ["High", "Mid", "Low"]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) { throw new Error("Invalid crew indices"); }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

function sortByPriorityDescending (crew) {
  for (let i = 0; i < crew.length - 1; i++) {
    for (let j = 0; j < crew.length - 1 - i; j++) {
      if (crew[j].priority < crew[j + 1].priority) {
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
      }
    }
  }
}

function getEVAReadyCrew (crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible)  eligible.push(astronaut);
  }
--fcc-editable-region--

--fcc-editable-region--

  return eligible;
}
```
