---
id: 695a4a70d5843b7f18c1050f
title: Step 30
challengeType: 1
dashedName: step-30
---

# --description--

Directly sort the local `sorted` array by invoking your helper function `sortByPriorityDescending()` on it.

# --hints--

You should invoke `sortByPriorityDescending` on the `sorted` array.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(printCrewSummary.toString()));
assert.match(cleaned, /(?:^|[^a-zA-Z])sortByPriorityDescending\(sorted\)/);
const viableSolutions = [`functionprintCrewSummary(crew){constsorted=crew.slice();sortByPriorityDescending(sorted);}`, `functionprintCrewSummary(crew){constsorted=crew.slice();sortByPriorityDescending(sorted)}`, `functionprintCrewSummary(crew){varsorted=crew.slice();sortByPriorityDescending(sorted);}`, `functionprintCrewSummary(crew){varsorted=crew.slice();sortByPriorityDescending(sorted)}`];
assert.ok(
  viableSolutions.includes(cleaned)
);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    throw new Error("Invalid crew indices");
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

function sortByPriorityDescending (crew) {
  for (let i = 0; i < crew.length - 1; i++) {
    for (let j = 0; j < crew.length - 1 - i; j++) {
      if (crew[j].priority < crew[j + 1].priority) {
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
      }
    }
  }
}

function getEVAReadyCrew (crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible)  eligible.push(astronaut);
  }
  sortByPriorityDescending(eligible); 

  return eligible;
}

const EVAReadySquad = getEVAReadyCrew(updatedSquad);
function chunkCrew(crew, size) {
  if (size < 1) {
    throw new Error("Chunk size must be >= 1");
  }

  const chunks = [];
  for (let i = 0; i < crew.length; i += size) {
    chunks.push(crew.slice(i, i + size));
  }

  return chunks;
}

const EVAChunks = chunkCrew(EVAReadySquad, 3);

function printCrewSummary(crew) {
  const sorted = crew.slice();
--fcc-editable-region--

--fcc-editable-region--
}
```
