---
id: 695a4b4afad548377ea2f146
title: Step 31
challengeType: 1
dashedName: step-31
---

# --description--

Loop through the members of the `sorted` array and log each of their names.

# --hints--

Use a `for` loop to iterate through and log the names of each astronaut in the `sorted` array.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(printCrewSummary.toString()))
assert.match(cleaned, /for\s*\(/);
```

You must correctly log the names of each astronaut in the `sorted` array, once each and in order.

```js
function captureConsole(fn) {
  const logs = [];
  const originalLog = console.log;
  console.log = (...args) => logs.push(args.join(" "));
  try {
    fn();
  } finally {
    console.log = originalLog;
  }
  return logs;
}
const testCrew = [
  { name: "Andy", priority: 3 },
  { name: "Bart", priority: 8 },
  { name: "Caroline", priority: 4 },
  { name: "Diego", priority: 1 },
  { name: "Elise", priority: 7 },
  { name: "Felix", priority: 6 },
  { name: "Gertrude", priority: 4 },
  { name: "Hank", priority: 2 },
  { name: "Irene", priority: 5 },
  { name: "Joan", priority: 1 }
];
const logs = captureConsole(() => printCrewSummary(testCrew));
const expectedOrder = [
  "Bart",
  "Elise",
  "Felix",
  "Irene",
  "Caroline",
  "Gertrude",
  "Andy",
  "Hank",
  "Diego",
  "Joan"
];
assert.strictEqual(
  logs.length,
  expectedOrder.length,
  "You must log each astronaut's name exactly once"
);
assert.deepEqual(
  logs,
  expectedOrder,
  "Names must be logged in descending priority order"
);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      throw new Error("Duplicate ID: " + astronaut.id);
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    throw new Error("Invalid crew indices");
  }

  const updatedCrew = crew.slice();
  updatedCrew[fromIndex] = updatedCrew.splice(toIndex, 1, updatedCrew[fromIndex])[0];

  return updatedCrew; 
}

const updatedSquad = swapCrewMembers(squad, 2, 5);

function sortByPriorityDescending (crew) {
  for (let i = 0; i < crew.length - 1; i++) {
    for (let j = 0; j < crew.length - 1 - i; j++) {
      if (crew[j].priority < crew[j + 1].priority) {
        const temp = crew[j];
        crew[j] = crew[j + 1];
        crew[j + 1] = temp;
      }
    }
  }
}

function getEVAReadyCrew (crew) {
  const eligible = [];
  for (const astronaut of crew) {
    if (astronaut.isEVAEligible)  eligible.push(astronaut);
  }
  sortByPriorityDescending(eligible); 

  return eligible;
}

const EVAReadySquad = getEVAReadyCrew(updatedSquad);
function chunkCrew(crew, size) {
  if (size < 1) {
    throw new Error("Chunk size must be >= 1");
  }

  const chunks = [];
  for (let i = 0; i < crew.length; i += size) {
    chunks.push(crew.slice(i, i + size));
  }

  return chunks;
}

const EVAChunks = chunkCrew(EVAReadySquad, 3);

function printCrewSummary(crew) {
  const sorted = crew.slice();
  sortByPriorityDescending(sorted); 
--fcc-editable-region--

--fcc-editable-region--
}
```
