---
id: 66c63bfa4685e4e3843fa0dc
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop you will continue to learn about arrays by building a grocery shopping list.

Start by adding a `console.log` statement that logs the string `"Let's create a grocery shopping list"` to the console.

**NOTE**: If you are going to use single quotes in for your string, remember to escape them with a backslash (`\`) like this: `'Let\'s`.

# --hints--

You should have a `console.log` in your code.

```js
assert.match(code, /console\.log(.*)/);
```

Your console statement should output the message `"Let's create a grocery shopping list"`.

```js
assert.match(code, /console\.log\((["'])(Let\'s\s+create\s+a\s+grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
