---
id: 66c645b345a39bf04864dc50
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

To see the updated shopping list, log the message `Current Shopping List: <shopping list variable goes here>` to the console.

You should use template literal syntax to construct the string. 

Make sure to replace `<shopping list variable goes here>` with the `shoppingList` array.

# --hints--

You should log the message `Current Shopping List: <shopping list variable goes here>` to the console. Remember to use correct template literal syntax.

```js
assert.match(code, /console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Let's create a grocery shopping list");
const shoppingList = [];
console.log("It will be nice to have some fruit to eat. Let's add some to our list.");
shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
