---
id: 66c64ad16796c7f2419b45c5
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now it is time to start adding items to the top of the grocery list.

Start by adding a `console.log` statement that logs the message `"It looks like we are out of cooking oil at home. Let's put that at the top of the list so we get that first."`

Remember the properly escape the string if you are using single quotes instead of double quotes.

# --hints--

You should log the message `"It looks like we are out of cooking oil at home. Let's put that at the top of the list so we get that first."` to the console.

```js
assert.match(code, /console\.log\(("')(It\s+looks\s+like\s+we\s+are\s+out\s+of\s+cooking\s+oil\s+at\s+home\.\s+Let's\s+put\s+that\s+at\s+the\s+top\s+of\s+the\s+list\s+so\s+we\s+get\s+that\s+first\.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Let's create a grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat. Let's add some to our list.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

--fcc-editable-region--

--fcc-editable-region--
```
