---
id: 66c6491fe4c8e0f16845425f
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now it is time to add another fruit to the list.

Using the same array method as earlier, add the string `"Grapes"` to the `shoppingList` array.

Then, use the same format from the previous step to log the message `Current Shopping List: ` followed by the `shoppingList` array to the console.

# --hints--

Your `shoppingList` array should now contain two items.

```js
assert.lengthOf(shoppingList, 2);
```

Your `shoppingList` array should contain the string `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

You should log the message `Current Shopping List: ` followed by the `shoppingList` array to the console. Remember to use correct template literal syntax.

```js
assert.lengthOf(code.match(/console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\);?/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

--fcc-editable-region--

--fcc-editable-region--
```
