---
id: 66c6491fe4c8e0f16845425f
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now it is time to add another fruit to the list.

Using the same array method as earlier, add the string `"Grapes"` to the `shoppingList` array.

Then add the same console statement you used in the previous step to log the updated shopping list to the console.

# --hints--

Your `shoppingList` array should now contain two items.

```js
assert.lengthOf(shoppingList, 2);
```

Your `shoppingList` array should contain the string `"Grapes"`.

```js
assert.includeMembers(shoppingList, ["Grapes"]);
```

You should log the message `Current Shopping List: <shopping list variable goes here>` to the console. Remember to use correct template literal syntax.

```js
assert.match(code, /console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Let's create a grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat. Let's add some to our list.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

--fcc-editable-region--

--fcc-editable-region--
```
