---
id: 66c645b345a39bf04864dc50
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

To see the updated shopping list, you will need to log the current shopping list and a short message to the console.

You will be doing this task repeatedly throughout the workshop so here is what the format will look like:

```js
console.log(`Message goes here: ${variableNameGoesHere}`);
```

Using template literal syntax and the format shown above, log the message `Current Shopping List: ` followed by the `shoppingList` array to the console.


# --hints--

You should log the message `Current Shopping List: ` followed by the `shoppingList` array to the console. Refer to the format shown in the example.

```js
assert.match(code, /console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
