---
id: 66c6491fe4c8e0f16845425f
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now it is time to add another fruit to the list.

Using the same array method as earlier, add the string `"Grapes"` to the end of the `shoppingList` array.

Then, add a `console.log` and call the `getShoppingListMsg` function inside of the `console.log` to see the updated list logged to the console.

# --hints--

Your `shoppingList` array should now contain two items.

```js
assert.lengthOf(shoppingList, 2);
```

Your `shoppingList` array should contain the string `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

You should call the `getShoppingListMsg` function inside of the `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
