---
id: 66c74079c30b1c2d166cb9a4
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

In the lecture videos, you learned how to remove an item from the beginning of the array using the `shift` method. 

Here is a reminder of how to use the `shift` method:

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

Use the `shift` method to remove the first item from the `shoppingList` array.

# --hints--

Your `shoppingList` array should have a total of `5` items.

```js
assert.lengthOf(shoppingList, 5);
```

You should use the `shift` method to remove the first item from the `shoppingList` array.

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg());

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg());

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
