---
id: 66c72a55418cc9247b710827
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now it is time to log the updated `shoppingList` array to the console.

Start by adding another `console.log()` statement to the code. Inside the `console.log()` statement, use template literal syntax to log the message `Current Shopping List:` followed by the `shoppingList` array.

# --hints--

You should use template literal syntax to log the message `Current Shopping List:` followed by the `shoppingList` array.

```js
assert.match(code, /console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Let's create a grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat. Let's add some to our list.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we are out of cooking oil at home. Let's put that at the top of the list so we get that first.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
