---
id: 66c72f4d0528bd268a82107b
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Now it is time to log the updated `shoppingList` array to the console.

Start by adding another `console.log()` statement to the code. Inside the `console.log()` statement, use template literal syntax to log the message `Current Shopping List:` followed by the `shoppingList` array.

# --hints--

You should use template literal syntax to log the message `Current Shopping List:` followed by the `shoppingList` array.

```js
assert.match(code, /console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
