---
id: 66c730183f4020275cbf0611
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

For this next portion of the workshop, you will review how to remove items from the end of the array.

Start by adding a `console.log` that logs the message `"This looks like too much junk food."`.

# --hints--

You should log the message `"This looks like too much junk food."`.

```js
assert.match(code, /console\.log\((['"])(This\s+looks\s+like\s+too\s+much\s+junk\s+food.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
