---
id: 66c72f4d0528bd268a82107b
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Now it is time to log the updated `shoppingList` array to the console.

Add another `console.log` and call the `getShoppingListMsg` function inside of the `console.log` to see the updated list logged to the console.

# --hints--

You should call the `getShoppingListMsg` function inside of the `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
