---
id: 66c72a55418cc9247b710827
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now it is time to log the updated `shoppingList` array to the console.

Use template literal syntax, like in the previous steps, to log the message `Current Shopping List: ` followed by the `shoppingList` array to the console.

# --hints--

You should log the message `Current Shopping List: ` followed by the `shoppingList` array to the console. Remember to use correct template literal syntax.

```js
assert.lengthOf(code.match(/console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\);?/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
