---
id: 66c726c34ecf1e238aa9d7d5
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

In the lecture videos, you learned how to add elements to the beginning of an array using the `unshift()` method. 

Here is a reminder of how to use the `unshift()` method:

```js
array.unshift(item1, item2, ..., itemX);
```

Use the `unshift()` method to add the string `"Vegetable Oil"` to the beginning of the `shoppingList` array.

# --hints--

You should have three items in your `shoppingList` array.

```js
assert.lengthOf(shoppingList, 3);
```

You should use the `unshift()` method to add `"Vegetable Oil"` to the beginning of the `shoppingList` array.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
