---
id: 66c73fa7433e082c4be096b1
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

In this last part of the workshop, you will review how to remove an item from the beginning of an array.

Start by adding a `console.log` statement that logs the message `"On second thought, maybe we should be more health conscious. Let's remove the dessert from the list."`.

# --hints--

You should have a `console.log` that logs the message `"On second thought, maybe we should be more health conscious. Let's remove the dessert from the list."`.

```js
assert.match(code, /console\.log\(("')(On\s+second\s+thought,\s+maybe\s+we\s+should\s+be\s+more\s+health\s+conscious\.\s+Let's\s+remove\s+the\s+dessert\s+from\s+the\s+list\.)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("This looks like to much junk food. Let's remove the last item from the list.");

shoppingList.pop();
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It might be nice to get a dessert. Let's add that to the top of the list so we get that first.");

shoppingList.unshift("Chocolate Cake");
console.log(`Current Shopping List: ${shoppingList}`);

--fcc-editable-region--

--fcc-editable-region--
```
