---
id: 66c73a7798f6f62b2ae58f22
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Now it is time to add more items to the beginning of the grocery list.

Start by adding a `console.log` statement that logs the message `"It might be nice to get a dessert. Let's add that to the top of the list so we get that first."`

Below that `console` statement, use the correct array method to add the string `"Chocolate Cake"` to the beginning of the `shoppingList` array.

Finally, add a `console.log` statement that logs the message `"Current Shopping List: "` followed by the `shoppingList` array. Remember to use template literal syntax for the message.

# --hints--

You should have a `console.log` that logs the message `"It might be nice to get a dessert. Let's add that to the top of the list so we get that first."`.

```js
assert.match(code, /console\.log\(("')(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.\s+Let's\s+add\s+that\s+to\s+the\s+top\s+of\s+the\s+list\s+so\s+we\s+get\s+that\s+first\.)\1\)/);
```

Your `shoppingList` array should have a total of `6` items.

```js
assert.lengthOf(shoppingList, 6);
```

You should use the `unshift` method to add the string `"Chocolate Cake"` to the beginning of the `shoppingList` array.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

You should use template literal syntax to log the message `Current Shopping List:` followed by the `shoppingList` array.

```js
assert.match(code, /console\.log\(`Current\s+Shopping\s+List:\s+\${shoppingList}`\);?/);
```


# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("This looks like to much junk food. Let's remove the last item from the list.");

shoppingList.pop();
console.log(`Current Shopping List: ${shoppingList}`);

--fcc-editable-region--

--fcc-editable-region--
```
