---
id: 66c730183f4020275cbf0611
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

For this next portion of the workshop, you will review how to remove items from the end of the array.

Start by adding a `console.log` that logs the message `"This looks like to much junk food. Let's remove the last item from the list."`.

# --hints--

You should log the message `"This looks like to much junk food. Let's remove the last item from the list."`.

```js
assert.match(code, /console\.log\(("')(This\s+looks\s+like\s+to\s+much\s+junk\s+food\.\s+Let's\s+remove\s+the\s+last\s+item\s+from\s+the\s+list\.?)\1\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(`Current Shopping List: ${shoppingList}`);

--fcc-editable-region--

--fcc-editable-region--
```
