---
id: 66c726c34ecf1e238aa9d7d5
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

In the lecture videos, you learned how to add elements to the beginning of an array using the `unshift()` method. 

Here is a reminder of how to use the `unshift()` method:

```js
const arr = [1, 2, 3];
arr.unshift(0);

// [0, 1, 2, 3]
console.log(arr); 
```

Use the `unshift()` method to add the string `"Vegetable Oil"` to the beginning of the `shoppingList` array.

# --hints--

You should have three items in your `shoppingList` array.

```js
assert.lengthOf(shoppingList, 3);
```

You should use the `unshift()` method to add `"Vegetable Oil"` to the beginning of the `shoppingList` array.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Let's create a grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat. Let's add some to our list.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we are out of cooking oil at home. Let's put that at the top of the list so we get that first.");

--fcc-editable-region--

--fcc-editable-region--
```
