---
id: 66c730ee6ae076281721d0b9
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

In the previous lecture videos, you learned how to remove items from the end of an array using the `pop` method. 

Here is reminder of how to use the `pop` method:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Use the `pop` method to remove the last item from the `shoppingList` array.

# --hints--

Your `shoppingList` array should now have `5` items in it.

```js
assert.lengthOf(shoppingList, 5);
```

Your `shoppingList` array should no longer have the string `"Potato Chips"` in it.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg());

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
