---
id: 66c72b0ffbc5522525768558
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

In earlier steps, you reviewed how to add an item to the end of the array using the `push` method. 

But the `push` method accepts multiple arguments, so you can add multiple items to the end of the array like this:

```js
array.push(item1, item2, item3);
```

In this step, use the `push` method to add the strings `"Popcorn", "Beef Jerky", "Potato Chips"` to the `shoppingList` array.

The order is important, so make sure to add the items in the order they are listed.

# --hints--

Your `shoppingList` array should have a total of `6` strings.

```js
assert.lengthOf(shoppingList, 6);
```

Your `shoppingList` should have the string `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Your `shoppingList` should have the string `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Your `shoppingList` should have the string `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg() {
  return `Current Shopping List: ${shoppingList}`;
}

console.log(getShoppingListMsg());

shoppingList.push("Grapes");
console.log(getShoppingListMsg());

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg());

--fcc-editable-region--

--fcc-editable-region--
```
