---
id: 66c742d045c9fc2e09fa64b1
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

The last change to make to the grocery list is to update the first item in the list.

In the previous lecture videos, you learned how to update an item using bracket notation and the index of the item you want to update.

Here is a reminder of how to update an item in an array:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

Update the first item in the `shoppingList` array to be `"Canola Oil"`.

# --hints--

You should not remove or add any items to the `shoppingList` array.

```js
assert.lengthOf(shoppingList, 6);
```

You should update the first item in the `shoppingList` array to be `"Canola Oil"`.

```js
assert.equal(shoppingList[0], "Canola Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Let's create a grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat. Let's add some to our list.");

shoppingList.push("Apples");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Grapes");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It looks like we are out of cooking oil at home. Let's put that at the top of the list so we get that first.");

shoppingList.unshift("Vegetable Oil");
console.log(`Current Shopping List: ${shoppingList}`);

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("This looks like to much junk food. Let's remove the last item from the list.");

shoppingList.pop();
console.log(`Current Shopping List: ${shoppingList}`);

console.log("It might be nice to get a dessert. Let's add that to the top of the list so we get that first.");

shoppingList.unshift("Chocolate Cake");
console.log(`Current Shopping List: ${shoppingList}`);

console.log("On second thought, maybe we should be more health conscious. Let's remove the dessert from the list.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
