---
id: 61696ef7ac756c829f9e4048
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

في هذا المشروع ستعمل مع ملف CSS خارجي لتصميم الصفحة. لقد قمنا بالفعل بإنشاء ملف `styles.css` لك. ولكن قبل أن تتمكن من استخدامه، ستحتاج إلى ربطه بالصفحة.

Nest a `link` element within the `head` element. اعطيها سمة `rel` بقيمة `stylesheet` و `href` بقيمة `styles.css`.

# --hints--

الكود الخاص بك يجب أن يحتوي على عنصر `link` مغلق ذاتيا.

```js
assert(document.querySelectorAll('link').length === 1);
```

عنصر `link` الخاص بك يجب أن يكون داخل عنصر `head` الخاص بك.

```js
assert.exists(document.querySelector('head > link'));
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `rel` بقيمة `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

عنصر `link` الخاص بك يجب أن يحتوي على سمة `href` بقيمة `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
