---
id: 616971b3cd990186b66c99a1
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

الآن بعد إعداد ملف CSS الخارجي الخاص بك، يمكنك البدء في تصميم الصفحة.

كتذكير، إليك كيفية استهداف عنصر paragraph و محاذاته الي اليمين:

```css
p {
  text-align: right;
}
```

قم بإنشاء قاعدة CSS جديدة تستهدف عنصر `h1` ، وقم بتعيين خاصية `text-align` إلى `center`.

# --hints--

يجب عليك استخدام منتقي `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

قاعدة CSS الخاصة بك لـ `h1` يجب أن تحتوي على خاصية `text-align` بقيمة `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
