---
id: 61695197ac34f0407e339882
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

كما رأيت في المشاريع السابقة، ينبغي أن تبدأ صفحات الويب بتعريف `DOCTYPE html` ، يليه عنصر `html`.

قم بإضافة تعريف `DOCTYPE html` في الجزء العلوي من الوثيقة، وعنصر `html` بعد ذلك. اعطي عنصر `html` سمة `lang` بقيمة `en`.

# --hints--

يجب أن يحتوي الكود الخاص بك على `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب عليك تضمين مسافة بعد `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب عليك تحديد نوع المستند ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب عليك إغلاق تعريف `DOCTYPE` مع `>` بعد type.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على علامة افتتاحية مع سمة `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` الخاص بك على closing tag.

```js
assert(code.match(/<\/html\s*>/gi));
```

يجب أن يكون تعريف `DOCTYPE` في بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
