---
id: 61695ab9f6ffe951c16d03dd
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

قم بإضافة عنصر `head` ضمن عنصر `html`. بعد عنصر `head` مباشرة، أضف عنصر `body`.

# --hints--

يجب أن يكون لديك الـ opening tag للـ `head`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك الـ opening tag للـ `body`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون هناك علامة فتح `head` واحدة فقط.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

يجب أن يكون هناك علامة أغلاق `head` واحدة فقط.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

يجب أن يكون هناك علامة فتح `body` واحدة فقط.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

يجب أن يكون هناك علامة أغلاق `body` واحدة فقط.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

يجب أن يكون عنصر `head` فارغاً.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

يجب أن يكون عنصر `body` فارغاً.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

يجب أن يأتي عنصر `body` الخاص بك بعد عنصر `head` الخاص بك.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
