---
id: 61695c4aad56f95497c19583
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

تذكر أن عنصر `title` يعطي محركات البحث معلومات إضافية عن الصفحة. It also displays the content of that `title` element in two more ways:

* in the title bar when the page is open
* in the browser tab for the page when you hover on it. Even if that tab is not active, once you hover on the tab, the `title` text is displayed.

داخل عنصر `head`، قم بدمج عنصر `title` مع النص `Colored Markers`.

# --hints--

يجب أن يكون لديك الـ opening tag للـ `title`.

```js
assert(code.match(/<title\s*>/i));
```

يجب أن يكون لديك الـ closing tag للـ `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

يجب أن يكون عنوان مشروعك `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

تذكر أن الـ casing، اي حالة الحرف سواء كبير او صغير (capital or small) والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
