---
id: 61695d1fbc003856628bf561
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

لأخبار المتصفحات عن كيفية ترميز الأحرف ( encode characters) على صفحتك، قم بتعيين `charset` إلى `utf-8`. `utf-8` هي مجموعة أحرف عالمية تتضمن تقريبًا كل حرف من جميع اللغات البشرية.

داخل عنصر `head` قم بدمج عنصر `meta` مع السمة `charset` بقيمة `utf-8`. تذكر أن `meta` عناصر مغلقة ذاتيا، ولا تحتاج إلى علامة إغلاق.

# --hints--

يجب أن يكون لديك عنصر `meta` واحد.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

يجب أن يكون عنصر `meta` عنصر مغلق ذاتيا (self-closing element).

```js
assert(code.match(/<\/meta>/i) === null);
```

يجب أن يحتوي عنصر `meta` على السمة `charset` بقيمة `utf-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
