---
id: 616965351e74d4689eb6de30
title: الخطوة ٥
challengeType: 0
dashedName: step-5
---

# --description--

ضف عنصر `meta` ذاتي الأغلاق داخل `head`. أعطيه سمة `name` بقيمة `viewport` و سمة `content` بقيمة `width=device-width, initial-scale=1.0` حتى تطابق مظهر صحفتك في كل الأجهزة.

# --hints--

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

يجب أن يكون عنصر `meta` مغلق ذاتيا (self-closing).

```js
assert(code.match(/<\/meta>/i) === null);
```

يجب أن يحتوي عنصر `meta` الجديد على السمة `name` بقيمة `viewport`، و السمة `content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
