---
id: 616d525007b8c5d8b3308b1c
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

سيبدو الـ marker الخاص بك أفضل إذا تم توسيطه على الصفحة. طريقة سهلة للقيام بذلك هي باستخدام الـ <dfn>shorthand property</dfn> للـ `margin`.

في المشروع الأخير، قمت بتعيين منطقة الـ margin للعناصر بشكل منفصل مع خصائص مثل `margin-top` و `margin-left`. خاصية اختصار `margin` تجعل من السهل تعيين مناطق margin متعددة في نفس الوقت.

لتوسيط marker الخاص بك في الصفحة، عيين خاصية `margin` إلى `auto`. هذا يعين `margin-top`, و `margin-right`, و `margin-bottom`, و `margin-left` إلى `auto`.

# --hints--

يجب أن تكون لقاعدة CSS الخاصة بك `.marker` خاصية `margin` بقيمة `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
