---
id: 616d595270d902f0e7086e18
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

مع أنّ لديك ثلاث عناصر `div` منفصلة من marker فإنها تبدو مثل مستطيل كبير واحد. يجب عليك إضافة بعض المساحة بينهم لجعل رؤية كل عنصر أسهل.

عندما تكون الخاصية المختصرة `margin` ذات قيمتين، يعين `margin-top` و `margin-bottom` إلى أول قيمة، و `margin-left` و `margin-right` إلى القيمة الثانية.

في قاعدة CSS الخاصة بك من `.marker`، عيّن الخاصية `margin` إلى `10px auto`.

# --hints--

يجب أن تكون لقاعدة CSS الخاصة بك `.marker` خاصية `margin` بقيمة `10px auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
