---
id: 616d3a67ccf800ad94ec89ae
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

الآن ستضيف بعض العناصر التي ستصممها في نهاية المطاف لتكون لون markers.

First, within the `body` element, add a `div` element and set its `class` attribute to `container`. تأكد من أن عنصر `div` تحت عنصر `h1`.

# --hints--

يجب أن يكون لديك العلامة مفتوحة (opening tag) الآتية `div`.

```js
assert(code.match(/<div.*>/i));
```

يجب أن يكون لديك العلامة المغلقة (closing tag) الآتية `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

عنصر `div` الخاص بك يجب أن يكون داخل عنصر `body` الخاص بك.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

عنصر `div` الخاص بك يجب أن يكون بعد عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

عنصر `div` الخاص بك يجب أن يحتوي على سمة `class` بقيمة `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
