---
id: 616d50b93ba424d6282c99cf
title: الخطوة ١٢
challengeType: 0
dashedName: step-12
---

# --description--

لاحظ أن marker الخاص بك لا يبدو أن له أي لون. تم بالفعل تطبيق لون الخلفية، ولكن بما أن عنصر `div` للـ marker فارغ، فليس لديه أي ارتفاع افتراضي.

في قاعدة CSS مسمى `.marker` الخاص بك، عيين خاصية `width` إلى `200px` و `height` إلى `25px`.

# --hints--

يجب أن تكون لقاعدة CSS الخاصة بك `.marker` خاصية `width` بقيمة `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

يجب أن تكون لقاعدة CSS الخاصة بك `.marker` خاصية `height` بقيمة `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
