---
id: 6176519636a76b810ab1219a
title: الخطوة 20
challengeType: 0
dashedName: step-20
---

# --description--

أنشئ قاعدة CSS جديدة تستهدف فئة `two` وتعين خاصية `background-color` إلى`green`.

وأنشئ قاعدة CSS أخرة, تستهدف فئة `three` وتعين خاصية `background-color` إلى `blue`.

# --hints--

يجب عليك استخدام مستهدف لفئة لاستهداف فئة `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

يجب أن يكون لقاعدة CSS مسمى `.two` خاصية `background-color` بقيمة `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

يجب عليك استخدام مستهدف لفئة لاستهداف فئة `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

يجب أن يكون لقاعدة CSS مسمى `.three` خاصية `background-color` بقيمة `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
