---
id: 61764c602bee6974e7790f35
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

لإعطاء markers ألوان مختلفة، ستحتاج إلى إضافة فئة فريدة لكل واحدة. يمكن إضافة فئات متعددة إلى عنصر بإدراجها في سمة `class` وفصلها بمسافة. على سبيل المثال، يضيف ما يلي class تسمى `animal` و `dog` إلى عنصر `div`.

```html
<div class="animal dog">
```

يجب عليك إضافة class يسمى `one` إلى العنصر `div` الأول ألى marker.


# --hints--

يجب عليك إضافة class يسمى `one` إلى العنصر `div` الأول ألى marker.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

يجب أن يحتوي عنصر الأول `div` من marker على class تسمى `marker` و `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
