---
id: 617bbb6b97a83f6d1f7d6e38
title: الخطوة 31
challengeType: 0
dashedName: step-31
---

# --description--

الآن بعد أن اصبحت على دراية بالألوان الثانوية، ستتعلم كيفية إنشاء <dfn>tertiary colors</dfn> (الألوان الثلاثية). يتم إنشاء الألوان الثلاثية بدمج لون أساسي مع لون ثانوي قريب.

لإنشاء اللون الثالث البرتقالي، قم بتحديث دالة `rgb` في قاعدة CSS لـ `.one` بحيث يكون الاحمر بقيمة الحد الأقصى، و الأخضر بقيمة `127`.

# --hints--

يجب أن يكون لقاعدة CSS لـ `.one` الخاصية `background-color` بقيمة `rgb(255, 127, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
