---
id: 617bc3386dc7d07d6469bf20
title: الخطوة 32
challengeType: 0
dashedName: step-32
---

# --description--

لاحظ أنه لإنشاء البرتقالي كان عليك زيادة كثافة اللون الأحمر وتقليل كثافة قيم `rgb` الخضراء. وذلك لأن البرتقالي هو مزيج من اللون الأحمر والأصفر، ويسقط بين اللونين على عجلة الألوان.

لإنشاء ربيع اللون الثالث أخضر، الجمع بين السيان والأخضر. قم بتحديث وظيفة `rgb` في قاعدة CSS مسمى `.two` بحيث تكون قيمة اللون الأخضر بالحد الأقصى، وقيمة اللون الأزرق `127`.

# --hints--

يجب أن يكون لقاعدة CSS مسمى `.two` الخاصية `background-color` بقيمة `rgb(0, 255, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
