---
id: 617b8e0d93a8d10d9a90e720
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

لاحظ أن `background-color` من marker الخاص بك لا يزل أحمر. هذا لأنك عيّنت القيمة الحمراء للوظيفة `rgb` إلى الحد الأقصى `255`، أو 100٪ أحمر، وتعيين القيم الخضراء والزرقاء على حد سواء إلى `0`.

الآن استخدم وظيفة `rgb` لتعيين الألوان الأخرى.

في قاعدة CSS مسمى `.two`، استخدم الوظيفة `rgb` لتعيين `background-color` إلى الحد الأقصى للقيمة الخضراء، و `0` للقيم الأخرى. وفي قاعدة CSS مسمى `.three`، استخدم الوظيفة `rgb` لتعيين `background-color` إلى القيمة القصوى للأزرق، و `0` للقيم الأخرى.

# --hints--

لا ينبغي لقاعدة CSS الخاصة بك `.two` أن تستخدم كلمة `green` لتحديد خاصية `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

يجب أن يكون لقاعدة CSS مسمى `.two` الخاصية `background-color` بقيمة `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

لا ينبغي لقاعدة CSS الخاصة بك `.three` أن تستخدم كلمة `blue` لتحديد خاصية `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

يجب أن يكون لقاعدة CSS مسمى `.three` الخاصية `background-color` بقيمة `rgb(0, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
