---
id: 617b97abd9f3f61d1590b815
title: الخطوة 26
challengeType: 0
dashedName: step-26
---

# --description--

في نموذج ألوان RGB الأضافية، يكون <dfn>ألوان الأساسية</dfn> هي الألوان التي عند الجمع بينها تنشئ لون ابيض نقي. ولكن لكي يحدث هذا ، يجب أن يكون كل لون بأقصى حد له.

قبل أن تجمع الألوان، عيّن marker الأخضر مرة أخرى إلى اللون الأخضر النقي. بالنسبة لوظيفة `rgb` في قاعدة CSS مسمى `.two`، عيّن الأخضر إلى القيمة القصوى `255`.

# --hints--

يجب أن يكون لقاعدة CSS مسمى `.two` الخاصية `background-color` بقيمة `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
