---
id: 617b8b38f32bf2080a140675
title: الخطوة ٢٢
challengeType: 0
dashedName: step-22
---

# --description--

الوظيفة (function) هي جزء من الكود البرمجي التي يمكن أن تأخذ مدخلات وتنفذ إجراء محددا. تقوم وظيفة CSS تدعي `rgb` بقبول القيم، أو <dfn>حجج</dfn>، للأحمر والأخضر والأزرق، وتنتج لون:

```css
rgb(red, green, blue);
```

كل قيمة حمراء وخضراء وزرقاء هي رقم من `0` إلى `255`. `0` يعني أن هناك 0% من هذا اللون، وهو أسود. `255` يعني أن هناك 100% من هذا اللون.

في قاعدة CSS مسمى `.one` استبدل كلمة اللون `red` بالوظيفة `rgb`. بالنسبة لوظيفة `rgb`، عيين قيمة الأحمر إلى `255`، قيمة اللون الأخضر إلى `0`، والقيمة للأزرق إلى `0`.

# --hints--

لا ينبغي لقاعدة CSS الخاصة بك `.one` أن تستخدم كلمة `red` لتحديد خاصية `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

يجب أن يكون لقاعدة CSS مسمى `.one` الخاصية `background-color` بقيمة `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
