---
id: 617b92b9de349513466f6156
title: الخطوة 24
challengeType: 0
dashedName: step-24
---

# --description--

بينما تبدو markers الحمراء والزرقاء متماثلة، فإن marker الأخصر أخف بكثير مما كان عليه من قبل. هذا لأن الكلمة `green` هي في الواقع ظل أغمق، وتقع في منتصف المسافة بين الأسود وأقصى قيمة للأخضر.

في قاعدة CSS لـ `.two`، قم بتعيين القيمة الخضراء في دالة `rgb` إلى `127` لتقليل شدتها.

# --hints--

يجب أن يكون لقاعدة CSS مسمى `.two` الخاصية `background-color` بقيمة `rgb(0, 127, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
