---
id: 617b954d9f4f6217a749380e
title: الخطوة 25
challengeType: 0
dashedName: step-25
---

# --description--

الآن أضف مساحة رأسية بين markers الخاصة بك وحافة عنصر `container` الموجودة فيه.

في قاعدة CSS مسمى `.container`، استخدم الخاصية القصيرة `padding` بقيمة `10px` إلى top padding و bottom padding، وعيّن left padding و right padding إلى `0`. يعمل هذا بشكل مماثل مع خاصية اختصار `margin` الذي استخدمتها من قبل.

# --hints--

لا يجب عليك إزالة خاصية `background-color` وقيمتها من قاعدة CSS مسمى `.container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

يجب أن تكون لقاعدة CSS الخاصة بك `.container` خاصية `padding` بقيمة `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
