---
id: 618a0927005553b74bfa05ae
title: الخطوة 36
challengeType: 0
dashedName: step-36
---

# --description--

عجلة الألوان هي دائرة تكون فيها الألوان المتشابهة أو <dfn>hues</dfn> (تدرجات) قريبة من بعضها البعض والألوان المختلفة متباعدة. على سبيل المثال ، يقع اللون الأحمر الخالص بين درجات اللون الوردي والبرتقالي.

يطلق على لونين متقابلين على عجلة الألوان <dfn>complementary colors</dfn> (ألوان تكميلية). وإذا تم الجمع بين لونين متكاملين، فإنهما ينتجان الرمادي. لكن عندما توضع جنبا إلى جنب هذه الألوان تنتج تناقضا بصريا قويا وتبدو أكثر إشراقا.

في دالة `rgb` في قاعدة CSS لـ `.one`، قم بتعيين القيمة الحمراء إلى الحد الأقصى من `255` لإنتاج اللون الأحمر النقي. في دالة `rgb` في قاعدة CSS لـ `.two`، قم بتعيين القيمة الخضراء و الزرقاء إلى الحد الأقصى من `255` لإنتاج اللون cyan.

# --hints--

يجب أن يكون لقاعدة CSS لـ `.one` الخاصية `background-color` بقيمة `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

يجب أن يكون لقاعدة CSS لـ `.two` الخاصية `background-color` بقيمة `rgb(0, 255, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
