---
id: 618a00ed1ca871a2b3aca0cb
title: الخطوة 35
challengeType: 0
dashedName: step-35
---

# --description--

الآن بعد أن مررت بجميع الألوان الابتدائية والثانوية والثلاثية على عجلة الألوان، سيكون من الأسهل فهم مفاهيم نظرية الألوان الأخرى وكيف تؤثر على التصميم.

First, in the CSS rules `.one`, `.two`, and `.three`, adjust the values in the `rgb` function so that the `background-color` of each element is set to pure black. تذكر أن دالة `rgb` تستخدم نموذج اللون المضاف (additive color model)، حيث تبدأ الألوان باللون الأسود وتتغير كقيم اللون الأحمر والأخضر والأزرق.

# --hints--

يجب أن يكون لقاعدة CSS لـ `.one` الخاصية `background-color` بقيمة `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

يجب أن يكون لقاعدة CSS لـ `.two` الخاصية `background-color` بقيمة `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

يجب أن يكون لقاعدة CSS لـ `.three` الخاصية `background-color` بقيمة `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
