---
id: 619b72a0db211f1c29afb906
title: الخطوة 42
challengeType: 0
dashedName: step-42
---

# --description--

Update the `.one` CSS rule to target the new `red` class.

# --hints--

لم يعد يجب أن يحتوي الكود الخاص بك على class selector الـ `.one`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.one'));
```

يجب عليك استخدام class selector لاستهداف class الـ `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red'));
```

يجب أن يكون لقاعدة CSS لـ `.red` الخاصية `background-color` بقيمة `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
