---
id: 619b7fd56aa2253778bcf5f7
title: الخطوة 48
challengeType: 0
dashedName: step-48
---

# --description--

نموذج الألوان <dfn>HSL</dfn> أو hue و saturation و lightness، طريقة أخرى لتمثيل الألوان.

The CSS hsl function accepts 3 values: a number from 0 to 360 for hue, a percentage from 0 to 100 for saturation, and a percentage from 0 to 100 for lightness.

إذا تخيلتم عجلة لون، فاللون الأحمر هو عند الدرجة صفر، والأخضر عند 120 درجة، والأزرق عند 240 درجة.

Saturation is the intensity of a color from 0%, or gray, to 100% for pure color. You must add the percent sign `%` to the saturation and lightness values.

Lightness (الضوء) هي مدى سطوع اللون الذي يظهر، من 0٪ أو أسود بالكامل، إلى 100٪ أبيض بالكامل و 50٪ حيادي.

في قاعدة CSS لـ `.blue`، استخدم دالة `hsl` لتغيير خاصية `background-color` إلى اللون الأزرق النقي. عيّن hue إلى `240`، و saturation إلى `100%`، و lightness إلى `50%`.

# --hints--

يجب أن يكون لقاعدة CSS لـ `.blue` الخاصية `background-color` بقيمة `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
