---
id: 619b7c3c83de403126b69c1e
title: الخطوة 47
challengeType: 0
dashedName: step-47
---

# --description--

قد تكون معتادًا بالفعل على القيم العشرية أو القيم باساس 10 ، والتي تبدأ من 0 إلى 9. القيم السداسية العشرية (Hexadecimal)، أو القيم باساس 16 ، تبدأ من 0 إلى 9، ثم A - F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

مع ألوان hex الـ `00` هو 0% من هذا اللون، و `FF` هو 100%. إذن `#00FF00` يترجم إلى 0% أحمر و 100% أخضر و 0% أزرق، و هو نفس `rgb(0, 255, 0)`.

Lower the intensity of green by setting the green value of the hex color to `7F`.

# --hints--

قاعدة CSS لـ `.green` يجب أن تحتوي علي الخاصية `background-color` بقيمة `#007F00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
