---
id: 619b74fa777a2b2473c94f82
title: الخطوة 45
challengeType: 0
dashedName: step-45
---

# --description--

قم بتحديث CSS class selector الـ `.two` بحيث أنه يستهدف class الـ `green` الجديد. And update the `.three` class selector so it targets the new `blue` class.

# --hints--

لم يعد يجب أن يحتوي الكود الخاص بك على class selector الـ `.two`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

يجب عليك استخدام class selector لاستهداف class الـ `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

يجب أن يكون لقاعدة CSS لـ `.green` الخاصية `background-color` بقيمة `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

لم يعد يجب أن يحتوي الكود الخاص بك على class selector الـ `.three`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

يجب عليك استخدام class selector لاستهداف class الـ `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

يجب أن يكون لقاعدة CSS لـ `.blue` الخاصية `background-color` بقيمة `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
