---
id: 61a49d15bdbb5e57cc6fd280
title: الخطوة 54
challengeType: 0
dashedName: step-54
---

# --description--

تتيح لك التوقفات اللونية (color-stops) ضبط مكان وضع الألوان على طول خط التدرج اللوني. إنها وحدة طول مثل `px` أو النسب المئوية التي تتبع اللون في وظيفة `linear-gradient`.

على سبيل المثال، في معامل التدرج الأحمر الأسود هذا (red-black gradient)، يتم الانتقال من الأحمر إلى الأسود عند نقطة 90 في المائة على طول خط التدرج، لذا الأحمر يأخذ معظم المساحة المتاحة:

```css
linear-gradient(90deg, red 90%, black);
```

في وظيفة `linear-gradient`، أضف لون `75%` يتوقف بعد الحجة للون الأحمر الأول. لا تضيف color stops إلى حجج الألوان الأخرى.

# --hints--

يجب أن تحتوي قاعدة CSS مسمى `.red` على خاصية `background` بقيمة `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
