---
id: 61a498399534644f59cff05e
title: الخطوة 53
challengeType: 0
dashedName: step-53
---

# --description--

كما تري، فإن دالة `linear-gradient` أنتجت معامل تدرج أحمر أخضر سلس (smooth red-green gradient). في حين أن دالة `linear-gradient` تحتاج إلى الحد الأدنى من اثنين arguments للالوان، فإنها يمكن أن تقبل العديد من الـ arguments للالوان.

استخدم دالة `rgb` لإضافة اللون الأزرق النقي كـ argument اللون الثالث إلى دالة `linear-gradient`.

# --hints--

يجب أن تحتوي قاعدة CSS لـ `.red` على خاصية `background` بقيمة `linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0),rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0),rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
