---
id: 61a5c906ab73313316e342f0
title: الخطوة 65
challengeType: 0
dashedName: step-65
---

# --description--

إذا لم يتم تقديم حجة `gradientDirection` إلى وظيفة `linear-gradient`، تقوم الوظيفة بترتيب الألوان من الأعلى إلى الأسفل، أو على طول خط 180 درجة، بشكل افتراضي.

نظف الكود أكثر قليلاً عن طريق إزالة الحجة `gradientDirection` من وظيفة `linear-gradient`.

# --hints--

يجب عليك إزالة حجة `gradientDirection` من وظيفة `linear-gradient` في قواعد CSS مسمى `.red` و `.green`.

```js
assert(!code.match(/linear-gradient\s*\(\s*180deg/gi));
```

يجب أن تحتوي قاعدة CSS مسمى `.red` على خاصية `background` بقيمة `linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

قاعدة CSS مسمى `.green` يجب أن تحتوي علي خاصية `background` بالقيمة `linear-gradient(#55680D, #71F53E, #116C31)`.

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
