---
id: 5a24c314108439a4d4036166
title: Compose React Components
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

As the challenges continue to use more complex compositions with React components and JSX, there is one important point to note. Rendering ES6 style class components within other components is no different than rendering the simple components you used in the last few challenges. You can render JSX elements, stateless functional components, and ES6 class components within other components.

# --instructions--

In the code editor, the `TypesOfFood` component is already rendering a component called `Vegetables`. Also, there is the `Fruits` component from the last challenge.

Nest two components inside of `Fruits` — first `NonCitrus`, and then `Citrus`. Both of these components are provided for you behind the scenes. Next, nest the `Fruits` class component into the `TypesOfFood` component, below the `h1` heading element and above `Vegetables`. The result should be a series of nested components, which uses two different component types.

# --hints--

The `TypesOfFood` component should return a single `div` element.

```js
assert(
  (function () {
    globalThis.NonCitrus = function NonCitrus() {
      return React.createElement('div', null);
    };
    globalThis.Citrus = function Citrus() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

The `TypesOfFood` component should return the `Fruits` component.

```js
assert(
  (function () {
    globalThis.NonCitrus = function NonCitrus() {
      return React.createElement('div', null);
    };
    globalThis.Citrus = function Citrus() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

The `Fruits` component should return the `NonCitrus` component and the `Citrus` component.

```js
assert(
  (function () {
    globalThis.NonCitrus = function NonCitrus() {
      return React.createElement('div', null);
    };
    globalThis.Citrus = function Citrus() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

The `TypesOfFood` component should return the `Vegetables` component below the `Fruits` component.

```js
assert(
  (function () {
    globalThis.NonCitrus = function NonCitrus() {
      return React.createElement('div', null);
    };
    globalThis.Citrus = function Citrus() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
