---
id: 6ac200f06b568b700b78fb5d
title: Implement a Stack
challengeType: 27
dashedName: implement-a-stack-python
---

# --description--

In this lab, you will implement a stack data structure using functions. A stack is a Last-In-First-Out (LIFO) data structure where elements are added and removed from the top.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should have an `init_stack` function that returns a dictionary with a `collection` key set to an empty list.
1. You should have a `push` function that adds an element to the top of the stack.
1. You should have a `pop` function that removes and returns the top element of the stack, or `None` if there isn't one.
1. You should have a `peek` function that returns the top element of the stack without removing it, or `None` if there isn't one.
1. You should have an `is_empty` function that returns `True` if the stack contains no elements, and `False` otherwise.
1. You should have a `clear` function that removes all elements from the stack.

**Note:** Most tests depend on the `init_stack` and `push` functions. Implement them first, as tests for `pop`, `peek`, `is_empty`, and `clear` require adding elements to the stack.

# --hints--

You should have an `init_stack` function.

```js
({ 
  test: () => assert(runPython(`
    _Node(_code).has_function("init_stack")
  `)) 
})
```

Your `init_stack` function should return a dictionary with a `collection` key set to an empty list.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    assert isinstance(test_stack, dict)
    assert test_stack.get("collection", None) == []
  `)
})
```

You should have a `push` function.

```js
({ 
  test: () => assert(runPython(`
    _Node(_code).has_function("push")
  `)) 
})
```

Your `push` function should add an element to the top of the stack.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "first")
    push(test_stack, "second")
    push(test_stack, "third")
    assert test_stack.get("collection", None) == ["first", "second", "third"]
  `)
})
```

You should have a `pop` function.

```js
({ 
  test: () => assert(runPython(`
    _Node(_code).has_function("pop")
  `)) 
})
```

Your `pop` function should remove and return the top element of the stack.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "first")
    push(test_stack, "second")
    push(test_stack, "third")
    assert pop(test_stack) == "third"
    assert test_stack.get("collection", None) == ["first", "second"]
    assert pop(test_stack) == "second"
    assert test_stack.get("collection", None) == ["first"]
  `)
})
```

Your `pop` function should return `None` if the stack is empty.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    assert pop(test_stack) is None
  `)
})
```

Your `pop` function should return falsy values correctly.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "")
    push(test_stack, [])
    push(test_stack, 0)
    push(test_stack, None)
    assert pop(test_stack) == None
    assert test_stack.get("collection", None) == ["", [], 0]
    assert pop(test_stack) == 0
    assert test_stack.get("collection", None) == ["", []]
    assert pop(test_stack) == []
    assert test_stack.get("collection", None) == [""]
    assert pop(test_stack) == ""
    assert test_stack.get("collection", None) == []
  `)
})
```

You should have a `peek` function.

```js
({ 
  test: () => assert(runPython(`
    _Node(_code).has_function("peek")
  `)) 
})
```

Your `peek` function should return the top element of the stack without removing it.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "first")
    push(test_stack, "second")
    push(test_stack, "third")
    assert peek(test_stack) == "third"
    assert test_stack.get("collection", None) == ["first", "second", "third"]
    assert peek(test_stack) == "third"
    assert test_stack.get("collection", None) == ["first", "second", "third"]
  `)
})
```

Your `peek` function should return `None` if the stack is empty.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    assert peek(test_stack) is None
  `)
})
```

Your `peek` function should return falsy values correctly.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "")
    push(test_stack, [])
    push(test_stack, 0)
    push(test_stack, None)
    assert peek(test_stack) == None
    pop(test_stack)
    assert peek(test_stack) == 0
    pop(test_stack)
    assert peek(test_stack) == []
    pop(test_stack)
    assert peek(test_stack) == ""
  `)
})
```

You should have an `is_empty` function.

```js
({ 
  test: () => assert(runPython(`
    _Node(_code).has_function("is_empty")
  `)) 
})
```

Your `is_empty` function should return `True` for an empty stack.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    assert is_empty(test_stack) is True
  `)
})
```

Your `is_empty` function should return `False` for a non-empty stack.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "element")
    assert is_empty(test_stack) is False
  `)
})
```

Your `is_empty` function should return `False` when the top element is falsy. 

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, False)
    assert is_empty(test_stack) is False
  `)
})
```

You should have a `clear` function.

```js
({ 
  test: () => assert(runPython(`
    _Node(_code).has_function("clear")
  `)) 
})
```

Your `clear` function should remove all elements from the stack.

```js
({ 
  test: () => runPython(`
    test_stack = init_stack()
    push(test_stack, "first")
    push(test_stack, "second")
    push(test_stack, "third")
    clear(test_stack)
    assert test_stack.get("collection", None) == []
  `)
})
```

# --seed--

## --seed-contents--

```py
```

# --solutions--

```py
def init_stack():
    return {"collection": []}

def print_stack(stack):
    print(stack["collection"])

def push(stack, element):
    stack["collection"].append(element)

def pop(stack):
    if is_empty(stack):
        return None
    return stack["collection"].pop()

def peek(stack):
    if is_empty(stack):
        return None
    return stack["collection"][-1]

def is_empty(stack):
    return not stack["collection"]

def clear(stack):
    stack["collection"].clear()
```
