---
id: 65821fcc010c3245718f2a06
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

The `expenses` parameter of your `add_expense` function will be a list of expenses. You want to be able to add items at the end of your list. For that you can use the `.append()` list method:

```py
my_list = [2, 4, 7]
my_list.append(3)
```

In the example above, after appending `3`, `my_list` would be `[2, 4, 7, 3]`. Replace `pass` with a call to the `.append()` method on the `expenses` list. Don't pass any arguments to `.append()` for now.

# --hints--

You should have `expenses.append()` in your function.

```js
assert.match(code, /^\s+expenses\.append\s*\(\s*\)/m)
```

# --seed--

## --seed-contents--

```py    
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
