---
id: 65822773f4c60a479ec89f66
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Inside the `print_expenses` function, create a `for` loop that iterates over each item in the `expenses` list. Use `expense` as the loop variable and move `pass` inside the loop body.

# --hints--

You should have `for expense in expenses:` in your `print_expenses` function. Remember to use `pass` within the loop body.

```js
({ test: () => assert.match(code, /for\s+expense\s+in\s+expenses\s*:/ ) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
