---
id: 65825a9520a0795c8afbef50
title: Paso 35
challengeType: 20
dashedName: step-35
---

# --description--

After the `print()` call, call the `print_expenses` function to display all the expenses that have been added so far. Pass the `expenses` list as the argument.

# --hints--

You should have `print_expenses(expenses)` in your `elif` statement.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')2\1\s*:\s*print\s*\(\s*("|')\\nAll\sExpenses:\2\s*\)\s*print\s*_expenses\s*\(\s*expenses\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
--fcc-editable-region--
        elif choice == '2':
            print('\nAll Expenses:')

--fcc-editable-region--
```
