---
id: 69cb589c3b9dafbc0c162d22
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Try submitting only spaces again. This time, you should see an `Uncaught Error` in the console.

This happens because the error you threw is not being handled, which causes the program to stop running.

In the previous lesson, you learned about the `try...catch` block, which allows you to catch errors and handle them gracefully instead of breaking the program. Wrap your `if` statements in a `try` block.

Add `error` as the parameter in your `catch` block so you can access the error object. Leave `catch` as an empty block for now.

# --hints--

Wrap your validation logic inside a `try...catch` block. The `catch` block should have `error` as its parameter.

```js
const match = code.match(
  __helpers.functionRegex('registerUser', [], { capture: true })
);
assert.exists(match);
assert.match(match[1], /try\s*\{[\s\S]*\}\s*catch\s*\(error\)\s*\{\s*\}/);
```

The `if` statements should be inside the `try` block.

```js
const match = code.match(
  __helpers.functionRegex('registerUser', [], { capture: true })
);
assert.exists(match);
assert.match(match[1], /try\s*\{[\s\S]*if\s*\(name\.trim\(\)/);
assert.match(match[1], /try\s*\{[\s\S]*if\s*\(age\.trim\(\)/);
```

# --seed--

## --seed-contents--

```js
const nameInput = document.getElementById("name-input");
const ageInput = document.getElementById("age-input");
const btn = document.getElementById("submit-btn");
const result = document.getElementById("result");

console.clear();

console.dir(nameInput);
console.dir(ageInput);

btn.addEventListener("click", registerUser);

function registerUser() {
  console.log("Button clicked");
  
  const name = nameInput.value;
  const age = ageInput.value;

  console.log("Name:", name);
  console.log("Age:", age);

--fcc-editable-region--

  if (name.trim() === "") {
    throw new Error("Name is required");
  }
  if (age.trim() === "") {
    throw new Error("Age is required");
  }
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>User Registration Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>

  <body>
    <main>
      <h1>User Registration</h1>

      <form aria-labelledby="registration-heading">
        <fieldset>
          <legend id="registration-heading">Register a new user</legend>

          <label for="name-input">Name</label>
          <input
            type="text"
            id="name-input"
            name="name"
            placeholder="Enter your name"
            autocomplete="name"
            required
          />

          <label for="age-input">Age</label>
          <input
            type="number"
            id="age-input"
            name="age"
            placeholder="Enter your age"
            inputmode="numeric"
            min="13"
            required
            aria-describedby="age-help"
          />
          <p id="age-help">Users must be at least 13 years old.</p>

          <button type="button" id="submit-btn">Register</button>
        </fieldset>
      </form>

      <p id="result" aria-live="polite" aria-atomic="true"></p>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 0;
  min-height: 100vh;
  display: grid;
  place-items: start center;
  padding: 40px 16px;
  background-color: #0a0a23;
  color: #f5f6f7;
}

main {
  width: min(100%, 28rem);
}

h1 {
  margin: 0 0 1rem;
  text-align: center;
}

form {
  text-align: left;
}

fieldset {
  margin: 0;
  padding: 2.5rem;
  border: 1px solid #3b3b4f;
  border-radius: 0.5rem;
  background-color: #1b1b32;
}

legend {
  padding: 0 0.5rem;
  font-weight: 700;
}

label {
  display: block;
  margin-top: 1rem;
  margin-bottom: 0.375rem;
  font-weight: 600;
}

label:first-of-type {
  margin-top: 0;
}

input {
  display: block;
  width: 100%;
  padding: 8px;
  border: 1px solid #d0d0d5;
  border-radius: 0.25rem;
  background-color: #0a0a23;
  color: #f5f6f7;
}

input::placeholder {
  color: #d0d0d5;
}

button {
  margin-top: 1rem;
  padding: 8px 12px;
  cursor: pointer;
  border: 2px solid #ffac33;
  background-image: linear-gradient(#fecc4c, #ffac33);
  transition: all 0.1s ease-in-out;
}

input:focus-visible,
button:focus-visible {
  outline: 3px solid #99c9ff;
  outline-offset: 2px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c);
}

button:active {
  box-shadow: inset 0 0 5px rgba(0, 0, 0, 0.4);
}

#age-help {
  margin: 0.5rem 0 0;
  font-size: 0.875rem;
  color: #dfdfe2;
}

#result {
  min-height: 1.5rem;
  margin-top: 20px;
  font-weight: 600;
  text-align: center;
}
```
