---
id: 5d5a813321b9e3db6c106a46
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript allows you to build interactive applications by manipulating the <dfn>Document Object Model (DOM)</dfn>.  The `DOM` is a representation of the HTML elements on your web page and you will learn how to manipulate the `DOM` to create a role playing game that is fun and interactive.

*Note*: For the remaining projects in this curriculum, you will need to have a basic level knowledge of HTML and CSS. If you are new to HTML and CSS, please go through the [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/) first.

All of the CSS and HTML boilerplate has been provided for you. But you will need to add the HTML for the game components.

Start, by creating a `div` element with `id` set to `game` within your `body`.

# --hints--

You should have a `div` element.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Your `div` element should have an `id` attribute with the value `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Your `div` element should be within the `body` element.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```
