---
id: 62a2509ba163e020bb9d84ea
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now you can start writing your JavaScript. To add JavaScript to your HTML, you need to use a <dfn>script</dfn> element. The `script` element is used to load JavaScript into your HTML file.

```html
<script>
  // Your JavaScript code goes here
</script>
```

Inside the `head` element, add a `script` element. You should use an opening `<script>` and closing `</script>` tag.

# --hints--

You should have a `script` element.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

Your `script` element should have an opening tag.

```js
assert.match(code, /<script\s*>/i);
```

Your `script` element should have a closing tag.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
