---
id: 62a2509ba163e020bb9d84ea
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

When building projects, it is best to keep your code organized and separate into HTML, CSS, and JavaScript files. 

You can use the <dfn>script</dfn> element to link your JavaScript code to your HTML file. You can then use the `src` attribute in the `script` element to specify the path to the JavaScript file that will be used to control the page.

```html
<script src="link-to-file-goes-here"></script>
```

Inside the `head` element, add a `script` element. Give your empty `script` element a `src` attribute set to `./script.js`.

# --hints--

You should have a `script` element with a `src` attribute set to `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```
