---
id: 62a2509ba163e020bb9d84ea
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

When building projects, it is best to keep your code organized and separate into HTML, CSS and JavaScript files. 

The <dfn>script</dfn> element can be used to link your JavaScript code to your HTML file. The `src` attribute in the `script` element is used to specify the path to the JavaScript file that will be used to control the page.

```html
<script src="link-to-file-goes-here"></script>
```

Inside the `head` element, add a `script` element. Give your empty `script` element a `src` attribute set to `./script.js`.

# --hints--

You should have a `script` element.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 2);
```

Your `script` element should have a `src` attribute set to `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
