---
id: 62a257659d0d1e2456f24ba2
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

When building out projects, it is best to keep your code organized and separate into HTML, CSS and JavaScript files. 

Give your empty `script` element a `src` attribute set to `./script.js`.

# --hints--

Your `script` element should have a `src` attribute set to `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script></script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
