---
id: 62a3b5843544ce3a77459c27
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Your player's inventory will hold multiple weapons. For now, you want the player to start with just a `"stick"`. 

Assign an array to the `inventory` variable that has the string `"stick"` as its only value.

Remember, that you worked with arrays in the previous project and this is the syntax.

```js
["element 1", "element 2", "element 3"]
```

# --hints--

Your `inventory` variable should still be an array.

```js
assert.isArray(inventory);
```

Your `inventory` variable should only have one value.

```js
assert.lengthOf(inventory, 1);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
