---
id: 62a3a7e4f1060e2fc5ffb34b
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

The player will be able to accumulate different weapons throughout the game. When the player switches to a new weapon, you will need to keep track of the player's current weapon.

Use `let` to create a variable called `currentWeapon` and set it to `0`. The number `0` represents the player's first weapon in the `weapons` array. 

*NOTE:* You will create the `weapons` array later on.

# --hints--

You should use `let` to declare a variable called `currentWeapon`.

```js
assert.match(code, /let\s+currentWeapon/i);
```

You should use camelCase to name your variable.

```js
assert.match(code, /currentWeapon/);
```

Your `currentWeapon` variable should be set to `0`.

```js
assert.equal(currentWeapon, 0);
```

You should initialize your variable to `0`.

```js
assert.match(code, /let\s+currentWeapon\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
