---
id: 62a3a0a3c0a4b32915d26a6e
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In the game, the player will be able to gain experience points (XP) when they win fights against monsters. To track the player's XP, you will need to create a variable.

Use the `let` keyword to declare a variable called `xp`. Assign it a value of the number `0`.

You want to use `let` here because the value of `xp` will change as the player gains experience points. If you used `const`, you would not be able to change the value of `xp` after it was initialized.

# --hints--

You should use the `let` keyword to declare a variable named `xp`.

```js
assert.match(code, /let\s+xp/);
```

You should initialize the `xp` variable to `0`. 

```js
assert.equal(xp, 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```


