---
id: 62a3b3eab50e193608c19fc6
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Throughout the game, the player will be able to fight different monsters. You will need to keep track of the type of monster the player is fighting.

Use `let` to declare a variable called `fighting` but do not initialize it with a value. You are using `let` here because the value of `fighting` will change throughout the game.

# --hints--

You should use `let` to declare a variable called `fighting`.

```js
assert.match(code, /let\s+fighting/);
```

Your `fighting` variable should not have a value.

```js
assert.isUndefined(fighting);
```

You should not assign a value to your `fighting` variable. 

```js
assert.match(code, /let\s+fighting;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
--fcc-editable-region--
```
