---
id: 62a3b41c9494f937560640ab
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

In the game, you will need to keep track of the monster's health and the player's current weapon inventory. 

Use `let` to declare two variables called `monsterHealth` and `inventory` but do not initialize them with a value.
# --hints--

You should use `let` to declare your `monsterHealth` variable.

```js
assert.match(code, /let\s+monsterHealth/i);
```

You should use camelCase to name your `monsterHealth` variable.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` should not have a value.

```js
assert.isUndefined(monsterHealth);
```

You should not assign a value to your `monsterHealth` variable.

```js
assert.match(code, /let\s+monsterHealth;?/);
```

You should use `let` to declare your `inventory` variable.

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` should not have a value.

```js
assert.isUndefined(inventory);
```

You should not assign a value to your `inventory` variable. 

```js
assert.match(code, /let\s+inventory\s*;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
