---
id: 62a3a0a3c0a4b32915d26a6e
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

In the game, the player will be able to gain experience points (XP) when they win fights against monsters. To track the player's XP, you will need to create a variable.

Use the `let` keyword to declare a variable called `xp`. Assign it a value of the number `0`.

You want to use `let` here because the value of `xp` will change as the player gains experience points. If you used `const`, you would not be able to change the value of `xp` after it was initialized.

# --hints--

You should use the `let` keyword to declare a variable named `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` should have a value of `0`.

```js
assert.equal(xp, 0);
```

You should initialize the `xp` variable to `0`. 

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
