---
id: 62a3b49686792938718b90d3
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Your player's inventory will hold multiple weapons. For now, you want the player to start with just a `"stick"`. 

Assign an array to the `inventory` variable that has the string `"stick"` as its only value.

Remember, that you worked with arrays in the previous project and this is the syntax.

```js
["element 1", "element 2", "element 3"]
```

# --hints--

You should assign an array to your `inventory` variable.

```js
assert.isArray(inventory);
```

Your `inventory` variable should only have one value.

```js
assert.lengthOf(inventory, 1);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

#game {
  max-width: 500px;
  max-height: 400px;
  background-color: #ffffff;
  color: #ffffff;
  margin: 30px auto 0px;
  padding: 10px;
}

#controls,
#stats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #0a0a23;
}

#monsterStats {
  border: 1px solid #0a0a23;
  padding: 5px;
  color: #ffffff;
  background-color: #c70d0d;
}

.stat {
  padding-right: 10px;
}

button {
  cursor: pointer;
  color: #0a0a23;
  background-color: #feac32;
  background-image: linear-gradient(#fecc4c, #ffac33);
  border: 3px solid #feac32;
}
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
let monsterHealth;
let inventory;
--fcc-editable-region--
```
