---
id: a8d97bd4c764e91f9d2bda01
title: 翻譯二進制字符串
challengeType: 1
forumTopicId: 14273
dashedName: binary-agents
---

# --description--

請實現一個函數，把傳入的二進制字符串轉換成英文句子。

二進制字符串會以空格分隔。

# --hints--

`binaryAgent("01000001 01110010 01100101 01101110 00100111 01110100 00100000 01100010 01101111 01101110 01100110 01101001 01110010 01100101 01110011 00100000 01100110 01110101 01101110 00100001 00111111")` 應返回 `Aren't bonfires fun!?`。

```js
assert.deepEqual(
  binaryAgent(
    '01000001 01110010 01100101 01101110 00100111 01110100 00100000 01100010 01101111 01101110 01100110 01101001 01110010 01100101 01110011 00100000 01100110 01110101 01101110 00100001 00111111'
  ),
  "Aren't bonfires fun!?"
);
```

`binaryAgent("01001001 00100000 01101100 01101111 01110110 01100101 00100000 01000110 01110010 01100101 01100101 01000011 01101111 01100100 01100101 01000011 01100001 01101101 01110000 00100001")` 應返回 `I love FreeCodeCamp!`。

```js
assert.deepEqual(
  binaryAgent(
    '01001001 00100000 01101100 01101111 01110110 01100101 00100000 01000110 01110010 01100101 01100101 01000011 01101111 01100100 01100101 01000011 01100001 01101101 01110000 00100001'
  ),
  'I love FreeCodeCamp!'
);
```

# --seed--

## --seed-contents--

```js
function binaryAgent(str) {
  return str;
}

binaryAgent("01000001 01110010 01100101 01101110 00100111 01110100 00100000 01100010 01101111 01101110 01100110 01101001 01110010 01100101 01110011 00100000 01100110 01110101 01101110 00100001 00111111");
```

# --solutions--

```js
function binaryAgent(str) {
  return str.split(' ').map(function(s) { return parseInt(s, 2); }).map(function(b) { return String.fromCharCode(b);}).join('');
}
```
