---
id: 663d18c574456976d11b6e87
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在游戏中，将会有多个回合。 谁先达到三分，谁就赢了。

在这一步，你将专注于确定玩家是否赢得了这一局。

完成 `hasPlayerWonTheRound` 函数。 这个函数有两个参数：`player` 和 `computer`。 如果玩家赢得了这一局，该函数应该返回 `true`。如果玩家输了或打平了这一局，该函数则应返回 `false`。

以下是玩家赢得一轮比赛的标准：

- 如果玩家选择 `"Rock"`，而电脑选择 `"Scissors"`
- 如果玩家选择 `"Scissors"`，而电脑选择 `"Paper"`
- 如果玩家选择 `"Paper"`，而电脑选择 `"Rock"`

已经为你提供了一些函数调用来测试你的函数。

# --hints--

你的 `hasPlayerWonTheRound` 函数应该返回一个布尔值。

```js
assert.isBoolean(hasPlayerWonTheRound("Rock", "Scissors"));
```

你的 `hasPlayerWonTheRound` 函数应该在玩家选择了 `"Rock"` 而电脑选择 `"Scissors"` 时返回 `true`。

```js
assert.isTrue(hasPlayerWonTheRound("Rock", "Scissors"));
```

你的 `hasPlayerWonTheRound` 函数应该在玩家选择 `"Scissors"` 且电脑选择 `"Paper"` 时返回 `true`。

```js
assert.isTrue(hasPlayerWonTheRound("Scissors", "Paper"));
```

你的 `hasPlayerWonTheRound` 函数应在玩家选择 `"Paper"` 而电脑选择 `"Rock"` 时返回 `true`。

```js
assert.isTrue(hasPlayerWonTheRound("Paper", "Rock"));
```

如果玩家和计算机选择了相同的选项，你的 `hasPlayerWonTheRound` 函数应返回 `false`。

```js
assert.isFalse(hasPlayerWonTheRound("Rock", "Rock"));
assert.isFalse(hasPlayerWonTheRound("Scissors", "Scissors"));
assert.isFalse(hasPlayerWonTheRound("Paper", "Paper"));
```

如果电脑赢了这一轮，你的 `hasPlayerWonTheRound` 函数应返回 `false`。

```js
assert.isFalse(hasPlayerWonTheRound("Scissors", "Rock"));
assert.isFalse(hasPlayerWonTheRound("Paper", "Scissors"));
assert.isFalse(hasPlayerWonTheRound("Rock", "Paper"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Rock, Paper, Scissors game</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Let's play Rock, Paper, Scissors!</h1>
    <main>
      <details class="rules-container">
        <summary>Rules to the game</summary>

        <p>You will be playing against the computer.</p>
        <p>You can choose between Rock, Paper, and Scissors.</p>
        <p>The first one to three points wins.</p>

        <p>Here are the rules to getting a point in the game:</p>
        <ul>
          <li>Rock beats Scissors</li>
          <li>Scissors beats Paper</li>
          <li>Paper beats Rock</li>
        </ul>
        <p>
          If the player and computer choose the same option (Ex. Paper and
          Paper), then no one gets the point.
        </p>
      </details>

      <div class="score-container">
        <strong
          >Player Score: <span class="score" id="player-score">0</span></strong
        >
        <strong
          >Computer Score:
          <span class="score" id="computer-score">0</span></strong
        >
      </div>

      <section class="options-container">
        <h2>Choose an option:</h2>
        <div class="btn-container">
          <button id="rock-btn" class="btn">Rock</button>
          <button id="paper-btn" class="btn">Paper</button>
          <button id="scissors-btn" class="btn">Scissors</button>
        </div>
      </section>

      <div class="results-container">
        <p id="results-msg"></p>
        <p id="winner-msg"></p>
        <button class="btn" id="reset-game-btn">Play again?</button>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --very-dark-blue: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--very-dark-blue);
  text-align: center;
  color: var(--white);
}

h1 {
  margin: 15px 0 20px;
}

.btn {
  cursor: pointer;
  width: 100px;
  margin: 10px;
  color: var(--very-dark-blue);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

.rules-container {
  padding: 10px 0;
  margin: auto;
  border-radius: 15px;
  border: 5px solid var(--yellow);
  background-color: var(--white);
  color: var(--very-dark-blue);
}

.rules-container ul {
  list-style-type: none;
}

.rules-container p {
  margin: 10px 0;
}

@media (min-width: 760px) {
  .rules-container {
    width: 60%;
  }
}

.score-container {
  display: flex;
  justify-content: space-around;
  margin: 30px 0;
  font-size: 1.2rem;
}

.score {
  font-weight: 500;
}

.results-container {
  font-size: 1.3rem;
  margin: 15px 0;
}

#winner-msg {
  margin-top: 25px;
}

#reset-game-btn {
  display: none;
  margin: 20px auto;
}

```

```js
function getRandomComputerResult() {
  const options = ["Rock", "Paper", "Scissors"];
  const randomIndex = Math.floor(Math.random() * options.length);
  return options[randomIndex];
}

--fcc-editable-region--
function hasPlayerWonTheRound(player, computer) {

}

console.log(hasPlayerWonTheRound("Rock", "Scissors")); 
console.log(hasPlayerWonTheRound("Scissors", "Rock")); 
--fcc-editable-region--
```
