---
id: 63ee3f71381756f9716727ef
title: CSS Foundations Exercise A
challengeType: 14
dashedName: css-foundations-exercise-a
---

# --description--
In this exercise, you're going to practice adding `CSS` to an `HTML` file using all three methods: external `CSS`, internal `CSS`, and inline `CSS`. You should only be using type selectors for this exercise when adding styles via the external and internal methods. You should also use keywords for colors (e.g. "blue") instead of using `RGB` or `HEX` values.

There are three elements for you to add styles to, each of which uses a different method of adding `CSS` to it, as noted in the outcome image below. All other exercises in this section will have a `CSS` file provided and linked for you, but for this exercise you will have to create the file and link it in the `HTML` file yourself. This is all about practicing using these different methods and getting the syntax right.

You should see a `div` element, with white text and a `font-size` of `32px`, finally its text should be aligned in the center and the `font-style` should be bold. Make sure to add some text to the container.




# --hints--

Test 1

```js
assert(document.getElementById('hello').length == 1)
```

Test 2

```js

```

Test 3

```js

```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <title>Exercise A</title>
  </head>
  <body>
    <!-- add your code here -->

    <!-- add your code here -->
  </body>
</html>
```

# --solutions--

```html
```
