---
id: 5daa813381b9e3db6c126b43
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

PYTHON

Defina a variável `hello` como "world". Em seguida, imprima o valor.

# --hints--

O código-fonte deve incluir `one = 2`

```js
assert(code.match(/one\s*=\s*2/))
```

A variável `hello` deve ser igual a "world".

```js
({ test: () => assert.equal(__userGlobals.get("hello"), "world") })
```

The `one` variable should be an integer.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_variable("one").is_integer()`))
}})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
one = 1
hello = "goodbye"
print(hello)
--fcc-editable-region--
```

# --solutions--

```py
one = 2
hello = "world"
print(hello)
```
