---
id: 5daa813381b9e3db6c126b43
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

PYTHON

Defina a variável `hello` como "world". Em seguida, imprima o valor.

# --hints--

A variável `hello` deve ser igual a "world".

```js
({ test: () => assert.equal(__userGlobals.get("hello"), "world") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
one = 1
hello = "goodbye"
def a_function():
    local_thing = "world"
    print(local_thing)

a_function()

print(hello)
--fcc-editable-region--
```

# --solutions--

```py
one = 1
hello = "world"
print(hello)
```
