---
id: 64a5229b99ff0e8250cd9a72
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

instruções do passo 2

# --hints--

Teste 1

```js
({ input: ["Beau", "Carnes"], test: () => {
  assert.equal( "Beau", __userGlobals.get("name"));
  assert.equal( "Carnes", __userGlobals.get("last_name"));
  } })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
name = input('What is your name?')
print('Hi ' + name)
--fcc-editable-region--
```

# --solutions--

```py
name = input('What is your name?')
last_name = input('What is your last name?')
print('Hi ' + name + ' ' + last_name)
```
