---
id: 64b171849f925b0773aa434c
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Crie uma lista de números de 1 a 3 e atribua a lista à variável `xs`.

# --hints--

`xs` deve ser a lista `[1,2,3]`.

```js
({ test: () => assert.deepEqual(Array.from(__userGlobals.get("xs")), [1,2,3])})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
xs = [1,2,3]
```
