---
id: 65df3afd233057f6a620a860
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Isso mostra como as funções auxiliares podem ser usadas para analisar as definições de classe de um aluno

# --instructions--

Nós podemos simplesmente encontrar a definição da classe dentro do Python, retornando-a para o JavaScript e verificando se ela contém a string "pass".

Se precisarmos de algo mais confiável, podemos expandir os ajudantes para verificar especificamente se existem instruções pass. Isso seria útil se, por exemplo, quiséssemos escrever classes contendo, por exemplo, "passengers".

# --hints--

`TreeNode` não deve conter instruções `pass`.

```js
({ test: () => {
  const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
  assert.notInclude(pyClassStr, "pass")
}})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass

class Trap:
    pass
--fcc-editable-region--
```

# --solutions--

```py
class TreeNode:
    def __init__(self, thing):
        self.thing = thing

class Trap:
    pass

```
