---
id: 615f38cabc64e3556f98cc1a
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Bordas podem ser usadas para agrupar e priorizar o conteúdo.

Crie um seletor `.label` e dê a ele um `border` definido como `2px solid black`.

# --hints--

Você deve ter um seletor `.label`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

O seletor `.label` deve ter a propriedade `border` definida como `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
