---
id: 615f2d4150fe0d4cbd0f2628
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Abaixo do elemento `h1`, adicione um elemento `p` com o texto `8 servings per container`.

# --hints--

Você deve adicionar um novo elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

O elemento `p` deve estar dentro do elemento `body`.

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

O elemento `p` deve vir depois do elemento `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

O elemento `p` deve ter o texto `8 servings per container`.

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
