---
id: 615f3ed16592445e57941aec
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Linhas podem ajudar a separar e agrupar conteúdos importantes, especialmente quando o espaço é limitado.

Crie um elemento `div` abaixo do elemento `h1` e dê a ele um atributo `class` definido como `divider`.

# --hints--

Você deve criar um novo `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

O novo elemento `div` deve ter o atributo `class` definido como `divider`.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

O elemento `.divider` deve estar dentro do elemento `.label`.

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

O elemento `.divider` deve vir depois do elemento `h1`.

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
