---
id: 615f34ecc1091b4fd5a8a484
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro do elemento `head`, adicione o elemento `link` com o atributo `rel` definido como `stylesheet` e o atributo `href` definido como `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

Isto importará a família de fontes `Open Sans`, com os valores de peso da fonte `400`, `700`, e `800`.

Adicione também um elemento `link` para vincular o arquivo `styles.css`.

# --hints--

O código deve ter dois elementos `link` de autofechamento.

```js
assert(document.querySelectorAll('link').length === 2);
```

Ambos os elementos `link` devem ter o atributo `rel` definido como `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Um dos elementos `link` deve ter o atributo `href` definido como `./styles.css`.

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Um dos elementos `link` deve ter o atributo `href` definido como `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
