---
id: 615f357957e370510f21ea16
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Crie um seletor `body` e dê a ele a propriedade `font-family` definida como `Open Sans` com `sans-serif` como alternativa.

Lembre-se de que fontes com espaços no nome devem estar entre aspas no CSS.

# --hints--

Você deve ter um seletor `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

O seletor `body` deve ter a propriedade `font-family` definida como `"Open Sans", sans-serif`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
