---
id: 615f3cafd794015aa9547a6d
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Lembre-se de que o uso de `h1`, `h2` e tags similares determinam a estrutura semântica do seu HTML. No entanto, você pode ajustar o CSS desses elementos para controlar o fluxo visual e a hierarquia.

Crie uma regra para `h1` e defina a propriedade `font-weight` para `800`. Isto deixará o texto do elemento `h1` em negrito.

# --hints--

Você deve criar o seletor `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

O seletor `h1` deve ter a propriedade `font-weight` com o valor de `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
