---
id: 615f3b091162165948e1cb33
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Se você inspecionar o elemento `.label` com as ferramentas de desenvolvedor do seu navegador, poderá notar que ele de fato possui 288 pixels de largura ao invés de 270. Isto porque, por padrão, o navegador inclui a borda e o padding quando determina o tamanho de um elemento.

Para resolver isso, redefina o modelo da caixa criando um seletor `*` e dê a ele o atributo de `box-sizing` que é propriedade de `border-box`.

# --hints--

Você deve criar o seletor `*`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

O seletor `*` deve ter a propriedade `box-sizing` com o valor de `border-box`.

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
