---
id: 615f51e4e5b24a6e80eccce1
title: Passo 30
challengeType: 0
dashedName: step-30
---

# --description--

Dentro do elemento `.calories-info`, crie um elemento `p`. Dê àquele elemento `p` um atributo de `class` definido como `bold sm-text` e o texto `Amount per serving`.

# --hints--

Você deve criar um elemento `p` dentro do elemento `.calories-info`.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.localName === 'p');
```

O novo elemento `p` deve ter o atributo `class` definido como `bold sm-text`.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.classList.contains('bold'));
assert(document.querySelector('.calories-info')?.children?.[0]?.classList.contains('sm-text'));
```

O novo elemento `p` deve ter o texto `Amount per serving`.

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.innerText === 'Amount per serving');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
--fcc-editable-region--
      <div class="calories-info">

      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}
```
