---
id: 615f378014c2da526a109c81
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

A fonte é um pouco pequena. Crie um seletor `html` e defina a fonte com o tamanho de `16px`.

# --hints--

Você deve ter um seletor `html`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

O seletor `html` deve ter a propriedade `font-size` definida como `16px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.fontSize === '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

body {
  font-family: 'Open Sans', sans-serif;
}
```
