---
id: 615f2abbe7d18d49a1e0e1c8
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Fornecemos um modelo de HTML básico para você.

Crie um elemento `h1` dentro do seu elemento `body` e dê a ele o texto `Nutrition Facts`.

# --hints--

Você deve adicionar um novo elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

O elemento `h1` deve estar dentro do elemento `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

O elemento `h1` deve ter o texto `Nutrition Facts`.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
