---
id: 615f5486b8fd4b71633f69b0
title: Passo 32
challengeType: 0
dashedName: step-32
---

# --description--

Abaixo do elemento `.sm-text`, crie um elemento `h1` com a classe `Calories 230`. Envolva a parte do texto que diz `230` em outro elemento `span` com a `class` definida como `right`.

# --hints--

Você deve ter um elemento `h1` dentro do elemento `.calories-info`.

```js
assert(document.querySelector('.calories-info')?.querySelector('h1'))
```

O elemento `h1` deve vir depois do elemento `.sm-text`.

```js
assert(document.querySelector('.sm-text')?.nextElementSibling?.localName === 'h1');
```

O elemento `h1` deve conter o texto `Calories 230`.

```js
// the float: right property does weird things to the inner text.
assert(document.querySelector('.calories-info')?.lastElementChild?.innerText.match(/Calories[\s|\n]*230/));
```

Você deve criar um elemento `span` dentro do elemento `h1`.

```js
assert(document.querySelector('.calories-info')?.lastElementChild?.querySelector('span'));
```

O elemento `span` deve ter a `class` definida como `right`.

```js
assert(document.querySelector('.calories-info')?.lastElementChild?.querySelector('span')?.classList?.contains('right'));
```

O elemento `span` deve ter o texto `230`.

```js
assert(document.querySelector('.calories-info')?.lastElementChild?.querySelector('span')?.innerText === '230');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
--fcc-editable-region--
      <div class="calories-info">
        <p class="bold sm-text">Amount per serving</p>
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}

.sm-text {
  font-size: 0.85rem;
}
```
