---
id: 615f34948891834dd77655a6
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Adicione um segundo elemento `p` com o texto `Serving size 2/3 cup (55g)`.

# --hints--

Você deve ter um segundo elemento `p`.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

O segundo elemento `p` deve estar dentro do elemento `body`.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

O segundo elemento `p` deve vir após o elemento `p` já existente.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

O segundo elemento `p` deve ter o texto `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
