---
id: 5d822fd413a79914d39e98ca
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте початковий та кінцевий теґи `html` під `DOCTYPE`, щоб ви мали місце для початку розміщення коду. Переконайтесь, що встановили мову на англійську.

# --hints--

Ваша декларація `DOCTYPE` повинна бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ваш елемент `html` повинен мати початковий теґ та атрибут `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` повинен мати кінцевий теґ.

```js
assert(code.match(/<\/html\s*>/));
```

Ваші теґи `html` повинні бути в правильному порядку.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Ви повинні мати тільки один елемент `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
