---
id: 5d822fd413a79914d39e98e7
title: Крок 31
challengeType: 0
dashedName: step-31
---

# --description--

Надайте шести новим елементам такі значення `width` та `height`: `10%` та `60%` для `.fb1`, `10%` та `40%` для `.fb2`, `10%` та `35%` для `.fb3`, `8%` та `45%` для `.fb4`, `10%` та `33%` для `.fb5`, `9%` та `38%` для `.fb6`.

# --hints--

Ви повинні створити селектор `.fb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

Ви повинні надати селектору `.fb1` властивість `width` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

Ви повинні надати селектору `.fb1` властивість `height` зі значенням `60%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

Ви повинні створити селектор `.fb2`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

Ви повинні надати селектору `.fb2` властивість `width` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

Ви повинні надати селектору `.fb2` властивість `height` зі значенням `40%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

Ви повинні створити селектор `.fb3`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

Ви повинні надати селектору `.fb3` властивість `width` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

Ви повинні надати селектору `.fb3` властивість `height` зі значенням `35%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

Ви повинні створити селектор `.fb4`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

Ви повинні надати селектору `.fb4` властивість `width` зі значенням `8%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

Ви повинні надати селектору `.fb4` властивість `height` зі значенням `45%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

Ви повинні створити селектор `.fb5`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

Ви повинні надати селектору `.fb5` властивість `width` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

Ви повинні надати селектору `.fb5` властивість `height` зі значенням `33%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

Ви повинні створити селектор `.fb6`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

Ви повинні надати селектору `.fb6` властивість `width` зі значенням `9%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

Ви повинні надати селектору `.fb6` властивість `height` зі значенням `38%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--

```

