---
id: 5d822fd413a79914d39e98e0
title: Крок 24
challengeType: 0
dashedName: step-24
---

# --description--

Створіть нову змінну під іншими та назвіть її `--building-color3` та надайте їй значення `#cc6699`. Потім використайте її як `background-color` класу `.bb3` та надайте їй резервне значення `pink`.

# --hints--

Ви повинні визначити нову змінну з назвою `--building-color3`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('--building-color3'));
```

Ви повинні надати `--building-color3` значення `#cc6699`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color3')?.trim(), '#cc6699');
```

Ви повинні встановити властивість `background-color` на `.bb3`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb3')?.backgroundColor);
```

Ви повинні встановити `background-color`, використовуючи змінну `--building-color3` з резервним `pink`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb3')?.getPropVal('background-color', true), 'var(--building-color3,pink)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2, green);
}

.bb3 {
  width: 10%;
  height: 55%;
}
--fcc-editable-region--
.bb4 {
  width: 11%;
  height: 58%;
}

```
