---
id: 5d822fd413a79914d39e98d0
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Створіть елемент `div` в `body` з класом `background-buildings`. Це буде контейнер для групи будівель.

# --hints--

Ви повинні створити елемент `div`.

```js
assert.exists(document.querySelector('div'));
```

Ваш `div` повинен бути в межах елемента `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Ваші елементи `div` повинні мати клас зі значенням `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

