---
id: 5d822fd413a79914d39e98cf
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ви можете побачити `body` (прихований блок вашої сторінки); блок навколо нього є елементом `html`. Зробіть так, щоб ваш `body` заповнював весь вюпорт, надавши йому `height` зі значенням `100vh`. Видаліть `margin` за замовчуванням із `body`, встановивши `margin` на `0`. Зрештою, встановіть властивість `overflow` на `hidden`, щоб приховати будь-які смуги прокрутки, які з’являються, коли щось виходить за вюпорт.

# --hints--

Ви повинні використати селектор `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Ваш `body` повинен мати `height` зі значенням `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

Ваш `body` повинен мати `margin` зі значенням `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

Ваш `body` повинен мати властивість `overflow` зі значенням `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

