---
id: 5d822fd413a79914d39e98d2
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Вкладіть `div`, що має клас `bb1` в контейнер фонових будівель. Відкрийте свій файл `styles.css` та надайте `.bb1` властивості `width` зі значенням `10%` та `height` зі значенням `70%`. «bb» означає «background building» (фонова будівля); це буде ваша перша будівля.

# --hints--

Ви повинні створити новий елемент `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Ви повинні надати новому `div` клас `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Ви повинні використати селектор класу `.bb1` для стилізації елемента.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Ви повинні надати елементу `.bb1` властивість `width` зі значенням `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Ви повинні надати елементу `.bb1` властивість `height` зі значенням `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
--fcc-editable-region--
    <div class="background-buildings"></div>
--fcc-editable-region--
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

```

