---
id: 5d822fd413a79914d39e98cc
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

В межах `head` вкладіть елемент `meta` з `charset` зі значенням `UTF-8`, елемент `title` з титулом `City Skyline` та елемент `link`, який прив’язує ваш файл `styles.css`.

# --hints--

Ви повинні створити елемент `meta` в межах елемента `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Теґу `meta` потрібно надати `charset` зі значенням `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Елемент `title` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Ваш проєкт повинен мати титул `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Пам’ятайте, що вживання великої літери та написання мають значення.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Ваш код повинен мати елемент `link`.

```js
assert.match(code, /<link/)
```

Ви не повинні змінювати наявні теґи `head`. Переконайтеся, що не видалили кінцевий теґ.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Ви повинні мати один самозакривний елемент `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
