---
id: 5d822fd413a79914d39e98d6
title: Крок 14
challengeType: 0
dashedName: step-14
---

# --description--

Тепер у вас є щось, що нагадує будівлю. Ви готові створити свою першу змінну. Декларації змінних починаються з двох дефісів (`-`) та отримують ім’я і значення, ось так: `--variable-name: value;`. У правилі для класу `bb1` створіть змінну з назвою `--building-color1` та надайте їй значення `#999`.

# --hints--

Ви повинні створити нову змінну з назвою `--building-color1`.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

Ви повинні визначити змінну `--building-color1` в `.bb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

Ви повинні надати `--building-color1` значення `#999`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}

```
