---
id: 5d822fd413a79914d39e98d3
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Вкладіть чотири елементи `div` в контейнер `.bb1`. Надайте їм класи `bb1a`, `bb1b`, `bb1c` та `bb1d`, за таким порядком. Ця будівля складатиметься з чотирьох секцій.

# --hints--

Ви повинні створити чотири нові елементи `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 6);
```

Ви повинні надати одному з нових елементів `div` клас `bb1a`.

```js
assert.exists(document.querySelector('div.bb1a'));
```

Ви повинні надати одному з нових елементів `div` клас `bb1b`.

```js
assert.exists(document.querySelector('div.bb1b'));
```

Ви повинні надати одному з нових елементів `div` клас `bb1c`.

```js
assert.exists(document.querySelector('div.bb1c'));
```

Ви повинні надати одному з нових елементів `div` клас `bb1d`.

```js
assert.exists(document.querySelector('div.bb1d'));
```

Ви повинні розмістити нові елементи `div` в правильному порядку.

```js
function __t(a, b) {
  return [...document.querySelector(a)?.nextElementSibling?.classList]?.includes(b);
}
assert(__t('div.bb1a','bb1b') && __t('div.bb1b','bb1c') && __t('div.bb1c','bb1d'));
```

Ви повинні розмістити нові елементи `div` в межах елемента `.bb1`.

```js
assert.equal(document.querySelectorAll('div.bb1 > div')?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      <div class="bb1"></div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}

```

