---
id: 5d822fd413a79914d39e98cd
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

В CSS ви можете націлити будь-що, використовуючи зірочку. Додайте кордон до всього, використовуючи селектор `*` та надавши йому `border` зі значенням `1px solid black`. Це трюк, який допомагає візуалізувати розташування та розмір елементів. Пізніше ви видалите його.

# --hints--

Ви повинні використати селектор `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Ви повинні використати властивість `border`, щоб стилізувати всі елементи.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Всі елементи повинні мати кордон `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
