---
id: 5f33294a6af5e9188dbdb8f3
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this workshop, you will practice the basics of CSS (Cascading Style Sheets) by building a cafe menu.

Let's start by adding some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

Your code should have an opening `<main>` tag.

```js
assert.match(code, /<main>/i);
```

Your code should have a closing `</main>` tag.

```js
assert.match(code, /<\/main>/i);
```

You should not change your `body` element. Make sure you don't accidentally delete your closing tag.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Your `main` tag should be within your `body` tag.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
