---
id: 5f7692f7c5b3ce22a57788b6
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--

To complete the styling, add the applicable class names `flavor` and `price` to all the remaining `p` elements.

# --hints--

You should have five `.flavor` elements.

```js
assert.lengthOf(document.querySelectorAll('p.flavor'), 5);
```

You should have five `.price` elements.

```js
assert.lengthOf(document.querySelectorAll('p.price'), 5);
```

Your `.flavor` elements should be your `p` elements with the text `French Vanilla`, `Caramel Macchiato`, `Pumpkin Spice`, `Hazelnut`, and `Mocha`.

```js
const flavor = document.querySelectorAll('p.flavor');
assert.match(flavor?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(flavor?.[1]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(flavor?.[2]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(flavor?.[3]?.innerText.trim(), /Hazelnut/i);
assert.match(flavor?.[4]?.innerText.trim(), /Mocha/i);
```

Your `.price` elements should be your `p` elements with the text `3.00`, `3.75`, `3.50`, `4.00`, and `4.50`.

```js
const price = document.querySelectorAll('p.price');
assert.match(price?.[0]?.innerText.trim(), /3\.00/);
assert.match(price?.[1]?.innerText.trim(), /3\.75/);
assert.match(price?.[2]?.innerText.trim(), /3\.50/);
assert.match(price?.[3]?.innerText.trim(), /4\.00/);
assert.match(price?.[4]?.innerText.trim(), /4\.50/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article class="item">
            <p>Caramel Macchiato</p><p>3.75</p>
          </article>
          <article class="item">
            <p>Pumpkin Spice</p><p>3.50</p>
          </article>
          <article class="item">
            <p>Hazelnut</p><p>4.00</p>
          </article>
          <article class="item">
            <p>Mocha</p><p>4.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
```

