---
id: 657ed53c19461d4b95c4757a
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Тепер створіть новий список під назвою `snake_cased_char_list` у функції. Використайте пару порожніх квадратних дужок, щоб створити новий список.

Цей список міститиме символи рядка після того, як ви перетворили їх в зміїний регістр.

# --hints--

Замініть `pass` на порожній список під назвою `snake_cased_char_list` в межах `convert_to_snake_case`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
