---
id: 657e928716b77b2277980276
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви дізнаєтесь про розуміння списків у Python, створивши програму, яка відформатовує рядок у верблюдячомуРегістрі (`camelCase`) або РегістріПаскаля (`PascalCase`) в зміїний_регістр (`snake_case`).

Розуміння списків у Python пропонує стислий спосіб побудови списків без використання циклів або методу `.append()`, що часто призводить до коротшого та швидшого виконання.

Для початку визначте нову функцію під назвою `convert_to_snake_case()`, яка приймає `pascal_or_camel_cased_string` як вхідні дані. Додайте інструкцію `pass` до тіла функції, щоб тимчасово заповнити її.

# --hints--

Визначте нову функцію під назвою `convert_to_snake_case()`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло функції.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Функція має приймати параметр під назвою `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
