---
id: 657ee28cefc4945568287673
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер ви маєте порожній список, тому можете почати перебирати вхідний рядок та перетворювати кожен символ в зміїний регістр.

Використайте цикл `for`, щоб ітерувати `pascal_or_camel_cased_string`. Переконайтесь, що цільова змінна має назву `char`, яка виступає скороченням від character (укр. символ). Наразі додайте інструкцію `pass` до тіла циклу.

# --hints--

Напишіть новий цикл `for` з цільовою змінною під назвою `char`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло циклу.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
