---
id: 657e928716b77b2277980276
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви дізнаєтесь про розуміння списків у Python, створивши програму, яка приймає рядок у верблюдячомуРегістрі (`camelCase`) або РегістріПаскаля (`PascalCase`) та перетворює його в зміїний_регістр (`snake_case`).

Розуміння списків у Python — це стислий спосіб створити список без використання циклів або методу `.append()`. Такий спосіб не лише стислий, але й часто виконується швидше.

Для початку визначте нову функцію під назвою `convert_to_snake_case()`, яка приймає рядок під назвою `pascal_or_camel_cased_string` як вхідні дані. Наразі додайте інструкцію `pass` до функції.

# --hints--

Визначте нову функцію під назвою `convert_to_snake_case()`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло функції.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get("convert_to_snake_case")
      inspect.isfunction(a)
    `));
  }
})
```

Функція має приймати параметр під назвою `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get('convert_to_snake_case')
      'pascal_or_camel_cased_string' in inspect.signature(a).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
