---
id: 657efce98e958b75df86b305
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Використайте метод списку `.append()` так само в межах тіла інструкції `if`, щоб додати змінений символ до списку, який ви створили раніше.

```py
snake_cased_char_list.append(converted_character)
```

Метод `.append()` додає наданий об’єкт в кінець списку, на якому його викликають.

# --hints--

Використайте метод `.append()`, щоб додати перетворений символ до `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
