---
id: 657ef2a86d4e545cec9a85fb
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Великі літери позначають початок нових слів як у верблюдячому регістрі, так і в регістрі Паскаля. Щоб перетворити вхідний рядок в зміїний регістр, потрібно перевірити, чи символи вхідного рядка знаходяться у верхньому регістрі.

Щоб перевірити, чи символ знаходиться у верхньому регістрі, можна використати метод `.isupper()`. Цей метод повертає `True`, якщо символ у верхньому регістрі, а якщо ні — повертає `False`.

В межах циклу `for` додайте інструкцію `if`, щоб перевірити, чи поточний символ знаходиться у верхньому регістрі. Перемістіть інструкцію `pass` у нову інструкцію `if`.

# --hints--

Напишіть нову інструкцію `if` з умовою `char.isupper()`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло `if`.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Замініть інструкцію `pass` в тілі циклу на інструкцію `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
