---
id: 657e928716b77b2277980276
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви створите програму, яка відформатовує рядок у верблюдячомуРегістрі (`camelCase`) або РегістріПаскаля (`PascalCase`) в зміїний_регістр (`snake_case`), використовуючи два підходи. Спочатку ви будете використовувати цикл `for`, а потім — розуміння списку. Ви побачите, як розуміння списку може зробити код більш лаконічним.

Для початку визначте нову функцію під назвою `convert_to_snake_case()`, яка приймає рядок під назвою `pascal_or_camel_cased_string` як вхідні дані. Наразі додайте інструкцію `pass` до функції.

# --hints--

Визначте нову функцію під назвою `convert_to_snake_case()`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло функції.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Функція має приймати параметр під назвою `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
