---
id: 657e928716b77b2277980276
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you are going to learn about list comprehensions in Python by building a program that converts a `camelCase` or `PascalCase` formatted string into a `snake_case` formatted string.

List comprehensions in Python offer a concise way of constructing lists without using loops or the `.append()` method, often resulting in a briefer and faster execution.

To begin, define a new function named `convert_to_snake_case()` that takes `pascal_or_camel_cased_string` as input. Within the function body, include a `pass` statement to temporarily fill the function body.

# --hints--

Визначте нову функцію під назвою `convert_to_snake_case()`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло функції.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Функція має приймати параметр під назвою `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
