---
id: 657effaa2a5e0277d71f9cbe
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Наразі змінна `snake_cased_char_list` містить список перетворених символів. Щоб об’єднати ці символи в один рядок, можна використати метод `.join()`.

Метод `join` об’єднує всі елементи списку в один рядок, де вони розділені визначеним рядком (відомий як роздільник).

```py
result_string = ''.join(characters)
```

Наведений вище приклад об’єднує елементи списку `characters` в один рядок, де елементи об’єднані за допомогою роздільника у вигляді порожнього рядка.

Тепер одразу після циклу `for` скористайтеся методом `.join()`, щоб об’єднати елементи в `snake_cased_char_list`, використовуючи порожній рядок як роздільник. Призначте результат до нової змінної під назвою `snake_cased_string`.


# --hints--

Утворіть рядок з символів зі списку `snake_cased_char_list` за допомогою методу `.join()` та призначте рядок до змінної під назвою `snake_cased_string` після циклу `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
