---
id: 657ef2a86d4e545cec9a85fb
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Великі літери у верблюдячому регістрі чи регістрі Паскаля вказують на початок нових слів.

Використайте інструкцію `if` разом з методом рядка `.isupper()` в тілі циклу, щоб знайти символи у верхньому регістрі та перемістіть `pass` до нової інструкції `if`.

# --hints--

Напишіть нову інструкцію `if` з умовою `char.isupper()`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло `if`.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Замініть інструкцію `pass` в тілі циклу на інструкцію `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
