---
id: 657f28a0482132aca51a9212
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Замініть `pass` на виклик функції `main()` в межах тіла інструкції `if`. Ви маєте побачити, що наданий рядок у верблюдячому регістрі або регістрі Паскаля перетворений на рядок в зміїному регістрі.

# --hints--

Додайте виклик `main()` до інструкції `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        assert.match(transformedCode, /\n +main\s*\(\s*\)/);
    }
})
```

Інструкція `if` не повинна містити `pass`.

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+__name__\s+==\s+'__main__'\s*/);

        assert.notMatch(block_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
if __name__ == '__main__':
    pass
--fcc-editable-region--
```
