---
id: 657f456223b8c1187b461987
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Якщо почати розуміння списку зі схожої інструкції `if`, то Python вимагатиме умову `else` в цьому виразі.

```py
snake_cased_char_list = [
    '_' + char.lower() if char.isupper()
    else char
]
```

Python інтерпретує оновлений вираз як «додати `'_' + char.lower()` до списку, якщо `char` у верхньому регістрі, в іншому випадку — додати `char`», що стосується як великих, так і малих літер у вхідному рядку.

Додайте умову `else` між квадратними дужками.

# --hints--

Додайте `else char` до квадратних дужок `snake_cased_char_list` після `'_' + char.lower() if char.isupper()`.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*("|')_\1\s*\+\s*char\.lower\(\s*\)\s+if\s+char\.isupper\(\s*\)\s*else\s+char\s*\]/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()

    ]
--fcc-editable-region--



def main():
    print(convert_to_snake_case('aLongAndComplexString'))

if __name__ == '__main__':
    main()
```
