---
id: 657f456223b8c1187b461987
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

<dfn>Розуміння списку</dfn> в Python — це конструкція, яка дозволяє створити новий список, застосовуючи вираз до кожного елемента в наявному ітерованому об’єкті та додатково фільтруючи елементи за умовою. Такий спосіб не лише стислий, але й часто виконується швидше.

Базове розуміння списку складається з виразу, за яким слідує умова `for`:

```py
spam = [i * 2 for i in iterable]
```

Вище використовується змінна `i`, щоб ітерувати над `iterable`. Кожен елемент вихідного списку отримано шляхом обчислення виразу `i * 2` на поточній ітерації.

На цьому кроці вам потрібно заповнити порожній список `snake_cased_char_list` за допомогою синтаксису розуміння списку.

Перетворіть порожній список на розуміння списку, яке перетворює кожен символ в `pascal_or_camel_cased_string` на символ у нижньому регістрі та додає знак підкреслення перед ним (код, який ви закоментували раніше, може допомогти написати вираз). Використайте `char`, щоб ітерувати над `pascal_or_camel_cased_string`.

# --hints--

Перетворіть `snake_cased_char_list` на розуміння списку, яке ітерує над `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

Розуміння списку має використовувати `char`, щоб ітерувати над `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

Розуміння списку має оцінювати `'_' + char.lower()` для кожного `char` у `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower()")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
