---
id: 657f0044be09db790b1eb1c5
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Рядки в регістрі Паскаля починаються з великої літери. Оскільки ви змінили регістр всіх таких символів на нижній і поставили перед ними підкреслення, є ймовірність, що перетворений рядок в зміїному регістрі має підкреслення напочатку.

Найпростіший спосіб видалити непотрібний символ — використати метод рядка `.strip()` та передати знак підкреслення як аргумент.

```py
snake_cased_string.strip('_')
```

Переконайтесь, що зберегли рядок в змінній під назвою `clean_snake_cased_string` на тому ж рівні, що й змінна `snake_cased_string`.

# --hints--

Призначте рядок, отриманий з `snake_cased_string.strip('_')`, до змінної під назвою `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)

--fcc-editable-region--
```
