---
id: 657f04ed0035f47ed04d0f1f
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Перш ніж запускати функцію `main()`, вам потрібно переконатися, що файл працює як сценарій. Додайте інструкцію `if` на тому ж рівні, що й дві наявні функції, та перевірте, чи `__name__ == '__main__'`.

Не забудьте використати `pass`, щоб заповнити тіло інструкції `if`.

# --hints--

Напишіть умову `if`, щоб перевірити, чи `__name__ == '__main__'` оцінюється як `True`. Не забудьте двокрапку вкінці та використайте `pass`, щоб заповнити тіло `if`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        assert.match(transformedCode, /\nif\s+__name__\s*==\s*("|')__main__\1\s*:/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

--fcc-editable-region--
```
