---
id: 657f4345abe7f2161f99f1ad
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Для початку замініть `pass` на змінну `snake_cased_char_list` та призначте до неї порожній список. Використайте квадратні дужки, щоб створити список, але не розміщуйте нічого між дужками.

Розмістіть дужки в окремих рядках, щоб між ними залишався простір, де ви зможете працювати над кодом для розуміння списків.

# --hints--

Оголосіть порожній список під назвою `snake_cased_char_list` в межах `convert_to_snake_case`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, /clean_snake_cased_string\s+    snake_cased_char_list\s*=\s*\[\n*\s*\n*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string
--fcc-editable-region--
    pass


--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

if __name__ == '__main__':
    main()
```
