---
id: 657f0044be09db790b1eb1c5
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Рядки в регістрі Паскаля починаються з великої літери. Після перетворення всіх символів у нижній регістр та додавання до них знаку підкреслення є ймовірність, що на початку рядка з’явиться зайвий знак підкреслення.

Найпростіший спосіб виправити це — використати метод `.strip()`, який видаляє з рядка будь-які символи напочатку чи вкінці, передані як аргумент. Наприклад:

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

Метод `strip()` застосовано до `original_string`. Це видаляє будь-яке підкреслення напочатку чи вкінці. Результатом наведеного вище прикладу буде рядок `'example_string'`.

Оголосіть нову змінну під назвою `clean_snake_cased_string` та призначте до неї результат методу `.strip()`, застосованого до `snake_cased_string`, передаючи `'_'` як аргумент методу.

# --hints--

Призначте рядок, отриманий з `snake_cased_string.strip('_')`, до змінної під назвою `clean_snake_cased_string`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
