---
id: 657f47b12c51e41b3149e584
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Вам досі потрібно об’єднувати елементи списку в рядок, видаляти будь-які зайві підкреслення та повертати рядок. Хоча це можна робити як і раніше, розгляньте коротшу альтернативу.

```py
return ''.join(snake_cased_char_list).strip('_')
```

Цей єдиний рядок коду об’єднає список символів у рядок, видалить будь-які зайві підкреслення та поверне отриманий рядок. Додайте цей рядок на тому ж рівні, що й змінна `snake_cased_char_list`, та в межах функції `convert_to_snake_case()`.

# --hints--

Поверніть `''.join(snake_cased_char_list).strip('_')` в кінці функції `convert_to_snake_case()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +return\s+('|")\1\.join\(\s*snake_cased_char_list\s*\)\.strip\(\s*("|')_\2\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

--fcc-editable-region--



def main():
    print(convert_to_snake_case('aLongAndComplexString'))

if __name__ == '__main__':
    main()
```
