---
id: 6852db755c9c1274d75aae9a
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this workshop, you will create a PIN extractor. The PIN digits are hidden in each line of a poem.

Start by creating a function called `pin_extractor` with a parameter `poem`. A function without a body raises an error, so add `pass` inside the function as a temporary placeholder to make it valid.

# --hints--

You should have a function named `pin_extractor`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

The `pin_extractor` function should have a parameter named `poem`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
