---
id: 6855918e87c33c6979d21d3f
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Now the function is doing a good job in extracting the code from the poem, but there is an edge case to consider: the line of the poem could be shorter than expected, and that would make the function error out (you can try this removing the word `bright` from the poem).

Put the line `secret_code += str(len(words[line_index]))` in an `if` statement that checks that there are enough words in the `words` list.

# --hints--

Inside the loop you should have an `if` statement.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

The if statement condition should be `len(words) > line_index`.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

Inside the `if` statement you should have `secret_code += str(len(words[line_index]))`.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))
--fcc-editable-region--

```
