---
id: 615f522dea4f776f64dc3e91
title: ステップ 31
challengeType: 0
dashedName: step-31
---

# --description--

`rem` という単位は `root em` を表し、`html` 要素のフォントサイズに比例します。

`.small-text` セレクターを作成し、`font-size` を `0.85rem` に設定してください。これはおおよそ `13.6px` に相当します。(`html` 要素の `font-size` を `16px` に設定したことを思い出してください。)

# --hints--

`.small-text` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text'));
```

`.small-text` セレクターの `font-size` プロパティを `0.85rem` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.small-text')?.fontSize === '0.85rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
      </div>
    </div>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
