---
id: 615f34ecc1091b4fd5a8a484
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

`head` 要素内に `link` 要素を追加し、`rel` 属性を `stylesheet` に、`href` 属性を `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` に設定してください。

これにより、`Open Sans` フォントファミリーのうち太さ (font weight) の値が `400`、`700`、`800` のフォントがインポートされます。

更に、`styles.css` ファイルにリンクする `link` 要素を追加してください。

# --hints--

自己終了要素 `link` が 2 つ必要です。

```js
assert(document.querySelectorAll('link').length === 2);
```

`link` 要素はどちらも `rel` 属性が `stylesheet` に設定されている必要があります。

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

`link` 要素のうち 1 つは `href` 属性を `./styles.css` に設定されている必要があります。

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

`link` 要素のうち 1 つは `href` 属性を `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` に設定されている必要があります。

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
