---
id: 615f357957e370510f21ea16
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`body` セレクターを作成し、`font-family` に `Open Sans` とフォールバックとして `sans-serif` を設定してください。

CSS では名前にスペースを含むフォントは引用符で囲まなければならないことに注意しましょう。

# --hints--

`body` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` セレクターの `font-family` プロパティを `"Open Sans", sans-serif` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
