---
id: 615f51e4e5b24a6e80eccce1
title: ステップ 30
challengeType: 0
dashedName: step-30
---

# --description--

`.calories-info` 要素内に `div` 要素を作成してください。 その `div` 要素に `class` 属性を追加して、`left-container` に設定してください。 またその `div` 要素の中に `h2` 要素を作成して、`Amount per serving` というテキストを追加してください。 その `h2` 要素に `class` 属性を追加して、`bold small-text` に設定してください。

# --hints--

`.calories-info` の要素内に、新しい `div` 要素を作成する必要があります。

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.localName === 'div');
```

新しく追加した `div` 要素の `class` 属性を `left-container` に設定してください。

```js
assert(document.querySelector('.calories-info')?.children?.[0]?.classList.contains('left-container'));
```

`.left-container` の要素内に、新しい `h2` 要素を作成する必要があります。

```js
assert(document.querySelector('.left-container')?.children?.[0]?.localName === 'h2');
```

新しく追加した `h2` 要素の `class` 属性を `bold small-text` に設定してください。

```js
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('bold'));
assert(document.querySelector('.left-container')?.children?.[0]?.classList.contains('small-text'));
```

新しく追加した `h2` 要素には `Amount per serving` というテキストが必要です。

```js
assert(document.querySelector('.left-container')?.children?.[0]?.innerText === 'Amount per serving');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
--fcc-editable-region--
    <div class="calories-info">

    </div>
--fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}
```
