---
id: 615f38279e5c3d53692ea441
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

`h1` 要素と `p` 要素を `div` 要素で囲みましょう。 その `div` 要素の `class` 属性を `label` に設定してください。

# --hints--

新しい `div` 要素を作成してください。

```js
assert(document.querySelector('div'));
```

新しい `div` 要素の `class` 属性を `label` に設定する必要があります。

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

`h1` 要素と `p` 要素は新しく追加した `.label` 要素の中にある必要があります。

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

--fcc-editable-region--
<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
--fcc-editable-region--
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
