---
id: 615f4172e9eec061d6456221
title: ステップ 19
challengeType: 0
dashedName: step-19
---

# --description--

栄養成分表示ラベルには、重要な情報に注意を引くために太字のテキストがたくさんあります。 太字にしたい要素を 1 つずつ選択するより、クラスを使って全要素に太字のスタイルを適用した方が効率的ですね。

2 つ目の `p` 要素の `class` 属性を `bold` に設定してください。

# --hints--

2 つ目の `p` 要素の `class` 属性を `bold` に設定する必要があります。

```js
assert(document.querySelectorAll('p')?.[1]?.classList?.contains('bold'));
```

1 つ目の `p` 要素の `class` 属性は `bold` に設定しないでください。

```js
assert(!document.querySelector('p')?.classList?.contains('bold'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}
```
