---
id: 615f2abbe7d18d49a1e0e1c8
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

基本的な HTML ボイラープレートはあらかじめ用意しておきました。

`body` 要素内に `h1` 要素を作成し、`Nutrition Facts` というテキストを設定してください。

# --hints--

新しい `h1` 要素を 1 つ追加してください。

```js
assert.exists(document.querySelector('h1'));
```

`h1` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

`h1` 要素のテキストは `Nutrition Facts` でなければなりません。

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
