---
id: 615f34948891834dd77655a6
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

2 つ目の `p` 要素と `Serving size 2/3 cup (55g)` というテキストを追加してください。

# --hints--

2 つ目の `p` 要素が必要です。

```js
assert(document.querySelectorAll('p')?.length === 2);
```

2 つ目の `p` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

2 つ目の `p` 要素は既存の `p` 要素の後に置く必要があります。

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

2 つ目の `p` 要素には `Serving size 2/3 cup (55g)` というテキストが必要です。

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
