---
id: 615f3ed16592445e57941aec
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

スペースが限られている場合は特に、重要なコンテンツを分離したりグループ化するのに直線が役立ちます。

`h1` 要素の下に `div` 要素を作成し、`class` 属性を `divider` に設定してください。

# --hints--

新しい `div` を作成してください。

```js
assert(document.querySelectorAll('div')?.length === 2);
```

新しい `div` の `class` 属性を `divider` に設定する必要があります。

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

`.divider` 要素は `.label` 要素の中にある必要があります。

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

`.divider` 要素は `h1` 要素の後に置く必要があります。

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
--fcc-editable-region--
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
--fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
