---
id: 615f2d4150fe0d4cbd0f2628
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`h1` 要素の下に `p` 要素を追加して、テキストを `8 servings per container` にしてください。

# --hints--

新しい `p` 要素を 1 つ追加してください。

```js
assert.exists(document.querySelector('p'));
```

`p` 要素は `body` 要素の中にある必要があります。

```js
assert(document.querySelector('p')?.parentElement?.localName === 'body');
```

`p` 要素を `h1` 要素の後に置く必要があります。

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

`p` 要素には `8 servings per container` というテキストが必要です。

```js
assert(document.querySelector('p')?.innerText === '8 servings per container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>

</body>
</html>
--fcc-editable-region--
```

```css

```
