---
id: 615f38cabc64e3556f98cc1a
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

境界線を使ってコンテンツをグループ化したり強調したりできます。

`.label` セレクターを作成して、`border` を `2px solid black` に設定してください。

# --hints--

`.label` セレクターが必要です。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label'));
```

`.label` セレクターの `border` プロパティを `2px solid black` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('.label')?.border === '2px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

--fcc-editable-region--

--fcc-editable-region--
```
