---
id: 615f51257a8a516d80b6c743
title: ステップ 29
challengeType: 0
dashedName: step-29
---

# --description--

`.lg` 要素の下に新しい `div` を作成して、`class` 属性を `calories-info` に設定してください。

# --hints--

`div` 要素をもう 1 つ作成してください。

```js
assert(document.querySelectorAll('div')?.length === 4);
```

新しい `div` の `class` 属性を `calories-info` に設定する必要があります。

```js
assert(document.querySelector('.label')?.lastElementChild?.classList?.contains('calories-info'));
```

新しい `div` は `.lg` 要素の後に置く必要があります。

```js
assert(document.querySelector('.label')?.lastElementChild?.previousElementSibling?.classList?.contains('lg'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Nutrition Label</title>
    <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="label">
      <header>
        <h1 class="bold">Nutrition Facts</h1>
        <div class="divider"></div>
        <p>8 servings per container</p>
        <p class="bold">Serving size <span class="right">2/3 cup (55g)</span></p>
      </header>
      <div class="divider lg"></div>
--fcc-editable-region--

--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.right {
  float: right;
}

.lg {
  height: 10px;
}

.lg, .md {
  background-color: black;
  border: 0;
}
```
