---
id: 615f3b091162165948e1cb33
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

ブラウザーの開発者ツールで `.label` 要素を調べると、幅が 270 ピクセルではなく実際は 288 ピクセルになっていることに気づくかもしれません。 これは、ブラウザーが要素のサイズを決定する際、デフォルトでは境界線とパディングが含まれるためです。

これを解消するには、`*` セレクターを作成して `box-sizing` プロパティを `border-box` に設定することで、ボックスモデルをリセットします。

# --hints--

`*` セレクターを作成する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` セレクターの `box-sizing` プロパティを `border-box` に設定する必要があります。

```js
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}
```
