---
id: 66faa60000000000000000e4
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--
Spawn coins at intervals; eating one adds 50 to the distance score. Animate the spritesheet's 8-frame spin at 12 fps.

At the end of `create`:

```js
this.coins = this.physics.add.group({ allowGravity: false });
this.anims.create({
  key: "spin",
  frames: this.anims.generateFrameNumbers("coin", { frames: [0, 1, 2, 3, 4, 5, 6, 7] }),
  frameRate: 12,
  repeat: -1
});
this.spawnCoin = () => {
  const c = this.coins.create(520, 200, "coin", 0);
  c.body.setVelocityX(-this.scrollSpeed);
  c.play("spin");
};
this.coinTimer = this.time.addEvent({
  delay: 2700,
  loop: true,
  callback: this.spawnCoin
});
this.physics.add.overlap(this.player, this.coins, (_p, coin) => {
  coin.destroy();
  this.distance += 50;
});
```

2700 ms makes coins less frequent than cacti — they're rewards, not noise.

# --hints--
You should create `this.coins` as a `physics.add.group({ allowGravity: false })`.

```js
assert.match(code, /this\.coins\s*=\s*this\.physics\.add\.group\(\s*\{\s*allowGravity:\s*false\s*\}\s*\)/);
```

You should register a `"spin"` animation with frames `[0, 1, 2, 3, 4, 5, 6, 7]` and `frameRate: 12`.

```js
assert.match(code, /this\.anims\.create\(\s*\{\s*key:\s*["\']spin["\'][\s\S]*frames:\s*\[\s*0\s*,\s*1\s*,\s*2\s*,\s*3\s*,\s*4\s*,\s*5\s*,\s*6\s*,\s*7\s*\][\s\S]*frameRate:\s*12/);
```

You should define `this.spawnCoin` to create a coin at `(520, 200)` with leftward velocity and play the spin animation.

```js
assert.match(code, /this\.spawnCoin\s*=[\s\S]*this\.coins\.create\(\s*520\s*,\s*200\s*,\s*["\']coin["\']\s*,\s*0\s*\)[\s\S]*setVelocityX\(\s*-\s*this\.scrollSpeed\s*\)[\s\S]*\.play\(\s*["\']spin["\']\s*\)/);
```

You should register a coin overlap callback that destroys the coin and adds 50 to `this.distance`.

```js
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.player\s*,\s*this\.coins[\s\S]*coin\.destroy\(\s*\)[\s\S]*this\.distance\s*\+=\s*50/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.image(
      "cactus",
      "/curriculum-assets/phaser/spritesheets/cactus.png"
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) this.lastGroundedAt = time;
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        this.lastJumpPressedAt = time;
      }
      const canJump = time - this.lastGroundedAt < 120;
      const buffered = time - this.lastJumpPressedAt < 120;
      if (canJump && buffered) {
        this.player.setVelocityY(-600);
        this.lastGroundedAt = -Infinity;
        this.lastJumpPressedAt = -Infinity;
      }
    };
    this.scrollSpeed = 200;
    this.distance = 0;
    this.obstacles = this.physics.add.group({ allowGravity: false });
    this.lastObstacleX = -Infinity;
    this.spawnObstacle = () => {
      if (520 - this.lastObstacleX < 200) return;
      const cactus = this.obstacles.create(520, 248, "cactus");
      cactus.body.setVelocityX(-this.scrollSpeed);
      cactus.setOrigin(0.5, 1);
      cactus.body.setSize(16, 28).setOffset(4, 4);
      this.lastObstacleX = 520;
    };
    this.spawnObstacle();
    this.obstacleTimer = this.time.addEvent({
      delay: 1500,
      loop: true,
      callback: this.spawnObstacle
    });
    this.gameOver = false;
    this.gameOverDeath = () => {
      if (this.gameOver) return;
      this.gameOver = true;
      this.obstacleTimer.remove(false);
      this.player.setVelocity(0, 0);
      this.player.anims.stop();
      const dist = Math.floor(this.distance);
      if (dist > this.bestDistance) {
        this.bestDistance = dist;
        localStorage.setItem("runner:best", String(dist));
        this.bestText.setText("Best: " + dist);
      }
      this.add
        .text(240, 140, "Game Over", {
          fontSize: "24px",
          color: "#ffffff"
        })
        .setOrigin(0.5)
        .setDepth(200);
    };
    this.physics.add.overlap(this.player, this.obstacles, () => {
      this.gameOverDeath();
    });
    this.distanceText = this.add
      .text(8, 8, "Dist: 0", {
        fontSize: "12px",
        color: "#ffffff"
      })
      .setScrollFactor(0)
      .setDepth(120);
    this.bestDistance = parseInt(
      localStorage.getItem("runner:best") || "0",
      10
    );
    this.bestText = this.add
      .text(8, 24, "Best: " + this.bestDistance, {
        fontSize: "10px",
        color: "#cccccc"
      })
      .setScrollFactor(0)
      .setDepth(120);
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
    if (!this.gameOver) {
      this.distance += (this.scrollSpeed * delta) / 1000;
      const speedLevel = Math.floor(this.distance / 1000);
      this.scrollSpeed = 200 + speedLevel * 50;
      this.obstacleTimer.delay = Math.max(
        500,
        1500 - speedLevel * 100
      );
    }
    this.distanceText.setText(
      "Dist: " + Math.floor(this.distance)
    );
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
