---
id: 66faa60000000000000000cc
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--
Stack three `tileSprite`s back-to-front and define a closure that advances each layer at its own speed. Add order is render order, so the sky goes in first (back), then mountains, then ground (front). Each `tileSprite(x, y, width, height, key)` is centred at `(x, y)` and tiles its texture across `(width, height)`. Chain `setScrollFactor(0)` so the layers stay locked to the camera even if you add a shake later. Then capture the per-frame parallax advance in a `this.tickParallax(dt)` closure — when you add a cloud layer in step 34 you'll just edit one block.

Inside `create`:

```js
this.sky = this.add
  .tileSprite(240, 160, 480, 320, "sky")
  .setScrollFactor(0);
this.mountains = this.add
  .tileSprite(240, 200, 480, 160, "mountains")
  .setScrollFactor(0);
this.groundTile = this.add
  .tileSprite(240, 288, 480, 64, "ground")
  .setScrollFactor(0);
this.tickParallax = (dt) => {
  this.sky.tilePositionX += 5 * dt;
  this.mountains.tilePositionX += 25 * dt;
  this.groundTile.tilePositionX += 240 * dt;
};
```

The sky drifts at 5 px/s, mountains at 25 px/s, ground at 240 px/s. The 48× ratio between sky and ground is what tells your eye the sky is far and the ground is close.

# --hints--
You should chain `.setScrollFactor(0)` after creating `this.sky` with the `"sky"` key.

```js
assert.match(code, /this\.sky\s*=\s*this\.add\s*\.tileSprite\(\s*240\s*,\s*160\s*,\s*480\s*,\s*320\s*,\s*["\']sky["\']\s*\)\s*\.setScrollFactor\(\s*0\s*\)/);
```

You should create `this.mountains` and chain `.setScrollFactor(0)`.

```js
assert.match(code, /this\.mountains\s*=\s*this\.add\s*\.tileSprite\(\s*240\s*,\s*200\s*,\s*480\s*,\s*160\s*,\s*["\']mountains["\']\s*\)\s*\.setScrollFactor\(\s*0\s*\)/);
```

You should create `this.groundTile` and chain `.setScrollFactor(0)`.

```js
assert.match(code, /this\.groundTile\s*=\s*this\.add\s*\.tileSprite\(\s*240\s*,\s*288\s*,\s*480\s*,\s*64\s*,\s*["\']ground["\']\s*\)\s*\.setScrollFactor\(\s*0\s*\)/);
```

You should define `this.tickParallax = (dt) => { ... }` advancing the sky at `5 * dt`, mountains at `25 * dt`, and ground at `240 * dt`.

```js
assert.match(code, /this\.tickParallax\s*=\s*\(\s*dt\s*\)\s*=>\s*\{[\s\S]*this\.sky\.tilePositionX\s*\+=\s*5\s*\*\s*dt[\s\S]*this\.mountains\.tilePositionX\s*\+=\s*25\s*\*\s*dt[\s\S]*this\.groundTile\.tilePositionX\s*\+=\s*240\s*\*\s*dt/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
  }

  create() {
--fcc-editable-region--
--fcc-editable-region--
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
