---
id: 66faa60000000000000000d4
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--
**Coyote time** lets the player jump for a brief grace period after leaving a ledge — named for cartoons where Wile E. Coyote runs off a cliff and stays in the air a beat too long. Without it, players feel cheated when they press SPACE the same frame they walk off.

Replace the entire `this.updateJump` closure with the coyote-aware version. Track the latest grounded time, and allow a jump if the press happens within 120 ms of leaving the ground:

```js
this.updateJump = (time) => {
  const grounded =
    this.player.body.blocked.down || this.player.body.touching.down;
  if (grounded) this.lastGroundedAt = time;
  if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
    if (time - this.lastGroundedAt < 120) {
      this.player.setVelocityY(-600);
      this.lastGroundedAt = -Infinity;
    }
  }
};
```

Resetting `this.lastGroundedAt = -Infinity` after firing is what blocks repeated jumps from a single grounded frame.

# --hints--
You should stamp `this.lastGroundedAt = time` whenever the player is grounded.

```js
assert.match(code, /if\s*\(\s*grounded\s*\)\s*this\.lastGroundedAt\s*=\s*time/);
```

You should allow a jump when `time - this.lastGroundedAt < 120`.

```js
assert.match(code, /if\s*\(\s*time\s*-\s*this\.lastGroundedAt\s*<\s*120\s*\)/);
```

After firing the jump you should set `this.lastGroundedAt = -Infinity` to prevent a second jump.

```js
assert.match(code, /this\.player\.setVelocityY\(\s*-600\s*\)\s*;\s*\n\s*this\.lastGroundedAt\s*=\s*-Infinity/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
--fcc-editable-region--
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (
        Phaser.Input.Keyboard.JustDown(this.jumpKey) &&
          grounded
      ) {
        this.player.setVelocityY(-600);
      }
    };
--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
