---
id: 66faa60000000000000000cf
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--
Create the physics sprite at `(80, 200)` using frame 12 (the first frame of the rightward-walk row). Phaser's arcade physics gives it a body so gravity can pull it down once the world has a floor. Tell it to collide with the world bounds, then build the four-frame `run` animation cycling frames 12-15 at 10 fps, and play it.

At the end of `create` (after the parallax tile sprites):

```js
this.player = this.physics.add.sprite(80, 200, "hero", 12);
this.player.setCollideWorldBounds(true);
this.anims.create({
  key: "run",
  frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
  frameRate: 10,
  repeat: -1
});
this.player.play("run");
```

The player is now in the world. Without a ground body the next step will add, gravity will drop the hero through the floor — that's expected for one frame of progress.

# --hints--
You should create `this.player` with `this.physics.add.sprite(80, 200, "hero", 12)`.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\(\s*80\s*,\s*200\s*,\s*["\']hero["\']\s*,\s*12\s*\)/);
```

You should call `this.player.setCollideWorldBounds(true)`.

```js
assert.match(code, /this\.player\.setCollideWorldBounds\(\s*true\s*\)/);
```

You should create the `"run"` animation with frames `[12, 13, 14, 15]` at `frameRate: 10` and `repeat: -1`.

```js
assert.match(code, /this\.anims\.create\(\s*\{\s*key:\s*["\']run["\']\s*,\s*frames:\s*this\.anims\.generateFrameNumbers\(\s*["\']hero["\']\s*,\s*\{\s*frames:\s*\[\s*12\s*,\s*13\s*,\s*14\s*,\s*15\s*\]\s*\}\s*\)\s*,\s*frameRate:\s*10\s*,\s*repeat:\s*-1\s*\}\s*\)/);
```

You should call `this.player.play("run")`.

```js
assert.match(code, /this\.player\.play\(\s*["\']run["\']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
