---
id: 66faa60000000000000000ce
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--
Time to add the runner. The hero spritesheet is 64 × 64 with four rows of 16 × 16 frames; the rightward-walk row sits at frames 12-15. Load the spritesheet in `preload`, then in `create` add a physics sprite at the fixed runner X position (`80`), build a four-frame `run` animation, and play it.

First, append the hero loader at the end of `preload`:

```js
this.load.spritesheet(
  "hero",
  "/curriculum-assets/phaser/spritesheets/hero-walk.png",
  { frameWidth: 16, frameHeight: 16 }
);
```

# --hints--
You should call `this.load.spritesheet("hero", "/curriculum-assets/phaser/spritesheets/hero-walk.png", { frameWidth: 16, frameHeight: 16 })`.

```js
assert.match(code, /this\.load\.spritesheet\(\s*["\']hero["\']\s*,\s*["\']\/curriculum-assets\/phaser\/spritesheets\/hero-walk\.png["\']\s*,\s*\{\s*frameWidth:\s*16\s*,\s*frameHeight:\s*16\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
