---
id: 66faa60000000000000000f1
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--
Add a **second** camera covering the same `(0, 0, 480, 320)` viewport. Phaser renders cameras in the order they were added — so the new `uiCam` draws *on top of* the main camera. Until you tell it what to ignore (next step), both cameras still draw the entire scene, so the world appears doubled. Don't panic — that's the diagnostic that proves you have two cameras live.

Right after the HUD container:

```js
this.uiCam = this.cameras.add(0, 0, 480, 320);
```

Lecture recap: `cameras.add(x, y, width, height)` creates and returns a new camera; it's pushed to the end of the manager's list, which controls render order.

# --hints--
You should add a second camera with `this.uiCam = this.cameras.add(0, 0, 480, 320)`.

```js
assert.match(code, /this\.uiCam\s*=\s*this\.cameras\.add\(\s*0\s*,\s*0\s*,\s*480\s*,\s*320\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "clouds",
      "/curriculum-assets/phaser/spritesheets/parallax-clouds.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.image(
      "cactus",
      "/curriculum-assets/phaser/spritesheets/cactus.png"
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "flyer",
      "/curriculum-assets/phaser/spritesheets/flying-enemy.png",
      { frameWidth: 24, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.clouds = this.add
      .tileSprite(240, 120, 480, 96, "clouds")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.clouds.tilePositionX += 12 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) {
        this.lastGroundedAt = time;
        this.didDoubleJump = false;
      }
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        this.lastJumpPressedAt = time;
      }
      const canJump = time - this.lastGroundedAt < 120;
      const buffered = time - this.lastJumpPressedAt < 120;
      if (canJump && buffered) {
        this.player.setVelocityY(-600);
        this.lastGroundedAt = -Infinity;
        this.lastJumpPressedAt = -Infinity;
      } else if (
        buffered &&
        this.canDoubleJump &&
        !this.didDoubleJump &&
        !grounded
      ) {
        this.player.setVelocityY(-520);
        this.didDoubleJump = true;
        this.lastJumpPressedAt = -Infinity;
      }
    };
    this.scrollSpeed = 200;
    this.distance = 0;
    this.obstacles = this.physics.add.group({ allowGravity: false });
    this.lastObstacleX = -Infinity;
    this.spawnObstacle = () => {
      if (520 - this.lastObstacleX < 200) return;
      const cactus = this.obstacles.create(520, 248, "cactus");
      cactus.body.setVelocityX(-this.scrollSpeed);
      cactus.setOrigin(0.5, 1);
      cactus.body.setSize(16, 28).setOffset(4, 4);
      this.lastObstacleX = 520;
    };
    this.spawnObstacle();
    this.obstacleTimer = this.time.addEvent({
      delay: 1500,
      loop: true,
      callback: this.spawnObstacle
    });
    this.gameOver = false;
    this.gameOverDeath = () => {
      if (this.gameOver) return;
      this.gameOver = true;
      if (!this.reduceMotion) this.cameras.main.shake(220, 0.012);
      this.obstacleTimer.remove(false);
      this.player.setVelocity(0, 0);
      this.player.anims.stop();
      const dist = Math.floor(this.distance);
      if (dist > this.bestDistance) {
        this.bestDistance = dist;
        localStorage.setItem("runner:best", String(dist));
        this.bestText.setText("Best: " + dist);
      }
      this.add
        .text(240, 140, "Game Over", {
          fontSize: "24px",
          color: "#ffffff"
        })
        .setOrigin(0.5)
        .setDepth(200);
    };
    this.physics.add.overlap(this.player, this.obstacles, () => {
      this.gameOverDeath();
    });
    this.distanceText = this.add
      .text(8, 8, "Dist: 0", {
        fontSize: "12px",
        color: "#ffffff"
      })
      .setScrollFactor(0)
      .setDepth(120);
    this.bestDistance = parseInt(
      localStorage.getItem("runner:best") || "0",
      10
    );
    this.bestText = this.add
      .text(8, 24, "Best: " + this.bestDistance, {
        fontSize: "10px",
        color: "#cccccc"
      })
      .setScrollFactor(0)
      .setDepth(120);
    this.hud = this.add.container(0, 0).setScrollFactor(0).setDepth(120);
    this.hud.add([this.distanceText, this.bestText]);
--fcc-editable-region--

--fcc-editable-region--
    this.coins = this.physics.add.group({ allowGravity: false });
    this.anims.create({
      key: "spin",
      frames: this.anims.generateFrameNumbers("coin", { frames: [0, 1, 2, 3, 4, 5, 6, 7] }),
      frameRate: 12,
      repeat: -1
    });
    this.spawnCoin = () => {
      const c = this.coins.create(520, 200, "coin", 0);
      c.body.setVelocityX(-this.scrollSpeed);
      c.play("spin");
    };
    this.coinTimer = this.time.addEvent({
      delay: 2700,
      loop: true,
      callback: this.spawnCoin
    });
    this.physics.add.overlap(this.player, this.coins, (_p, coin) => {
      coin.destroy();
      this.distance += 50;
    });
    this.canDoubleJump = false;
    this.didDoubleJump = false;
    this.powerups = this.physics.add.group({ allowGravity: false });
    this.spawnDoubleJump = () => {
      const p = this.add.rectangle(520, 200, 14, 14, 0x66ddff);
      this.physics.add.existing(p);
      this.powerups.add(p);
      p.body.setVelocityX(-this.scrollSpeed);
      p.kind = "double";
    };
    this.powerupTimer = this.time.addEvent({
      delay: 6500,
      loop: true,
      callback: this.spawnDoubleJump
    });
    this.physics.add.overlap(this.player, this.powerups, (_p, item) => {
      if (item.kind === "double") {
        this.canDoubleJump = true;
      }
      item.destroy();
      if (!this.reduceMotion) {
        this.cameras.main.flash(160, 102, 221, 255);
      }
    });
    this.flyers = this.physics.add.group({ allowGravity: false });
    this.spawnFlyer = () => {
      const f = this.flyers.create(520, 200, "flyer", 0);
      f.body.setVelocityX(-this.scrollSpeed);
    };
    this.anims.create({
      key: "fly",
      frames: this.anims.generateFrameNumbers("flyer", { frames: [0, 1, 2, 3] }),
      frameRate: 8,
      repeat: -1
    });
    this.flyerTimer = this.time.addEvent({
      delay: 4500,
      loop: true,
      callback: this.spawnFlyer
    });
    this.physics.add.overlap(this.player, this.flyers, () => {
      this.gameOverDeath();
    });
    this.duckKey = this.input.keyboard.addKey("DOWN");
    this.applyDuck = (down) => {
      if (down) {
        this.player.body.setSize(12, 8).setOffset(2, 8);
      } else {
        this.player.body.setSize(16, 16).setOffset(0, 0);
      }
    };
    const dustGfx = this.make.graphics({ x: 0, y: 0, add: false });
    dustGfx.fillStyle(0xffffff, 0.6);
    dustGfx.fillCircle(2, 2, 2);
    dustGfx.generateTexture("dust", 4, 4);
    dustGfx.destroy();
    this.dust = this.add.particles(0, 0, "dust", {
      lifespan: 400,
      speedX: { min: -120, max: -60 },
      speedY: { min: -10, max: 10 },
      scale: { start: 1, end: 0 },
      quantity: 1,
      frequency: 90,
      follow: this.player,
      followOffset: { x: -4, y: 6 }
    });
    this.reduceMotion = window.matchMedia(
      "(prefers-reduced-motion: reduce)"
    ).matches;
    if (this.reduceMotion && this.dust) this.dust.stop();
    this.input.on("pointerdown", () => {
      if (this.gameOver) return;
      this.lastJumpPressedAt = this.time.now;
    });
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
    this.applyDuck(this.duckKey.isDown);
    if (!this.gameOver) {
      this.distance += (this.scrollSpeed * delta) / 1000;
      const speedLevel = Math.floor(this.distance / 1000);
      this.scrollSpeed = 200 + speedLevel * 50;
      this.obstacleTimer.delay = Math.max(
        500,
        1500 - speedLevel * 100
      );
    }
    this.distanceText.setText(
      "Dist: " + Math.floor(this.distance)
    );
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
