---
id: 66faa60000000000000000d7
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--
Build the obstacle system as a closure-driven group. Initialise the scroll speed and distance score, create a physics group with `allowGravity: false` (we want obstacles flying along the ground, not arcing), define `this.spawnObstacle` to drop a cactus at the right edge with a leftward velocity, and call it once for the first obstacle on screen.

At the end of `create`:

```js
this.scrollSpeed = 200;
this.distance = 0;
this.obstacles = this.physics.add.group({ allowGravity: false });
this.spawnObstacle = () => {
  const cactus = this.obstacles.create(520, 248, "cactus");
  cactus.body.setVelocityX(-this.scrollSpeed);
  cactus.setOrigin(0.5, 1);
  cactus.body.setSize(16, 28).setOffset(4, 4);
};
this.spawnObstacle();
```

`setOrigin(0.5, 1)` puts the cactus's *base* at `y = 248`, so the foot of the cactus sits on the ground surface. `setSize(16, 28)` shrinks the collision box (the PNG has a transparent border).

# --hints--
You should initialise `this.scrollSpeed = 200` and `this.distance = 0`.

```js
assert.match(code, /this\.scrollSpeed\s*=\s*200[\s\S]*this\.distance\s*=\s*0/);
```

You should create `this.obstacles` as `this.physics.add.group({ allowGravity: false })`.

```js
assert.match(code, /this\.obstacles\s*=\s*this\.physics\.add\.group\(\s*\{\s*allowGravity:\s*false\s*\}\s*\)/);
```

You should define `this.spawnObstacle` to call `this.obstacles.create(520, 248, "cactus")` and set the leftward velocity.

```js
assert.match(code, /this\.spawnObstacle\s*=[\s\S]*this\.obstacles\.create\(\s*520\s*,\s*248\s*,\s*["\']cactus["\']\s*\)[\s\S]*setVelocityX\(\s*-\s*this\.scrollSpeed\s*\)/);
```

You should call `this.spawnObstacle()` once after defining it to seed the first obstacle.

```js
assert.match(code, /this\.spawnObstacle\(\)\s*;/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.image(
      "cactus",
      "/curriculum-assets/phaser/spritesheets/cactus.png"
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) this.lastGroundedAt = time;
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        this.lastJumpPressedAt = time;
      }
      const canJump = time - this.lastGroundedAt < 120;
      const buffered = time - this.lastJumpPressedAt < 120;
      if (canJump && buffered) {
        this.player.setVelocityY(-600);
        this.lastGroundedAt = -Infinity;
        this.lastJumpPressedAt = -Infinity;
      }
    };
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
