---
id: 66faa60000000000000000d5
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--
**Jump buffer** is the symmetric forgiveness window: a press registered up to 120 ms *before* landing fires the jump on touchdown. Combined with coyote time, both timing failures of the naive check are gone — the jump feels right at every BPM.

Replace the `updateJump` closure once more. Stamp the press time in `this.lastJumpPressedAt`, then fire only when *both* the coyote and buffer windows are open:

```js
this.updateJump = (time) => {
  const grounded =
    this.player.body.blocked.down || this.player.body.touching.down;
  if (grounded) this.lastGroundedAt = time;
  if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
    this.lastJumpPressedAt = time;
  }
  const canJump = time - this.lastGroundedAt < 120;
  const buffered = time - this.lastJumpPressedAt < 120;
  if (canJump && buffered) {
    this.player.setVelocityY(-600);
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
  }
};
```

Resetting both timestamps after firing prevents double-fire on a single press.

# --hints--
You should stamp `this.lastJumpPressedAt = time` on `JustDown(this.jumpKey)`.

```js
assert.match(code, /if\s*\(\s*Phaser\.Input\.Keyboard\.JustDown\(\s*this\.jumpKey\s*\)\s*\)\s*\{\s*\n\s*this\.lastJumpPressedAt\s*=\s*time/);
```

You should compute `canJump` and `buffered` from the two timestamps.

```js
assert.match(code, /const\s+canJump\s*=\s*time\s*-\s*this\.lastGroundedAt\s*<\s*120[\s\S]*const\s+buffered\s*=\s*time\s*-\s*this\.lastJumpPressedAt\s*<\s*120/);
```

You should only jump when both `canJump && buffered` are true.

```js
assert.match(code, /if\s*\(\s*canJump\s*&&\s*buffered\s*\)/);
```

After firing you should reset *both* `this.lastGroundedAt` and `this.lastJumpPressedAt` to `-Infinity` inside the `canJump && buffered` block.

```js
assert.match(code, /if\s*\(\s*canJump\s*&&\s*buffered\s*\)\s*\{\s*[\s\S]*this\.lastGroundedAt\s*=\s*-Infinity\s*;\s*\n\s*this\.lastJumpPressedAt\s*=\s*-Infinity/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
--fcc-editable-region--
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) this.lastGroundedAt = time;
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        if (time - this.lastGroundedAt < 120) {
          this.player.setVelocityY(-600);
          this.lastGroundedAt = -Infinity;
        }
      }
    };
--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
