---
id: 66faa60000000000000000d2
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--
Wire the jump in two parts. This step sets up the **closure**; the next step calls it from `update`. Add a `SPACE` key, two timestamp trackers (you'll need them for coyote time and jump buffering soon), and a `this.updateJump(time)` closure that — for now — does the naive thing: jumps when grounded.

At the end of `create`:

```js
this.jumpKey = this.input.keyboard.addKey("SPACE");
this.lastGroundedAt = -Infinity;
this.lastJumpPressedAt = -Infinity;
this.updateJump = (time) => {
  const grounded =
    this.player.body.blocked.down || this.player.body.touching.down;
  if (
    Phaser.Input.Keyboard.JustDown(this.jumpKey) &&
      grounded
  ) {
    this.player.setVelocityY(-600);
  }
};
```

The closure isn't called yet — that's step 10. With gravity at `1500` and a launch velocity of `-600`, the arc lasts roughly 800 ms; tune both numbers if your obstacles need different spacing.

# --hints--
You should add the SPACE key as `this.jumpKey = this.input.keyboard.addKey("SPACE")`.

```js
assert.match(code, /this\.jumpKey\s*=\s*this\.input\.keyboard\.addKey\(\s*["\']SPACE["\']\s*\)/);
```

You should initialise `this.lastGroundedAt = -Infinity` (you'll use it for coyote time soon).

```js
assert.match(code, /this\.lastGroundedAt\s*=\s*-Infinity/);
```

You should initialise `this.lastJumpPressedAt = -Infinity`.

```js
assert.match(code, /this\.lastJumpPressedAt\s*=\s*-Infinity/);
```

Inside `this.updateJump = (time) => { ... }` you should jump only when grounded — call `setVelocityY(-600)`.

```js
assert.match(code, /this\.updateJump\s*=\s*\(\s*time\s*\)\s*=>\s*\{[\s\S]*Phaser\.Input\.Keyboard\.JustDown\(\s*this\.jumpKey\s*\)[\s\S]*setVelocityY\(\s*-600\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
