---
id: 66faa60000000000000000cb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--
The world has gravity but no scenery. Start by loading the three parallax PNGs you'll stack in the next step. The asset bundle ships with the curriculum, so all you need are three loader calls.

Inside `preload`:

```js
this.load.image(
  "sky",
  "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
);
this.load.image(
  "mountains",
  "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
);
this.load.image(
  "ground",
  "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
);
```

The images are tileable — left edge meets right edge — so they'll loop forever in a `tileSprite` without visible seams.

# --hints--
You should call `this.load.image("sky", "/curriculum-assets/phaser/spritesheets/parallax-sky.png")`.

```js
assert.match(code, /this\.load\.image\(\s*["\']sky["\']\s*,\s*["\']\/curriculum-assets\/phaser\/spritesheets\/parallax-sky\.png["\']\s*\)/);
```

You should call `this.load.image("mountains", "/curriculum-assets/phaser/spritesheets/parallax-mountains.png")`.

```js
assert.match(code, /this\.load\.image\(\s*["\']mountains["\']\s*,\s*["\']\/curriculum-assets\/phaser\/spritesheets\/parallax-mountains\.png["\']\s*\)/);
```

You should call `this.load.image("ground", "/curriculum-assets/phaser/spritesheets/parallax-ground.png")`.

```js
assert.match(code, /this\.load\.image\(\s*["\']ground["\']\s*,\s*["\']\/curriculum-assets\/phaser\/spritesheets\/parallax-ground\.png["\']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
--fcc-editable-region--
--fcc-editable-region--
  }

  create() {
  }

  update(time, delta) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
