---
id: 66faa60000000000000000d8
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--
One cactus is a hazard; an endless conveyor of cacti is a game. `Phaser.Time.Clock.addEvent` fires a callback on a delay; with `loop: true` it repeats forever.

At the end of `create`:

```js
this.obstacleTimer = this.time.addEvent({
  delay: 1500,
  loop: true,
  callback: this.spawnObstacle
});
```

1500 ms gives the player ~3 seconds of clear runway between cacti at the starting scroll speed — enough room to learn the jump timing. Step 24 will tighten this as difficulty ramps.

# --hints--
You should create `this.obstacleTimer` with `this.time.addEvent({ delay: 1500, loop: true, callback: this.spawnObstacle })`.

```js
assert.match(code, /this\.obstacleTimer\s*=\s*this\.time\.addEvent\(\s*\{\s*delay:\s*1500\s*,\s*loop:\s*true\s*,\s*callback:\s*this\.spawnObstacle\s*\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.image(
      "cactus",
      "/curriculum-assets/phaser/spritesheets/cactus.png"
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) this.lastGroundedAt = time;
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        this.lastJumpPressedAt = time;
      }
      const canJump = time - this.lastGroundedAt < 120;
      const buffered = time - this.lastJumpPressedAt < 120;
      if (canJump && buffered) {
        this.player.setVelocityY(-600);
        this.lastGroundedAt = -Infinity;
        this.lastJumpPressedAt = -Infinity;
      }
    };
    this.scrollSpeed = 200;
    this.distance = 0;
    this.obstacles = this.physics.add.group({ allowGravity: false });
    this.spawnObstacle = () => {
      const cactus = this.obstacles.create(520, 248, "cactus");
      cactus.body.setVelocityX(-this.scrollSpeed);
      cactus.setOrigin(0.5, 1);
      cactus.body.setSize(16, 28).setOffset(4, 4);
    };
    this.spawnObstacle();
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
