---
id: 66faa60000000000000000d9
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--
Define how the run ends — but don't wire the trigger yet. `this.gameOverDeath` flips the `gameOver` flag, stops the obstacle timer, freezes the player (kills velocity + animation), and shows a Game Over text. Idempotent: if `gameOver` is already true, return immediately.

At the end of `create`:

```js
this.gameOver = false;
this.gameOverDeath = () => {
  if (this.gameOver) return;
  this.gameOver = true;
  this.obstacleTimer.remove(false);
  this.player.setVelocity(0, 0);
  this.player.anims.stop();
  this.add
    .text(240, 140, "Game Over", { fontSize: "24px", color: "#ffffff" })
    .setOrigin(0.5)
    .setDepth(200);
};
```

`obstacleTimer.remove(false)` cancels the timer without firing one last callback — without that, you'd see one extra cactus spawn after death.

# --hints--
You should initialise `this.gameOver = false`.

```js
assert.match(code, /this\.gameOver\s*=\s*false/);
```

Inside `this.gameOverDeath` you should guard with `if (this.gameOver) return;` then set `this.gameOver = true`.

```js
assert.match(code, /this\.gameOverDeath\s*=[\s\S]*if\s*\(\s*this\.gameOver\s*\)\s*return\s*;[\s\S]*this\.gameOver\s*=\s*true/);
```

You should call `this.obstacleTimer.remove(false)`.

```js
assert.match(code, /this\.obstacleTimer\.remove\(\s*false\s*\)/);
```

You should freeze the player with `this.player.setVelocity(0, 0)` and `this.player.anims.stop()`.

```js
assert.match(code, /this\.player\.setVelocity\(\s*0\s*,\s*0\s*\)[\s\S]*this\.player\.anims\.stop\(\s*\)/);
```

You should show a `"Game Over"` text at `(240, 140)`.

```js
assert.match(code, /this\.add\s*\.text\(\s*240\s*,\s*140\s*,\s*["\']Game Over["\']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.image(
      "cactus",
      "/curriculum-assets/phaser/spritesheets/cactus.png"
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
    this.groundBody = this.add
      .rectangle(240, 264, 480, 8, 0x000000, 0)
      .setOrigin(0.5, 0);
    this.physics.add.existing(this.groundBody, true);
    this.physics.add.collider(this.player, this.groundBody);
    this.jumpKey = this.input.keyboard.addKey("SPACE");
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.updateJump = (time) => {
      const grounded =
        this.player.body.blocked.down || this.player.body.touching.down;
      if (grounded) this.lastGroundedAt = time;
      if (Phaser.Input.Keyboard.JustDown(this.jumpKey)) {
        this.lastJumpPressedAt = time;
      }
      const canJump = time - this.lastGroundedAt < 120;
      const buffered = time - this.lastJumpPressedAt < 120;
      if (canJump && buffered) {
        this.player.setVelocityY(-600);
        this.lastGroundedAt = -Infinity;
        this.lastJumpPressedAt = -Infinity;
      }
    };
    this.scrollSpeed = 200;
    this.distance = 0;
    this.obstacles = this.physics.add.group({ allowGravity: false });
    this.spawnObstacle = () => {
      const cactus = this.obstacles.create(520, 248, "cactus");
      cactus.body.setVelocityX(-this.scrollSpeed);
      cactus.setOrigin(0.5, 1);
      cactus.body.setSize(16, 28).setOffset(4, 4);
    };
    this.spawnObstacle();
    this.obstacleTimer = this.time.addEvent({
      delay: 1500,
      loop: true,
      callback: this.spawnObstacle
    });
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
    this.updateJump(time);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
