---
id: 66faa60000000000000000cd
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--
Closures don't run themselves. Convert `delta` (milliseconds) into seconds and call `this.tickParallax(dt)` once per frame.

Inside `update`:

```js
const dt = delta / 1000;
this.tickParallax(dt);
```

Reload the preview — the sky drifts almost imperceptibly, the mountains a little faster, and the ground races past. The speed-ratio gives the eye depth even though every layer is 2D.

# --hints--
You should compute `const dt = delta / 1000` so the parallax is frame-rate independent.

```js
assert.match(code, /const\s+dt\s*=\s*delta\s*\/\s*1000/);
```

You should call `this.tickParallax(dt)` to advance each tile sprite.

```js
assert.match(code, /this\.tickParallax\(\s*dt\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
  }

  update(time, delta) {
--fcc-editable-region--
--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
