---
id: 66faa60000000000000000ca
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--
Welcome to the endless-runner workshop. Across 45 steps you will build a complete endless runner: a player auto-scrolling past a four-layer parallax background, jumping over cacti, ducking flying birds, eating coins, persisting a personal best in `localStorage`, and ramping in difficulty until you crash.

The world has gravity baked in. The `config` object below sets `gravity.y: 1500` — strong gravity is what makes a satisfying jump arc; a wimpy `300` value floats forever. Endless runners traditionally pull a player down hard.

Declare a `MainScene` class extending `Phaser.Scene` with a constructor that calls `super("MainScene")` and empty `preload`, `create`, and `update(time, delta)` methods.

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {}

  create() {}

  update(time, delta) {}
}
```

The `config` object is already filled in — keep it as-is.

# --hints--
You should declare a `MainScene` class extending `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*["\']MainScene["\']\s*\)/);
```

You should declare a `preload` method.

```js
assert.match(code, /\bpreload\s*\(\s*\)\s*\{/);
```

You should declare a `create` method.

```js
assert.match(code, /\bcreate\s*\(\s*\)\s*\{/);
```

You should declare an `update` method that takes `time` and `delta`.

```js
assert.match(code, /\bupdate\s*\(\s*time\s*,\s*delta\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
