---
id: 66faa60000000000000000d0
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--
Gravity pulls the hero down forever — there's no floor. Add an **invisible static rectangle** as a physics body the player will collide with. The visible `groundTile` is decorative; this rectangle is the actual collision surface. Place it just below the grass strip of the ground tile so the player's feet line up with the visible ground.

At the end of `create`:

```js
this.groundBody = this.add
  .rectangle(240, 264, 480, 8, 0x000000, 0)
  .setOrigin(0.5, 0);
this.physics.add.existing(this.groundBody, true);
```

`add.rectangle(x, y, width, height, color, alpha)` with `alpha: 0` is invisible. `setOrigin(0.5, 0)` puts the rectangle's *top* at `y = 264`. `physics.add.existing(obj, true)` registers it as a **static** body — immovable, no gravity, perfect for floors.

# --hints--
You should create `this.groundBody` as `this.add.rectangle(240, 264, 480, 8, 0x000000, 0)` with origin `(0.5, 0)`.

```js
assert.match(code, /this\.groundBody\s*=\s*this\.add\s*\.rectangle\(\s*240\s*,\s*264\s*,\s*480\s*,\s*8\s*,\s*0x000000\s*,\s*0\s*\)\s*\.setOrigin\(\s*0\.5\s*,\s*0\s*\)/);
```

You should register the ground rectangle as a static physics body via `this.physics.add.existing(this.groundBody, true)`.

```js
assert.match(code, /this\.physics\.add\.existing\(\s*this\.groundBody\s*,\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Endless Runner</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image(
      "sky",
      "/curriculum-assets/phaser/spritesheets/parallax-sky.png"
    );
    this.load.image(
      "mountains",
      "/curriculum-assets/phaser/spritesheets/parallax-mountains.png"
    );
    this.load.image(
      "ground",
      "/curriculum-assets/phaser/spritesheets/parallax-ground.png"
    );
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.sky = this.add
      .tileSprite(240, 160, 480, 320, "sky")
      .setScrollFactor(0);
    this.mountains = this.add
      .tileSprite(240, 200, 480, 160, "mountains")
      .setScrollFactor(0);
    this.groundTile = this.add
      .tileSprite(240, 288, 480, 64, "ground")
      .setScrollFactor(0);
    this.tickParallax = (dt) => {
      this.sky.tilePositionX += 5 * dt;
      this.mountains.tilePositionX += 25 * dt;
      this.groundTile.tilePositionX += 240 * dt;
    };
    this.player = this.physics.add.sprite(80, 200, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.anims.create({
      key: "run",
      frames: this.anims.generateFrameNumbers("hero", { frames: [12, 13, 14, 15] }),
      frameRate: 10,
      repeat: -1
    });
    this.player.play("run");
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time, delta) {
    const dt = delta / 1000;
    this.tickParallax(dt);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 320,
  parent: "game-container",
  backgroundColor: "#1a1d2e",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 1500 }, debug: false }
  },
  scene: [MainScene]
};

new Phaser.Game(config);
```
