---
id: 646c47867800472a4ed5d2ea
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Inizia a scrivere la struttura HTML di base. Aggiungi un riferimento `DOCTYPE` `html` e un elemento `html` con l'attributo `lang` impostato su `en`. Aggiungi anche degli elementi `head` e `body` all'interno dell'elemento `html`.

# --hints--

Dovresti avere la dichiarazione `DOCTYPE` `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio dell'HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag `html` di apertura con un attributo `lang` con il valore `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dovresti avere un tag `html` di chiusura.

```js
assert(code.match(/<\/html>/i));
```

Dovresti avere un tag `head` di apertura.

```js
assert(code.match(/<head>/i));
```

Dovresti avere un tag `head` di chiusura.

```js
assert(code.match(/<\/head>/i));
```

Dovresti avere un tag di apertura `body`.

```js
assert(code.match(/<body>/i));
```

Dovresti avere un tag `body` di chiusura.

```js
assert(code.match(/<\/body>/i));
```

L'elemento `body` dovrebbe trovarsi dopo l'elemento `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Gli elementi `head` e `body` dovrebbero essere all'interno dell'elemento `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
