---
id: 646c59e942f35541923104bf
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Tornando all'HTML, crea un elemento `main`. All'interno di questo elemento `main`, aggiungi un elemento `div` con la classe `cat-head`.

# --hints--

Dovresti avere un elemento `main`.

```js
assert(document.querySelectorAll('main').length === 1);
```

Dovresti avere un elemento `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

L'elemento `div` dovrebbe avere la classe `cat-head`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

L'elemento `div` dovrebbe essere all'interno del tag `main`.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
