---
id: 646cf1206cac5f51804f49cf
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Per posizionare correttamente l'orecchio sinistro, è possibile utilizzare una trasformazione CSS per ruotarlo di un certo grado.

La proprietà `transform` consente di modificare la forma, la posizione e la dimensione di un elemento senza cambiare il layout o influenzare gli elementi circostanti. Ha funzioni come `translate()`, `rotate()`, `scale()`, `skew()` e `matrix()`.

Imposta la proprietà `transform` su `rotate(-45deg)` e vedi che succede.

# --hints--

Dovresti impostare la proprietà `transform` dell'elemento `.cat-left-ear` su `rotate(-45deg)`. Non dimenticare di aggiungere il punto e virgola.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform === 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
