---
id: 646c48df8674cf2b91020ecc
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Aggiungi un elemento `link` all'interno dell'elemento `head`. Imposta l'attributo `rel` dell'elemento `link` su `stylesheet` e `href` su `./styles.css`.

# --hints--

Dovresti avere un elemento `link`.

```js
assert.match(code, /<link/)
```

Your `link` element should have a `rel` and `href` attribute set to `stylesheet` and `./styles.css`.

```js
assert.match(code, /<link\s+(rel=("|')stylesheet\2\s+href=("|')\.\/styles\.css\3|href=("|')\.\/styles\.css\4\s+rel=("|')stylesheet\5)\s*\/?>/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```

```css

```
