---
id: 646c5e727e56e743c9aed4a1
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Le proprietà di posizionamento CSS ti consentono di impostare la posizione di un elemento nel browser come desideri. Puoi impostare la proprietà `position` su `static`, `absolute`, `relative`, `sticky` o `fixed`.

Una volta impostata la proprietà `position` di un elemento, puoi spostare l'elemento impostando un valore in pixel o percentuale per una o più proprietà tra `top`, `right`, `left` e `bottom`.

`static` è il posizionamento predefinito per tutti gli elementi. Se lo assegni a un elemento, non sarai in grado di muoverlo con `top`, `right`, `left` o `bottom`.

Dai a `.cat-head` una proprietà `position` impostata su `static`, poi imposta le proprietà `top` e `left` su `100px` ciascuna.


# --hints--

Il selettore `.cat-head` dovrebbe avere una proprietà `position` impostata su `static`. Assicurati di aggiungere un punto e virgola.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su `100px`. Assicurati di aggiungere un punto e virgola.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `left` impostata su `100px`. Assicurati di aggiungere un punto e virgola.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
