---
id: 646c48df8674cf2b91020ecc
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Aggiungi un elemento `link` all'interno dell'elemento `head`. Imposta l'attributo `rel` dell'elemento `link` su `stylesheet` e `href` su `./styles.css`.

# --hints--

Dovresti avere un elemento `link`.

```js
assert.match(code, /<link/)
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1/)
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `./styles.css`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1\s+href=\1\.\/styles\.css\1\s*\/?>/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
