---
id: 646c5ffef5598d449b52ec12
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Ora dovrai lavorare alle orecchie del gatto. Ci saranno un orecchio destro e sinistro, e dentro ciascuno ci sarà una parte interna.

All'interno dell'elemento `.cat-head`, crea un elemento `div` con la classe `cat-ears`.

# --hints--

Non dovresti cambiare l'elemento `div` esistente con la classe `cat-head`.

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

Dovresti creare un elemento `div` all'interno dell'elemento `.cat-head`.

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

L'elemento `div` dovrebbe avere la classe `cat-ears`.

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">

      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
