---
id: 646c48df8674cf2b91020ecb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno dell'elemento `head`, aggiungi un tag `meta` con l'attributo `charset` impostato su `utf-8`. Aggiungi anche un elemento `title` con il testo `fCC Cat Painting`.

# --hints--

Dovresti aggiungere un singolo elemento `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

L'elemento `meta` dovrebbe avere un attributo `charset` con il valore `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Dovresti aggiungere un singolo elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

L'elemento `title` dovrebbe avere il testo `fCC Cat Painting`. Fai attenzione all'ortografia e alle maiuscole/minuscole.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
