---
id: 646dd556d524bc61c0139bd6
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

La maggior parte dei gatti ha le orecchie e la parte interna delle orecchie di colori diversi. Puoi fare la stessa cosa. Ecco perché prima hai definito un elemento `div` per la parte interna sia dell'orecchio destro che sinistro.

Usando un selettore di classe, dai all'elemento `.cat-left-inner-ear` un bordo sinistro e un bordo destro ciascuno di `20px solid transparent`. Inoltre dagli un bordo inferiore di `40px solid #3b3b4f`.

# --hints--

Dovresti avere un selettore `.cat-left-inner-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear'))
```

Il selettore `.cat-left-inner-ear` dovrebbe avere la proprietà `border-left` impostata su `20px solid transparent`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderLeft === '20px solid transparent')
```

Il selettore `.cat-left-inner-ear` dovrebbe avere la proprietà `border-right` impostata su `20px solid transparent`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderRight === '20px solid transparent')
```

Il selettore `.cat-left-inner-ear` dovrebbe avere la proprietà `border-bottom` impostata su `40px solid #3b3b4f`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-inner-ear')?.borderBottom === '40px solid rgb(59, 59, 79)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  z-index: 1;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  position: absolute;
  top: -26px;
  left: 163px;
  z-index: 1;
  transform: rotate(45deg);
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

--fcc-editable-region--

--fcc-editable-region--
```
