---
id: 6476f7a4827bcc61682f2347
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

La prossima proprietà di posizione è `absolute`. Quando si usa il valore `absolute` per la proprietà `position`, l'elemento viene tolto dal normale flusso del documento e la sua posizione è determinata dalle proprietà `top`, `right`, `bottom` e `left`.

Imposta la proprietà di posizione dell'elemento `.cat-head` su `absolute`, quindi imposta le proprietà `top` e `left` su qualsiasi valore in pixel.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

Il selettore `.cat-head` dovrebbe avere una proprietà `position` con il valore `absolute`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su un numero qualsiasi di `px`. Make sure you add a semi-colon.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Il selettore `.cat-head` dovrebbe avere una proprietà `left` impostata su un numero qualsiasi di `px`. Make sure you add a semi-colon.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
