---
id: 6476fc5cf14b276e6d04e82a
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Usa un selettore di classe per dare all'elemento `.box` una larghezza di `200px`, un'altezza di `600px` e un colore di sfondo del valore `#000`. Inoltre, imposta `position` su `absolute`, `top` su `800px` e `left` su `650px`.

# --hints--

Dovresti avere un selettore `.box`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

Il selettore `.box` dovrebbe avere una proprietà `width` impostata su `200px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

Il selettore `.box` dovrebbe avere una proprietà `height` con il valore `600px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

Il selettore `.box` dovrebbe avere una proprietà `background-color` impostata su `#000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

Il selettore `.box` dovrebbe avere una `position` proprietà impostata su `absolute`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

Il selettore `.box` dovrebbe avere una proprietà `left` con il valore `650px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

Il selettore `.box` dovrebbe avere una proprietà `top` impostata su `800px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

