---
id: 6477062778c85972eb648030
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Ora dovresti centrare la testa del gatto.

Give the `.cat-head` element a `position` property set to `absolute`. Imposta il valore `0` per le proprietà `right`, `left`, `top` e `bottom`, quindi imposta la sua proprietà `margin` su `auto` per tutti i lati. Questo è un modo per centrare un elemento verticalmente e orizzontalmente usando il posizionamento CSS.

# --hints--

Il selettore `.cat-head` dovrebbe avere una proprietà `position` con il valore `absolute`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `top` impostata su `0`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `left` impostata su `0`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `right` impostata su `0`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `bottom` impostata su `0`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

Il selettore `.cat-head` dovrebbe avere una proprietà `margin` impostata su `auto`. Make sure you add a semi-colon.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
