---
id: 68e6aff5325410118e76e7be
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will practice working with ARIA by building an accessible audio controller.

To begin the project, add the `<!DOCTYPE html>`, and an `html` element with a `lang` attribute of `en`.

Inside the `html` element, add a `head` element.

# --hints--

You should have the `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

You should have an opening `html` tag with the language set to english.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

You should have a closing `html` tag.

```js
assert.match(code, /<\/html>/i);
```

Your `DOCTYPE` should come before the `html` element.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im);
```

You should have an opening `head` tag.

```js
assert.match(code, /<head>/i);
```

You should have a closing `head` tag.

```js
assert.match(code, /<\/head>/i);
```

Your opening `head` tag should come before the closing `head` tag.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

Your `head` element should be inside the `html` element.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
