---
id: 68eb7f2397da251a40e04aa2
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now, add an `h1` element with the text `Audio Controls` inside the `body` element.

Also you need to add a Play button, for this add a `button` element with the type attribute set to the value `"button"`.

The `button` element's text should be `Play`.

# --hints--

You should have an `h1` element.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` element should be inside your `body` element.

```js
assert.exists(document.querySelector('body > h1'));
```

Your `h1` element should have the text `Audio Controls`.

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

You should have a `button` element.

```js
assert.exists(document.querySelector('button'));
```

Your `button` element should have a `type` attribute.

```js
const button = document.querySelector('button');
assert.exists(button?.getAttribute('type'));
```

Your `button` element should have the `type` attribute set to `"button"`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

Your `button` element should have the text `Play`.

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

Your `h1` element should come before your `button` element.

```js
const h1 = document.querySelector('h1');
const button = document.querySelector('button');
assert.ok(h1.compareDocumentPosition(button) & Node.DOCUMENT_POSITION_FOLLOWING);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
