---
id: 68ed3a96e0ebcc36565b39e5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Next, you need to add an `input` element with the `type` attribute set to `"range"` to create the volume slider. Set the `min` attribute to `"0"`, the `max` attribute to `"100"`, and the `value` attribute to `"50"` to define the default volume level.

As you have learnt in the previous lectures, the `aria-labelledby` attribute is used when there is existing text on the page that can be used as a label. In this case, you have added the `id` attribute to your `span` elements which will now be used here.

So, add an `aria-labelledby` attribute to the `input` element set to `"volume-label volume-description"`.

# --hints--

You should have an `input` element below your second `span` element.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

Your `input` element should have a `type` attribute.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('type'));
```

Your `input` element should have a `type` attribute with the value `range`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

Your `input` element should have a `min` attribute.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

Your `input` element should have a `min` attribute with the value `0`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

Your `input` element should have a `max` attribute.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('max'));
```

Your `input` element should have a `max` attribute with the value `100`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

Your `input` element should have a `value` attribute.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('value'));
```

Your `input` element should have a `value` attribute with the value `50`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

Your `input` element should have a `aria-labelledby` attribute.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('aria-labelledby'));
```

Your `input` element should have a `aria-labelledby` attribute with the value `volume-label volume-description`.

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert(ariaLabelledbyAttr === 'volume-label volume-description' || ariaLabelledbyAttr === 'volume-description volume-label');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
