---
id: 660eebd83100d37862268781
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

目前你的 `character` 变量还没有一个值。 你可以利用<dfn>赋值</dfn>运算符 `=` 来赋值。 举个例子：

```js
let hello = "Hello";
```

在声明变量的同时给它赋值，叫做<dfn>初始化</dfn>。

在声明变量的阶段，通过将 `"Hello"` 赋值给 `character` 变量来实现变量的初始化。

# --hints--

你应该使用赋值运算符 `=` 。

```js
assert.match(code, /let\s+character\s*=/);
```

你应该使用字符串 `"Hello"` 。

```js
assert.match(code, /('|")Hello\1/);
```

你应该对你的 `"Hello"` 字符串使用双引号。

```js
assert.match(code, /"Hello"/);
```

你应该将 `"Hello"` 赋值给变量 `character` 。

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
