---
id: 660ef31a5be625914a0102cd
title: 步骤 5
challengeType: 1
dashedName: step-5
---

# --description--

当一个变量使用关键字 `let` 来声明时，你可以在之后<dfn>重新赋值</dfn>。 在此示例中，`programmer` 的值就从 `"Naomi"` 被重新赋值为了 `"CamperChan"`。

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Note that when reassigning a variable, you do **not** use the `let` keyword again.

在 `console.log` 语句后，将 `"World"` 赋值给变量 `character`。

# --hints--

你应该在你的代码中两次使用 `character`。

```js
assert.lengthOf(code.match(/character/g), 3);
```

你应该使用赋值运算符来重新赋值变量 `character`。

```js
assert.lengthOf(code.match(/character\s*\=/g), 2);
```

你应该在你的代码中使用字符串 `"World"`。

```js
assert.match(code, /("|')World\1/);
```

你的 `character` 变量的值应该为 `"World"`。

```js
assert.equal(character, "World");
```

你的重新赋值语句不需要使用 `let`。

```js
assert.isBelow(code.match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
