---
id: 660ef19b95d3308e7dd31bb6
title: 步骤 4
challengeType: 1
dashedName: step-4
---

# --description--

<dfn>console</dfn>（控制台指令）允许你打印并查看 JavaScript 的输出。 你可以使用 `console.log()` 来向控制台发送信息。 举个例子，下面这串代码将会在控制台打印 `"Naomi"`。

```js
let developer = "Naomi";
console.log(developer);
```

上面的代码通过在 `console.log()` 中的名称读取到了变量 `developer` 。 请注意，括号内的值是将要被打印的内容。

在控制台打印变量 `character` 的值。 然后点击 “Console”（控制台）按钮来查看 JavaScript 控制台。

# --hints--

你应该在你的代码中使用 `console` （控制台指令）。

```js
assert.match(code, /console/);
```

你应该使用 `console` 的 `log` 属性。

```js
assert.match(code, /console\.log/);
```

你应该使用括号来调用 `.log()` 方法。

```js
assert.match(code, /console\.log\(/);
```

你应该在控制台打印变量 `character` 的值。

```js
assert.match(code, /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
