---
id: 660ef0f7c4b8e68ccd1f0786
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript 有七种原始数据类型，`String`（字符串类型）就是其中之一。 在 JavaScript 中，<dfn>字符串</dfn>代表的是一串字符内容，被包裹在单引号 （`'`）或者双引号（`"`）之中。

注意，字符串是 <dfn>不可变的</dfn>，这意味着一旦被创建，它们就无法被更改。 仍然可以给该变量重新分配另一个值。

试着来使用单引号包裹你的 `"Hello"` 字符串。

# --hints--

你的代码中不应该使用双引号。

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

你应该对你的 `"Hello"` 字符串使用单引号。

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

你仍然应该在你的代码中使用 `let` 关键字。

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

你仍然应该声明一个 `character` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

你的 `character` 变量的值应该仍是字符串 `"Hello"`。

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
