---
id: 660ef55dd468079679ee0092
title: 步骤 11
challengeType: 1
dashedName: step-11
---

# --description--

你现在可以声明下一个变量了。 移除 `console.log` 语句和为 `character` 重新赋值的这两类语句。

同时移除变量 `secondCharacter`，但是保留 `character` 初始化语句。

# --hints--

现在你的代码中不应该有任何 `console.log()` 语句。

```js
assert.notMatch(code, /console/);
```

你不应该将 `character` 重新赋值。

```js
assert.isAtMost(code.match(/character/).length, 1);
```

你的变量 `character` 的值应该为 `"Hello"`。

```js
assert.equal(character, "Hello");
```

You should not have a `secondCharacter` variable.

```js
assert.notMatch(code, /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
