---
id: 660ee6e3a242da6bd579de69
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

JavaScript 是一门用以驱动网页的编程语言。 不同于你之前学习的 HTML 和 CSS，JavaScript 最常用于编写逻辑而不是标记。

变量是编程中最重要的概念之一。 一个<dfn>变量</dfn>指向一个用来存储值的特定内存地址。 变量被赋予一个名称，可以在整个代码中用来访问该值。

声明一个变量意味着为它命名。 在 JavaScript 中，通常使用 `let` 关键字来实现。 举个例子，你会这样声明一个 `hello` 变量。

```js
let hello;
```

变量的命名遵循特定的规则：名称能够包含字母、数字、美元符号以及下划线，但是不能包含空格，且一定不能以数字开头。

Use the `let` keyword to declare a variable called `character`.

_注意_ ：通常我们会以分号来结束JavaScript中的语句， `;`

# --hints--

你应该在你的代码中使用 `let` 关键字。

```js
assert.match(code, /let/);
```

你应该在你的代码中使用 `character` 。

```js
assert.match(code, /character/);
```

你应该使用 `let` 关键字来声明一个 `character` 变量。

```js
assert.match(code, /let\s+character/);
```

你的声明应该以分号结尾。

```js
assert.match(code, /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
