---
id: 660ef0f7c4b8e68ccd1f0786
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

JavaScript 有七种原始数据类型，`String`（字符串类型）就是其中之一。 在 JavaScript 中，<dfn>字符串</dfn>代表的是一串字符内容，被包裹在单引号 （`'`）或者双引号（`"`）之中。

请注意，字符串是<dfn>不可变的</dfn>，这意味着它们被创建后就不能再被更改了。

试着来使用单引号包裹你的 `"Hello"` 字符串。

# --hints--

你的代码中不应该使用双引号。

```js
assert.notMatch(code, /"/);
```

你应该对你的 `"Hello"` 字符串使用单引号。

```js
assert.match(code, /'Hello'/);
```

你仍然应该在你的代码中使用 `let` 关键字。

```js
assert.match(code, /let/);
```

你仍然应该声明一个 `character` 变量。

```js
assert.match(code, /let\s+character/);
```

你的 `character` 变量的值应该仍是字符串 `"Hello"`。

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
