---
id: 660ef6355e8f5a9e67fe5f46
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

通过 `number` 数据类型，你可以进行数学运算，比如加法运算。 试着在控制台打印 `count + 1`。

# --hints--

你应该在你的代码中使用 `console`。

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

你应该使用 `console` 的 `.log()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你应该打印 `count + 1` 到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
