---
id: 660ef6355e8f5a9e67fe5f46
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

With the `number` data type, you can perform mathematical operations, like addition. Try printing `count + 1` to the console.

# --hints--

You should access the `console` in your code.

```js
assert.match(code, /console/);
```

You should access the `.log()` of the `console`.

```js
assert.match(code, /console\.log/);
```

You should log `count + 1` to the console.

```js
assert.match(code, /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
