---
id: 660ef857f2806aa626d29d17
title: Step 17
challengeType: 1
dashedName: step-17
---

# --description--

同样你也可以使用减法（`-`）、乘法（`*`）以及除法（`/`）。 请用这些运算符和数字在你的 `console.log` 中尽情的做实验吧。 当你准备好进入下一节时，移除 `console.log`。

# --hints--

你的代码中不应该有 `console.log()`。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你的 `character` 变量的值应该仍是 `"Hello"`。

```js
assert.equal(character, "Hello");
```

你的 `count` 变量的值应该仍是 `8`。

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
