---
id: 660ef5c1904955978a986a5c
title: 步骤 12
challengeType: 1
dashedName: step-12
---

# --description--

使用 `let` 来声明一个变量 `count`。 给它赋值<dfn>数字类型</dfn> `8`。 当你要使用数字类型作为值的时候，你不再需要用引号去包裹它。 举个例子：

```js
let money = 100;
```

# --hints--

你的代码中应该有第二个 `let` 关键字。

```js
assert.lengthOf(code.match(/let/g), 2);
```

你应该用 `let` 来声明一个变量 `count`。

```js
assert.match(code, /let\s+count/);
```

你应该把数字类型的 `8` 赋值给变量 `count`。

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
