---
id: 660f5179b3b0ca558f6b4d4f
title: Step 113
challengeType: 1
dashedName: step-113
---

# --description--

使用 `if` 语句检查 `inverted` 是否为真。 请记住，这里不需要使用相等运算符。

# --hints--

你应该使用 `if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

你的 `if` 语句应检查 `inverted` 是否为 `true`。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

你的 `if` 条件不应使用任何比较运算符。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
