---
id: 660f39b444fd6f16d1e49c1f
title: Step 73
challengeType: 1
dashedName: step-73
---

# --description--

你不必将 `i + 1` 传递给 `padRow` 调用，而是可以从 `1` 开始循环。 这将允许你创建一个单索引循环。

更新你的迭代器以从 `1` 开始，而不是从 `0` 开始。

# --hints--

你的 `for` 循环应该将 `i` 初始化为 `1`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
