---
id: 660f415b76859a2736771607
title: Step 78
challengeType: 1
dashedName: step-78
---

# --description--

你的金字塔又消失了。 没关系——这是意料之中的事。

在创建新循环之前，你需要了解 `if` 语句。 <dfn>`if` 语句</dfn>允许你仅在满足条件时运行代码块。 它们使用以下语法：

```js
if (condition) {
  logic
}
```

创建一个 `if` 语句，以布尔值 `true` 作为条件。 在主体中，打印字符串 `"Condition is true"`。

# --hints--

你应该创建一个 `if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

你的 `if` 语句应以 `true` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

你的 `if` 主体应记录 `"Condition is true"`。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
