---
id: 660f48a419b40238e2b8b4d5
title: Step 93
challengeType: 1
dashedName: step-93
---

# --description--

你的循环不再依赖于 `continueLoop` 变量。 这使得该变量成为<dfn>未使用的声明</dfn>。 一般来说，你需要避免未使用的声明以防止将来的混淆。

删除你的 `continueLoop` 变量。

# --hints--

你应该不再有 `continueLoop` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
