---
id: 660f4efcb8068e4cb470dca1
title: Step 109
challengeType: 1
dashedName: step-109
---

# --description--

现在你已经尝试了这些方法，你可以采用另一种倒金字塔方法。 但首先你需要清理你的实验。

删除你的 `numbers` 数组以及方法调用和日志调用。

# --hints--

你不应该有 `numbers` 数组。

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

你不应该有 `unshifted` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

你不应该有 `shifted` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

你不应该有 `console.log` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
