---
id: 660f44f10ea40f300b896a5e
title: Step 84
challengeType: 1
dashedName: step-84
---

# --description--

现在你已经练习了使用 `if...else if...else` 语句，你可以将它们从代码中删除。

完成后，使用 `let` 声明 `continueLoop` 变量并为其分配布尔值 `false`。 然后使用 `let` 声明一个 `done` 变量并为其赋值 `0`。

# --hints--

你不应该有 `if` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

你不应该有 `else if` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

你不应该有 `else` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

你应该使用 `let` 来声明 `continueLoop` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

你的 `continueLoop` 变量应具有值 `false`。

```js
assert.isFalse(continueLoop);
```

你应该使用 `let` 来声明 `done` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

你的 `done` 变量应具有值 `0`。

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
