---
id: 660f0da9bf1035c9097af20a
title: Step 28
challengeType: 1
dashedName: step-28
---

# --description--

你应该看到了 `"freeCodeCamp"` 被打印到了控制台。 这是因为 `.pop()` 返回了从数组中移除的那个值 - 这个值也正是你早先向数组末尾推入的那个值 `"freeCodeCamp"` 。

那么 `.push()` 方法会返回什么内容呢？ 试着为你现有的 `rows.push()` 重新声明一个新的 `pushed` 变量，并且打印它。

# --hints--

你应该声明一个 `pushed` 变量。

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

你应该用 `let` 来声明变量 `pushed`。

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

你应该用 `rows.push("freeCodeCamp")` 来声明你的 `pushed` 变量。

```js
assert.equal(pushed, 4);
```

你应该打印你的变量 `pushed` 。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
