---
id: 660f34626216270c682e2f7b
title: Step 66
challengeType: 1
dashedName: step-66
---

# --description--

<dfn>函数调用</dfn>允许你实际使用函数。 你可能没有意识到，但你一直在使用的 `.push()` 之类的方法一直是函数调用。

通过引用函数的名称并添加 `()` 来调用函数。 调用 `test` 函数的方法如下：

```js
test();
```

将 `.push()` 调用中的 `character.repeat(i + 1)` 替换为 `padRow` 函数的函数调用。

# --hints--

你不应该在 `push` 调用中使用 `i + 1`。

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

你不应该在 `.push()` 调用中使用 `character.repeat`。

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

你应该在 `.push()` 调用中调用 `padRow`。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*\)\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
