---
id: 660f4a83373de83ea101685f
title: Step 100
challengeType: 1
dashedName: step-100
---

# --description--

因为你要朝相反方向循环，所以你的循环需要在 `i` 大于 `0` 时运行。 你可以为此使用 <dfn>大于</dfn> 运算符 `>`。

将循环条件设置为当 `i` 大于 `0` 时运行。

# --hints--

当 `i` 大于 `0` 时，你的 `for` 循环应该运行。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
