---
id: 660f0a55847d6cc485f29ba5
title: Step 26
challengeType: 1
dashedName: step-26
---

# --description--

JavaScript 中的 <dfn>方法</dfn> 是一个与特定值或特定对象相关联的函数。 你已经遇到的一个例子是 `.log()` 方法，它是 `console` 对象的一部分。

数组也有自己的方法，你将首先探索的是 `.push()` 方法。 这个方法允许你去“推入”一个值到一个数组的末尾。 这里有一个将数字 `12` 添加到数组末尾的例子。

```js
array.push(12);
```

使用 `.push()` 添加字符串 `"freeCodeCamp"` 到你的 `rows` 数组末尾。 在 `console.log` 前添加此代码，以便你可以看到你对数组所做的更改。

# --hints--

你应该在你的代码中使用 `.push()`。

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

你应该对你的 `rows` 数组使用 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

你应该将字符串 `"freeCodeCamp"` 传递给 `.push()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
