---
id: 660f38c34a4de6141c0c369f
title: Step 71
challengeType: 1
dashedName: step-71
---

# --description--

你的金字塔生成器现在可以按预期运行。 但这是一个进一步探索你所编写的代码的绝佳机会。

加法运算符不是向变量添加值的唯一方法。 <dfn>加法赋值</dfn>运算符可以简写为“取变量的原始值，添加该值，然后将结果赋给变量”。 例如，这两个语句将产生相同的结果：

```js
test = test + 1;
test += 1;
```

更新 `for` 循环中的迭代器语句以使用加法赋值。

# --hints--

你的 `for` 循环不应该使用 `i = i + 1`;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

你的 `for` 循环应该使用带有 `i` 的加法赋值。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
