---
id: 660f229d2dbe09ef2954a4a1
title: Step 42
challengeType: 1
dashedName: step-42
---

# --description--

回想一下，在之前的循环中，你使用了加法运算符 `+` 将 `i` 的值增加 `1`。

你可以对字符串值执行类似的操作，即将新字符串附加到现有字符串。 例如，`hello = hello + " World";` 会将字符串 `" World"` 添加到 `hello` 变量中存储的现有字符串中。 这被称为<dfn>连接</dfn>。

在 `for...of` 循环中，使用加法运算符将 `row` 值连接到 `result` 值。

# --hints--

你应该在 `result` 变量上使用连接运算符。

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

你应该将 `row` 连接到 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

你应该将连接的结果分配回 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
