---
id: 660f3ba3cceef11b6ba08b59
title: Step 75
challengeType: 1
dashedName: step-75
---

# --description--

不幸的是，现在金字塔的底层已经消失了。 这是因为你又创建了另一个 <dfn>off-by-one 错误</dfn>。

你的原始循环适用于 `i` 值从 `0` 到 `7`，因为 `count` 为 `8` 并且你的条件要求 `i` 小于 `count`。 你的循环现在正在运行 `i` 个值，从 `1` 到 `7`。

当 `i` 为 `8` 时，你的循环也需要更新才能运行。 查看你的逻辑，这意味着当 `i` <dfn>小于或等于</dfn> `count` 时，你的循环应该运行。 你可以为此使用小于或等于运算符 `<=`。

更新你的循环条件，使其在 `i` 小于或等于 `count` 时运行。

# --hints--

你的 `for` 循环不应检查 `i` 是否小于 `count`。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

你的 `for` 循环应检查 `i` 是否小于或等于 `count`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
