---
id: 660f17294346b7d69e79db3d
title: Step 32
challengeType: 1
dashedName: step-32
---

# --description--

你现在可以开始构建金字塔生成器了。 你的 `character` 变量将作为金字塔的构建块。

`"Hello"` 可能不太适合这种情况。 将 `character` 的值更改为井号 `"#"`。

# --hints--

你的 `character` 变量应该是一个井号。

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
