---
id: 660f280dda5040f707c76b4a
title: Step 45
challengeType: 1
dashedName: step-45
---

# --description--

现在你有一系列 `#` 字符，但仍然缺少金字塔形状。 幸运的是，`i` 变量代表循环中的当前“行”号，使你能够使用它来制作类似金字塔的结构。

为了实现这一点，你将使用字符串可用的 `.repeat()` 方法。 该方法接受一个数字作为参数，指定重复目标字符串的次数。 例如使用 `.repeat()` 生成字符串 `"Code! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

在你的 `character` 上使用 `.repeat()` 方法，并为其提供 `i` 作为数字。

# --hints--

你应该使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

你应该在 `character` 变量上使用 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

你应该将 `i` 传递给你的 `.repeat()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
