---
id: 660f447efc0e722f016c1be0
title: Step 81
challengeType: 1
dashedName: step-81
---

# --description--

文字又出现了！ 这是因为 `"false"` 是一个字符串，当将其计算为布尔值时变为 `true`。 这意味着 `"false"` 是一个真值。

<dfn>真值</dfn>是作为布尔值评估时被视为真的值。 你在 JavaScript 中遇到的大多数值都是真值。

<dfn>假值</dfn>则相反 - 当作为布尔值评估时被视为假的值。 JavaScript 有一个定义的假值列表。 其中包括 `false`、`0`、`""`、`null`、`undefined` 和 `NaN`。

尝试将你的 `if` 条件更改为空字符串 `""`，这是一个假值。

# --hints--

你的 `if` 语句应以 `""` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
