---
id: 660f45ccf4ca5c31f253005a
title: Step 86
challengeType: 1
dashedName: step-86
---

# --description--

现在，如果你将 `continueLoop` 更改为 true，则 `while` 循环将永远运行。 这被称为<dfn>无限循环</dfn>，你应该小心避免这种情况。 无限循环可能会锁死你的系统，需要完全重新启动才能摆脱。

为了避免这种情况，首先使用自增运算符来增加循环内 `done` 变量的值。

# --hints--

你的 `while` 循环应该增加 `done` 变量。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
