---
id: 660f383d4c772c12ff59904b
title: Step 70
challengeType: 1
dashedName: step-70
---

# --description--

你可以将完整的表达式作为参数传递。 该函数将接收该表达式的求值结果。 例如，这两个函数调用将产生相同的结果：

```js
test(2 * 3 + 1);
test(7);
```

再次查看模式：

```md
··#··
·###·
#####
```

将 `character` 值更新为重复 `2 * rowNumber - 1` 次。

再次打开控制台以查看更新的结果。

# --hints--

你的函数应该将 `2 * rowNumber - 1` 传递给你的 `.repeat()` 调用。

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
