---
id: 660f23b53db70af0f2620e78
title: Step 43
challengeType: 1
dashedName: step-43
---

# --description--

现在，你的所有数字都出现在同一行上。 这对于创建金字塔不起作用。

你需要在每一行中添加一个新行。 但是，在 JavaScript 中按回车键在引号之间插入换行符将导致解析错误。 相反，你需要使用特殊的 <dfn>转义序列</dfn> `\n`，它在记录字符串时被解释为新行。 例如：

```js
lineOne = lineOne + "\n" + lineTwo;
```

使用第二个加法运算符在现有的 `result` 值和添加的 `row` 值之间连接一个新行。

# --hints--

你应该使用 `\n` 转义序列。 请记住，它需要是一个字符串，因此用引号括起来。

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

你应该将 `\n` 转义序列连接到 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1/);
```

你应该将 `row` 变量连接到 `\n` 转义序列。

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*('|")\\n\1\s*\+\s*row/);
```

你应该将整个连接分配回你的 `result` 变量。 Don't forget your semicolon.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*result\s*\+\s*('|")\\n\1\s*\+\s*row;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
