---
id: 660f17d4e9f227d86e834abd
title: Step 33
challengeType: 1
dashedName: step-33
---

# --description--

要生成金字塔，你需要创建多行。 当你必须重复执行某项任务直到满足某个条件时，你将使用<dfn>循环</dfn>。 编写循环的方法有很多种。

你将从一个基本的 `for` 循环开始。 `for` 循环使用以下语法：

```js
for (iterator; condition; iteration) {
  logic;
}
```

在接下来的步骤中，你将详细探索循环的每个组件。 现在，构建一个 `for` 循环，其中包含三个组件的术语 `"iterator"`、`"condition"` 和 `"iteration"`。 保持循环 <dfn>主体</dfn>（花括号 `{}` 内的部分）为空。


# --hints--

你应该有一个 `for` 循环。

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

`for` 循环的第一个组件应该是字符串 `"iterator"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

`for` 循环的第二个组件应该是字符串 `"condition"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

`for` 循环的第三个部分应该是字符串 `"iteration"`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

`for` 循环的主体应该是空的。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
