---
id: 660f540c2176ea5dec01306d
title: Step 118
challengeType: 1
dashedName: step-118
---

# --description--

做得好！ 尝试使用 `character`、`count` 和 `inverted` 变量的不同值。

当你准备好继续下一个项目时，将 `character` 设置为 `"!"`、将 `count` 设置为 `10`，并将 `inverted` 设置为 `false` 以继续。

恭喜你完成你的第一个 JavaScript 项目！

# --hints--

你应该将 `character` 设置为 `"!"`。

```js
assert.equal(character, "!");
```

你应该将 `count` 设置为 `10`。

```js
assert.equal(count, 10);
```

你应该将 `inverted` 设置为 `false`。

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
