---
id: 660f2fbd45b520046cac68e8
title: Step 65
challengeType: 1
dashedName: step-65
---

# --description--

记住在之前的步骤中，你了解了返回值。 函数可以<dfn>返回</dfn>一个值，供你的应用程序单独使用。

在函数中，`return` 关键字用于指定返回值。 例如，此函数将返回赋予第一个参数的值：

```js
function name(parameter) {
  return parameter;
}
```

使用 `return` 关键字返回 `character` 变量的值，重复 `rowNumber` 次。

# --hints--

你应该使用 `.repeat()` 方法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

你应该在 `character` 变量上使用 `.repeat()` 方法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

你应该将 `rowNumber` 传递给 `.repeat()` 调用。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

你应该使用 `return` 关键字。

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

你应该返回 `.repeat()` 调用的结果。

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
