---
id: 660f20473aef47e9b8c9afc6
title: Step 40
challengeType: 1
dashedName: step-40
---

# --description--

添加 log 语句来打印 `result` 的值。 根据你使用的控制台，你可能看不到任何打印内容。

# --hints--

你应该添加一个 `console.log` 调用。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

你应该记录你的 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
