---
id: 660f4cffb1459d45e34902d1
title: Step 106
challengeType: 1
dashedName: step-106
---

# --description--

你实际上可以构建倒金字塔，而不需要像以前那样“向后”循环。

为此，你需要学习一些新的数组方法。 首先使用 `const` 声明一个 `numbers` 变量。 为其分配一个包含元素 `1`、`2` 和 `3` 的数组。 然后记录 `numbers` 数组。

# --hints--

你应该使用 `const` 来声明 `numbers` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

你的 `numbers` 变量应该是一个数组。

```js
assert.isArray(numbers);
```

你的 `numbers` 数组应按顺序包含元素 `1`、`2` 和 `3`。

```js
assert.deepEqual(numbers, [1,2,3]);
```

你应该记录你的 `numbers` 数组。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
