---
id: 660f1cedf3676fe26122ebf6
title: Step 38
challengeType: 1
dashedName: step-38
---

# --description--

你应该会看到控制台中打印出数字 0 到 7，每行一个。 这将作为生成金字塔的基础。

将你的日志语句替换为把 `i` 推送到 `rows` 数组的语句。

# --hints--

你不应该有 `console.log` 调用。

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

你应该在 `rows` 数组上调用 `.push()`。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

你应该将 `i` 推送到你的 `rows` 数组。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

你的 `.push()` 应该发生在你的 `for` 循环中。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
