---
id: 660f3915b41a441537ec9f5e
title: Step 72
challengeType: 1
dashedName: step-72
---

# --description--

因为你仅将 `i` 增加 `1`，所以你可以使用 <dfn>自增运算符</dfn> `++`。 该运算符将变量的值增加 1，更新该变量的赋值。 例如，`test` 在这里将变成 `8`：

```js
let test = 7;
test++;
```

将循环迭代的加法赋值替换为自增运算符。

# --hints--

你的 `for` 循环不应使用带有 `i` 的加法赋值。

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

你的 `for` 循环应该在 `i` 上使用自增运算符。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
