---
id: 660f4e74f7fd3f4a99ac2e50
title: Step 108
challengeType: 1
dashedName: step-108
---

# --description--

数组还有一个 `.shift()` 方法。 这将删除数组的 **第一个** 元素，而 `.pop()` 则会删除最后一个元素。 以下是 `.shift()` 方法的示例：

```js
const numbers = [1, 2, 3];
numbers.shift();
```

`numbers` 数组将是 `[2, 3]`。

在你的 `numbers` 数组正下方，声明一个 `shifted` 变量，并将在 `numbers` 数组上调用 `.shift()` 的结果分配给它。 在下一行，将 `shifted` 变量记录到控制台。

# --hints--

你应该使用 `const` 来声明 `shifted` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

你应该在 `numbers` 数组上调用 `.shift()`。

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

你应该将 `.shift()` 调用的结果分配给 `shifted` 变量。

```js
assert.equal(shifted, 1);
```

你应该记录你的 `shifted` 变量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
