---
id: 660f0f980e98e8cf77f1ce31
title: Step 30
challengeType: 1
dashedName: step-30
---

# --description--

更改你的 `rows` 声明，再次给它分配一个空数组。

同时，更改 `'Hello'` 字符串，再次使用双引号。 一般来说，用单引号或者双引号都没问题，但是你终会发现在整个项目中，统一使用一种方式是更好的选择。

# --hints--

你的 `rows` 数组应该是空的。

```js
assert.empty(rows);
```

你的 `"Hello"` 字符串应该用双引号。

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
