---
id: 660f3b664421471aa595170f
title: Step 74
challengeType: 1
dashedName: step-74
---

# --description--

金字塔现在看起来有点滑稽。 由于你从 `1` 而不是 `0` 开始循环，因此在将 `i` 传递给 `padRow` 时无需为其加一。

将 `padRow` 调用的第一个参数更新为 `i`。

# --hints--

你不应将 `i + 1` 传递给 `padRow`。

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

你应该将 `i` 传递给 `padRow`。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
