---
id: 660f48e1d3682f39e81843c4
title: Step 94
challengeType: 1
dashedName: step-94
---

# --description--

你的金字塔生成器仍在运转。 然而，它可能会再次陷入无限循环。

因为你仅检查 `done` 是否不等于 `count`，如果 `done` **大于** `count`，你的循环将永远持续下去。

更新循环的条件来检查 `done` 是否小于或等于 `count`。

# --hints--

你的 `while` 循环应检查 `done` 是否小于或等于 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
