---
id: 660f18f059fe0fda192ce394
title: Step 34
challengeType: 1
dashedName: step-34
---

# --description--

你的循环现在需要一个合适的迭代器。 <dfn>迭代器</dfn> 是一个变量，你可以在 `for` 循环中专门声明它，以控制循环如何迭代或执行逻辑。

在循环中使用 `i` 作为迭代器变量是一种常见的惯例。 `for` 循环允许你在括号 `()` 中声明这一点。 例如，这里有一个 `for` 循环，它声明一个 `index` 变量并为其赋值 `100`。

```js
for (let index = 100; "second"; "third") {

}
```

将字符串 `"iterator"` 替换为用 `let` 声明变量 `i`。 为其分配值 `0` 以开始。 这将在循环**第一次**运行时赋予 `i` 变量值 `0`。

# --hints--

你应该使用 `let` 来声明 `i` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

你应该将 `0` 分配给你的 `i` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

你的 `for` 循环应从 `0` 处启动 `i` 迭代器。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
