---
id: 660f165270622fd4ec0da3f7
title: Step 31
challengeType: 1
dashedName: step-31
---

# --description--

使用 `let` 关键字声明变量时，可以对其进行重新赋值。 这意味着你稍后可以将 `character` 更改为完全不同的值。

对于这个项目，你不会想要更改这些变量值。 因此，你应该使用 `const` 来声明它们。 `const` 变量很特殊。

首先，`const` 变量不能像 `let` 变量那样被重新分配。 这段代码将引发错误：

```js
const firstName = "Naomi";
firstName = "Jessica";
```

`const` 变量也不能未初始化。 这段代码将引发错误：

```js
const firstName;
```

用 `const` 替换 `let` 关键字。

# --hints--

你应该使用 `const` 来声明你的 `character` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

你应该使用 `const` 来声明你的 `count` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

你应该使用 `const` 来声明你的 `rows` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

你不应该在代码中使用 `let`。

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
