---
id: 660f4ae5b3924c3fc3373973
title: Step 101
challengeType: 1
dashedName: step-101
---

# --description--

你的迭代语句也会有所不同。 你不需要在每次循环时将 `1` 添加到 `i`，而是需要减去 `1`。

就像你之前对 `i = i + 1` 所做的那样，更新迭代语句，赋予 `i` 从其自身中减去 `1` 的值。

# --hints--

你的 `for` 循环应使用 `i = i - 1` 作为迭代。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
