---
id: 660f207334fabaeac3269c38
title: Step 41
challengeType: 1
dashedName: step-41
---

# --description--

要操作 `result` 字符串，你将使用不同类型的循环。 具体来说，`for...of` 循环遍历可迭代对象中的每个项目并将其临时分配给变量。

`for...of` 循环的语法如下：

```js
for (const value of iterable) {

}
```

请注意，你可以使用 `const`，因为该变量仅在一次迭代中存在，而不是在整个循环期间存在。

创建一个 `for...of` 循环来遍历 `rows` 数组，并将每个值分配给 `row` 变量。

# --hints--

你应该使用另一个 `for` 关键字。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

你应该声明一个 `row` 变量。

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

你应该使用 `const` 来声明你的 `row` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

你的 `for...of` 循环应该声明你的 `row` 变量。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

你应该使用 `of` 关键字从 `rows` 中提取 `row` 变量。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

你的 `for...of` 循环体应该是空的。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
