---
id: 660f1bf673487ae0bb25b900
title: Step 37
challengeType: 1
dashedName: step-37
---

# --description--

你的循环现在应该运行八次。 在循环主体内，打印 `i` 迭代器的值并观察会发生什么。

# --hints--

你应该使用 `console.log()`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

你应该记录 `i` 的值。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

你应该在 `for` 循环中记录 `i` 的值。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
