---
id: 660f3ce51f70571e1c5227c8
title: Step 76
challengeType: 1
dashedName: step-76
---

# --description--

注释有助于解释为什么你的代码采用某种方法，或者为你未来的自己留下待办事项。

在 JavaScript 中，你可以使用 `//` 在代码中留下单行注释。

在函数上方添加单行注释，提醒自己将代码更改为不同类型的循环。

# --hints--

你应该以 `//` 开始单行注释。

```js
assert.match(code, /\/\//);
```

你的单行注释至少应有五个字符长。

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
