---
id: 660f1e3f047bf4e403268713
title: Step 39
challengeType: 1
dashedName: step-39
---

# --description--

不幸的是，现在你无法看到循环正在做什么。

使用 `let` 声明一个 `result` 变量，并为其分配一个空字符串。 空字符串用引号表示，引号之间没有任何内容，例如 `""`。

# --hints--

你应该声明一个 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

你应该使用 `let` 来声明你的 `result` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s*result/);
```

你的 `result` 变量应该是一个空字符串。

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
