---
id: 660f4934fb48f63abd5ae371
title: Step 95
challengeType: 1
dashedName: step-95
---

# --description--

使用 `done` 来跟踪已生成的行数是可行的，但实际上你可以进一步清理逻辑。

数组具有特殊的 `length` 属性，可让你查看数组中有多少个值或 <dfn>元素</dfn>。 你可以使用类似 `myArray.length` 的语法访问此属性。

更新你的条件以检查 `rows.length` 是否小于或等于 `count`。

# --hints--

你的 `while` 循环应检查 `rows.length` 是否小于或等于 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
