---
id: 660f2b6fd54ac1fc142804dd
title: Step 47
challengeType: 1
dashedName: step-47
---

# --description--

格式化这个金字塔的逻辑可能会变得复杂，这意味着现在是将该代码提取到函数中的好时机。

<dfn>函数</dfn> 是可以在整个应用程序中重复使用的代码块。 函数使用以下语法声明：

```js
function name(parameter) {

}
```

`function` 关键字告诉 JavaScript `name` 变量将是一个函数。 `parameter` 是一个变量，表示在使用函数时传递给函数的值。 一个函数可以具有任意数量的 <dfn>参数</dfn>。 类似于 `for` 循环，花括号之间的空间是 <dfn>函数主体</dfn>。

声明一个 `padRow` 函数。 暂时不要创建任何参数变量。 函数体应该是空的。 请记住，你需要使用驼峰式命名约定。

# --hints--

你应该使用 `function` 关键字。

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

你应该声明一个 `padRow` 函数。

```js
assert.isFunction(padRow);
```

你的 `padRow()` 函数不应有任何参数。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

你的 `padRow()` 函数应该有一个空的主体。

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
