---
id: 660f17294346b7d69e79db3d
title: Step 32
challengeType: 1
dashedName: step-32
---

# --description--

You are now ready to start building your pyramid generator. Your `character` variable will serve as the building block for the pyramid.

`"Hello"` might not work very well for that. Change the value of `character` to be the hash character `"#"`.

# --hints--

Your `character` variable should be a hash symbol.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
