---
id: 660f2a70ad6225fa503e71c3
title: Step 46
challengeType: 1
dashedName: step-46
---

# --description--

你越来越近了！ 此时，你会遇到所谓的 <dfn>差一错误</dfn>，这是 JavaScript 等基于零的索引语言中经常出现的问题。

`rows` 数组的第一个索引是 `0`，这就是为什么你用 `i = 0` 启动 `for` 循环。 但重复字符串零次将导致不打印任何内容。

为了解决此问题，请在 `.repeat()` 调用中给 `i` 的值添加 `1`。 不要像在循环条件中那样将其分配回 `i`。

# --hints--

你应该在 `.repeat()` 方法中给 `i` 添加 `1`。

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
