---
id: 660f0c34aad72dc712b97624
title: Step 27
challengeType: 1
dashedName: step-27
---

# --description--

在这个项目中另一个必不可少的方法是 `.pop()`。 这个方法会将一个数组中的最后一个元素移除，并且 <dfn>returns</dfn>（返回）该元素。

当一个方法返回一个值时，你可以把它看作是把这个值还给你，让你能在代码的其他部分使用它。

创建一个新变量叫作 `popped` ，并将 `rows.pop()` 的结果分配给它。 然后，在控制台记录 `popped`。

# --hints--

你应该声明一个名为 `popped` 的变量。

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

你应该用关键字 `let` 来声明你的 `popped` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

你应该调用 `.pop()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

你应该在你的数组 `rows` 上，调用 `.pop()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

你应该记录你的变量 `popped`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
