---
id: 660f09a2694b59c3a10ee304
title: Step 23
challengeType: 1
dashedName: step-23
---

# --description--

现在，删除你的第一个控制台日志和 `rows[rows.length - 1]` 赋值。 将第二个 `rows` 日志语句留待以后使用。

# --hints--

你应该删除 `console.log(rows[0])` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

你应该删除 `rows[rows.length - 1]` 重新赋值。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

你不应删除 `console.log(rows)` 语句。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
