---
id: 660f4a1472f8e63d76162ce5
title: Step 99
challengeType: 1
dashedName: step-99
---

# --description--

如果将金字塔反过来，或者说 <dfn>倒置</dfn>，会怎样？ 是时候尝试一下了！

首先创建一个新的 `for` 循环。 声明你的迭代器 `i` 并为其分配 `count` 的值，然后使用布尔值 `false` 作为条件和迭代语句。

# --hints--

你的代码应该有一个 `for` 循环。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(/);
```

你的 `for` 循环应该使用 `count` 的值初始化 `i`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count/);
```

你的 `for` 循环应使用 `false` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false/);
```


你的 `for` 循环应使用 `false` 作为迭代。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
