---
id: 660f07d231941bc11719f664
title: Step 21
challengeType: 1
dashedName: step-21
---

# --description--

数组的特殊之处在于它们被认为是<dfn>可变的</dfn>。 这意味着你可以直接更改索引处的值。

例如，此代码将数字 `25` 分配给数组中的第二个元素：

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

将 `rows` 数组的 **第三** 元素更新为数字 `10`。 然后将 `rows` 数组打印到控制台。

# --hints--

你应该再次在 `rows` 数组上使用括号表示法。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

你应该访问 `rows` 数组的第三个元素。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

你应该对 `rows` 数组的第三个元素使用赋值运算符。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

你应该将值 `10` 分配给 `rows` 数组的第三个元素。

```js
assert.equal(rows[2], 10);
```

你的代码中应该有第二个 `console.log` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你应该记录 `rows` 数组。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
