---
id: 660f4f79e2a82a4e92290f44
title: Step 110
challengeType: 1
dashedName: step-110
---

# --description--

有时你可能希望恢复以前注释掉的代码。 你可以通过删除该代码周围的 `/*` 和 `*/` 来实现此目的。 这叫作<dfn>取消注释</dfn>。

仅取消注释第一个 `for` 循环。 保留单行注释和另外两个多行注释。

# --hints--

你应该取消注释第一个 `for` 循环。

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

你不应该删除你的单行注释。

```js
assert.match(code, /\/\/\sTODO:/);
```

你不应取消注释 `while` 循环。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

你不应取消注释第二个 `for` 循环。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
