---
id: 660f3dd626be3a1ffe27e5d1
title: Step 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript 还支持多行注释。 多行注释以 `/*` 开头，以 `*/` 结尾。

与单行注释不同，多行注释将封装多行。

使用 `/*` 和 `*/` 将当前的 `for` 循环（包括主体）转换为多行注释。

# --hints--

你应该以 `/*` 开始多行注释。

```js
assert.match(code, /\/\*/);
```

你应该以 `*/` 结束多行注释。

```js
assert.match(code, /\*\//);
```

你的整个 `for` 循环应该被注释掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
