---
id: 660f49e32001983c90b75850
title: Step 98
challengeType: 1
dashedName: step-98
---

# --description--

这是一个非常干净并且实用的循环。 做得好！ 但仍有更多的内容有待探索。

使用多行注释来注释掉你的 `while` 循环。

# --hints--

你的 `while` 循环应该被注释掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

你的 `while` 循环体应该被注释掉。

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /rows\.push/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
