---
id: 660f34e99571070d56d2f231
title: Step 67
challengeType: 1
dashedName: step-67
---

# --description--

你的 `padRow` 函数有两个你定义的参数。 当调用函数时，会为这些参数提供值。

你向函数调用提供的值叫作<dfn>参数</dfn>，并且你将<dfn>传递</dfn>参数给函数调用。 以下是将 `"Hello"` 作为参数传递的函数调用：

```js
test("Hello");
```

将 `i + 1` 和 `count` 作为参数传递给 `padRow` 调用。 与参数一样，传递的参数也用逗号分隔。

# --hints--

你应该将 `i + 1` 传递给 `padRow()` 调用。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

`i + 1` 参数后面应该有一个逗号。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

你应该传递 `count` 作为第二个参数。

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
