---
id: 660f061d259bbebc37461080
title: Step 20
challengeType: 1
dashedName: step-20
---

# --description--

你可以使用值的 <dfn>index</dfn> 访问数组内的值。 索引是一个数字，表示值在数组中的位置，从 `0` 开始表示第一个值。

你可以使用<dfn>括号表示法</dfn>访问该值，例如 `array[0]`。

使用 `console.log` 和括号表示法打印 `rows` 数组中的第一个值。

# --hints--

你的代码中应该有一个 `console.log()` 语句。

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

你应该访问你的 `rows` 数组。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

你应该在 `rows` 数组中使用括号表示法。

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

你应该使用括号表示法来访问 `rows` 数组的第一个元素。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

你应该记录 `rows` 数组的第一个元素。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
