---
id: 660f4de78f775e480ba2e451
title: Step 107
challengeType: 1
dashedName: step-107
---

# --description--

数组的 `.unshift()` 方法允许你将值添加到数组的 **开头**，这与 `.push()` 不同，后者将值添加到数组的末尾。 `.unshift()` returns the new length of the array it was called on.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

使用 `const` 声明一个 `unshifted` 变量，并将在 `numbers` 数组上调用 `.unshift()` 的结果分配给它。 传递 `5` 作为参数。 然后打印你的 `unshifted` 变量。

# --hints--

你应该声明你的 `unshifted` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

你应该在 `numbers` 数组上调用 `.unshift()` 方法。

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

你应该将 `5` 作为参数传递给 `.unshift()` 调用。

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

你应该将 `.unshift()` 调用的结果分配给 `unshifted` 变量。

```js
assert.equal(unshifted, 4);
```

You should log your `unshifted` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
