---
id: 660f1b6e60bd9edf902c81fd
title: Step 36
challengeType: 1
dashedName: step-36
---

# --description--

你的 <dfn>iteration</dfn> 语句将告诉你的循环在每次运行后如何处理迭代器。

重新分配变量时，可以使用该变量引用重新分配之前的先前值。 这使你可以执行一些操作，比如在已有的数字上加上 3。 例如，`bees = bees + 3;` 会将 `bees` 的值增加 3。

使用该语法将 `"iteration"` 字符串替换为给 `i` 增加 1 的新的赋值语句。

# --hints--

你应该给 `i` 变量加上 1。

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

你应该给你的 `i` 变量赋值 `i + 1`。

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

你的 `for` 循环应在每次迭代中将 `i` 增加 `1`。

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
