---
id: 660f359af3e32e0f1a6880b7
title: Step 68
challengeType: 1
dashedName: step-68
---

# --description--

现在你应该在控制台中看到同一组字符。 你的 `padRow` 函数正在执行与你之前完全相同的操作，但现在它位于其自己的可重复使用部分中。

使用加法运算符将单个空格 `" "` 连接到重复 `character` 字符串的开头和结尾。

请记住，你可以使用 `+` 运算符来连接字符串，如下所示：

```js
" " + "string"
```

# --hints--

你应该将一个空格连接到返回值的开头。

```js
assert.match(padRow(1, 1), /^\s/);
```

你应该将一个空格连接到返回值的末尾。

```js
assert.match(padRow(1, 1), /\s$/);
```

你的 `padRow()` 函数应返回重复的 `character` 系列，并且在系列前后都有空格。

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
