---
id: 660f033cf051ebb50ea3bf48
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

在编程中，你经常需要处理大量数据。 有许多数据结构可以帮助你组织和管理数据。 最基本的数据结构之一是<dfn>数组</dfn>。

<dfn>数组</dfn>是一种非原始数据类型，可以保存一系列值。 非原始数据类型与原始数据类型的不同之处在于前者可以保存更复杂的数据。 字符串和数字等原始数据类型一次只能保存一个值。

数组使用方括号 (`[]`) 表示。 下面是一个值为空数组的变量的示例：

```js
let array = [];
```

声明一个 `rows` 变量并为其分配一个空数组。

# --hints--

你应该有一个 `rows` 变量。

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

你应该使用 `let` 来声明你的 `rows` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

你应该为 `rows` 变量分配一个空数组。

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
