---
id: 660f50a21fe7645252804f2b
title: Step 112
challengeType: 1
dashedName: step-112
---

# --description--

如果你有办法在倒金字塔和标准金字塔之间切换会怎样？

首先声明一个 `inverted` 变量，并为其赋值 `true`。 你不会在代码中更改此变量，但你需要使用 `let` 以便我们的测试可以稍后修改它。

# --hints--

你应该使用 `let` 声明一个 `inverted` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

你应该使用值 `true` 初始化 `inverted`。

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
