---
id: 660f039ff313dbb696b007ca
title: Step 19
challengeType: 1
dashedName: step-19
---

# --description--

当数组包含值或 <dfn>元素</dfn> 时，这些值以逗号分隔。 这是一个包含两个字符串的数组：

```js
let array = ["first", "second"];
```

将 `rows` 声明更改为包含字符串 `"Naomi"`、`"Quincy"` 和 `"CamperChan"` 的数组。 数组中值的顺序很重要，因此请遵循该顺序。 请记住，字符串区分大小写。

# --hints--

数组中的第一个元素应该是字符串 `"Naomi"`。

```js
assert.equal(rows[0], "Naomi");
```

数组中的第二个元素应该是字符串 `"Quincy"`。

```js
assert.equal(rows[1], "Quincy");
```

数组中的第三个元素应该是字符串 `"CamperChan"`。

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
