---
id: 660f255022991ef34ed0ee88
title: Step 44
challengeType: 1
dashedName: step-44
---

# --description--

打印数字不会产生具有视觉吸引力的金字塔。 现在你正在输出 `rows` 数组的格式化内容，是时候更新原始循环了。

不要将 `i` 加到数组，而是添加 `character` 变量的值。

# --hints--

你不应再添加 `i` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

你应该添加你的 `character` 变量。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
