---
id: 660f487dc0c8fa38084f9754
title: Step 92
challengeType: 1
dashedName: step-92
---

# --description--

由于你已将比较移至 `while` 条件，因此你可以删除整个 `if` 语句。

# --hints--

你不应再有 `if` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(\s*done\)/);
```

你不应再将 `continueLoop` 设置为 `false`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/continueLoop\s*=\s*false/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (done !== count) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
