---
id: 660f2eccfe3f820304af1b39
title: Step 64
challengeType: 1
dashedName: step-64
---

# --description--

要知道如何格式化一行，你的 `padRow` 函数需要知道你所在的行号，以及总共生成了多少行。

最好的方法是给它们创建函数参数。 为 `padRow` 函数提供 `rowNumber` 和 `rowCount` 参数。 多个参数以逗号分隔：

```js
function name(first, second) {

}
```

# --hints--

你的 `padRow` 函数应该有一个 `rowNumber` 参数。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

你应该在 `rowNumber` 参数后添加一个逗号。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

你的 `padRow` 函数应该有一个 `rowCount` 参数。

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
