---
id: 660f4455f457ef2e3ec6920f
title: Step 80
challengeType: 1
dashedName: step-80
---

# --description--

现在字符串不再打印，因为 `false` 不是 `true`。 但其他值会怎样呢？

尝试将条件更改为字符串 `"false"`。

# --hints--

你的 `if` 语句应以字符串 `"false"` 作为条件。

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*('|")false\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (false) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
