---
id: 660f46460f9c36330ebc07d8
title: Step 88
challengeType: 1
dashedName: step-88
---

# --description--

相等运算符可能会导致 JavaScript 中出现一些奇怪的行为。 例如，`"0" == 0` 为真，即使一个是字符串，一个是数字。

<dfn>严格相等</dfn>运算符 `===` 用于检查两个值是否相等且具有相同的类型。 一般来说，这是你应该始终使用的相等运算符。 使用严格相等运算符时，`"0" === 0` 变为假，因为虽然它们可能具有相同的零值，但它们不是同一类型。

更新你的 `done == count` 条件以使用严格相等运算符。

# --hints--

你的 `if` 条件应使用严格相等。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
