---
id: 660f374d532dc41189cc9cc2
title: Step 69
challengeType: 1
dashedName: step-69
---

# --description--

现在该进行一些数学运算了。 考虑一个三排金字塔。 如果我们想让它居中，它看起来应该像这样：

```md
··#··
·###·
#####
```

为了便于阅读，空格已被中间点取代。 如果你研究该图案，你会发现行首和行末的空格遵循一种模式。

更新空白字符串以重复 `rowCount - rowNumber` 次。

打开控制台查看结果。

# --hints--

你应该在 `" "` 字符串上调用 `.repeat()` 来重复 `rowCount - rowNumber` 次。

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
