---
id: 6610b8017d1671f2814e8c77
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

当你声明一个变量但没有初始化它时，该变量被视为<dfn>未初始化</dfn>。 目前，你的 `secondCharacter` 变量未初始化。

添加 `console.log()` 来查看 `secondCharacter` 变量的值。

# --hints--

你应该添加第二个 `console.log`。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

你应该记录你的 `secondCharacter` 变量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
