---
id: 6610c71600966a2191d3a64a
title: Step 59
challengeType: 1
dashedName: step-59
---

# --description--

函数返回的值可通过调用该函数来使用。 你可以直接使用函数调用作为其返回的值，或者在变量中捕获返回的值。 这样，你就可以在创建该变量的函数之外使用分配给本地作用域变量的值。

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

要使用 `"Testing"` 值，请通过更新 `return` 语句以仅返回 `test` 变量，将其从 `padRow` 函数中返回。


# --hints--

你的 `padRow` 函数应该返回 `test` 变量。

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
