---
id: 6610bbed59bc2a0194d85533
title: Step 22
challengeType: 1
dashedName: step-22
---

# --description--

注意到 `rows` 数组内的值是如何直接改变的？ 这被称为<dfn>突变</dfn>。 当你对数组了解得更多时，你就会了解何时应该改变数组，何时不应该改变数组。

在继续之前，这是一个学习常见数组用法的好机会。 目前，你的代码使用 `rows[2]` 访问数组中的最后一个元素。 但是当你想要最后一个元素时，你可能不知道数组中有多少个元素。

你可以使用数组的 `.length` 属性 - 这将返回数组中元素的数量。 要获取任何数组的最后一个元素，可以使用以下语法：

```js
array[array.length - 1]
```

`array.length` 返回数组中元素的数量。 通过减去 `1`，你可以得到数组中最后一个元素的索引。 你可以将相同的概念应用到 `rows` 数组。

更新 `rows[2]` 以动态访问 `rows` 数组中的最后一个元素。 请参阅上面的例子来获得帮助。

你将不会在控制台中看到任何变化。

# --hints--

你应该使用 `rows` 数组的 `.length` 属性。

```js
assert.match(__helpers.removeJSComments(code), /rows\.length/);
```

你应该在括号符号中使用 `rows.length`。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length/);
```

你应该从括号符号的长度中减去 `1`。

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--
rows[2] = 10;
--fcc-editable-region--
console.log(rows);
```
