---
id: 6610c424b7119919b62932f4
title: Step 53
challengeType: 1
dashedName: step-53
---

# --description--

函数不必返回硬编码值。 它可以返回存储在变量中的值。 参数是函数的特殊变量，因此也可以返回它们。

将你的 `padRow` 函数更改为直接 `return` `name` 参数。

# --hints--

你的 `padRow` 函数应该返回 `name` 参数的值。

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
