---
id: 6610c48c4ea0891afa7c4696
title: Step 54
challengeType: 1
dashedName: step-54
---

# --description--

如果你再次打开控制台，你会看到 `padRow` 函数返回 `undefined`，即使你定义了返回值！ 这是因为当你**调用**该函数时，需要给参数赋值。

当你将一个值传递给函数调用时，该值被称为<dfn>参数</dfn>。 下面是调用 `demo` 函数并将 `"Naomi"` 作为 `name` 参数传递的示例。

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

将你自己的名字作为 `padRow` 调用中 `name` 参数传递。 请记住，你的名字是一个字符串，因此你需要使用引号。

# --hints--

你应该将一个字符串传递给你的 `padRow()` 调用。

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
