---
id: 6610c83b52583e245a079217
title: Step 61
challengeType: 1
dashedName: step-61
---

# --description--

现在你的 `call` 变量具有值 `"Testing"`。 但是你的函数不再使用 `name` 参数。

从函数声明中删除 `name` 参数，然后从 `padRow` 调用中删除 `"CamperChan"` 字符串。

另外，从 `padRow` 函数中删除两个 `console.log`。
# --hints--

你的 `padRow` 函数不应该有 `name` 参数。

```js
assert.notMatch(padRow.toString(), /name/);
```

你不应将 `"CamperChan"` 传递给你的 `padRow` 调用。

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

你仍然应该调用 `padRow` 函数。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

在 `return` 关键字之前不应该有 `console.log`。

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

`return` 关键字之后不应该有 `console.log`。

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
