---
id: 6610c1d97b1671140f95cfbb
title: Step 50
challengeType: 1
dashedName: step-50
---

# --description--

现在添加一个 log 语句来打印 `call` 变量的值。

# --hints--

你应该添加一个 `console.log` 调用。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

你应该记录你的 `call` 变量。 Don't forget the semicolon.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
