---
id: 6610c2d8d67563174fcf96dc
title: Step 52
challengeType: 1
dashedName: step-52
---

# --description--

当你的代码中有一个明确写出的值时，例如函数中的 `"Hello!"` 字符串，它被视为<dfn>硬编码</dfn>。 在函数内部对值进行硬编码可能无法使其像你希望的那样具有可重用性。

相反，你可以为该函数定义 <dfn>参数</dfn>。 参数是调用函数时被赋予值的特殊变量，并且可以在函数中使用来动态地改变函数代码的结果。

要向函数添加参数，需要在括号内添加变量名。 例如，此 `demo` 函数有一个 `name` 参数：

```js
function demo(name) {

}
```

`name` 听起来像是一个有用的参数，所以继续将其添加到你的 `padRow` 函数中。

# --hints--

你的 `padRow` 函数应该有一个 `name` 参数。

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
