---
id: 6610c538372aa61cc0f5b122
title: Step 57
challengeType: 1
dashedName: step-57
---

# --description--

JavaScript 中的变量在特定的<dfn>范围</dfn>内可用。 换句话说，变量的声明位置决定了它在代码中的使用位置。

第一个范围是全局范围。 在任何“块”（例如函数或 `for` 循环）之外声明的变量都位于<dfn>全局范围</dfn>内。 你的 `character`、`count` 和 `rows` 变量均位于全局范围内。

当变量位于全局范围内时，函数可以在其定义中访问它。 以下是使用全局 `title` 变量的函数示例：

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

此示例将返回 `"Professor Naomi"`。 更新你的 `padRow` 函数以返回将你的 `character` 变量连接到 `name` 参数开头的值。

# --hints--

你的 `padRow` 函数应该将 `character` 连接到 `name` 的开头。

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

你的 `padRow` 函数应该返回 `character + name` 的结果。

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
