---
id: 6610c16c4fa0df12c0e30675
title: Step 49
challengeType: 1
dashedName: step-49
---

# --description--

你正在调用 `padRow` 函数，但没有对该函数调用执行任何操作。 JavaScript 中的所有函数都<dfn>返回</dfn>一个值，这意味着它们提供了调用它们的定义结果供你在其他地方使用。

要查看调用 `padRow` 函数的结果，请声明一个 `call` 变量并将现有的 `padRow` 调用分配给该变量。

# --hints--

你应该声明一个 `call` 变量。

```js
assert.match(__helpers.removeJSComments(code), /(const|let)\s+call/);
```

你应该使用 `const` 来声明你的 `call` 变量。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

你应该将 `padRow` 调用的结果分配给 `call`。

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
