---
id: 6610c77d50636722e5b6be17
title: Step 60
challengeType: 1
dashedName: step-60
---

# --description--

在 `return` 语句下方，将字符串 `"This works!"` 记录到控制台。

执行此操作后，你将看到字符串 `"This works!"` 没有显示在控制台中，并且 `console.log("This works!")` 行显示为灰色。

复制控制台日志并将其粘贴到 `return` 语句上方。 现在，字符串 `"This works!"` 应该出现在控制台中。

关于 `return` 关键字，需要了解的一件重要事情是，它不仅定义从函数返回的值，还会停止函数或块语句内代码的执行。 这意味着 `return` 语句之后的任何代码都不会运行。

# --hints--

你的 `padRow` 函数应该返回 `test` 变量。

```js
assert.equal(padRow("Naomi"), "Testing");
```

你的第一个 `console.log` 应该位于 `return` 关键字之后。

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

你的第二个 `console.log` 应该位于 `return` 关键字之前。

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";

  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
