---
id: 6610bf6fa14d700beed1b109
title: Step 87
challengeType: 1
dashedName: step-87
---

# --description--

<dfn>相等</dfn>运算符 `==` 用于检查两个值是否相等。 要比较两个值，你可以使用类似 `value == 8` 的语句。

在循环中添加 `if` 语句。 该语句应使用相等运算符检查 `done` 是否等于 `count`。

# --hints--

你应该在循环中使用 `if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

你的 `if` 语句应使用相等运算符来比较条件中的 `done` 和 `count`。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
