---
id: 6610c87eac0f0b256d7b037e
title: Step 62
challengeType: 1
dashedName: step-62
---

# --description--

因为你的函数不再使用该参数，所以更改调用中传递的值不会对其产生影响。

继续从 `padRow` 函数中删除 `test` 声明。 另外，删除 `return` 语句，这样你的函数再次变为空。

# --hints--

你的 `padRow` 函数不应该有 `test` 变量。

```js
assert.notMatch(padRow.toString(), /test/);
```

你的 `padRow` 函数不应返回值。

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
