---
id: 6610c105bbdacc114d6cdc44
title: Step 48
challengeType: 1
dashedName: step-48
---

# --description--

为了使用函数，你需要调用它。 <dfn>函数调用</dfn> 告诉你的应用程序在你选择调用该函数的任何地方运行该函数的代码。 函数调用的语法是函数名称后跟括号。 例如，此代码定义并调用 `test` 函数。

```js
function test() {

}
test();
```

调用你的 `padRow` 函数。

# --hints--

你应该调用 `padRow` 函数。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

你的 `padRow` 函数体应该是空的。

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
