---
id: 6610c8cfe4cf4d278e35c156
title: Step 63
challengeType: 1
dashedName: step-63
---

# --description--

正如预期的那样，你的函数现在再次返回 `undefined`。 你的 `call` 变量不再需要，因此请删除 `call` 声明和 `call` 变量的 `console.log`。

# --hints--

你不应该有 `call` 声明。

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s*call/);
```

你不应该记录你的 `call` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
