---
id: 6610c6541c82551f95e765ab
title: Step 58
challengeType: 1
dashedName: step-58
---

# --description--

变量也可以在函数内部声明。 这些变量被认为处于<dfn>局部作用域</dfn>或<dfn>块级作用域</dfn>内。 在函数内部声明的变量只能在该函数内部使用。 如果你尝试在函数外部访问它，则会收到引用错误。

要查看其实际效果，请使用 `const` 在 `padRow` 函数中声明一个 `test` 变量。 使用值 `"Testing"` 对其进行初始化。

然后，在你的函数下方，尝试将 `test` 记录到控制台。 你将看到一个错误，因为它没有在函数的局部作用域之外定义。 删除 `console.log` 以通过测试并继续。

# --hints--

你的函数应该声明一个 `test` 变量。

```js
assert.match(padRow.toString(), /var\s+test/);
```

你应该使用值 `"Testing"` 初始化 `test`。 Don't forget the semicolon.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

你的 `test` 变量应位于 `return` 关键字之前。

```js
const str = padRow.toString();
assert.isBelow(str.indexOf("test"), str.indexOf("return"));
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
