---
id: 6610b8f6a98d25f4d485a94d
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

未初始化变量的默认值是 `undefined`。 这是一种特殊的数据类型，表示尚未定义的值。

你仍然可以为未初始化的变量赋值。 以下是一个例子：

```js
let uninitialized;
uninitialized = "assigned";
```

将字符串 `"Test"` 分配给声明下方的 `secondCharacter` 变量。 打开控制台查看你的日志是如何变化的。

# --hints--

你不应该初始化 `secondCharacter`。 请记住，初始化意味着在声明变量时分配一个值。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

你应该在 `secondCharacter` 上使用赋值运算符。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

你应该将字符串 `"Test"` 分配给 `secondCharacter` 变量。

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
