---
id: 6610c21b3ef82015573ffbbe
title: Step 51
challengeType: 1
dashedName: step-51
---

# --description--

你的 `call` 变量具有 `undefined` 值，即使你定义了它！ 这是因为你的 `padRow` 函数当前没有返回值。 默认情况下，函数返回 `undefined` 作为其值。

为了返回其他内容，你需要使用 `return` 关键字。 下面是一个返回字符串 `"Functions are cool!"` 的函数示例：

```js
function demo() {
  return "Functions are cool!";
}
```

使用 `return` 关键字让你的函数返回字符串 `"Hello!"`。

# --hints--

Your `padRow` function should use the `return` keyword.

```js
assert.match(padRow.toString(), /return/);
```

Your `padRow` function should return the string `"Hello!"`.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
