---
id: 6610b9f7619764fad5fd516d
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

你还可以将一个变量的值分配给另一个变量。 例如：

```js
let first = "One";
let second = "Two";
second = first;
```

`second` 变量现在的值为 `"One"`。

要查看实际效果，请将 `secondCharacter` 赋值从 `"Test"` 更改为 `character` 变量。

然后打开控制台查看记录的内容。

# --hints--

你不应将值 `"Test"` 分配给 `secondCharacter` 变量。

```js
assert.notEqual(secondCharacter, "Test");
```

你应该将 `character` 变量的值分配给 `secondCharacter` 变量。 Don't forget your semicolon.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

你的 `secondCharacter` 变量现在应该具有值 `"World"`。


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
