---
id: 6610b741b54b90f0c0fb3d58
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

当变量名超过一个单词时，对于如何将单词大写，有特定的命名约定。 在 JavaScript 中，使用的约定是 <dfn>驼峰式命名法</dfn>。

驼峰式大小写是指名称中的第一个单词全部小写，但后续单词均采用首字母大写。 以下是驼峰式命名法的一些示例：

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

使用驼峰式命名法声明一个新的 `secondCharacter` 变量。

# --hints--

你应该声明一个 `secondCharacter` 变量。

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

你应该使用 `let` 来声明你的 `secondCharacter` 变量。

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

你不应该为 `secondCharacter` 变量赋值。 Do not forget your semicolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
