---
id: 661483051820c3c1ab4595e0
title: Step 96
challengeType: 1
dashedName: step-96
---

# --description--

你还可以在 `padRow` 调用中替换 `done` 引用。

请注意，此处的 `rows.length` 会出现差一错误，因为 `done` 在被调用之前 *已经* 递增了。

因此你需要将此处的 `done` 替换为 `rows.length + 1`。 当你执行此操作时，你可能会看到 `Range Error`，因为我们又创建了另一个差一错误。

你需要更改 `while` 条件以使用小于运算符，而不是小于或等于运算符。

# --hints--

你应该将 `rows.length + 1` 作为 `padRow` 调用的第一个参数传递。

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(padRow\s*\(\s*rows\.length\s*\+\s*1/);
```

当 `rows.length` 小于 `count` 时，你的 `while` 循环应该运行。

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
