---
id: 664599653fcd6e97104f9261
title: Step 83
challengeType: 1
dashedName: step-83
---

# --description--

有时，你会希望在所有 `if...else if` 条件均为 `false` 时运行不同的代码。 你可以通过添加 `else` 块来实现这一点。

仅当 `if` 和 `else if` 块中的条件不满足时，才会评估 `else` 块。

这里将 `else` 块添加到 `else if` 块。

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

向 `else if` 块添加 `else` 块。 在 `else` 块内，将字符串 `"This is the else block"` 记录到控制台。

要在控制台中查看结果，可以手动将 `else if` 语句中的 `<` 更改为 `>`。 这将使得条件为 `false` 并且 `else` 块将运行。

# --hints--

你应该有一个 `else` 块。

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

你的 `else` 块应该将字符串 `"This is the else block"` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

