---
id: 66458f0a05df478aa627629e
title: Step 82
challengeType: 1
dashedName: step-82
---

# --description--

文字又不见了！ 空字符串的计算结果为 `false`，使其成为 <dfn>假值</dfn>。 你将在未来的项目中了解有关真值和假值的更多信息。

除了 `if` 语句之外，JavaScript 还有 <dfn>else if</dfn> 语句。 `else if` 语句允许你在单个代码块中检查多个条件。

以下是 `else if` 语句的语法：

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

如果第一个条件为 `false`，JavaScript 将检查链中的下一个条件。 如果第二个条件为 `false`，JavaScript 将检查第三个条件，依此类推。

在 `if` 语句下方，添加一个 `else if` 语句，检查 `5` 是否小于 `10`。 然后在 `else if` 语句的主体内，将字符串 `"5 is less than 10"` 记录到控制台。

检查控制台来查看结果。

# --hints--

你应该有一个 `else if` 语句。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

你的 `else if` 语句应检查 `5` 是否小于 `10`。

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

你应该将字符串 `"5 is less than 10"` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```

