---
id: 664ed97c55a99e5ffab759dc
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Before moving forward, you should take a moment to review the concepts you have learned.

Use the `let` keyword to declare a `profession` variable and an `age` variable. Initialize `profession` with the string `"teacher"`, but do not initialize `age` with any value.

Log both of your variables to the console to see the results.

# --hints--

You should declare a `profession` variable.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

You should use `let` to declare the `profession` variable.

```js
assert.match(code, /let\s+profession/);
```

You should assign the string `"teacher"` to the `profession` variable.

```js
assert.strictEqual(profession, "teacher");
```

You should declare an `age` variable.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

You should use `let` to declare the `age` variable.

```js
assert.match(code, /let\s+age/);
```

You should not assign a value to the `age` variable.

```js
assert.isUndefined(age);
```

You should log `age` to the console.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

You should log `profession` to the console.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
