---
id: 664ed9344ba3495e1f67f8ea
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

现在你已经检查了声明和初始化，删除你为检查而编写的代码。

不要删除你的 `character` 变量。

# --hints--

你不应该有 `profession` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

你不应该有 `age` 变量。

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

你不应该对 `age` 有 `console.log` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

你不应该对 `profession` 有 `console.log` 语句。

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

你应该仍然有你的 `character` 变量。

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
