---
id: 664ed97c55a99e5ffab759dc
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

在继续之前，你应该花点时间复习一下已经学过的概念。

使用 `let` 关键字声明 `profession` 变量和 `age` 变量。 使用字符串 `"teacher"` 初始化 `profession`，但不要使用任何值初始化 `age`。

将两个变量记录到控制台以查看结果。

# --hints--

你应该声明一个 `profession` 变量。

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

你应该使用 `let` 来声明 `profession` 变量。

```js
assert.match(code, /let\s+profession/);
```

你应该将字符串 `"teacher"` 分配给 `profession` 变量。

```js
assert.strictEqual(profession, "teacher");
```

你应该声明一个 `age` 变量。

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

你应该使用 `let` 来声明 `age` 变量。

```js
assert.match(code, /let\s+age/);
```

你不应该为 `age` 变量赋值。

```js
assert.isUndefined(age);
```

你应该将 `age` 记录到控制台。

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

你应该将 `profession` 记录到控制台。

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
