---
id: 665732da4815b70bb083915e
title: Step 24
challengeType: 1
dashedName: step-24
---


# --description--

在前面几个步骤中，你学习了有关使用数组的所有知识。 花一点时间复习一下你所学到的知识。

首先声明一个 `cities` 变量，并将其初始化为字符串 `"London"`、`"New York"` 和 `"Mumbai"` 的数组。 然后将该变量记录到控制台。

记录完成后，将 `cities` 的最后一个元素更改为字符串 `"Mexico City"`，然后再次记录 `cities` 变量。

正确完成后，你应该在控制台中看到此输出。

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

你应该使用 `let` 来声明 `cities` 变量。

```js
assert.match(code, /let\s+cities/);
```

你应该将字符串 `"London"`、`"New York"` 和 `"Mumbai"` 的数组分配给 `cities` 变量。

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*\]/);
```

你应该使用 `console.log()` 将整个 `cities` 数组记录到控制台。

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

你应该将 `cities` 数组的最后一个元素更新为字符串 `"Mexico City"`。 请记住，你可以使用 `array[array.length - 1]` 访问数组的最后一个元素。

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

你的代码中应该有两个 `console.log(cities)` 语句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
