---
id: 66643c9a9e0ad1c787f5fe2b
title: Step 56
challengeType: 1
dashedName: step-56
---

# --description--

完成快速回顾后，你应该删除 `addTwoNumbers` 函数、`sum` 变量和日志语句。

# --hints--

你的代码中不应该有 `addTwoNumbers` 函数。

```js
assert.notMatch(code, /addTwoNumbers/);
```

你的代码中不应该有 `sum` 变量。

```js
assert.notMatch(code, /sum/);
```

你不应该记录 `sum` 变量。

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
