---
id: 66643c93e05093c728abdbe9
title: Step 55
challengeType: 1
dashedName: step-55
---

# --description--

在继续之前，请花点时间回顾一下函数的工作原理。

声明一个名为 `addTwoNumbers` 的函数。 该函数应该接受两个参数并返回这两个参数的总和。

然后声明一个 `sum` 变量并为其分配使用 `5` 和 `10` 作为参数调用 `addTwoNumbers` 函数的值。 将 `sum` 变量记录到控制台。


# --hints--

你应该有一个名为 `addTwoNumbers` 的函数。

```js
assert.isFunction(addTwoNumbers);
```

你的函数应该返回两个数字的总和。

```js
assert.strictEqual(addTwoNumbers(2,3), 5);
```

你应该声明一个 `sum` 变量。

```js
assert.isDefined(sum);
```

你应该将以 `5` 和 `10` 作为参数调用 `addTwoNumbers` 函数得到的值分配给 `sum`。

```js
assert.strictEqual(sum, 15);
```

你应该记录你的 `sum` 变量。

```js
assert.match(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--

--fcc-editable-region--

const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
