---
id: bad87fee1348bd9aedf08806
title: 更改元素的字体大小
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 1
---

# --description--

字体大小由`font-size`属性控制，如下所示：

```css
h1 {
  font-size: 30px;
}
```

# --instructions--

在包含`red-text`class 选择器的`<style>`声明区域的里，创建一个`p`元素样式规则，并设置`font-size`为`16px`。

# --hints--

在`style`样式声明区域里，`p`元素的`font-size`的值应为`16px`，浏览器和文本缩放应设置为 100％。

```js
assert(code.match(/p\s*{\s*font-size\s*:\s*16\s*px\s*;\s*}/i));
```

# --solutions--

