---
id: 660ee6e3a242da6bd579de69
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Uno de los conceptos más importantes en la programación son las variables. Una <dfn>variable </dfn> apunta a una dirección de memoria específica que almacena un valor. A las variables se les asigna un nombre que puede ser usado a través de tu código para acceder a ese valor.

Declarar una variable significa darle un nombre. En JavaScript, esto a menudo se hace con la palabra clave `let`. Por ejemplo, así es cómo declararías una variable `hola`:

```js
let hello;
```

Debes serguir reglas específicas para nombrar una variable: los nombres pueden incluir letras, números, símbolo de dólar y guiones bajos, pero no pueden contener espacios y no deben comenzar con un número.

Usa la palabra clave `let` para declarar una variable llamada `character`.

_Nota_: Terminar las declaraciones en JavaScript con punto y coma es una práctica común. `;`

# --hints--

Deberías usar `let` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Deberías usar `character` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Deberías usar `let` para declarar la variable `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Your declaration should end with a semicolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
