---
id: 660ef19b95d3308e7dd31bb6
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

La <dfn>consola</dfn> te permite imprimir y ver salida de JavaScript. Puedes enviar información a la consola usando `console.log()`. Por ejemplo, este código imprimirá `"Naomi"` en la consola:

```js
let developer = "Naomi";
console.log(developer);
```

El código anterior accede a la variable `developer` con su nombre en el `console.log() `. Ten en cuenta que el valor entre paréntesis es el valor que se imprimirá.

Imprime el valor de la variable `character` en la consola. Luego, haz clic en el botón "Consola" para ver la consola JavaScript.

# --hints--

Deberías acceder a la `consola` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Deberías acceder a la propiedad `log` de `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Deberías usar paréntesis para llamar al método `.log()`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Deberías imprimir la variable `character` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
