---
id: 660ef5c1904955978a986a5c
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Usa `let` para declarar una variable `count`. Asígnale el <dfn>número</dfn> `8`. Cuando se utiliza un valor numérico, no se utilizan comillas. Por ejemplo:

```js
let money = 100;
```

# --hints--

Deberías tener una segunda palabra reservada `let` en tu código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Deberías usar `let` para declarar la variable `count`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Deberías asignar el número `8` a tu variable `count`.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
