---
id: 660ef0f7c4b8e68ccd1f0786
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript tiene siete tipos de datos primitivos, siendo `String` uno de ellos. En JavaScript un <dfn>string</dfn> representa una secuencia de caracteres y puede encerrarse entre comillas simples (`'`) o dobles (`"`).

Ten en cuenta que las cadenas son <dfn>inmutables</dfn>, lo que significa que una vez creadas, no se pueden cambiar. La variable aún puede ser reasignada con otro valor.

Cambia tu string `"Hello"` para usar comillas simples.

# --hints--

No deberías tener comillas dobles en tu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Deberías usar comillas simples para tu cadena de texto `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Deberías seguir usando `let` en tu código.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Aún debes declarar una variable `character`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Tu variable `character` debe tener la cadena de texto `"Hello"` como valor.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
