---
id: 660eebd83100d37862268781
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Tu variable `character` actualmente no tiene ningún valor. Puedes asignarle un valor usando el operador de <dfn>asignación</dfn> `=`. Por ejemplo:

```js
let hello = "Hello";
```

Asignar un valor a una variable en el momento de su declaración es conocido como <dfn>inicialización</dfn>.

Inicializa tu variable `character` asignándole el valor `"Hola"` durante su declaración.

# --hints--

Debes utilizar el operador de asignación `=`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Debes utilizar la cadena de texto `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Deberías usar comillas dobles para tu cadena de texto `"Hello"`.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Debes asignar `"Hello"` a la variable `character`.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
