---
id: 6610c71600966a2191d3a64a
title: Step 52
challengeType: 1
dashedName: step-52
---

# --description--

There is a bit of a problem, here. Your `test` declaration is currently after your `return` statement. An important thing to know about the `return` keyword is that it does not just define a value to be returned from your function, it also stops the execution of your function code. This means that any code after a `return` statement will not run.

Move your `test` initialization to the line above your `return` statement.

# --hints--

Your `test` variable should come before your `return` keyword.

```js
const str = padRow.toString();
assert.isBelow(str.indexOf("test"), str.indexOf("return"));
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
  const test = "Testing";
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
