---
id: 6551f5756c20146fc94f8675
title: 6° passo
challengeType: 20
dashedName: step-6
---

# --description--

You can access the number of characters in a string with the built-in `len()` function.

Modify your existing `print()` call to print `len(text)`.

# --hints--


You should call the `len()` function.

```js
({ test: () => assert.match(code, /len\s*\(.*\)/) })

```

You should pass the variable `text` to the `len()` function.

```js
({ test: () => assert.match(code, /len\s*\(\s*text\s*\)/) })

```

You should pass `len(text)` to the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
