---
id: 65524b790ba8558a2f1c9fe5
title: Passo 32
challengeType: 20
dashedName: step-32
---

# --description--

Clean the output a bit. Delete `print(char, index)`, and turn the last `print()` call into `print('char:', char, 'new char:', new_char)`.

# --hints--

You should delete the `print(char, index)` call.

```js
({ test: () => assert.isFalse(/print\s*\(\s*char\s*,\s*index\s*\)/.test(code)) })
```

You should turn `print(new_char)` into `print('char:', char, 'new char:', new_char)`.

```js
({
  test: () => {
    assert.isFalse(/print\s*\(\s*new_char\s*\)/.test(code));
    assert.match(code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')new\schar:\2\s*,\s*new_char\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
    print(new_char)
--fcc-editable-region--
```
