---
id: 60a3e3396c7b40068ad69978
title: ステップ 15
challengeType: 0
dashedName: step-15
---

# --description--

要素外の余白を調整するには、マージンを使用します。

`margin` プロパティを使用して、`.frame` 要素の垂直方向のマージンを `20px` に、水平方向のマージンを `auto` に設定してください。 これにより、額縁が 20 ピクセル下に移動し、ページの水平方向の中央に配置されます。

# --hints--

`margin` プロパティを `20px auto` に設定する必要があります。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

`.frame` 要素には値が `20px auto` の `margin` が必要です。

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
