---
id: 60a3e3396c7b40068ad69972
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

`.canvas` ルールに、`height` プロパティと値 `600px` を追加してください。

# --hints--

`height` プロパティを `600px` に設定する必要があります。

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

`.canvas` 要素の `height` が `600px` である必要があります。

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
