---
id: 60a3e3396c7b40068ad6996f
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`body` 内に `div` 要素を追加してください。

`class` 属性を `canvas` に設定してください。 例えば `<div class="my-div">` のようにして設定します。

この要素が、あなたの絵画のキャンバスとなります。

# --hints--

コードには `div` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('div').length === 1)
```

`div` 要素の `class` を `canvas` に設定する必要があります。

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
