---
id: 60a3e3396c7b40068ad6996c
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

CSS ボックスモデルでは、すべての HTML 要素は 4 つの領域を持つボックスとして扱われます。

あなたのお気に入りのオンラインショップから荷物を受け取ったと想像してみてください。右図の content (コンテンツ) は箱の中の商品にあたります。HTML においては、ヘッダー、段落、画像などの要素です。

では、`<img>` の `src` 属性を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` から `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png` に変更してください。

# --hints--

`img` 要素の `src` を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png` に設定する必要があります。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

