---
id: 60a3e3396c7b40068ad69980
title: ステップ 23
challengeType: 0
dashedName: step-23
---

# --description--

`one` の要素の下に別の `div` を追加して、`class` の値を `two` に設定してください。 これが 2 つ目の長方形になります。

# --hints--

既存の `.one` 要素を変更しないでください。

```js
assert(document.querySelectorAll('.one').length === 1);
```

`.canvas` 要素の中に 2 つ目の `div` 要素をネストする必要があります。

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

2 つ目の `div` 要素の `class` の値を `two` に設定する必要があります。

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

`.two` の要素は `.one` の要素の後に置く必要があります。

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
