---
id: 60a3e3396c7b40068ad69970
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

追加した `div` のスタイル設定を始める前に、CSS を HTML にリンクする必要があります。

`styles.css` ファイルにリンクする `link` 要素を追加してください。 `href` を `styles.css` に設定し、`rel` 属性も忘れずに `stylesheet` に設定してください。

# --hints--

コードには `link` 要素が 1 つ必要です。

```js
assert(/<link/.test(code))
```

`link` 要素は値が `stylesheet` に設定されている `rel` 属性を持つ必要があります。

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
