---
id: 60a3e3396c7b40068ad69974
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

絵画には額縁が必要です。

`.canvas` 要素を別の `div` で囲んでください。 その `div` に `frame` というクラスを設定してください。
# --hints--

新しい `div` 要素を 1 つ追加してください。

```js
assert(document.querySelectorAll('div').length === 2)
```

`.canvas` 要素は新しい `div` 要素の中にネストされる必要があります。

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

新しい `div` には `frame` という値の `class` が必要です。

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

新しい `div` は `body` 要素内にある必要があります。

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
