---
id: 60a3e3396c7b40068ad6996b
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`head` 要素内に、`charset` を `UTF-8` に設定する `meta` タグと、`Rothko Painting` というタイトルの `title` 要素を追加しましょう。

`body` 要素内には、`src` を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` に設定した `img` 要素を追加してください。
# --hints--

コードには `meta` タグが 1 つ必要です。

```js
assert(document.querySelectorAll('meta').length === 1);
```

`meta` タグは `charset` を `UTF-8` に設定されている必要があります。

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

コードには `title` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('title').length === 1);
```

`title` は、`Rothko Painting` でなければなりません。

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

コードには `img` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 要素の `src` を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` に設定する必要があります。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
