---
id: 60a3e3396c7b40068ad6997e
title: ステップ 21
challengeType: 0
dashedName: step-21
---

# --description--

`.one` を水平方向の中央に配置できましたが、上のマージンがキャンバスを越えて額縁の境界線を押しやって、キャンバス全体を 20 ピクセル下に移動させています。

`.one` 要素が押しやることのできる実体を追加するため、`.canvas` 要素に `1px` の `padding` を追加しましょう。

# --hints--

`padding` プロパティを `1px` に設定する必要があります。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '1px');
assert(hasPadding);
```

`.canvas` 要素の `padding` の値が `1px` である必要があります。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('padding');
assert(canvasPadding === '1px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--

--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
