---
id: 60a3e3396c7b40068ad69973
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

キャンバスの `background-color` を `#4d0f00` に変更してください。

# --hints--

`background-color` プロパティを `#4d0f00` に設定する必要があります。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

`.canvas` 要素の `background-color` は `#4d0f00` である必要があります。

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
