---
id: 60a3e3396c7b40068ad69975
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

`.frame` というクラスセレクターを使用して、新しいルールを書きましょう。

`border` の一括指定を使用して、`.frame` 要素に、実線 (solid)、黒色 (black)、幅 `50px` の境界線を設定してください。

# --hints--

コード内に `.frame` セレクターが必要です。

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

`border` プロパティを `50px solid black` に設定する必要があります。

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

`.frame` 要素は `50px solid black` の `border` を持つ必要があります。

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
