---
id: 60a3e3396c7b40068ad69986
title: ステップ 28
challengeType: 0
dashedName: step-28
---

# --description--

`.two` の要素の下に新しい `div` を追加して、`class` の値を `three` に設定してください。 これが 3 つ目の長方形になります。

# --hints--

既存の `.one` と `.two` の要素は変更しないでください。

```js
assert(document.querySelectorAll('.one').length === 1);
assert(document.querySelectorAll('.two').length === 1);
```

新しい `div` は `.canvas` 要素内にネストされる必要があります。

```js
assert(document.querySelector('.canvas').children[2].tagName === 'DIV');
```

新しい `div` は `.two` 要素の後に置く必要があります。

```js
assert(document.querySelector('.two').nextElementSibling.tagName === 'DIV');
```

新しい `div` には `three` という値の `class` が必要です。

```js
assert(document.querySelector('.canvas').children[2].className.split(' ').includes('three'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
