---
id: 60a3e3396c7b40068ad6996e
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

マージン (margin) は箱の外側の領域で、他の箱、つまり他の要素との間隔を制御するために使えます。

右の図では下にある要素の上部マージンが大きく設定されているため、ページの下の方へ押し下げられています。

CSS ボックスモデルが理解できたところで、ロスコ風の絵に取り掛かりましょう。

`<img>` 要素を削除してください。

# --hints--

コード内に `img` 要素が無いようにしてください。

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
