---
id: 60a3e3396c7b40068ad6997a
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

`.one` を選択する新しいルールを作成し、その `width` を 425 ピクセルに設定してください。

# --hints--

`.one` セレクターが必要です。

```js
const hasOne = new __helpers.CSSHelp(document).getStyle('.one');
assert(hasOne);
```

`width` プロパティを `425px` に設定する必要があります。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '425px');
assert(hasWidth);
```

`.one` 要素の `width` の値が `425px` である必要があります。

```js
const oneWidth = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('width');
assert(oneWidth === '425px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
