---
id: 60a3e3396c7b40068ad6998d
title: ステップ 35
challengeType: 0
dashedName: step-35
---

# --description--

`.one` と `.two` の両方を選択するルールを作成し、さらに 1 ピクセルの `blur` 効果をかけてください。

# --hints--

`.one, .two` のセレクターリストが必要です。

```js
const oneTwo = new __helpers.CSSHelp(document).getStyle('.one, .two');
assert(oneTwo);
```

`filter` プロパティを `blur(1px)` に設定する必要があります。

```js
const hasFilter = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.filter === 'blur(1px)');
assert(hasFilter)
```

`.one` 要素の `filter` の値が `blur(1px)` である必要があります。

```js
const one = document.querySelector('.one');
const oneFilter = getComputedStyle(one).filter;
assert(oneFilter === 'blur(1px)');
```

`.two` 要素の `filter` の値が `blur(1px)` である必要があります。

```js
const two = document.querySelector('.two');
const twoFilter = getComputedStyle(two).filter;
assert(twoFilter === 'blur(1px)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
  filter: blur(2px);
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
}

--fcc-editable-region--

--fcc-editable-region--

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
