---
id: 60a3e3396c7b40068ad6996a
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

ここまで来れば、HTML ページが持つべき基本的な要素には見覚えがあるかと思います。

`DOCTYPE` 宣言、`html` 要素、`head` 要素、`body` 要素でコードを書き始めましょう。

# --hints--

コードには `<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

コードには `html` 要素が 1 つ必要です。

```js
assert(document.querySelectorAll('html').length === 1);
```

コードの `html` 要素内に `head` 要素が必要です。

```js
assert(document.querySelectorAll('head').length === 1);
```

コードの `html` 要素内に `body` 要素が必要です。

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 要素は `body` 要素の前に置く必要があります。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
