---
id: 60a3e3396c7b40068ad69981
title: ステップ 24
challengeType: 0
dashedName: step-24
---

# --description--

`.two` のセレクターを使って新しい CSS ルールを作成し、その `width` を 475 ピクセルに設定してください。

# --hints--

`.two` セレクターが必要です。

```js
const hasTwo = new __helpers.CSSHelp(document).getStyle('.two');
assert(hasTwo);
```

`width` プロパティを `475px` に設定する必要があります。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '475px');
assert(hasWidth);
```

`.two` 要素の `width` の値が `475px` である必要があります。

```js
const twoWidth = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('width');
assert(twoWidth === '475px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
      </div>
    </div>
  </body>
</html>
```
