---
id: 60a3e3396c7b40068ad69979
title: ステップ 16
challengeType: 0
dashedName: step-16
---

# --description--

`.canvas` 要素内に新しい `div` 要素を追加してください。

新しい `div` の `class` 属性には `one` という値を設定してください。 これが 1 つ目の長方形になります。

# --hints--

新しい `div` 要素を作成してください。

```js
assert(document.querySelectorAll('div').length === 3);
```

`.canvas` 要素の中に新しい `div` 要素をネストする必要があります。

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

新しい `div` には `class` 属性と値 `one` が必要です。

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
