---
id: 60a3e3396c7b40068ad6997d
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

マージンを利用して、`.one` 要素のキャンバス上の位置を調整しましょう。

`margin` 一括指定プロパティを追加して、垂直方向のマージンを `20px`、水平方向のマージンを `auto` に設定してください。

# --hints--

`margin` プロパティを `20px auto` に設定する必要があります。

```js
const marginFilter = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.style.margin === '20px auto');
assert(marginFilter.length === 2);
```

`.one` 要素の `margin` の値が `20px auto` である必要があります。

```js
const oneMargin = new __helpers.CSSHelp(document).getStyle('.one')?.getPropertyValue('margin');
assert(oneMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
