---
id: 60a3e3396c7b40068ad6998b
title: ステップ 33
challengeType: 0
dashedName: step-33
---

# --description--

マージンが押しやる方向を一方向にしておくと便利です。

ここでは、`.one` 要素の下部マージンが `.two` を下方向に 20 ピクセル押しやっています。

`.two` セレクターの中で、`margin` 一括指定プロパティを使用して上部マージンを `0` に、水平方向のマージンを `auto` に、下部マージンを `20px` に設定してください。 これにより上部マージンが取り除かれ、水平方向の中央に配置され、下部マージンが 20 ピクセルに設定されます。

# --hints--

`margin` プロパティを `0 auto 20px` に設定する必要があります。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '0px auto 20px');
assert(hasMargin);
```

`.two` 要素の `margin` の値が `0 auto 20px` である必要があります。

```js
const twoMargin = new __helpers.CSSHelp(document).getStyle('.two')?.getPropertyValue('margin');
assert(twoMargin === '0px auto 20px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
