---
id: 60a3e3396c7b40068ad69977
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

要素内の余白を調整するには、パディングを使用します。

`.frame` について、`padding` 一括指定プロパティを使用して `.frame` 要素と `.canvas` 要素の間の余白を `50px` に増やしてください。 一括指定により、要素の境界線と内側のキャンバスの間隔が上下左右すべて広がります。

# --hints--

`padding` プロパティを `50px` に設定する必要があります。

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

`.frame` 要素には値が `50px` の `padding` が必要です。

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
