---
id: 60a3e3396c7b40068ad6996a
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

ここまで来れば、HTML ページが持つべき基本的な要素には見覚えがあるかと思います。

`DOCTYPE` 宣言、言語を英語に設定した `html` 要素、`head` 要素、および `body` 要素をコードに準備してください。

# --hints--

コードには `<!DOCTYPE html>` 宣言が必要です。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

`lang` 属性を `en` にした開始タグ `<html>` が必要です。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

コードの `html` 要素内に `head` 要素が必要です。

```js
assert(document.querySelectorAll('head').length === 1);
```

コードの `html` 要素内に `body` 要素が必要です。

```js
assert(document.querySelectorAll('body').length === 1);
```

`head` 要素は `body` 要素の前に置く必要があります。

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
