---
id: 60a3e3396c7b40068ad6997f
title: ステップ 22
challengeType: 0
dashedName: step-22
---

# --description--

キャンバスの上下左右に 1 ピクセルのパディングを追加したことで、寸法が 502 ピクセル × 602 ピクセルに変わりました。

`padding` プロパティを `overflow` プロパティに置き換えて `hidden` に設定し、キャンバスの寸法を元に戻しましょう。

# --hints--

`.canvas` セレクターから `padding` プロパティを削除する必要があります。

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

`overflow` プロパティを `hidden` に設定する必要があります。

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

`.canvas` 要素の `overflow` の値が `hidden` である必要があります。

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
