---
id: 60a3e3396c7b40068ad6996d
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

コンテンツは、パディング (padding) と呼ばれる余白で囲まれています。これは商品と周りの箱の間に入っている緩衝材に似ています。

境界線 (border) は、商品が送られてきた段ボール箱のようなものだと考えてください。

では、`src` 属性を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png` に変更してください。

# --hints--

`img` 要素の `src` を `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png` に設定する必要があります。

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
