---
id: 60a3e3396c7b40068ad69971
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

では CSS を書きましょう。

`<div>` にはテキストがありませんが、それでもコンテンツの入ったボックスとして扱われます。 `.canvas` クラスセレクターを使って CSS ルールを作成し、その `width` を 500 ピクセルに設定してください。 下記は、`card` クラスの幅を 300 ピクセルに設定する CSS ルールの例です。

```css
.card {
  width: 300px;
}
```

# --hints--

コード内に `.canvas` セレクターが必要です。

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

`width` プロパティを `500px` に設定する必要があります。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

`.canvas` 要素の `width` が `500px` である必要があります。

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
