---
id: 657df49b46e49a4f73aff391
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, you're going to build a simple yet functional expense tracker in Python, focusing on lambda functions.

A list is a collection of items that can be modified. It's like a shopping list where you can add, remove, or change the items. In Python, you create a list by putting items inside square brackets `[]`, with each item separated from the following one by a comma. `[]` signifies that the list is empty. In the code below, you can see an example of a list.

```py
names = []  # Empty list
numbers = [1, 2, 3, 4]
```

For your expense tracker, you need a list to store each expense. Each expense will be a piece of data that we can later manipulate.

Create an empty list named `expense`. 

# --hints--

You should declare a new variable named `expense`.

```js
({ test: () => assert(__userGlobals.has("expense")) })
```

You should assign an empty list to the `expense` variable.

```js
({ test: () => assert.equal(__userGlobals.get("expense").length, 0) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
