---
id: 657df49b46e49a4f73aff391
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

In this project, we're going to build a simple yet functional expense tracker in Python, focusing on lambda functions. Let's start by understanding what a list is in Python.

A list is a collection of items that can be modified. It's like a shopping list where you can add, remove, or change the items. In Python, you create a list by putting items inside square brackets `[]`, separated by commas. `[]` signifies that the list is empty. In the code below, you can see an example of a list.

```py
list_name = []
```

For our expense tracker, we need a list to store each expense. Each expense will be a piece of data that we can later manipulate. Let's create our first list.

Create an empty list called `expense`. 

# --hints--

You should declare a new list named `expense`.

```js
({ test: () => assert(__userGlobals.has("expense")) })
```

The contents of the `expense` list should be empty

```js
({ test: () => assert.equal(__userGlobals.get("expense").length, 0) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
