---
id: 657eb8e0d19d7d50509c1723
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

To add an expense, we use an `if` statement to check the user's choice. If the choice is '1', it means the user wants to add an expense.

Still in the `while` loop, under the `choice` input, write an `if` statement to check if choice equals '1'. If it's true, it will be the starting point for adding a new expense. The `if` statement should contain a `pass` statement for now.

# --hints--

1 should be a string('1'), not an integer (1).

```js
({
    test: () => assert.match(code, /while\s+True:.*choice\s*=\s*input\(\s*['"]Enter\s+your\s+choice:\s*['"]\s*\).*if\s+choice\s*==\s*['"]1['"]\s*:\s*pass/s,)
})

```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        
--fcc-editable-region--
```
