---
id: 657ebb82de7b425221a30305
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

Once the user chooses to add an expense, you need to ask for the expense's amount and category. You will use the input() function for this and convert the amount to a float. This is how you can convert the input to a float:

```py
flat_number = float(input('Enter your number: '))
```

Inside the if statement, ask the user to enter the amount and category of the expense. Store the amount in a variable named `amount` and in the `input()`, use the prompt "Enter amount: ". Also store the category as a string in a variable named `category` and in the `input()`, use the prompt "Enter category: ". Remember to convert the amount to a float.


# --hints--

Use the `input()` function and assign it to a variable.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/amount\s*=\s*float\(input\(['"]Enter\s+amount:\s*['"]\)\)/ms) &&
           function_body.match(/category\s*=\s*input\(['"]Enter\s+category:\s*['"]\)/ms),);
  }
})

```
# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            
--fcc-editable-region--
```
