---
id: 657e895c6ab1b62db4b3438e
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Now, let's enhance the `total_expenses` function with a lambda function. Lambda functions are brief, anonymous functions in Python, ideal for simple, one-time tasks. It can be used with `sum` and `map` to calculate the total expenses:

- Lambda Function: A concise, anonymous function, e.g `lambda x: x * 2` is a lambda function that doubles its input.
- Map Function: Applies a function to each item in a list. Example: map(lambda x: x * 2, [1, 2, 3]) doubles each number.
- Sum Function: Calculates the total of a list of numbers.

For example, consider a list that consists of dictionaries which represents employees names and worked hours:

```py
employees = [
    {"name": "Alice", "hours": 40},
    {"name": "Bob", "hours": 35},
]
```

A lambda function can be used to extract a specific value (like 'hours') and calculate the sum:

```py
total_hours = sum(map(lambda employee: employee['hours'], employees))
print(total_hours)  # Outputs total hours worked
```

Now, inside the `total_expenses` function, use the sum and map functions along with a lambda function to calculate and return the total amount of expenses. The lambda function should extract the amount from each expense in the expenses list.

# --hints--

Make sure the lambda function is written properly, look at the example above

```js
({ 
  test: () => assert.match(code, /lambda\s+expense\s*:\s*expense\['amount'\]/, "Ensure the lambda function is correctly written to extract the 'amount' from each expense.") 
})

```

Make sure the sum function is used

```js
({ 
  test: () => assert.match(code, /sum\(/, "Make sure the 'sum' function is used to calculate the total of the amounts.") 
})

```

Make sure the map function is used

```js
({ 
  test: () => assert.match(code, /map\(/, "Ensure the 'map' function is used to apply the lambda function to each element in the 'expenses' list.") 
})

```

Make sure the `expenses` list is present in the map function

```js
({ 
  test: () => assert.match(code, /sum\(map\(lambda\s+expense\s*:\s*expense\['amount'\]\s*,\s*expenses\)\)/, "Verify that 'expenses' is used correctly in the 'sum(map(lambda...))' expression.") 
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    
--fcc-editable-region--
```
