---
id: 657eea43f3853d6be988b939
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--
To provide a way for users to exit the program, we'll use an elif statement in our main function. This statement will check if the user's choice is '5', which we'll designate as the option to exit the program.

Add an `elif` statement to your main function to handle the exit option. The statement should check if choice equals '5'. The content of the `elif` statement can be a `pass` statement for now.

# --hints--

Ensure there is an `elif` statement checking if `choice` equals '5'.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/elif\s+choice\s*==\s*'5'/ms),);
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []  # Create a new expenses list for each instance
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            add_expense(expenses, amount, category)

        elif choice == '2':
            print("\nAll Expenses:")
            print_expenses(expenses)
            
        elif choice == '3':
            print("\nTotal Expenses: ", total_expenses(expenses))

        elif choice == '4':
            category = input("Enter category to filter: ")
            print(f"\nExpenses for {category}:")
            expenses_from_category = filter_expenses_by_category(expenses, category)
            print_expenses(list(expenses_from_category))

        
--fcc-editable-region--
```
