---
id: 657eae63a49c7f4a3c82896e
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Now, inside the while loop, use the input() function to get the user's choice. This allows the user to interact with the menu and select an option.

For example, suppose you're creating a simple program that asks the user for their name and then stores their name in a variable. Here's how you could do it:

```py
user_name = input("Please enter your name: ")
```

In this example, when the program runs, it will display the message "Please enter your name: ". The user can then type their name and press enter. The entered name gets stored in the variable `user_name`

Now, apply this concept to the expense tracker program. Inside the `while True: loop`, after printing the menu options, call `input()` to ask the user to enter their choice and store this input in a variable named `choice`. The prompt in the `input()` should be "Enter your choice: "

# --hints--

Ensure the `input()` function is correctly used to get the user's choice and store it in a variable named `choice`, placed after the print statements.

```js
({
    test: () => assert.match(code, /choice\s*=\s*input\(\s*['"]Enter\s+your\s+choice:\s*['"]\s*\)/, "Ensure the `input()` function is used with the prompt 'Enter your choice: ' in the code.")
})

```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        
--fcc-editable-region--
```
