---
id: 657eca6e8965365ab6e57a79
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

In the `elif` block for choice '3', we'll call the `total_expenses` function to calculate the sum of all expenses and print the result.

Now inside your `elif` statement, print the total expenses by calling the `total_expenses` function with `expenses` as its argument and using the `print()` function to print the result. A sample syntax for this is:

```py
print('\nTotal Employees: ', total_employees(employees))
```

Use a similar syntax for calling the `total_expenses` function. 

# --hints--

Ensure it is similar to the syntax specified above.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/print\("\\nTotal Expenses: ",\s*total_expenses\(expenses\)\)/ms),);
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []  # Create a new expenses list for each instance
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            add_expense(expenses, amount, category)

        elif choice == '2':
            print("\nAll Expenses:")
            print_expenses(expenses)
            
        elif choice == '3':

--fcc-editable-region--
```
