---
id: 657e78c3451d69249549ad52
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

In this step, you will define a new function called `print_expenses`. This function will later be used to display each expense in our list.

Start by defining a function named `print_expenses` that takes one parameter: `expenses`. At this stage, you can fill the body of the function with a `pass` statement.

# --hints--


The function name should be `print_expenses`

```js
({ test: () => assert.match(code, /def\s+print_expenses\s*\(/) })

```

The `print_expenses` function should have a parameter named `expenses`

```js
({ test: () => assert.match(code, /def\s+print_expenses\s*\(\s*expenses\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})


--fcc-editable-region--
```
