---
id: 657e967c93048c3c7c0ce143
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

In this step, you will create a function called `filter_expenses_by_category`. This function will be enhanced to filter expenses by category using the filter function.

Define a function named `filter_expenses_by_category` that takes two parameters: `expenses` and `category`. The function body can be a pass statement for now.



# --hints--


Define a function named `filter_expenses_by_category` with two parameters `expenses` and `category`.

```js
({ test: () => assert.match(code, /def\s+filter_expenses_by_category\s*\(\s*expenses\s*,\s*category\s*\)\s*:/,) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))


--fcc-editable-region--
```
