---
id: 657e8090acd31028fc8b8751
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

In this step, you will use the `print()` function and F-strings to display the details of each expense. F-strings in Python allow you to include variables directly in strings for easy formatting.

For instance, with a dictionary like this:

```py
person_profile = {"name": "First Last", "age": 28,}
```

You can use an f-string to format the output:

```py
print(f"Name: {person_profile['name']}, Age: {person_profile['age']}")
```

This displays: `Name: First Last, Age: 28`

Now, inside the for loop in the `print_expenses` function, add a print statement that uses an f-string to display each expense's amount and category. Format the output like this: 
`Amount: 10.0, Category: Transport`

# --hints--

Use an f-string in the `print` statement to display each expense's `amount` and `category`.

```js
({ test: () => assert.match(code, /print\(f"Amount:\s*\{\s*expense\['amount'\]\s*\},\s*Category:\s*\{\s*expense\['category'\]\s*\}"\)/,) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        
--fcc-editable-region--
```
