---
id: 657e8565afb64e2b8efb97c3
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

In this step, a function called `total_expenses` will be created and it will be used to calculate the total amount of all expenses. This function will be enhanced with lambda functions in the next step.

Define a function named `total_expenses` that takes one parameter: `expenses`. The function body should contain a `pass` statement for now.


# --hints--


Define a function named `total_expenses` with one parameter `expenses`.

```js
({ test: () => assert.match(code, /def\s+total_expenses\s*\(\s*expenses\s*\)\s*:/, ) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")


--fcc-editable-region--
```
