---
id: 657eab36af77b0482d2a8cd6
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Within the while loop, a couple of print statements should be added to display the menu options for the expense tracker. This step makes the program user-friendly by providing clear choices to the user.

In the code below, the first 3 print statements has been added. The first statement shows the title of the program (Expense Tracker). The `\n` character in the first print statement adds a newline, this creates a space above the menu for better readability in the output. The second print statement shows the option for adding an expense, and the third print statement shows the option to list all the expenses.

Complete the code by adding the remaining 3 options using a similar print statement. Add these three options:

- 3. Show total expenses
- 4. Filter expenses by category
- 5. Exit

# --hints--


Include the remaining three menu options: '3. Show total expenses', '4. Filter expenses by category', and '5. Exit' in the correct order and format using `print()` statements.

```js
({ 
  test: () => assert.match(code, /def\s+main\s*\(\)\s*:\s*[\s\S]*expenses\s*=\s*\[\]\s*[\s\S]*while\s+True\s*:\s*[\s\S]*print\("\\nExpense Tracker"[\s\S]*print\("1. Add an expense"[\s\S]*print\("2. List all expenses"[\s\S]*print\("3. Show total expenses"[\s\S]*print\("4. Filter expenses by category"[\s\S]*print\("5. Exit"\)/,) 
})


```
# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []  # Create a new expenses list for each instance
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")

--fcc-editable-region--
```
