---
id: 657e6231622cd919e2d7094b
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now it's time to be able to append each expense as a dictionary to the `expenses` list. This illustrates how to create functions that manipulate lists by adding new items.


To append a value to a list, the `append` methos is used. This is the syntax used to append a value to a list:
```python
def function_name(parameter_name):
    list_name = []
    list_name.append(parameter_name)
```

But to append a dictionary to a list, this syntax is used:
```python
def function_name(parameter_name):
    list_name = []
    list_name.append({"sample_key": parameter_name})
```

Your task is to give the `add_expense` function the functionality to append a new expense in form of a dictionary to the `expense` list. The dictionary should have two keys, `"amount"` and `"category"`. The values of the keys should be the values of the arguments passed into the `add_expense` function.



# --hints--


The `add_expense` should use the `append` method to append to the `expense` list
```js
({
  test: () => assert.match(code, /expenses\.append\(/)
})

```
 
A dictionary containing `amount` and `category` keys should be appended

```js
({
  test: () => assert.match(code, /expenses\.append\(\{\s*"amount"\s*:\s*amount\s*,\s*"category"\s*:\s*category\s*\}\)/, "The line 'expenses.append({\"amount\": amount, \"category\": category})' should be present in the code.")
})

```
# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):

--fcc-editable-region--
```
