---
id: 657ec30272cf4255edaccf32
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

To list all expenses, we use an `elif` statement to check if the user's choice is '2'. This is how we provide different functionalities based on user input.

This is the syntax of an `elif` statement:

```py
elif variable_name == 'value':
    pass
```

Now write an `elif` statement to check if `choice` equals '2'. If it's true, it indicates the user wants to list all expenses. For now, add a pass statement under the `elif` statement.


# --hints--



Use the `elif` syntax correctly, and make sure you add a `pass` statement under it.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/elif\s+choice\s*==\s*['"]2['"]\s*:/ms),);
  }
})


```
# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            add_expense(expenses, amount, category)

        
--fcc-editable-region--
```
