---
id: 657ea3f2e2e649443416d863
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

The next step is to create the main function, which will be the entry point of the interactive expense tracker program. Inside this function, an empty list named `expenses` will be initialized and it will be used to store the expense records.

In the code, define a function named `main`, the function should have no parameter. Inside the function, create an empty list named `expenses`


# --hints--


Define the `main` function with correct syntax.

```js
({ test: () => assert.match(code, /def\s+main\s*\(\)\s*:/,) })

```

Create an expenses list inside the `main` function.

```js
({ test: () => assert.match(code, /def\s+main\s*\(\)\s*:\s*[\s\S]*?expenses\s*=/,) })

```

Initialize the `expenses` list as an empty list.

```js
({ test: () => assert.match(code, /def\s+main\s*\(\)\s*:\s*[\s\S]*?expenses\s*=\s*\[\]/,) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


--fcc-editable-region--
```
