---
id: 657ed635f40df661219dddbe
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

Inside this `elif` block, first ask the user to enter the category they want to filter by using the `input()` function.

Write the code using the `input()` function to ask a user to input their category and store that value in a variable named `category`. Use this prompt for the `input()` function: `Enter category to filter: `

# --hints--

Make sure input() is used.

```js
({
    test: () => assert.match(code, /elif\s+choice\s*==\s*['"]4['"]\s*:\s*category\s*=\s*input\(['"]Enter category to filter: ['"]\)/,)
})

```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []  # Create a new expenses list for each instance
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            add_expense(expenses, amount, category)

        elif choice == '2':
            print("\nAll Expenses:")
            print_expenses(expenses)
            
        elif choice == '3':
            print("\nTotal Expenses: ", total_expenses(expenses))

--fcc-editable-region--
        elif choice == '4':

--fcc-editable-region--
```
