---
id: 657ea808f82a1946e08f22b8
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Inside the main function, after the `expenses` list, a `while True: ` loop will be set up. This will create an infinite loop, allowing the program to continuously present menu options until the user decides to exit.

This is the syntax for implementing a `while True: ` loop:

```py
while True:
    # operation goes here
```

Now, implement a `while True:` loop inside the main function, under the     `expenses` list. The content of this loop can be a pass statement for now.

# --hints--


Implement an infinite `while` loop inside the `main` function. Make sure to add correct indentation, and add a pass statement inside the loop.

```js
({ 
  test: () => assert.match(code, /def\s+main\s*\(\)\s*:\s*[\s\S]*expenses\s*=\s*\[\]\s*[\s\S]*while\s+True\s*:/,) 
})

```
# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
    
--fcc-editable-region--
```
