---
id: 657e7d9701d40d2792aac88d
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

In this step, we'll enhance the `print_expenses` function by adding a for loop. The for loop in Python is used to iterate over each item in a collection, such as a list. We'll use it to go through each expense in our expenses list.

This is the syntax for writing a basic for loop:

```py
for variable in iterable:
     # Code to execute in each iteration
```

In the code above: 
- `variable` is the loop variable that takes the value of each element in the iterable during each iteration of the loop.
- `iterable` is a collection of items (like a list, tuple, set, or any other iterable object) that the loop will iterate over.
- The indented block under the for loop (# Code to execute in each iteration) is the code that will be executed for each item in the iterable.

Now, inside the `print_expenses` function, add a for loop that iterates over each item in the list named `expenses` (the iterable). Each item during the iteration should be referred to as `expense` (the variable). At this stage, you can fill the body of the loop with a `pass` statement.

# --hints--


You should define a for loop with a loop variable named `expense` that iterates over an iterable named `expenses`.

```js
({ test: () => assert.match(code, /for\s+expense\s+in\s+expenses\s*:/, ) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    

--fcc-editable-region--
```
