---
id: 657e984243d8f13de95792a6
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

In this step, we'll use the `filter()` function, a built-in Python function that allows you to create a new iterator from elements of an existing iterable (like a list) for which a specified function returns `True`. An iterator is a special object that enables traversal through the elements of a collection, like a list, based on a condition.

For example, if we have a list of dictionaries representing books, and we want to filter out books of a particular genre, we can use `filter()` with a lambda function, like this:

```py
books = [
    {"title": "Book A", "genre": "Fiction"},
    {"title": "Book B", "genre": "Non-Fiction"},
    {"title": "Book C", "genre": "Fiction"}
]

def fiction_books(book_genre):
    return filter(lambda book: book['genre'] == 'book_genre', books)
```

Now, this same concept can be applied in this expense tracker project. Under the `filter_expenses_by_category` function, use `filter()` with a lambda function to return an iterator of expenses where the category matches the specified category that was passed into the function as a parameter.

# --hints--

Ensure the lambda function is correctly written to compare the 'category' of each expense. Look at the example above.

```js
({ 
  test: () => assert.match(code, /lambda\s+expense\s*:\s*expense\['category'\]\s*==\s*category/,) 
})

```

Verify that the 'filter' function is used with a lambda function.

```js
({ 
  test: () => assert.match(code, /filter\(\s*lambda/,) 
})

```

Ensure that the `expenses` list is correctly passed as an argument to the 'filter' function.

```js
({ 
  test: () => assert.match(code, /filter\(\s*.*,\s*expenses\s*\)/,) 
})

```

Make sure the function returns the result of the 'filter(lambda...)' expression correctly.

```js
({
  test: () => assert.match(code, /return\s+filter\(lambda\s+expense\s*:\s*expense\['category'\]\s*==\s*category,\s*expenses\)/, "Make sure the function returns the result of the 'filter(lambda...)' expression correctly.")
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    
--fcc-editable-region--
```
