---
id: 657e7d9701d40d2792aac88d
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

You can enhance the `print_expenses` function using a for loop. In Python, a for loop iterates over items in a collection, like a list. Use it to go through and process each expense in the expenses list.

This is the syntax for writing a basic for loop:

```py
for item in collection:
    # Actions for each item
```

In the code above:

- `item` represents each element from `collection` during the loop.
- `collection` is the list or other iterable you're looping over.
- The indented code under the for loop is executed for each item.

Now, inside the `print_expenses` function, add a for loop that iterates over each item in the list named `expenses`. Each element during the iteration should be referred to as `expense`. At this stage, you can fill the body of the loop with a `pass` statement.

# --hints--

You should define a for loop with a loop variable named `expense` that iterates over an iterable named `expenses`.

```js
({ test: () => assert.match(code, /for\s+expense\s+in\s+expenses\s*:/, ) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    

--fcc-editable-region--
```
