---
id: 657e895c6ab1b62db4b3438e
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Now, let's enhance the `total_expenses` function using a lambda function. Lambda functions in Python are small anonymous functions defined with the lambda keyword. They are concise and useful for short, one-time operations. We'll use a lambda function with the sum and map functions to calculate the total expenses.

In the context of our total_expenses function, we'll employ a lambda function in combination with the sum and map functions. Here's a brief explanation:

Lambda Function: A lambda function is a small anonymous function. For example, `lambda x: x * 2` is a lambda function that doubles its input.

Map Function: The map function applies a function to all items in a list. For example, `map(lambda x: x * 2, [1, 2, 3])` will apply the lambda function to each element in the list [1, 2, 3].

Sum Function: The sum function computes the sum of a list of numbers.

Consider this simple example, where we use a lambda function to add 10 to each number in a list and then sum them up:

```py
numbers = [1, 2, 3]
print(sum(map(lambda x: x + 10, numbers)))  # Outputs: 36
```

Let's look at an example where a lambda function is used with a dictionary, similar to what we need for the total_expenses function in our expense tracker project.
For example, suppose you have a list of employees, where each employee is represented as a dictionary with their name and the number of hours they worked:

```py
employees = [
    {"name": "Alice", "hours": 40},
    {"name": "Bob", "hours": 35},
]
```

Let's say you want to calculate the total hours worked by all employees. You can achieve this using a lambda function with map and sum:

```py
total_hours = sum(map(lambda employee: employee['hours'], employees))
print(total_hours)  # Outputs the total hours worked by all employees
```


Now, inside the `total_expenses` function, use the sum and map functions along with a lambda function to calculate and return the total amount of expenses. The lambda function should extract the amount from each expense in the expenses list.


# --hints--


Make sure the lambda function is written properly, look at the example above
```js
({ 
  test: () => assert.match(code, /lambda\s+expense\s*:\s*expense\['amount'\]/, "Ensure the lambda function is correctly written to extract the 'amount' from each expense.") 
})

```

Make sure the sum function is used
```js
({ 
  test: () => assert.match(code, /sum\(/, "Make sure the 'sum' function is used to calculate the total of the amounts.") 
})

```

Make sure the map function is used
```js
({ 
  test: () => assert.match(code, /map\(/, "Ensure the 'map' function is used to apply the lambda function to each element in the 'expenses' list.") 
})

```

Make sure the `expenses` list is present in the map function
```js
({ 
  test: () => assert.match(code, /sum\(map\(lambda\s+expense\s*:\s*expense\['amount'\]\s*,\s*expenses\)\)/, "Verify that 'expenses' is used correctly in the 'sum(map(lambda...))' expression.") 
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    
--fcc-editable-region--
```
