---
id: 657ebe128041f1539d04693f
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

After getting the expense details, we call the `add_expense` function to add this new expense to the `expenses` list.

After getting the `amount` and `category` using `input()`, add a line to call the `add_expense` function, passing three arguments: `expenses`, `amount` and `category`. 

- `expenses` is the empty list created in the main function earlier in this project.
- `amount` is the amount of the expense.
- `category` is the category of the expense.

# --hints--

Ensure the `add_expense` function is called with the `expenses` list, `amount`, and `category` as arguments in the `main` function.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/add_expense\(expenses,\s*amount,\s*category\)/ms),);
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            
--fcc-editable-region--
```
