---
id: 657e6231622cd919e2d7094b
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now it's time to be able to append each expense as a dictionary to the `expenses` list. This illustrates how to create functions that manipulate lists by adding new items.


In Python, the `.append()` method is used to append an item to a list. This is an example:
```python
def function_name(parameter_name):
    list_name = []
    list_name.append(parameter_name)
```

This is an example of how you can append a dictionary to a list:

```python
def function_name(parameter_name):
    list_name = []
    list_name.append({"sample_key": parameter_name})

Your task is to give the `add_expense()` function the functionality to append a new expense dictionary to the `expenses` list. The dictionary should have two keys: `"amount"` and `"category"`. The values of the keys should be the values of the arguments passed into the `add_expense` function.



# --hints--


The `add_expense()` function should call the `.append()` method to append the dictionary to the `expenses` list.
```js
({
  test: () => assert.match(code, /expenses\.append\(/)
})

```
 
You should append a dictionary with the keys `amount` and `category` to the `expenses` list.

```js
({
  test: () => assert.match(code, /expenses\.append\(\{\s*"amount"\s*:\s*amount\s*,\s*"category"\s*:\s*category\s*\}\)/, "The line 'expenses.append({\"amount\": amount, \"category\": category})' should be present in the code.")
})

```
# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):

--fcc-editable-region--
```
