---
id: 657e8090acd31028fc8b8751
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

Now that we have our for loop set up, it's time to add functionality to the `print_expenses` function. In this step, we'll use the print function to display the details of each expense. We will format our output using f-strings, a feature in Python that allows you to interpolate expressions inside string literals.

For example, if you have the variables `name = "Alice"` and `age = 30`, you can use an f-string to interpolate the values of these variables in a message:

```py
print(f"{name} is {age} years old.")
```
This will output: Alice is 30 years old.

To give a clearer idea of how to apply f-strings for dictionaries, look at the example below. 
Suppose you have a dictionary representing a person's profile:

```py
person_profile = {
    "name": "First Last",
    "age": 28,
}
```

You can use an f-string to print out details from this dictionary like this:
```py
print(f"Name: {person_profile['name']}, Age: {person_profile['age']}")
```

This will output: `Name: John Doe, Age: 28`.


Now, inside the for loop in the `print_expenses` function, add a print statement that uses an f-string to display each expense's amount and category. Format the output like this: 
`Amount: 10.0, Category: Transport`

# --hints--


Use an f-string in the `print` statement to display each expense's `amount` and `category`.

```js
({ test: () => assert.match(code, /print\(f"Amount:\s*\{\s*expense\['amount'\]\s*\},\s*Category:\s*\{\s*expense\['category'\]\s*\}"\)/,) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        
--fcc-editable-region--
```
