---
id: 657eebd507b1bc6d492131e8
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

Once the user selects the option to exit, we need to print a farewell message and then break out of the while loop to terminate the program.

Inside the `elif` block you just created, add the following two lines:

1. A print statement to inform the user that the program is exiting. Use this statement: "Exiting the program."
2. The break statement to exit the while loop and therefore end the program. Just add the code `break` under the print statement.

# --hints--

Ensure there is a print statement with the message 'Exiting the program.'

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/print\("Exiting the program."\)/ms),);
  }
})

```

Ensure the `break` statement is present under the print statement in the 'elif' block for exiting the program.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/print\(\s*"Exiting the program."\s*\)\s*;?\s*break\s*;?(\n|\r\n|\r)?/ms),);
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []  # Create a new expenses list for each instance
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            add_expense(expenses, amount, category)

        elif choice == '2':
            print("\nAll Expenses:")
            print_expenses(expenses)
            
        elif choice == '3':
            print("\nTotal Expenses: ", total_expenses(expenses))

        elif choice == '4':
            category = input("Enter category to filter: ")
            print(f"\nExpenses for {category}:")
            expenses_from_category = filter_expenses_by_category(expenses, category)
            print_expenses(list(expenses_from_category))

        elif choice == '5':
            
--fcc-editable-region--
```
