---
id: 657ec4994640ba56e67fc2bc
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

Inside the `elif` block, we'll call the `print_expenses` function to display all the expenses that have been added so far.

This code has already been added: `print("\nAll Expenses:")`
It uses the `\n` character to add a new line for better readability, and it prints "All Expenses" as the title.

Complete the code by calling the `print_expenses` function with the `expenses` list as its argument.

# --hints--

Ensure `expenses` is passed as an argument.

```js
({
  test: () => {
    const main_func = __helpers.python.getDef(e.code.original["main.py"], "main");
    const { function_body } = main_func;
    assert(function_body.match(/print_expenses\(expenses\)/ms),);
  }
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({"amount": amount, "category": category})

def print_expenses(expenses):
    for expense in expenses:
        print(f"Amount: {expense['amount']}, Category: {expense['category']}")

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
    while True:
        print("\nExpense Tracker")
        print("1. Add an expense")
        print("2. List all expenses")
        print("3. Show total expenses")
        print("4. Filter expenses by category")
        print("5. Exit")

        choice = input("Enter your choice: ")

        if choice == '1':
            amount = float(input("Enter amount: "))
            category = input("Enter category: ")
            add_expense(expenses, amount, category)

        elif choice == '2':
            print("\nAll Expenses:")

--fcc-editable-region--
```
