---
id: 61537bb9b1a29430ac15ad38
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Start this project by linking your `styles.css` file to the page.

# --hints--

You should have a `link` element within your `head` element.

```js
const link = document.querySelector("head > link");
assert.isNotNull(link);
```

Your `link` element should have an `href` attribute.

```js
const link = document.querySelector("head > link");
const hasHref = link?.hasAttribute("href");
assert.isTrue(link?.hasAttribute("data-href") || hasHref);
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const link = document.querySelector("head > link");
const href = link?.getAttribute("data-href");
const hasHref = link?.hasAttribute("href");
assert.isTrue(assert.oneOf(href, ["./styles.css", "styles.css"]) || !hasHref);
```

Your `link` element should have a `rel` attribute.

```js
const link = document.querySelector("head > link");
assert.isNotNull(link?.getAttribute("rel"));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const link = document.querySelector("head > link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
