---
id: 66fbcf750a62784cf11f5630
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

A third `recipe3` object has been filled in for you. It has the same properties as `recipe1` and `recipe2`.

You should now push the three objects into the `recipes` array. To do thus, you can use the `push()` method.

Here's the syntax of the `push()` method:

```js
array.push(element1, element2, ..., elementN);
```

And here's an example:

```js
const fruits = ['apple', 'banana'];

fruits.push('orange', 'mango');
console.log(fruits); // ['apple', 'banana', 'orange', 'mango']
```

Use the `push()` method to push all the recipes objects into the `recipes` array.

# --hints--

You should push all the recipes objects into the `recipes` array.

```js
assert.match(code, /recipes\.push\((?:recipe1\s*,\s*recipe2\s*,\s*recipe3|recipe1\s*,\s*recipe3\s*,\s*recipe2|recipe2\s*,\s*recipe1\s*,\s*recipe3|recipe2\s*,\s*recipe3\s*,\s*recipe1|recipe3\s*,\s*recipe1\s*,\s*recipe2|recipe3\s*,\s*recipe2\s*,\s*recipe1)\);?/)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

--fcc-editable-region--

--fcc-editable-region--
```
