---
id: 66fbcf750a62784cf11f5631
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now, you should work on calculating the `averageRating`, `totalIngredients`, and the `difficultyLevel` for each recipe in the `recipes` array.

Start by creating a `getAverageRating` function that takes a single argument, which is an array with ratings. Inside the function, calculate the average rating from the array passed to the function.

Your `getAverageRating` function must return a number.

# --hints--

You should create a `getAverageRating` function.

```js
assert.isFunction(getAverageRating)
```

Your `getAverageRating` function should have a single parameter.

```js
assert.lengthOf(getAverageRating, 1);
```

Your `getAverageRating` function should return a number.

```js
assert.isNumber(getAverageRating(recipe1.ratings))
```

`getAverageRating(recipe1.ratings)` should return `4.5`.

```js
assert.strictEqual(getAverageRating(recipe1.ratings), 4.5)
```

`getAverageRating(recipe2.ratings)` should return `4.75`.

```js
assert.strictEqual(getAverageRating(recipe2.ratings), 4.75)
```

`getAverageRating(recipe3.ratings)` should return `4`.

```js
assert.strictEqual(getAverageRating(recipe3.ratings), 4)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
