---
id: 66fbcf750a62784cf11f5633
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Create `getDificultyLevel` function with `cookingTime` as the parameter. Inside the function. Create the conditional to help you assign an `easy`, `medium`, and `hard` difficulty levels based on the `cookingTime` of each recipe.

If the `cookingTime` is less than `30`, your `getDificultyLevel` function should return `"easy"`, if it is less than or equal to `60`, the function should return `"medium"`, otherwise, the function should return `"hard"`.

# --hints--

You should create a `getDificultyLevel` function.

```js
assert.isFunction(getDificultyLevel)
```

You `getDificultyLevel` function should have a `cookingTime` parameter.

```js
assert.match(getDificultyLevel.toString(), /cookingTime/);
```

Your `getDificultyLevel` function should return `"easy"` when the `cookingTime` is less than or equal to `30`.

```js
assert.strictEqual(getDificultyLevel(10), "easy")
assert.strictEqual(getDificultyLevel(20), "easy")
assert.strictEqual(getDificultyLevel(29), "easy")
assert.strictEqual(getDificultyLevel(30), "easy")
```

Your `getDificultyLevel` function should return `"medium"` when the `cookingTime` is greater than `31` and less than `60`.

```js
assert.strictEqual(getDificultyLevel(31), "medium")
assert.strictEqual(getDificultyLevel(40), "medium")
assert.strictEqual(getDificultyLevel(50), "medium")
assert.strictEqual(getDificultyLevel(60), "medium")
```

Your `getDificultyLevel` function should return `"hard"` when the `cookingTime` is greater than `60`.

```js
assert.strictEqual(getDificultyLevel(61), "hard")
assert.strictEqual(getDificultyLevel(65), "hard")
assert.strictEqual(getDificultyLevel(70), "hard")
assert.strictEqual(getDificultyLevel(75), "hard")
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length.toFixed(2);
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
