---
id: 66fbcf750a62784cf11f5630
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

A third `recipe3` object has been filled in for you. It has the same properties as `recipe1` and `recipe2`.

You should now push the three objects into the `recipes` array. To do this, you can use the `push()` method.

Use the `push()` method to push all the recipe objects into the `recipes` array. Make sure to push `recipe1`, `recipe2`, and `recipe3` in that order.

Also delete the `recipe1Name`, `recipe2Name`, `recipe1CookingTime`, and `recipe2CookingTime` variables, and the `console.log` statements which log those variables.

# --hints--

You should remove the `recipe1Name` variable.

```js
try {
  recipe1Name;
  assert.fail('Variable is defined/declared');
} catch (e) {
  assert.include(e?.message, 'recipe1Name is not defined');
}
```

You should remove the `console.log(recipe1Name);` statement.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

You should remove the `recipe2Name` variable.

```js
try {
  recipe2Name;
  assert.fail('Variable is defined/declared');
} catch (e) {
  assert.include(e?.message, 'recipe2Name is not defined');
}
```

You should remove the `console.log(recipe2Name);` statement.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2Name\s*\)\s*;?/);
```

You should remove the `recipe1CookingTime` variable.

```js
try {
  recipe1CookingTime;
  assert.fail('Variable is defined/declared');
} catch (e) {
  assert.include(e?.message, 'recipe1CookingTime is not defined');
}
```

You should remove the `console.log(recipe1CookingTime);` statement.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1CookingTime\s*\)\s*;?/);
```

You should remove the `recipe2CookingTime` variable.

```js
try {
  recipe2CookingTime;
  assert.fail('Variable is defined/declared');
} catch (e) {
  assert.include(e?.message, 'recipe2CookingTime is not defined');
}
```

You should remove the `console.log(recipe2CookingTime);` statement.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

You should push all the recipes objects into the `recipes` array. Make sure to push them in the order they are declared.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2Name = recipe2.name;
console.log(recipe2Name);

const recipe1CookingTime = recipe1.cookingTime;
console.log(recipe1CookingTime);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);
--fcc-editable-region--
```
