---
id: 66fbcf750a62784cf11f5637
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Now, log the `recipes` array to the console to see all items of the `recipes` array filled up.

With that, your recipes tracker project is complete.

# --hints--

You should log the `recipes` array to the console.

```js
assert.match(code, /console\.log\(recipes\);?/)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length.toFixed(2);
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return 'easy';
  } else if (cookingTime <= 60) {
    return 'medium';
  } else {
    return 'hard';
  }
}

const recipe1AverageRating = getAverageRating(recipe1.ratings);
console.log(`Recipe 1 Average Rating ${recipe1AverageRating}`);

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(`Recipe 1 Total Ingredients: ${recipe1TotalIngredients}`);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(`Recipe 1 Difficulty: ${recipe1DifficultyLevel}`);

for (const recipe of recipes) {
  recipe.averageRating = getAverageRating(recipe.ratings);
  recipe.totalIngredients = getTotalIngredients(recipe.ingredients);
  recipe.difficultyLevel = getDifficultyLevel(recipe.cookingTime);
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return (total / ratings.length).toFixed(2);
}

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return 'easy';
  } else if (cookingTime <= 60) {
    return 'medium';
  } else {
    return 'hard';
  }
}

const recipe1AverageRating = getAverageRating(recipe1.ratings);
console.log(`Recipe 1 Average Rating ${recipe1AverageRating}`);

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(`Recipe 1 Total Ingredients: ${recipe1TotalIngredients}`);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(`Recipe 1 Difficulty: ${recipe1DifficultyLevel}`);

for (const recipe of recipes) {
  recipe.averageRating = getAverageRating(recipe.ratings);
  recipe.totalIngredients = getTotalIngredients(recipe.ingredients);
  recipe.difficultyLevel = getDifficultyLevel(recipe.cookingTime);
}

console.log(recipes);
```
