---
id: 66fbcf750a62784cf11f562f
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Other properties and values the `recipe2` object should have are:

| Key | Value   |
| ----------- | ------- |
| `"cookingTime"`       | `42` |
| `"totalIngredients"`   | `null`   |
| `"difficultyLevel"`   | `''`   |
| `"averageRating"`   | `null`   |


Go ahead and create those.

# --hints--

Your `recipe2` object should have a `"cookingTime"` property.

```js
assert.property(recipe2, 'cookingTime');
```

Your `"cookingTime"` property value should be a number.

```js
assert.isNumber(recipe2['cookingTime']);
```

Your `"cookingTime"` property should be set to `42`.

```js
assert.equal(recipe2['cookingTime'], 42);
```

Your `recipe2` object should have a `"totalIngredients"` property.

```js
assert.property(recipe2, 'totalIngredients');
```

Your `"totalIngredients"` property should be set to `null`.

```js
assert.deepEqual(recipe2['totalIngredients'], null);
```

Your `recipe2` object should have a `"difficultyLevel"` property.

```js
assert.property(recipe2, 'difficultyLevel');
```

Your `"difficultyLevel"` property should be set to an empty string.

```js
assert.deepEqual(recipe2['difficultyLevel'], '');
```

Your `recipe2` object should have an `"averageRating"` property.

```js
assert.property(recipe2, 'averageRating');
```

Your `"averageRating"` property should be set to `null`.

```js
assert.deepEqual(recipe2['averageRating'], null);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  ratings: [4, 5, 5, 5],
--fcc-editable-region--

--fcc-editable-region--
};
```
