---
id: 66fbcf750a62784cf11f562d
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Create an object named `recipe1`. Inside the `recipe1` object, create a `name` property with the value `"Spaghetti Carbonara"`.

Also inside the `recipe1` object, create an `ingredients` property with an array as the value. The array should have `spaghetti`, `Parmesan cheese`, `pancetta`, and `black pepper` inside it.

Create another `ratings` property with array value. The array should have `4`, `5`, `4`, and `5` inside it.

# --hints--

You should create an object named `recipe1`.

```js
assert.isObject(recipe1)
```

Your `recipe1` object should have a `name` property.

```js
assert.property(recipe1, 'name');
```

Your `name` property should be set to `Spaghetti Carbonara`.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

Your `recipe1` object should have an `ingredients` property.

```js
assert.property(recipe1, 'ingredients');
```

Your `ingredients` property should have an array value.

```js
assert.isArray(recipe1.ingredients)
```

The array value of your `ingredients` property should have `spaghetti`, `Parmesan cheese`, `pancetta`, and `black pepper` in it.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

Your `recipe1` object should have a `ratings` property.

```js
assert.property(recipe1, "ratings");
```

Your `ratings` property should have an array value.

```js
assert.isArray(recipe1.ratings)
```

The array value of your `ratings` property should have `4`, `5`, `4`, and `5` in it.

```js
assert.deepEqual(recipe1.ratings, [4, 5, 4, 5]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
