---
id: 66fbcf750a62784cf11f5632
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Create a `getTotalIngredients` function that takes an `ingredients` parameter. This function should return the number of ingredients for each recipe by returning the length of the `ingredients` array.

# --hints--

You should create a `getTotalIngredients` function.

```js
assert.isFunction(getTotalIngredients)
```

You `getTotalIngredients` function should have an `ingredients` parameter.

```js
assert.match(getTotalIngredients.toString(), /ingredients/);
```

Your `getTotalIngredients` function should return a number.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients))
```

Your `getTotalIngredients` function should return a number when the total ingrendients is `5`.

```js
assert.isNumber(getTotalIngredients(recipe2.ingredients))
```

Your `getTotalIngredients` function should return a number when the total ingrendients is `3`.

```js
assert.isNumber(getTotalIngredients(recipe3.ingredients))
```

Your `getTotalIngredients` function should return a number when the total ingrendients is `4`.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients))
```

Your `getTotalIngredients` function should return the correct number of ingredients.

```js
assert.equal(getTotalIngredients(recipe1.ingredients), 4)
assert.equal(getTotalIngredients(recipe2.ingredients), 5)
assert.equal(getTotalIngredients(recipe3.ingredients), 3)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

const recipe3 = {
  name: 'Vegetable Stir Fry',
  ingredients: ['broccoli', 'carrot', 'bell pepper'],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 3, 4, 5],
  averageRating: null,
};

recipes.push(recipe1, recipe2, recipe3);

function getAverageRating(ratings) {
  const total = ratings[0] + ratings[1] + ratings[2] + ratings[3];
  return total / ratings.length.toFixed(2);
}

--fcc-editable-region--

--fcc-editable-region--
```
