---
id: 670e4f45f7116c0f216a5177
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Before you move on, you should practice how to access properties from an object.

You can use either dot (`.`) or bracket (`[]`) notation to do this. Here's an example:

```js
const person = {
  name: 'John',
  age: 30,
  job: 'Software Engineer'
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

Access the `name` properties of both `recipe1` and `recipe2`, and assign them to the variables `recipe1Name` and `recipe2Name`, respectively.

Next, access the `cookingTime` properties of both recipes and assign them to the variables `recipe1CookingTime` and `recipe2CookingTime`, respectively.

Make sure all the variables you created are logged to the console.

# --hints--

You should create a `recipe1Name` variable.

```js
assert.isNotNull(recipe1Name)
```

You should set your `recipe1Name` variable to the accessment of the `name` property of `recipe1`.

```js
assert.match(code, /recipe1\.name|recipe1\[("|'|`)name("|'|`)\]/)
```

You should log `recipe1Name` to the console.

```js
assert.match(code, /console\.log\(recipe1Name\)/)
```

You should create a `recipe2Name` variable.

```js
assert.isNotNull(recipe2Name)
```

You should set your `recipe2Name` variable to the accessment of the `name` property of `recipe2`.

```js
assert.match(code, /recipe2\.name|recipe2\[("|'|`)name("|'|`)\]/)
```

You should log `recipe2Name` to the console.

```js
assert.match(code, /console\.log\(recipe2Name\)/)
```

You should create a `recipe1CookingTime` variable.

```js
assert.isNotNull(recipe1CookingTime)
```

You should set your `recipe1CookingTime` variable to the accessment of the `cookingTime` property of `recipe1`.

```js
assert.match(code, /recipe1\.cookingTime|recipe1\[("|'|`)cookingTime("|'|`)\]/)
```

You should log `recipe1CookingTime` to the console.

```js
assert.match(code, /console\.log\(recipe1CookingTime\)/)
```

You should create a `recipe2CookingTime` variable.

```js
assert.isNotNull(recipe2CookingTime)
```

You should set your `recipe2CookingTime` variable to the accessment of the `cookingTime` property of `recipe2`.

```js
assert.match(code, /recipe2\.cookingTime|recipe2\[("|'|`)cookingTime("|'|`)\]/)
```

You should log `recipe2CookingTime` to the console.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/)
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: 'Spaghetti Carbonara',
  ingredients: ['spaghetti', 'Parmesan cheese', 'pancetta', 'black pepper'],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 4, 5],
  averageRating: null,
};

const recipe2 = {
  name: 'Chicken Curry',
  ingredients: ['chicken breast', 'coconut milk', 'curry powder', 'onion', 'garlic'],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: '',
  ratings: [4, 5, 5, 5],
  averageRating: null,
};

--fcc-editable-region--

--fcc-editable-region--
```
