---
id: cf1111c1c11feddfaeb6bdef
title: Dividir um número por outro com JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cqkbdAr'
forumTopicId: 17566
dashedName: divide-one-number-by-another-with-javascript
---

# --description--

Nós também podemos dividir um número por outro.

JavaScript usa o símbolo `/` para divisão.

**Exemplo**

```js
const myVar = 16 / 2;
```

`myVar` agora possui o valor `8`.
# --instructions--

Altere `0` para que a variável `quotient` seja igual a `2`.

# --hints--

A variável `quotient` deve ser igual a 2.

```js
assert(quotient === 2);
```

Você deve usar o operador `/`.

```js
assert(/\d+\s*\/\s*\d+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'quotient = '+z;})(quotient);
```

## --seed-contents--

```js
const quotient = 66 / 0;
```

# --solutions--

```js
const quotient = 66 / 33;
```
