---
id: 56533eb9ac21ba0edf2244a9
title: Inicializar variáveis com o operador de atribuição
challengeType: 1
videoUrl: 'https://scrimba.com/c/cWJ4Bfb'
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

É comum <dfn>inicializar</dfn> a variável com um valor inicial na mesma linha em que é declarada.

```js
var myVar = 0;
```

Crie uma nova variável chamada `myVar` e atribua o seu valor inicial como `0`.

# --instructions--

Defina uma variável `a` com `var` e a inicialize com o valor de `9`.

# --hints--

Você deve inicializar `a` para o valor de `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(code));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
