---
id: bd7993c9c69feddfaeb7bdef
title: Multiplicar dois decimais com JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/ce2GeHq'
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

Em JavaScript, você também pode realizar cálculos com números decimais, assim como com números inteiros.

Vamos multiplicar dois decimais para obter o produto deles.

# --instructions--

Altere o `0.0` para que o produto seja igual a `5.0`.

# --hints--

A variável `product` deve ser igual `5.0`.

```js
assert(product === 5.0);
```

Você deve usar o operador `*`

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(y){return 'product = '+y;})(product);
```

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
