---
id: 5f33071498eb2472b87ddee4
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

正如你在 Cat Photo App 的最后几个步骤中所学习的那样，开始构建你的网页需要一个基本结构。

添加 `<!DOCTYPE html>` 标记和具有值为 `en` 的 `lang` 属性的 `html` 元素。

# --hints--

需要先声明 `DOCTYPE`。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你应该有一个带有值为 `en` 的 `lang` 属性的 `<html>` 开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你需要一个 `<html>` 结束标签。 别忘了结束标签是在开始的 `<` 后有一个 `/`。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
