---
id: 5f331e55dfab7a896e53c3a1
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` 是网页上几个不可见附加信息之一，但它对搜索引擎和网页如何显示很有用。

在`head`元素中，嵌套一个 `meta` 元素，其属性名为 `charset`，值为`utf-8`，以告诉浏览器如何对页面进行编码。

# --hints--

你应该有一个 `meta` 标签。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 元素是一个空元素，它不应该有一个结束标签 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 标签应该有一个 `charset` 属性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 属性的值应该为 `utf-8`。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
