---
id: 5f33071498eb2472b87ddee4
title: 第1步：
challengeType: 0
dashedName: step-1
---

# --description--

之前在Cat Photo App的最后几个步骤中提过，构建网页前需要一个基本的结构。

添加标签 `<!DOCTYPE html>` 和元素 `html` 。

# --hints--

需要先声明`DOCTYPE`

```js
assert(code.match(/<!DOCTYPE html>/i));
```

需要一个起始标签 `<html>`

```js
assert(code.match(/<html>/i));
```

需要一个闭合标签 `<html>` 别忘了闭合标签的尖括号 `<` 后需要有一个 `/`

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
