---
id: 5f332a88dc25a0fd25c7687a
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

咖啡厅的名称是 `CAMPER CAFE`。 在 `main` 元素中添加一个 `h1` 元素。 给它命名为 “cafe”，用大写字母，使其突出显示。

# --hints--

你应该有一个 `<h1>` 开始标签。

```js
assert(code.match(/<h1>/i));
```

你应该有一个 `</h1>` 结束标签。

```js
assert(code.match(/<\/h1>/i));
```

你不应该改变你现有的 `main` 元素。

```js
assert($('main').length === 1);
```

你的 `h1` 元素应该嵌套在你的 `main` 元素中。

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

你的 `h1` 元素应该有文字 `CAMPER CAFE`。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

