---
id: 5f332b23c2045fb843337579
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

为了让游客知道这家咖啡馆成立于 2020 年，在 `h1` 元素下面添加一个 `p`元素，文字为 `Est. 2020`。

# --hints--

你应该有一个 `<p>` 开始标签。

```js
assert(code.match(/<p>/i));
```

你应该有一个 `</p>` 结束标签。

```js
assert(code.match(/<\/p>/i));
```

你不应该改变你现有的 `h1` 元素。 确认你没有删除结束标签。

```js
assert($('h1').length === 1);
```

你的 `p` 元素应该在 `h1` 元素的下面。

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

你的 `p` 元素应该有文本 `Est. 2020`。

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

