---
id: 5f33310c1851c6c4da013250
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

是时候添加一些菜单内容。 在现有的 `header` 元素下面添加一个 `main`元素。 它最终将包含有关咖啡馆提供的咖啡和甜点的价格信息。

# --hints--

你应该有一个 `<main>` 开始标签。

```js
assert(code.match(/<main>/i));
```

你应该有一个 `</main>` 结束标签。

```js
assert(code.match(/<\/main>/i));
```

你不应该改变你的 `header` 元素。 确认你没有删除结束标签。

```js
assert($('header').length === 1);
```

你的 `main` 元素应该在你的 `header` 元素之后。

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
