---
id: 5f3313e74582ad9d063e3a38
title: 步骤 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素中添加一个 `head` 元素，接下来你可以添加 `title` 元素。 `title` 元素的文本应该是 `Cafe Menu`。

# --hints--

你需要一个开始标签 `<head>`

```js
assert(code.match(/<head>/i));
```

需要一个结束标签 `</head>`

```js
assert(code.match(/<head>/i));
```

需要一个开始标签 `<title>`

```js
assert(code.match(/<title>/i));
```

需要一个结束标签 `</title>`

```js
assert(code.match(/<\/title>/i));
```

所有的 `title` 元素都应嵌套在 `head` 元素中。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

你的 `title` 元素内应该有文本 `Cafe Menu`。 您可能需要检查您的拼写。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
