---
id: 5f33294a6af5e9188dbdb8f3
title: 步骤 5
challengeType: 0
dashedName: step-5
---

# --description--

咖啡厅的名称是 `CAMPER CAFE`。 你的 `body` 元素中应该有一个 `h1` 元素。 用大写字母显示咖啡店名，使其突出显示。

# --hints--

你应该有一个 `<h1>` 开始标签。

```js
assert(code.match(/<h1>/i));
```

你应该有一个 `</h1>` 结束标签。

```js
assert(code.match(/<\/h1>/i));
```

你不应该改变你现有的 `body`元素。

```js
assert($('body').length === 1);
```

你的 `h1` 元素应该嵌套在你的 `body` 元素中。

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

你的 `h1` 元素应该有文字 `CAMPER CAFE`。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
