---
id: 5f332b23c2045fb843337579
title: 步骤 7
challengeType: 0
dashedName: step-7
---

# --description--

由于上一步添加的 `p` 元素提供了关于咖啡馆的补充信息，因此将 `h1` 和 `p` 元素都嵌套在一个 `header`元素中。

# --hints--

你应该有一个 `<header>` 开始标签。

```js
assert(code.match(/<header>/i));
```

你应该有一个 `</header>` 结束标签。

```js
assert(code.match(/<\/header>/i));
```

你的 `h1` 元素应该嵌套在你的 `header` 元素中。

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

你的 `p` 元素应该嵌套在你的 `header`元素中。

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
</html>
```

