---
id: 5f33071498eb2472b87ddee4
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS 是一种用于设置 HTML 文档样式的语言。 它描述了 HTML 元素在屏幕上的显示方式。

正如你在 Cat Photo App 的最后几个步骤中所学习的那样，开始构建你的网页需要一个基本结构。 每个 HTML 文档都应包含 `DOCTYPE` 声明和 `html` 元素。 `DOCTYPE` 告诉浏览器文档的 HTML 版本。 而 `html` 元素代表根元素，它包含所有其他元素。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

添加 `<!DOCTYPE html>` 标记和具有值为 `en` 的 `lang` 属性的 `html` 元素。

# --hints--

你应该有 `<!DOCTYPE html>` 声明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你应该有一个 `<html lang="en">` 开始标签，语言设置为英语。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你应该有一个 `</html>` 结束标签。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
