---
id: 5f3477aefa51bfc29327200b
title: 步骤 14
challengeType: 0
dashedName: step-14
---

# --description--

你已经通过在 `style` 标签内编写 CSS 来设置了三个元素的样式。 这样做可以，但由于会有更多的样式，最好把所有的样式放在一个单独的文件中并链接到它。

我们已经为你创建了一个单独的 `styles.css` 文件，并将编辑器视图切换到该文件。 你可以通过编辑器顶部的标签在不同的文件之间转换。

首先，将你创建的样式改写到 `styles.css` 文件中。 请确保排除开头和结尾的 `style` 标签。

# --hints--

你的 `styles.css` 文件应该有 `h1, h2, p` 类型选择器。

```js
(getUserInput) => {
  assert(getUserInput('editableContents').replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
}
```

你的选择器应将 `text-align` 属性设置为 `center`。

```js
(getUserInput) => {
  assert(getUserInput('editableContents').match(/text-align:\s*center;?/));
}
```

代码不应该包含 `<style>` 标签。

```js
  assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
