---
id: 5f3477ae9675db8bb7655b30
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

在上一步中，你使用 <dfn>类型选择器</dfn> 来设置 `h1` 元素的样式。 Center the `h2` and `p` elements by adding a new type selector for each one to the existing `style` element.

# --hints--

你不应更改现有的 `h1` 选择器。

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

You should not add a new `style` tag. Add the new CSS rules to the existing `style` tag.

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

You should add a new `h2` selector.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

You should add a new `p` selector.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Your `h1` element should have a `text-align` of `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Your `h2` element should have a `text-align` of `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Your `p` element should have a `text-align` of `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

