---
id: 5f344fad8bf01691e71a30eb
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

到目前为止，在你创建的内容的演示和外观方面，你一直受到限制。 为了开始控制，在 `head` 元素中添加一个 `style` 元素。

# --hints--

你应该有一个 `<style>` 开始标签。

```js
assert(code.match(/<style\s*>/i));
```

你应该有一个 `</style>` 结束标签。

```js
assert(code.match(/<\/style\s*>/));
```

你的 `style` 元素都应嵌套在 `head` 元素中。

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
