---
id: 5f3477cb2e27333b1ab2b955
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

现在你需要链接 `styles.css` 文件，以便再次应用这些样式。 在 `head` 元素中嵌套一个自闭合 `link` 元素。 给它一个 `rel` 属性，值为 `stylesheet` 和一个 `href` 属性，值为 `styles.css`。

# --hints--

你的代码应该有一个 `link` 元素。

```js
const link = document.querySelector('link');
assert(link);
```

你不应该改变你现有的 `head` 元素。 请确保你没有删除结束标记。

```js
assert($('head').length === 1);
```

你应该有一个自闭合的 `link` 元素。

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
const link = document.querySelector('head > link');
assert(link);
```

你的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
