---
id: 5f344fc1520b6719f2e35605
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

菜单上将有两个部分，一个是咖啡，一个是甜点。 在 `main` 元素中添加一个 `section` 元素，这样你就有一个地方来放置所有可用的咖啡。

# --hints--

你应该有一个 `<section>` 开始标签。

```js
assert(code.match(/<section\s*>/i));
```

你应该有一个 `</section>` 结束标签。

```js
assert(code.match(/<\/section\s*>/i));
```

你不应该改变你现有的 `main` 元素。 确认你没有删除结束标签。

```js
assert($('main').length === 1);
```

你的 `section`元素应该在你的 `main` 元素中。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
