---
id: 5f3477cbcb6ba47918c1da92
title: 步骤 17
challengeType: 0
dashedName: step-17
---

# --description--

为了使页面样式在移动端看起来与在桌面或笔记本电脑上相似，你需要添加一个带有特殊 `content` 属性的 `meta` 元素。

在 `head` 元素中添加以下内容：

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

你应该有两个 `meta` 元素。

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

你的 `meta` 元素应该有一个 `name` 属性，其值为 `viewport`。

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

你的 `meta` 元素应该有一个 `content` 属性，其值为 `width=device-width, initial-scale=1.0`。

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
