---
id: 5f344f9c805cd193c33d829c
title: 步骤 11
challengeType: 0
dashedName: step-11
---

# --description--

你可以在 `style` 元素中指定一个样式，并为此设置一个属性，像这样：

```css
element {
 property: value;
}
```

通过设置 `text-align` 属性值为 `center`，让你的 `h1` 元素居中，。

# --hints--

你的 `style` 元素中应该有一个 `h1` 元素。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

你的 `text-align` 属性的值应该是 `center`。

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

你的 `h1` 选择器应该将 `text-align` 属性设置为 `center`。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

