---
id: 5f3477ae8466a9a3d2cc953c
title: 步骤 15
challengeType: 0
dashedName: step-15
---

# --description--

现在你在 `styles.css` 文件中有了 CSS，继续删除 `style` 元素及其所有内容。 删除后，居中的文本将移回左侧。

# --hints--

你的代码中不应该有任何 `style` 标签。

```js
assert(!code.match(/style/i));
```

你的 HTML 文件中不应有任何 CSS 选择器。

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

