---
id: 5f3477ae9675db8bb7655b30
title: 步骤 12
challengeType: 0
dashedName: step-12
---

# --description--

在上一步中，你使用 <dfn>类型选择器</dfn> 来设置 `h1` 元素的样式。 将 `h2` 和 `p` 元素居中，通过在现有的 `style` 元素中给它们添加一个新的选择器来实现。

# --hints--

你不应更改现有的 `h1` 选择器。

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

你不应该添加一个新的 `style` 标签。 将新的 CSS 规则添加到现有的 `style` 标签中。

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

你应该添加一个新的 `h2` 选择器。

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

你应该添加一个新的 `p` 选择器。

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

你的 `h1` 元素应该有一个 `text-align` 值为 `center`。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

你的 `h2` 元素应该有一个 `text-align` 值为 `center`。

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

你的 `p` 元素应该有一个 `text-align` 值为 `center`。

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

