---
id: 5f344fbc22624a2976425065
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

在 `section` 元素中创建一个 `h2` 元素，并在它里面填写文字 `Coffee`。

# --hints--

你应该有一个 `<h2>` 开始标签。

```js
assert(code.match(/<h2\s*>/i));
```

你应该有一个 `</h2>` 结束标签。

```js
assert(code.match(/<\/h2\s*>/i));
```

你不应该改变你现有的 `section` 元素。 确认你没有删除结束标签。

```js
assert($('section').length === 1);
```

你的 `h2` 元素应该在 `section` 元素中。

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

你的 `h2` 元素内应该有文本 `Coffee`。

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
