---
id: 5f35e5c4321f818cdc4bed30
title: 步骤 29
challengeType: 0
dashedName: step-29
---

# --description--

看起来很好。 是时候开始添加一些菜单项了。 在 `Coffee` 标题下添加一个空的 `article` 元素。 它将包含你当前提供的每种咖啡的风味和价格。

# --hints--

你应该有一个 `<article>` 开始标签。

```js
assert(code.match(/<article>/i));
```

你应该有一个 `</article>` 结束标签。

```js
assert(code.match(/<\/article>/i));
```

你不应更改现有的 `h2` 元素。

```js
assert($('h2').length === 1);
```

你的 `article` 元素应该在你的 `h2` 元素之后。

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

