---
id: 5f356ed69db0a491745e2bb6
title: 步骤 27
challengeType: 0
dashedName: step-27
---

# --description--

要给 `div` 元素应用类属性，需要移除 `id` 属性，并在 `div` 元素的开始标签中添加 `class` 属性。 确保将 class 的值设置为 `menu`。

# --hints--

你的 `div` 仍然应该渲染。 确保你没有改变 `<div>` 标签。

```js
assert($('div').length === 1);
```

你的 `div` 元素应该有 `menu` 类。

```js
assert($('div').attr('class').includes('menu'));
```

你的 `div` 元素不应该再有 `menu` id。

```js
assert(!$('div#menu').length);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

