---
id: 5f356ed6199b0cdef1d2be8f
title: 步骤 26
challengeType: 0
dashedName: step-26
---

# --description--

到目前为止，你一直在使用类型和 id 选择器来设置元素的样式。 不过，更常见的做法是使用不同的选择器来设置元素的样式。

<dfn>class 选择器</dfn>是名称前面有一个点，像这样：

```css
.class-name {
  styles
}
```

将现有的 `#menu` 选择器更改为类选择器，将 `#menu` 替换为名为 `.menu` 的类。

# --hints--

你应该有一个 `.menu` 类选择器。

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

你不应该有一个 `#menu` 选择器。

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

