---
id: 5f35e5c44359872a137bd98f
title: 步骤 28
challengeType: 0
dashedName: step-28
---

# --description--

由于咖啡馆主要销售的产品是咖啡，因此你可以使用咖啡豆的图像作为页面背景。

删除 `body` 类型选择器中的注释及其内容。 现在添加一个 `background-image` 属性并将其值设置为 `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`。

# --hints--

你应该删除注释掉的 `background-color` 属性。

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

你的 `body` 选择器不应有任何注释。

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

你应该将 `background-image` 属性设置为 `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

你的 `body` 元素应该有咖啡豆的背景图片。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

