---
id: 5f356ed60a5decd94ab66986
title: 步骤 22
challengeType: 0
removeComments: false
dashedName: step-22
---

# --description--

CSS 中的注释看起来像这样：

```css
/* comment here */
```

在你的样式表中，注释掉包含 `background-color` 属性和值的行，这样你就可以看到仅设置 `div` 元素样式的效果。 这将使背景再次变成白色。

# --hints--

你应该在你的 CSS 中注释掉 `background-color: burlywood;` 一行。

```js
assert(code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i));
```


你的 `body` 应该有一个白色的背景。

```js
const bodyCSS = $('body').css('background-color');
assert(bodyCSS === "rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div>
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

div {
  width: 300px;
}
```

