---
id: 5f3c866d28d7ad0de6470505
title: 步骤 32
challengeType: 0
dashedName: step-32
---

# --description--

口味和价格目前堆叠在一起，并以各自的 `p` 元素居中。 如果口味在左边，价格在右边，那就太好了。

给 `French Vanilla` `p` 元素添加 `flavor` class。

# --hints--

你应该给 `p` 元素添加 `flavor` class。

```js
assert(code.match(/<p\s*class=('|")flavor\1\s*>/i));
```

应该只有一个元素有 `flavor` class。

```js
assert($('.flavor').length === 1);
```

你的 `flavor` class 应该在 `p` 元素上，并带有文本 `French Vanilla`。

```js
assert($('.flavor')[0].innerText.match(/French Vanilla/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            <p>French Vanilla</p>
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
