---
id: 5f3c866de7a5b784048f94b1
title: 步骤 36
challengeType: 0
dashedName: step-36
---

# --description--

这正是你想要的，但如果口味和价格在同一条线上，那就太好了。 `p` 元素是 <dfn>块级</dfn> 元素，因此它们占据了其父元素的整个宽度。

要将它们放在同一行，你需要对 `p` 元素应用一些样式，因此它们更像<dfn>内联</dfn>元素。 将值为 `item` 的 `class` 属性添加到 `Coffee` 标题下的第一个 `article` 元素。

# --hints--

你应该将 `item` class 应用于你的 `article` 元素。

```js
assert(code.match(/<article\s*class=('|")item\1\s*>/i))
```

你应该只有一个 `item` class 元素。

```js
assert($('.item').length === 1);
```

你的第一个 `article` 元素应该具有 `item` class。

```js
assert($('article')[0].className === 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
