---
id: 5f3c866daec9a49519871816
title: 步骤 30
challengeType: 0
dashedName: step-30
---

# --description--

`article` 元素通常包含多个具有相关信息的元素。 在这个示例里，它将包含咖啡风味和该风味的价格。 在 `article` 元素中嵌套两个 `p` 元素。 第一个的文本应该是 `French Vanilla`，第二个的文本应该是 `3.00`。

# --hints--

你不应更改现有的 `article` 元素。

```js
assert($('article').length === 1);
```

你的 `article` 元素应该有两个 `p` 元素。

```js
assert($('article').children('p').length === 2);
```

你的第一个 `p` 元素应该包含文本 `French Vanilla`。

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

你的第二个 `p` 元素应该包含文本 `3.00`。

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
