---
id: 5f716ad029ee4053c7027a7a
title: 步骤 49
challengeType: 0
dashedName: step-49
---

# --description--

在 `article` 元素中嵌套两个 `p` 元素。 第一个的文本应该是 `Donut`，第二个的文本应该是 `1.50`。 把它们放在同一行，确保它们之间没有空格。

# --hints--

你不应更改现有的 `article` 元素。

```js
assert($('article').length === 6);
```

你的新 `article` 元素应该有两个 `p` 元素。

```js
assert($('article').last().children('p').length === 2);
```

你的第一个 `p` 元素应该包含文本 `Donut`。

```js
assert($('article').last().children('p')[0].innerText.match(/Donut/i));
```

你的第二个 `p` 元素应该包含文本 `1.50`。

```js
assert($('article').last().children('p')[1].innerText.match(/1\.50/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

