---
id: 5f76967fad478126d6552b0d
title: 步骤 34
challengeType: 0
dashedName: step-34
---

# --description--

接着，你想要将价格（price）右对齐。 给你的 `p` 元素添加一个名为 `price` 的类，其文本为 `3.00`。

# --hints--

你应该把 `price` 类添加到你的 `p` 元素中。

```js
assert(code.match(/<p\s*class=('|")price\1\s*>/i));
```

具有 `price` 类应该只有一个元素。

```js
assert($('.price').length === 1);
```

你的 `price` 类应该在 `p` 元素上，文字是 `3.00`。

```js
assert($('.price')[0].innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
--fcc-editable-region--
            <p>3.00</p>
--fcc-editable-region--
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}
```
