---
id: 5f7b87422a560036fd03ccff
title: 步骤 50
challengeType: 0
dashedName: step-50
---

# --description--

对于刚刚添加的两个 `p` 元素，添加 `dessert` 作为第一个 `p` 元素的 `class` 属性的值，然后添加 `price` 作为第二个 `p` 元素的 `class` 属性的值。

# --hints--

你应该有一个带有 `dessert` class 的 `p` 元素。

```js
assert($('.dessert').length === 1);
```

带有文本 `Donut` 的 `p` 元素应该具有 `dessert` class。

```js
assert($('.dessert')[0].innerText.match(/donut/i));
```

带有文本 `1.50` 的 `p` 元素应该具有 `price` class。

```js
assert($('.price').last().text().match(/1\.50/));
```

`p` 元素之间不应有任何空格。

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
--fcc-editable-region--
            <p>Donut</p><p>1.50</p>
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

