---
id: 66faab00000000000000000b
title: Build a Platformer Level 1-1
challengeType: 26
dashedName: lab-platformer-level-1-1
---

# --description--

**Objective:** Build a complete side-scrolling platformer level using Phaser 3 that fulfills the user stories below.

The Phaser library is loaded for you on the global `Phaser` object. The hero spritesheet lives at `/curriculum-assets/phaser/spritesheets/hero-walk.png` (4×4 grid of 16×16 frames; frames 12–15 are the right-facing walk cycle).

**User Stories:**

1. You should declare a class named `MainScene` that extends `Phaser.Scene`.
1. `MainScene` should have a constructor that calls `super("MainScene")` and `preload`, `create`, and `update` methods.
1. You should declare a `config` object with a `type` of `Phaser.AUTO`, a `width` of `480`, a `height` of `192`, a `parent` of `"game-container"`, and arcade physics with world gravity of `y: 800`.
1. Inside `MainScene.preload`, you should load the `"hero"` spritesheet from `/curriculum-assets/phaser/spritesheets/hero-walk.png` with 16×16 frames.
1. You should declare `this.player` as `this.physics.add.sprite(...)` using the `"hero"` key, and call `setCollideWorldBounds(true)` on it.
1. You should declare `this.platforms` as a `physics.add.staticGroup()` populated from a `LEVEL` (or `LEVELS`) array of segment tuples — at least 8 segments at varied heights — and register a `physics.add.collider` between the player and the platforms.
1. You should jump when `cursors.up` is pressed and the player's body is on the ground (`body.touching.down` or `body.blocked.down`), giving the body a negative `setVelocityY(...)`.
1. You should implement variable jump height by listening for `Phaser.Input.Keyboard.JustUp(cursors.up)` while the player is rising and clamping `velocity.y` to a smaller-magnitude negative value.
1. You should implement a **coyote-time** grace window of at least 80ms by stamping `this.lastGroundedAt = time` while on ground and allowing a jump if `time - this.lastGroundedAt` is below the threshold.
1. You should implement a **jump buffer** of at least 80ms by stamping `this.lastJumpPressedAt = time` on the jump press and consuming it when the player lands.
1. You should track a `canDoubleJump` flag — reset it to `true` on landing, and consume it for a mid-air jump when the coyote window has expired.
1. You should declare `this.movingPlatforms` as a `physics.add.group({ allowGravity: false, immovable: true })`, animate at least one moving platform with `this.tweens.add(...)`, register a `physics.add.collider` between player and moving platforms, and ride the platform by adding `plat.body.deltaX()` to `this.player.x` while standing on it.
1. You should declare two enemy groups — `this.enemies` (patrolling, walks left/right and flips at walls) and `this.flyingEnemies` (gravity-immune, sine-wave Y motion) — and register `physics.add.overlap` for each that calls `handleEnemyContact`.
1. You should resolve enemy contact in `handleEnemyContact(player, enemy)` so that landing on the enemy from above (`player.body.bottom <= enemy.body.top + N` while falling) destroys the enemy and bounces the player upward, and side contact calls `takeDamage()` (which decrements `this.hp` and grants ≈800ms of i-frames).
1. You should declare `this.coins` as a `physics.add.staticGroup()` containing at least 5 coins, and register a `physics.add.overlap` that destroys the coin and increments `this.score` on pickup.
1. You should declare a `this.goal` rectangle/sprite at the level end and register a `physics.add.overlap` that calls `this.scene.start("LevelCompleteScene", { ... })` when the player picks up `this.key` first (`this.hasKey === true`).
1. You should declare a class `GameOverScene extends Phaser.Scene` that calls `super("GameOverScene")`, render a "GAME OVER" message, and restart on `keydown-R`. Trigger it from `MainScene` when `this.hp` drops to `0`.
1. You should declare a class `LevelCompleteScene extends Phaser.Scene` that calls `super("LevelCompleteScene")` and renders the elapsed time and best time.
1. You should persist the player's best level-completion time across page reloads using `localStorage.getItem` and `localStorage.setItem`.
1. You should register all four scenes via `scene: [MainScene, GameOverScene, LevelCompleteScene, ...]` (the order does not matter as long as `MainScene` is first).

# --hints--

You should declare a class named `MainScene` that extends `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{/);
```

`MainScene` should call `super("MainScene")` in its constructor and define `preload`, `create`, and `update` methods.

```js
assert.match(code, /super\(\s*(['"`])MainScene\1\s*\)/);
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?\bpreload\s*\(\s*\)\s*\{/);
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?\bcreate\s*\(\s*\)\s*\{/);
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?\bupdate\s*\(/);
```

You should declare a `config` object with the required Phaser settings (Phaser.AUTO, width 480, height 192, parent "game-container", arcade physics with world gravity y of 800).

```js
assert.match(code, /const\s+config\s*=\s*\{/);
assert.match(code, /type\s*:\s*Phaser\.AUTO/);
assert.match(code, /width\s*:\s*480\b/);
assert.match(code, /height\s*:\s*192\b/);
assert.match(code, /parent\s*:\s*(['"`])game-container\1/);
assert.match(code, /physics\s*:\s*\{[\s\S]*?default\s*:\s*(['"`])arcade\1[\s\S]*?arcade\s*:\s*\{[\s\S]*?gravity\s*:\s*\{\s*y\s*:\s*800\s*\}/);
```

`MainScene.preload` should load the `"hero"` spritesheet from `/curriculum-assets/phaser/spritesheets/hero-walk.png` with 16×16 frames.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])hero\1\s*,\s*(['"`])\/curriculum-assets\/phaser\/spritesheets\/hero-walk\.png\2\s*,\s*\{\s*frameWidth\s*:\s*16\s*,\s*frameHeight\s*:\s*16\s*\}\s*\)/);
```

You should declare `this.player` with `this.physics.add.sprite(...)` keyed by `"hero"`, and call `setCollideWorldBounds(true)` on it.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\([\s\S]*?(['"`])hero\1/);
assert.match(code, /this\.player\.setCollideWorldBounds\(\s*true\s*\)/);
```

You should declare `this.platforms` as a `physics.add.staticGroup()`, declare a `LEVEL` (or `LEVELS`) array of segment tuples with at least 8 entries, and register a collider between the player and the platforms.

```js
assert.match(code, /this\.platforms\s*=\s*this\.physics\.add\.staticGroup\(/);
assert.match(code, /(?:const|let|var)\s+LEVEL[A-Z]*\s*=\s*\[\s*(?:\[[^\]]*\]\s*,\s*){7,}\[[^\]]*\]\s*\]/);
assert.match(code, /this\.physics\.add\.collider\(\s*this\.player\s*,\s*this\.platforms/);
```

You should jump when the up cursor is pressed and the player is on the ground (`body.touching.down` or `body.blocked.down`), giving the body a negative `setVelocityY(...)`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(/);
assert.match(code, /Phaser\.Input\.Keyboard\.JustDown\([\s\S]{0,40}cursors\.up/);
assert.match(code, /(?:onGround\s*=\s*[\s\S]{0,100}?touching\.down\s*\|\|[\s\S]{0,40}?blocked\.down)/);
assert.match(code, /canJump\s*&&\s*buffered\s*\)\s*\{[\s\S]{0,80}?this\.player\.setVelocityY\(\s*-\s*\d+/);
```

You should implement variable jump height by clamping `velocity.y` when the player releases `up` early.

```js
assert.match(code, /Phaser\.Input\.Keyboard\.JustUp\([\s\S]{0,40}cursors\.up/);
```

You should implement a coyote-time grace window of at least 80ms.

```js
assert.match(code, /this\.lastGroundedAt\s*=\s*time/);
assert.match(code, /time\s*-\s*this\.lastGroundedAt\s*<\s*(?:8\d|9\d|1\d{2,}|[2-9]\d{2,})/);
```

You should implement a jump buffer of at least 80ms.

```js
assert.match(code, /this\.lastJumpPressedAt\s*=\s*time/);
assert.match(code, /time\s*-\s*this\.lastJumpPressedAt\s*<\s*(?:8\d|9\d|1\d{2,}|[2-9]\d{2,})/);
```

You should track a `canDoubleJump` flag, reset on landing and consumed in mid-air.

```js
assert.match(code, /onGround[\s\S]{0,200}?this\.canDoubleJump\s*=\s*true/);
assert.match(code, /this\.canDoubleJump\s*=\s*false/);
```

You should declare `this.movingPlatforms` as a gravity-disabled, immovable group, animate at least one platform with a tween, collide it with the player, and add `plat.body.deltaX()` to the player's `x` while standing on it.

```js
assert.match(code, /this\.movingPlatforms\s*=\s*this\.physics\.add\.group\(\s*\{[\s\S]*?allowGravity\s*:\s*false[\s\S]*?immovable\s*:\s*true[\s\S]*?\}\s*\)/);
assert.match(code, /this\.tweens\.add\(/);
assert.match(code, /this\.physics\.add\.collider\(\s*this\.player\s*,\s*this\.movingPlatforms\b/);
assert.match(code, /\.body\.deltaX\(/);
```

You should declare two enemy groups — `this.enemies` (ground patrollers) and `this.flyingEnemies` (sine-wave flyers, `allowGravity: false`) — and overlap each with the player.

```js
assert.match(code, /this\.enemies\s*=\s*this\.physics\.add\.group\(/);
assert.match(code, /this\.flyingEnemies\s*=\s*this\.physics\.add\.group\(\s*\{[\s\S]*?allowGravity\s*:\s*false[\s\S]*?\}\s*\)/);
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.player\s*,\s*this\.enemies\s*,\s*this\.handleEnemyContact/);
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.player\s*,\s*this\.flyingEnemies\s*,\s*this\.handleEnemyContact/);
```

You should resolve enemy contact in `handleEnemyContact(player, enemy)` — stomp from above destroys the enemy and bounces the player; side contact calls `takeDamage`, decrements `this.hp`, and grants i-frames.

```js
assert.match(code, /handleEnemyContact\s*\(\s*player\s*,\s*enemy\s*\)\s*\{/);
assert.match(code, /player\.body\.bottom\s*<=\s*enemy\.body\.top/);
assert.match(code, /enemy\.destroy\(\s*\)/);
assert.match(code, /takeDamage\s*\(\s*\)\s*\{/);
assert.match(code, /this\.hp\s*-=\s*1/);
assert.match(code, /this\.invincibleUntil/);
```

You should declare `this.coins` as a static group with at least 5 coin positions stored in a `COINS` array, and overlap each pickup with the player to destroy the coin and increment `this.score`.

```js
assert.match(code, /this\.coins\s*=\s*this\.physics\.add\.staticGroup\(/);
assert.match(code, /this\.physics\.add\.overlap\([\s\S]*?this\.coins/);
assert.match(code, /this\.score\s*\+=/);
assert.match(code, /(?:const|let|var)\s+COINS\s*=\s*\[\s*(?:\[[^\]]*\]\s*,\s*){4,}\[[^\]]*\]\s*\]/);
```

You should declare `this.goal`, gate it on `this.hasKey`, and start `LevelCompleteScene` on overlap.

```js
assert.match(code, /this\.goal\s*=\s*/);
assert.match(code, /this\.hasKey/);
assert.match(code, /this\.scene\.start\(\s*(['"`])LevelCompleteScene\1/);
```

You should declare `class GameOverScene extends Phaser.Scene` calling `super("GameOverScene")`, render a game-over message, and restart on `keydown-R`. Trigger it from `MainScene` when `this.hp <= 0`.

```js
assert.match(code, /class\s+GameOverScene\s+extends\s+Phaser\.Scene\s*\{/);
assert.match(code, /super\(\s*(['"`])GameOverScene\1\s*\)/);
assert.match(code, /class\s+GameOverScene\s+extends\s+Phaser\.Scene\s*\{(?:(?!class\s+\w)[\s\S])*?keydown-R/);
assert.match(code, /this\.hp\s*<=\s*0/);
assert.match(code, /this\.scene\.start\(\s*(['"`])GameOverScene\1/);
```

You should declare `class LevelCompleteScene extends Phaser.Scene` calling `super("LevelCompleteScene")`.

```js
assert.match(code, /class\s+LevelCompleteScene\s+extends\s+Phaser\.Scene\s*\{/);
assert.match(code, /super\(\s*(['"`])LevelCompleteScene\1\s*\)/);
```

You should persist the best level-completion time across reloads using `localStorage.getItem` and `localStorage.setItem`.

```js
assert.match(code, /localStorage\.getItem\(/);
assert.match(code, /localStorage\.setItem\(/);
```

You should register all the scenes via `scene: [MainScene, ...]` in `config`, with `MainScene` first.

```js
assert.match(code, /scene\s*:\s*\[\s*MainScene\s*,/);
assert.match(code, /scene\s*:\s*\[[\s\S]*?\bGameOverScene\b/);
assert.match(code, /scene\s*:\s*\[[\s\S]*?\bLevelCompleteScene\b/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const SAVE_KEY = "phaser-platformer-progress";

const LEVEL = [
  [0, 192, 176],
  [256, 480, 176],
  [560, 768, 176],
  [832, 1280, 176],
  [96, 192, 128],
  [256, 320, 96],
  [384, 480, 128],
  [560, 624, 80],
  [720, 832, 128],
  [928, 1056, 96],
  [1136, 1216, 64]
];

const COINS = [
  [120, 140],
  [260, 100],
  [420, 60],
  [620, 130],
  [880, 80],
  [1100, 50]
];

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.platforms = this.physics.add.staticGroup();
    LEVEL.forEach(([xs, xe, y]) => {
      for (let x = xs; x < xe; x += 16) {
        this.platforms.add(this.add.rectangle(x + 8, y, 16, 16, 0x66ddff));
      }
    });
    this.player = this.physics.add.sprite(40, 60, "hero", 12);
    this.player.setCollideWorldBounds(true);
    this.player.body.setSize(10, 14);
    this.player.body.setOffset(3, 2);
    this.player.body.setMaxVelocity(140, 380);
    this.physics.add.collider(this.player, this.platforms);
    this.physics.world.setBounds(0, 0, 1280, 192);
    this.cameras.main.setBounds(0, 0, 1280, 192);
    this.cameras.main.startFollow(this.player, true, 0.1, 0.1);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.anims.create({
      key: "walk",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.coins = this.physics.add.staticGroup();
    COINS.forEach(([x, y]) => {
      this.coins.add(this.add.rectangle(x, y, 8, 8, 0xffee44));
    });
    this.score = 0;
    this.physics.add.overlap(
      this.player,
      this.coins,
      (player, coin) => {
        coin.destroy();
        this.score += 1;
      },
      null,
      this
    );
    this.hp = 3;
    this.invincibleUntil = 0;
    this.canDoubleJump = true;
    this.lastGroundedAt = -Infinity;
    this.lastJumpPressedAt = -Infinity;
    this.lastEnemyHitAt = 0;
    this.prevVelocityY = 0;
    this.hud = this.add.text(8, 8, "", {
      fontSize: "10px",
      color: "#ffffff"
    });
    this.hud.setScrollFactor(0);
    this.movingPlatforms = this.physics.add.group({
      allowGravity: false,
      immovable: true
    });
    const moving = this.add.rectangle(240, 140, 48, 8, 0xffaa44);
    this.movingPlatforms.add(moving);
    this.tweens.add({
      targets: moving,
      x: 320,
      duration: 1500,
      yoyo: true,
      repeat: -1,
      ease: "Sine.easeInOut"
    });
    this.physics.add.collider(this.player, this.movingPlatforms);
    this.spikes = this.physics.add.staticGroup();
    [[400, 168], [800, 168], [1024, 168]].forEach(([x, y]) => {
      this.spikes.add(this.add.rectangle(x, y, 16, 8, 0xff4444));
    });
    this.physics.add.overlap(
      this.player,
      this.spikes,
      () => this.takeDamage(),
      null,
      this
    );
    this.enemies = this.physics.add.group();
    [{ x: 320, y: 160 }, { x: 700, y: 64 }].forEach(({ x, y }) => {
      const e = this.add.rectangle(x, y, 12, 12, 0xff66aa);
      this.enemies.add(e);
      e.body.setVelocityX(-40);
    });
    this.physics.add.collider(this.enemies, this.platforms);
    this.physics.add.overlap(
      this.player,
      this.enemies,
      this.handleEnemyContact,
      null,
      this
    );
    this.flyingEnemies = this.physics.add.group({
      allowGravity: false
    });
    const flyer = this.add.circle(900, 96, 8, 0xaa66ff);
    this.flyingEnemies.add(flyer);
    flyer.baseY = 96;
    this.physics.add.overlap(
      this.player,
      this.flyingEnemies,
      this.handleEnemyContact,
      null,
      this
    );
    this.hasKey = false;
    this.key = this.add.rectangle(1120, 48, 10, 10, 0xffdd00);
    this.physics.add.existing(this.key, true);
    this.physics.add.overlap(
      this.player,
      this.key,
      () => {
        this.key.destroy();
        this.hasKey = true;
      },
      null,
      this
    );
    this.goal = this.add.rectangle(1240, 152, 12, 32, 0x44ee88);
    this.physics.add.existing(this.goal, true);
    this.physics.add.overlap(
      this.player,
      this.goal,
      () => {
        if (!this.hasKey || this.won) return;
        this.won = true;
        const elapsed = (this.time.now - this.startTime) / 1000;
        const stored = parseFloat(localStorage.getItem(SAVE_KEY + ":best"));
        const previousBest = isNaN(stored) ? Infinity : stored;
        const best = Math.min(previousBest, elapsed);
        if (best === elapsed) {
          localStorage.setItem(SAVE_KEY + ":best", elapsed.toFixed(2));
        }
        this.scene.start("LevelCompleteScene", {
          time: elapsed,
          best: best
        });
      },
      null,
      this
    );
    this.startTime = this.time.now;
  }

  update(time, delta) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-140);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(140);
    } else {
      this.player.setVelocityX(0);
    }
    const onGround =
      this.player.body.touching.down || this.player.body.blocked.down;
    if (onGround && !this.player.body.blocked.up) {
      this.lastGroundedAt = time;
      this.canDoubleJump = true;
    }
    if (Phaser.Input.Keyboard.JustDown(this.cursors.up)) {
      this.lastJumpPressedAt = time;
    }
    const canJump = time - this.lastGroundedAt < 120;
    const buffered = time - this.lastJumpPressedAt < 120;
    if (canJump && buffered) {
      this.player.setVelocityY(-310);
      this.lastGroundedAt = -Infinity;
      this.lastJumpPressedAt = -Infinity;
    } else if (
      Phaser.Input.Keyboard.JustDown(this.cursors.up) &&
      !canJump &&
      this.canDoubleJump
    ) {
      this.player.setVelocityY(-260);
      this.canDoubleJump = false;
    }
    if (
      Phaser.Input.Keyboard.JustUp(this.cursors.up) &&
      this.player.body.velocity.y < -120
    ) {
      this.player.setVelocityY(-120);
    }
    const vx = this.player.body.velocity.x;
    if (Math.abs(vx) > 1) {
      this.player.anims.play("walk", true);
      this.player.setFlipX(vx < 0);
    } else {
      this.player.anims.stop();
      this.player.setFrame(12);
    }
    this.movingPlatforms.getChildren().forEach((plat) => {
      if (
        this.player.body.touching.down &&
        this.player.y < plat.y &&
        Math.abs(this.player.x - plat.x) < plat.width / 2 + 8
      ) {
        this.player.x += plat.body.deltaX();
      }
    });
    this.enemies.getChildren().forEach((e) => {
      if (e.body.blocked.left) e.body.setVelocityX(40);
      else if (e.body.blocked.right) e.body.setVelocityX(-40);
    });
    this.flyingEnemies.getChildren().forEach((f) => {
      f.y = f.baseY + Math.sin(time / 300) * 24;
    });
    this.hud.setText("HP " + this.hp + "  COINS " + this.score);
    this.prevVelocityY = this.player.body.velocity.y;
  }

  takeDamage() {
    if (this.time.now < this.invincibleUntil) return;
    this.invincibleUntil = this.time.now + 800;
    this.hp -= 1;
    this.player.setTint(0xff4444);
    this.time.delayedCall(150, () => this.player.clearTint());
    this.cameras.main.flash(150, 200, 50, 50);
    if (this.hp <= 0) this.scene.start("GameOverScene");
  }

  handleEnemyContact(player, enemy) {
    if (player.body.bottom <= enemy.body.top + 4 && player.body.velocity.y > 0) {
      enemy.destroy();
      player.setVelocityY(-220);
      return;
    }
    if (this.time.now - this.lastEnemyHitAt < 200) return;
    this.lastEnemyHitAt = this.time.now;
    this.takeDamage();
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create() {
    this.add.rectangle(240, 96, 480, 192, 0x000000, 0.85);
    this.add
      .text(240, 70, "GAME OVER", { fontSize: "20px", color: "#ff6644" })
      .setOrigin(0.5);
    this.add
      .text(240, 110, "Press R to retry", {
        fontSize: "10px",
        color: "#ffee44"
      })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-R", () => {
      this.scene.start("MainScene");
    });
  }
}

class LevelCompleteScene extends Phaser.Scene {
  constructor() {
    super("LevelCompleteScene");
  }

  create(data) {
    this.add.rectangle(240, 96, 480, 192, 0x000000, 0.85);
    this.add
      .text(240, 60, "LEVEL COMPLETE", { fontSize: "18px", color: "#44ee88" })
      .setOrigin(0.5);
    this.add
      .text(240, 92, "Time: " + data.time.toFixed(2) + "s", {
        fontSize: "12px",
        color: "#ffffff"
      })
      .setOrigin(0.5);
    const bestLabel =
      data.best === Infinity ? "--" : data.best.toFixed(2) + "s";
    this.add
      .text(240, 112, "Best: " + bestLabel, {
        fontSize: "12px",
        color: "#aaaaaa"
      })
      .setOrigin(0.5);
    this.add
      .text(240, 140, "Press R to play again", {
        fontSize: "10px",
        color: "#ffee44"
      })
      .setOrigin(0.5);
    this.input.keyboard.once("keydown-R", () => {
      this.scene.start("MainScene");
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 192,
  parent: "game-container",
  backgroundColor: "#1a1a2e",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 800 }, debug: false }
  },
  scene: [MainScene, GameOverScene, LevelCompleteScene]
};

new Phaser.Game(config);
```
