---
id: 5d5a813321b9e3db6c106a46
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript 是一種強大的語言，它允許你構建交互式網站。 首先，使用 `DOCTYPE`、`html`、`head` 和 `body` 創建你的標準 HTML 樣板。 包括一個 `title` 元素和一個 `link` 用於樣式表，以及 `meta` 標記用於 `charset`。 然後，在 `body` 中創建一個 ID 爲 `game` 的 `div` 元素。 給你的 `title` 元素添加文本 `RPG - Dragon Repeller`。

# --hints--

代碼應包含 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

應該在 `DOCTYPE` 聲明之後包含一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

應該將文檔類型定義爲 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

應該在類型之後使用 `>` 關閉 `DOCTYPE` 聲明。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 元素應該有一個開始標籤。 不要忘記 `lang` 屬性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

`DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

應該有一個 `head` 開始標籤。

```js
assert(code.match(/<head\s*>/i));
```

應該有一個 `head` 結束標籤。

```js
assert(code.match(/<\/head\s*>/i));
```

應該有一個 `body` 開始標籤。

```js
assert(code.match(/<body\s*>/i));
```

應該有一個 `body` 結束標籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是兄弟元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該在 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

代碼應該有一個 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 元素應該有一個 `charset` 屬性，其值爲 `UTF-8`。

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

代碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

代碼應該有一個 `link` 元素。

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

代碼應該有一個 `div` 元素。

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` 元素應該有一個 `id` 屬性，其值爲 `game`。

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

`div` 元素應該在 `body` 元素內。

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
