---
id: 5d5a813321b9e3db6c106a46
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

JavaScript 是一種強大的語言，它允許你構建交互式網站。

*Note*: For all remaining projects in this curriculum, you will need a basic level of knowledge in HTML and CSS. If you are new to HTML and CSS, please go through the [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

首先，使用 `DOCTYPE`、`html`、`head` 和 `body` 標籤創建你的標準 HTML 模版，然後爲 `meta` 添加一個 `charset` 屬性。 添加 `title` 元素，將它的文本設置爲 `RPG - Dragon Repeller`。 添加一個 `link` 標籤，鏈接到 `styles.css` 文件，以添加樣式表。

最後，在你的 `body` 中創建一個 `div` 元素，將其 `id` 設置爲 `game`。

# --hints--

Your code should contain the `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE/gi));
```

You should include a space after the `DOCTYPE` reference.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

You should define the document type to be `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

You should close the `DOCTYPE` declaration with a `>` after the type.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Your `html` element should have an opening tag. Don't forget the `lang` attribute.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Your `html` element should have a closing tag.

```js
assert(code.match(/<\/html\s*>/));
```

你的 `DOCTYPE` 聲明應位於 HTML 的開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

You should have an opening `head` tag.

```js
assert(code.match(/<head\s*>/i));
```

You should have a closing `head` tag.

```js
assert(code.match(/<\/head\s*>/i));
```

You should have an opening `body` tag.

```js
assert(code.match(/<body\s*>/i));
```

You should have a closing `body` tag.

```js
assert(code.match(/<\/body\s*>/i));
```

The `head` and `body` elements should be siblings.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

The `head` element should be within the `html` element.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

The `body` element should be within the `html` element.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

You should have a closing `title` tag.

```js
assert(code.match(/<\/title\s*>/i));
```

Your code should have a `link` element.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

You should have a `div` element.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Your `div` element should have an `id` attribute with the value `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Your `div` element should be within the `body` element.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
