---
id: 62a23d1c5f1c93161f3582ae
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

與 `#stats` 元素類似，`#monsterStats` 元素需要兩個 `span` 元素。 給它們類 `stat`，並給第一個元素文本 `Monster Name:`，給第二個元素文本 `Health:`。 在每個文本之後，添加一個 `strong` 元素和一個空的嵌套 `span` 元素。

# --hints--

Your `#monsterStats` element should have two `span` elements.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

新的 `span` 元素都應該具有 `stat` 的 `class` 值。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

第一個 `span` 元素應該包含文本 `Monster Name:`。 請確保間隔正確。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].innerText, 'Monster Name: ');
```

第二個 `span` 元素應包含文本 `Health:`。 請確保間隔正確。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[1].textContent, 'Health: ');
```

第一個 `span` 元素應該有一個 `strong` 元素和一個空的嵌套 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

第二個 `span` 元素應該有一個 `strong` 元素和一個空的嵌套 `span` 元素。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

`strong` 和 `span` 元素應該在文本之後。

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name: <strong>/);
assert.match(spans[1].innerHTML, /Health: <strong>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
--fcc-editable-region--
        <div id="monsterStats">

        </div>
--fcc-editable-region--
        <div id="text"></div>
    </div>
</body>
</html>
```
