---
id: 62a255dae245b52317da824a
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

你的開發者控制檯是最強大的工具之一。 根據你的瀏覽器不同，可能需要通過按 `F12` 或 `Ctrl+Shift+I` 來打開它。 On Mac, you can press `Option + ⌘ + C` and select "Console". 也可以點擊預覽窗口上方的“控制檯”按鈕來查看我們的內置控制檯。

開發者控制檯將包含你的代碼產生的錯誤，但你也可以使用它來查看代碼中變量的值，這有助於調試。

Add a `console.log("Hello World");` line between your `script` tags. Then click the "Console" button to open the console. You should see the text `Hello World`.

請注意該行如何以分號結尾。 JavaScript 中的常見做法是用分號結束代碼行。

# --hints--

應該在代碼中添加 `console.log("Hello World");` 行。 不要忘記分號。

```js
assert.match(code, /console\.log\("Hello World"\);/);
```

`console.log("Hello World");` 行應該在 `script` 標記之間。

```js
assert.match(code, /<script>\s*console\.log\("Hello World"\);\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```
