---
id: 62a255dae245b52317da824a
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

你的開發者控制檯是最強大的工具之一。 根據你的瀏覽器不同，可能需要通過按 `F12` 或 `Ctrl+Shift+I` 來打開它。 On Mac, you can press `Option + ⌘ + C` and select "Console". 也可以點擊預覽窗口上方的“控制檯”按鈕來查看我們的內置控制檯。

開發者控制檯將包含你的代碼產生的錯誤，但你也可以使用它來查看代碼中變量的值，這有助於調試。

Add a `console.log("Hello World");` line between your `script` tags. Then click the "Console" button to open the console. You should see the text `"Hello World"`.


# --hints--

應該在代碼中添加 `console.log("Hello World");` 行。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

`console.log("Hello World");` 行應該在 `script` 標記之間。

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
