---
id: 62a3b3eab50e193608c19fc6
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

Declare a variable called `fighting` but do not initialize it with a value. 記得用分號結束你這一行的代碼。

# --hints--

你應該使用 `let` 來聲明變量 `fighting` 。

```js
assert.match(code, /let\s+fighting/);
```

你的 `fighting` 變量不應該有值。

```js
assert.isUndefined(fighting);
```

你不應該爲你的 `fighting` 變量分配一個值。 不要忘記末尾的分號。

```js
assert.match(code, /let\s+fighting\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
--fcc-editable-region--
```
