---
id: 62a3a0a3c0a4b32915d26a6e
title: 步驟6
challengeType: 0
dashedName: step-6
---

# --description--

刪除 `console.log("Hello World");` ，然後開始寫你自己的項目代碼。

在 JavaScript 中，使用<dfn>變量</dfn>來保存一個值。 若要使用變量，你必須先<dfn>聲明</dfn>。 例如，要聲明一個名爲 `camperbot` 的變量，你應該寫：

```js
let camperbot;
```

The `let` keyword tells JavaScript you are declaring a variable. 聲明一個名爲 `xp` 的變量。

# --hints--

你的代碼中不應該有 `console.log("Hello World");`。

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

You should use the `let` keyword to declare your variable.

```js
assert.match(code, /let/);
```

你應該聲明一個名爲 `xp` 的變量。

```js
assert.match(code, /let\s+xp/);
```

你不應將一個值分配給你的變量。

```js
assert.notMatch(code, /let\s+xp\s*=/);
```

別忘了代碼末尾的分號。

```js
assert.match(code, /let\s+xp\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
