---
id: 62a3b41c9494f937560640ab
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

聲明另外兩個變量，名爲 `monsterHealth` 和 `inventory`，但不要初始化。

# --hints--

你應該使用 `let` 關鍵詞聲明你的 `monsterHealth` 變量。

```js
assert.match(code, /let\s+monsterHealth/i);
```

You should use camelCase to name your `monsterHealth` variable.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 不應該有值。

```js
assert.isUndefined(monsterHealth);
```

你不應該爲你的 `monsterHealth` 變量分配一個值。 別忘了分號。

```js
assert.match(code, /let\s+monsterHealth\s*;/);
```

你應該使用 `let` 關鍵詞聲明你的 `inventory` 變量。

```js
assert.match(code, /let\s+inventory/i);
```

`inventory` 不應該有值。

```js
assert.isUndefined(inventory);
```

你不應該爲你的 `inventory` 變量分配一個值。 別忘了分號。

```js
assert.match(code, /let\s+inventory\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeapon = 0;
let fighting;
--fcc-editable-region--
```
