---
id: 62a3b41c9494f937560640ab
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Declare two more variables named `monsterHealth` and `inventory`.

For your `inventory` variable, assign it the value of an array containing the string `"stick"`.

Remember that you worked with arrays in the previous project like this:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

你應該使用 `let` 關鍵詞聲明你的 `monsterHealth` 變量。

```js
assert.match(code, /let\s+monsterHealth/i);
```

你應該使用駝峯命名法（camelCase）來命名 `monsterHealth` 變量。

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` 不應該有值。

```js
assert.isUndefined(monsterHealth);
```

你不應該爲你的 `monsterHealth` 變量分配一個值。

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

你應該使用 `let` 關鍵詞聲明你的 `inventory` 變量。

```js
assert.match(code, /let\s+inventory/i);
```

Your `inventory` variable should still be an array.

```js
assert.isArray(inventory);
```

Your `inventory` variable should only have one value.

```js
assert.lengthOf(inventory, 1);
```

Your `inventory` variable should include the string `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
