---
id: 62a39f5a5790eb27c1e5d4bf
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

你的視圖已切換到新的 `script.js` 文件。 你可以使用上面的選項卡在文件之間切換。

在文件中添加 `console.log("Hello World");`，然後就能在控制檯看到它。

# --hints--

你的代碼中應該有 `console.log("Hello World");`。

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
