---
id: 62a3a0a3c0a4b32915d26a6e
title: 步驟 19
challengeType: 0
dashedName: step-19
---

# --description--

刪除 `console.log("Hello World");` ，然後開始寫你自己的項目代碼。

在 JavaScript 中，使用<dfn>變量</dfn>來保存一個值。 若要使用變量，你必須先<dfn>聲明</dfn>。 例如，要聲明一個名爲 `camperbot` 的變量，你應該寫：

```js
var camperbot;
```

`var` 關鍵字告訴 JavaScript 你正在聲明一個變量。 聲明一個名爲 `xp` 的變量。

# --hints--

你的代碼中不應該有 `console.log("Hello World");`。

```js
assert.notMatch(code, /console\.log\("Hello World"\);/);
```

你應該使用 `var` 關鍵字聲明變量。

```js
assert.match(code, /var/);
```

你應該聲明一個名爲 `xp` 的變量。

```js
assert.match(code, /xp/);
```

你不應將一個值分配給你的變量。

```js
assert.notMatch(code, /var xp =/);
```

別忘了代碼末尾的分號。

```js
assert.match(code, /var xp;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
body {
    background-color: darkblue;
}

#text {
    background-color: black;
    color: white;
    padding: 10px;
}

#game {
    max-width: 500px;
    max-height: 400px;
    background-color: lightgray;
    color: white;
    margin: 0 auto;
    padding: 10px;
}

#controls, #stats {
    border: 1px solid black;
    padding: 5px;
    color: black;
}

#monsterStats {
    display: none;
    border: 1px solid black;
    padding: 5px;
    color: white;
    background-color: red;
}

.stat {
    padding-right: 10px;
}
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
