---
id: 62a3b79d520a7f3d0e25afd6
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript 使用<dfn>文檔對象模型</dfn>（又稱 DOM）與 HTML 交互。 DOM 是一個代表 HTML 的對象樹。 你可以使用 `document` 對象來訪問 HTML，其代表你的整個HTML文檔。

在你的 HTML 中找到特定元素的一個方法是使用 `querySelector()` 函數。 `querySelector()` 方法需要一個 CSS 選擇器作爲參數，並返回與該選擇器匹配的第一個元素。 例如，要在你的 HTML 中找到 `<h1>` 元素，你應該寫下：

```js
let h1 = document.querySelector("h1");
```

請注意， `h1` 是一個字符串，它匹配你將要使用的 CSS 選擇器。

創建一個 `button1` 變量。並使用 `querySelector()` 來給它賦值 `id` 爲 `button1` 的元素。 請記住，CSS 的 `id` 選擇器前綴爲 `#`。

# --hints--

你應該使用 `let` 關鍵詞來聲明 `button1` 變量。

```js
assert.match(code, /let\s+button1/);
```

你應該使用 `document.querySelector()` 。

```js
assert.match(code, /document\.querySelector/);
```

You should use the `"#button1"` selector.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
