---
id: 68e6172975008fea742a63cd
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now let's explore the `slice()` method, which extracts a portion of a string.

Create a `const` variable named `message` and assign it the string `"Welcome to freeCodeCamp!"`.

Then use `console.log()` to print `"Here are some examples of the slice() method:"`.

# --hints--

You should declare a variable named `message`.

```js
assert.match(code, /const\s+message/);
```

Your `message` variable should be `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

You should log the heading for slice examples.

```js
assert.match(code, /console\.log\(\s*(['"`])Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
