---
id: 68e6172975008fea742a63d1
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

The `slice()` method can use negative indices to count from the end of the string. `-1` refers to the last character.

Create a `const` variable named `endPunctuation` and use `message.slice(-1)` to extract the last character.

Then log: `The ending punctuation mark is a "${endPunctuation}"`

# --hints--

You should declare a variable named `endPunctuation`.

```js
assert.match(code, /const\s+endPunctuation/);
```

You should use `message.slice(-1)`.

```js
assert.match(code, /message\.slice\(\s*-1\s*\)/);
```

Your `endPunctuation` variable should equal `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

You should log the ending punctuation with the correct message.

```js
assert.match(code, /console\.log\(\s*`The ending punctuation mark is a "\$\{endPunctuation\}"`\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
