---
id: 68e6172975008fea742a63cb
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now demonstrate what happens when the substring is NOT found.

Create a `const` variable named `hasJavaScript` and check if `fccSentence` includes `"JavaScript"`.

Then log the result with: `fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`

# --hints--

You should declare a variable named `hasJavaScript`.

```js
assert.match(code, /const\s+hasJavaScript/);
```

You should use `fccSentence.includes("JavaScript")`.

```js
assert.match(code, /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Your `hasJavaScript` variable should be `false`.

```js
assert.strictEqual(hasJavaScript, false);
```

You should log the result with a template literal.

```js
assert.match(code, /console\.log\(\s*`.*\$\{hasJavaScript\}.*not in the sentence.*`\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
