---
id: 68e6172975008fea742a63cc
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

The `includes()` method is case-sensitive, meaning it distinguishes between uppercase and lowercase letters.

Create a `const` variable named `hasLowercaseFCC` and check if `fccSentence` includes `"freecodecamp"` (all lowercase).

Then log: `fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`

# --hints--

You should declare a variable named `hasLowercaseFCC`.

```js
assert.match(code, /const\s+hasLowercaseFCC/);
```

You should use `fccSentence.includes("freecodecamp")`.

```js
assert.match(code, /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Your `hasLowercaseFCC` variable should be `false`.

```js
assert.strictEqual(hasLowercaseFCC, false);
```

You should log the result mentioning case-sensitivity.

```js
assert.match(code, /console\.log\(\s*`.*\$\{hasLowercaseFCC\}.*case-sensitive.*`\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
