---
id: 68e6172975008fea742a63ca
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now log a message that includes the result.

You can use a template literal or string concatenation to output: `fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`

# --hints--

You should log a message that includes `fccSentence.includes("freeCodeCamp")`, the `hasFreeCodeCamp` result, and an explanation.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(.*fccSentence\.includes.*freeCodeCamp.*hasFreeCodeCamp.*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
