---
id: 68e6172875008fea742a63c9
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Use `console.log()` to log `"Here are some examples of the includes() method:"` to the console.

The `includes()` method checks if a string contains a specific substring and returns `true` or `false`.

Here is an example using the `includes()` method:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Create a variable named `hasFreeCodeCamp` and use the `includes()` method on `fccSentence` to check if it contains `"freeCodeCamp"`.

# --hints--

You should log `"Here are some examples of the includes() method:"` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the includes\(\) method:\1\s*\)/);
```

You should declare a variable named `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

You should use `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(code, /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Your `hasFreeCodeCamp` variable should be `true`.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
