---
id: 68e6172975008fea742a63d0
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

You can use `slice()` to extract from the beginning of a string.

Create a `const` variable named `greetingWord` and use `message.slice(0, 7)` to extract the first word `"Welcome"`.

Then log: `The first word is "${greetingWord}".`

# --hints--

You should declare a variable named `greetingWord`.

```js
assert.match(code, /const\s+greetingWord/);
```

You should use `message.slice(0, 7)`.

```js
assert.match(code, /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Your `greetingWord` variable should equal `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

You should log the greeting word with the correct message.

```js
assert.match(code, /console\.log\(\s*`The first word is "\$\{greetingWord\}"\.`\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
