---
id: 68e6172975008fea742a63cb
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Now it is time to see what happens when the substring is NOT found.

Create a variable named `hasJavaScript`. Then assign it the result of using the `includes()` method to check if `fccSentence` includes `"JavaScript"`.

Then log a message. You can use a template literal or string concatenation to output: `fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`

# --hints--

You should declare a variable named `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

You should use `fccSentence.includes("JavaScript")`.

```js
assert.match(code, /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

You should assign the result of `fccSentence.includes("JavaScript")` to your `hasJavaScript` variable.

```js
assert.match(code, /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Your `hasJavaScript` variable should be `false`.

```js
assert.isFalse(hasJavaScript);
```

You should log a message that includes `fccSentence.includes("JavaScript")`, the `hasJavaScript` result, and an explanation.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(.*fccSentence\.includes.*JavaScript.*hasJavaScript.*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
