---
id: 68e6172975008fea742a63ca
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now log the result with an explanation.

Use a template literal in `console.log()` to output: `fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`

# --hints--

You should use `console.log()` with a template literal.

```js
assert.match(code, /console\.log\(\s*`.*\$\{hasFreeCodeCamp\}.*`\s*\)/);
```

Your message should explain why the result is true.

```js
assert.match(code, /console\.log\(\s*`.*fccSentence\.includes\("freeCodeCamp"\).*returns.*\$\{hasFreeCodeCamp\}.*because.*in the sentence.*`\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
