---
id: 68e6172975008fea742a63ca
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Now log the result with an explanation.

You can use a template literal or string concatenation to output: `fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`

# --hints--

You should log the `hasFreeCodeCamp` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*.*hasFreeCodeCamp.*\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
