---
id: 68e6172975008fea742a63cf
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now it is time to log the result from using the `slice()` method.

You can use a template literal or string concatenation to output: `The word "${platform}" was sliced from the message.`

# --hints--

You should log a message that includes the word "sliced" and the `platform` variable.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(.*sliced.*platform.*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
