---
id: 68e6172975008fea742a63cb
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now it is time to see what happens when the substring is NOT found.

Create a variable named `hasJavaScript` and check if `fccSentence` includes `"JavaScript"`.

Then log the result. You can use a template literal or string concatenation to output: `fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`

# --hints--

You should declare a variable named `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

You should use `fccSentence.includes("JavaScript")`.

```js
assert.match(code, /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Your `hasJavaScript` variable should be `false`.

```js
assert.isFalse(hasJavaScript);
```

You should log the `hasJavaScript` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*.*hasJavaScript.*\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
