---
id: 68e6172875008fea742a63c9
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Before demonstrating the `includes()` method, you should add a heading to the console.

Use `console.log()` to print `"Here are some examples of the includes() method:"`.

Then, create a `const` variable named `hasFreeCodeCamp` and use the `includes()` method on `fccSentence` to check if it contains `"freeCodeCamp"`.

# --hints--

You should use `console.log()` to print the heading.

```js
assert.match(code, /console\.log\(\s*(['"`])Here are some examples of the includes\(\) method:\1\s*\)/);
```

You should declare a variable named `hasFreeCodeCamp`.

```js
assert.match(code, /const\s+hasFreeCodeCamp/);
```

You should use `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(code, /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Your `hasFreeCodeCamp` variable should be `true`.

```js
assert.strictEqual(hasFreeCodeCamp, true);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
