---
id: 61fd719788899952e67692b9
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

The `thead` and `tbody` elements are used to indicate which portion of your table is the header, and which portion contains the primary data or content.

Add a `thead` and `tbody` to your first `table`, below the `caption` element.

# --hints--

Your first `table` element should have a `thead` element.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Your first `table` element should have a `tbody` element.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

Your `thead` element should be immediately below your `caption` element.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

Your `tbody` element should be immediately below your `thead` element.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
