---
id: 699b70b10ae344c802d106b0
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Now it is time to add a couple more properties to the `userRoles` object.

Using dot notation or bracket notation, add a `moderator` property with a value of `"medium-access"`. Then below that, add a `guest` property with a value of `3` to the `userRoles` object.

# --hints--

Your `userRoles` object should have a `moderator` property.

```js
assert.property(userRoles, "moderator");
```

Your `moderator` property should have a value of `"medium-access"`.

```js
assert.propertyVal(userRoles, "moderator", "medium-access");
```

Your `userRoles` object should have a `guest` property.

```js
assert.property(userRoles, "guest");
```

Your `guest` property should have a value of the number `3`.

```js
assert.propertyVal(userRoles, "guest", 3);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

--fcc-editable-region--

--fcc-editable-region--

console.log(userRoles);
```
