---
id: 699b62c99ec8c63470313365
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now it is time to add another property to your `profile` object. 

Start by updating the inline object type to include a `bio` property with a type set to `string`.

# --hints--

You should update the inline object type to include a `bio` property set to `string`.

```js
assert.match(
  code,
  /const\s+profile\s*:\s*{[^}]*\bbio\s*:\s*string\b[^}]*}/
);
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
