---
id: 65821c1c40ccac44adf38e6b
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

A list in Python is a built-in data type that allows you to store many items in a single data structure. In Python, you create a list by putting items inside square brackets (`[]`), with each item separated from the following one by a comma.

```py
numbers = [1, 2, 3, 4]
```

Use a pair of square brackets to create an empty list named `expenses`. You will use it to store each of your expenses.

# --hints--

You should declare a new variable named `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

You should assign an empty list to the `expenses` variable.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
