---
id: 6582352cbdbcde4d0c4f7b0b
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

To call a lambda function you can use the usual function syntax with a pair of parentheses after the variable name.

Call your `test` lambda function and pass `3` as the argument. Then, print the result.

# --hints--

You should have `print(test(3))` in your code.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
