---
id: 666aae6a5d30a71f1fd7749f
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

The `insert` method can add an element at any position in a list, the first argument is the position at which the element has to be added, and the second argument is the element to add. For example, let's add a new element in the third position of `example_list`.

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

The `my_list` list is missing the `1` now. Add it using `insert()` in the right position, then print the list.

# --hints--

You should add a new element to `my_list` using `insert()`.

```js
assert(false);
```

`my_list` should have 4 elements.

```js
assert(false);
```

`my_list` should equal `[0, 1, 2, 3]`.

```js
assert(false);
```

You should print the list after adding the new element.

```js
assert(false);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
