---
id: 66a83601cd819e37f0dccd14
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

For the introductory text, you will want to display the main title followed by a short note about leaving feedback. 

Inside your `body` element, add a `header` element.

Inside the `header` element, add an `h1` element. The `h1` element's text should be `Hotel Feedback Form`.

Below your `h1` element, add a `p` element. The `p` element's text should be `Thank you for staying with us. Please provide feedback on your recent stay.`

# --hints--

You should have an opening `header` tag.

```js
assert.match(code, /<header>/i);
```

You should have a closing `header` tag.

```js
assert.match(code, /<\/header>/i);
```

You should have an opening `h1` tag.

```js
assert.match(code, /<h1>/i);
```

You should have a closing `h1` tag.

```js
assert.match(code, /<\/h1>/i);
```

Your `h1` element should have the text `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1.innerText, 'Hotel Feedback Form');
```

Your `h1` element should be inside the `header` element.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1.parentElement.tagName, 'HEADER');
```

You should have an opening `p` tag.

```js
assert.match(code, /<p>/i);
```

You should have a closing `p` tag.

```js
assert.match(code, /<\/p>/i);
```

Your paragraph element should have the text `Thank you for staying with us. Please provide feedback on your recent stay.`

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement.innerText, 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Your paragraph element should be inside your `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement.parentElement.tagName, 'HEADER');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
