---
id: 66a83bdcf425e7446900b7c4
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

In the previous lecture videos, you learned how to work with the `form` element like this:

```html
<form method="value-goes-here" action="url-goes-here">
  <!-- inputs go inside here -->
</form>
```

The `action` attribute is used to specify where the form data should be sent when the form is submitted.

The `method` attribute is used to specify the HTTP method to use when sending the form data. The most common methods are `GET` and `POST`. 

**NOTE**: You will learn about how HTTP methods work in later modules.

Inside your `main` element, add a `form` element with an `action` attribute set to `"https://hotel-feedback.freecodecamp.org"` and a `method` attribute set to `"POST"`.

# --hints--

You should have an opening `form` tag.

```js
assert.match(code, /<form.*>/i);
```

You should have a closing `form` tag.

```js
assert.match(code, /<\/form>/i);
```

Your `form` element should have an `action` attribute set to `"https://hotel-feedback.freecodecamp.org"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement.getAttribute('action'), 'https://hotel-feedback.freecodecamp.org');
```

Your `form` element should have a `method` attribute set to `"POST"`.

```js
const formElement = document.querySelector('form');
assert.strictEqual(formElement.getAttribute('method'), 'POST');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
    --fcc-editable-region--

    --fcc-editable-region--
    </main>
  </body>
</html>
```
