---
id: 66a8380d911e3f4270d5cadc
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Now, it is time to add the `main` element which represents the main content of the page.

# --hints--

You should have an opening `main` tag.

```js
assert.match(code, /<main>/i);
```

You should have a closing `main` tag.

```js
assert.match(code, /<\/main>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
