---
id: 66a8347f97a33e36ffc81b9f
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

You should have an opening `body` tag.

```js
assert.match(code, /<body>/i);
```

You should have a closing `body` tag.

```js
assert.match(code, /<\/body>/i);
```

You should not change your `head` element. Make sure you did not delete your closing tag.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Your `body` element should come after your `head` element.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>

  --fcc-editable-region--
</html>
```
