---
id: 66a93c95bc58e26a8fe95818
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Next, add an `input` with the type of `"email"` to the form. This `input` should have an `id` attribute set to the value of `"email"`. Also, set the `name` attribute to the value of `"email"`. 

This `input` is also required, so make sure to add the `required` attribute.

Finally, add a `placeholder` attribute set to `"example@email.com"`.

# --hints--

You should have an `input` element.

```js
assert.strictEqual(document.querySelectorAll('input').length, 2);
```

Your `input` element should have a `type` of `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Your `input` element should have an `id` attribute set to `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Your `input` element should have a `name` attribute set to `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Your `input` element should have a `required` attribute.

```js
assert.strictEqual(document.querySelector('input#email').getAttribute('required'), "");
```

Your `input` element should have a `placeholder` attribute set to `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email').getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>
        --fcc-editable-region--
          <label for="email">Email address*:</label>

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
