---
id: 66a954b2bcddba72076c1857
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Below your first radio button, add another `label` element with the `for` attribute set to `"no-option"`. The text for the `label` should be `No`.

Below your second `label` element, add a `radio` button with the `id` set to `"no-option"`, and the `name` attribute set to `"hotel-stay"`.

When you are finished, you can now try out the radio buttons by selecting one option at a time.

# --hints--

You should have a `label` element inside your `fieldset` element.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(2) label').length, 2);
```

Your `label` should have a `for` attribute set to `"no-option"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(2) label[for="no-option"]').length, 1);
```

Your `label` should have the text of `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label[for="no-option"]').textContent, 'No');
```

You should have a second `input` inside your `fieldset` element.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(2) input').length, 2);
```

Your `input` should have a `type` of `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2)').getAttribute('type'), 'radio');
```

Your `input` should have an `id` of `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2)').id, 'no-option');
```

Your `input` should have a `name` of `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2)').name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
