---
id: 66a93730719e1f68410cce54
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

The `name` attribute is used to identify form data after it has been submitted to the server.

Here is an example of how to use the `name` attribute:

```html
<input type="email" name="email">
```

Add a `name` attribute to the `input` element with the value of `"name"`.

# --hints--

You should not remove the existing `input` element.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

You should not create a new `input` element.

```js
assert.strictEqual(document.querySelectorAll('fieldset input').length, 1);
```

Your `input` element should have a `name` attribute with the value of `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
