---
id: 66a9689b1bf24b7750898a1b
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

For the final `label` and `input` inside this `fieldset`, you will add a `label` element with the text of `Price` and the `for` attribute set to `"price"`. 

The checkbox `input` should have an `id` set to `"price"` and a `name` set to `"price"`.

Now you can test out your `form` by selecting the various checkboxes.

# --hints--

You should have a `label` element with the text of `Price`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[4].textContent, 'Price');
```

Your `label` should have the `for` attribute set to `"price"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[4].getAttribute('for'), 'price');
```

You should have an checkbox input with an `id` of `"price"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) input[type="checkbox"]')[4].getAttribute('id'), 'price');
```

You should have a checkbox input with a `name` of `"price"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) input[type="checkbox"]')[4].getAttribute('name'), 'price');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form action="example-url"> 
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="10" max="99" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes </label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No </label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input type="checkbox" id="recommendation" name="recommendation" />

          <label for="location">Location</label>
          <input type="checkbox" id="location" name="location" />

          <label for="reputation">Reputation</label>
          <input checked type="checkbox" id="reputation" name="reputation" />
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
