---
id: 66a9401c9d660d6bb15993e2
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Your hotel feedback form should have an option for users to add their age. 

Start by adding a `label` element with the text `Age(optional):` to the form.

The `for` attribute should be set to `"age"`.

# --hints--

You should have a `label` element.

```js
assert.strictEqual(document.querySelectorAll('fieldset label').length, 3);
```

Your `label` element should have a `for` attribute set to `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

You `label` element should have the text `Age(optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]').textContent, "Age(optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
        --fcc-editable-region--

        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
