---
id: 66a937e74920ba68ebe5e86d
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

In the previous lecture videos, you learned how to work with the `placeholder` and `required` attributes like this:

```html
<input type="text" placeholder="Ex. John Doe" required>
```

For your existing `input` element, add a `placeholder` attribute with the value of `"Ex. John Doe"`.

Also, add the `required` attribute to the `input` element.

# --hints--

You should not remove the existing `input` element.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

Your `input` should have a `placeholder` attribute with the value of `"Ex. John Doe"`.

```js
assert.strictEqual(document.querySelector('fieldset input').getAttribute('placeholder'), 'Ex. John Doe');
```

Your `input` should have a `required` attribute.

```js
assert.isNotNull(document.querySelector('fieldset input[required]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
