---
id: 66a96127422411756204bc1b
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Add another `label` with the text of `Personal Recommendation`. The `for` attribute should be set to `"recommendation"`.

Below the `label` element, add another checkbox `input` with the `id`, `name` and `value` attributes set to `"recommendation"`.

# --hints--

You should have a `label` element.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label").length, 2);
```

Your `label` element should have the text of `Personal Recommendation`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label")[1].textContent, "Personal Recommendation");
```

Your `label` element should have a `for` attribute set to `"recommendation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label")[1].getAttribute("for"), "recommendation");
```

You should have a checkbox `input`.

```js
assert(document.querySelector("fieldset:nth-of-type(3) input:nth-of-type(2)[type='checkbox']"));
```

Your checkbox should have an `id` set to `"recommendation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input:nth-of-type(2)[type='checkbox']").id, "recommendation");
```

Your checkbox should have a `name` attribute set to `"recommendation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input:nth-of-type(2)[type='checkbox']").name, "recommendation");
```

Your checkbox should have a `value` attribute set to `"recommendation"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input:nth-of-type(2)[type='checkbox']").value, "recommendation");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
