---
id: 66a962954f4e0d76223b37ed
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Next, add another `label` element with the text of `Location` and the `for` attribute set to `"location"`. 

For the checkbox `input`, both the `id`, `name` and `value` attributes should be set to `"location"`.

Below that `input` element, add another `label` element with the text of `Reputation` and the `for` attribute set to `"reputation"`. 

For the checkbox `input`, both the `id`, `name` and `value` attributes should be set to `"reputation"`.

# --hints--

You should have a `label` element with the text of `Location`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2].textContent, 'Location');
```

Your `label` should have the `for` attribute set to `"location"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2].getAttribute('for'), 'location');
```

You should have a third checkbox `input`.

```js 
assert(document.querySelector("fieldset:nth-of-type(3) input:nth-of-type(3)[type='checkbox']"));
```

Your third checkbox input should have an `id` set to `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) input[type='checkbox']")[2].getAttribute('id'), 'location');
```

Your third checkbox should have a `name` of `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) input[type='checkbox']")[2].getAttribute('name'), 'location');
```

Your third checkbox should have a `value` of `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) input[type='checkbox']")[2].getAttribute('value'), 'location');
```

You should have a `label` element with the text of `Reputation`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[3].textContent, 'Reputation');
```

Your `label` should have the `for` attribute set to `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[3].getAttribute('for'), 'reputation');
```

You should have a fourth checkbox `input`.

```js 
assert(document.querySelector("fieldset:nth-of-type(3) input:nth-of-type(4)[type='checkbox']"));
```

You should have a fourth checkbox `input` with an `id` of `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) input[type='checkbox']")[3].getAttribute('id'), 'reputation');
```

Your fourth checkbox should have a `name` of `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) input[type='checkbox']")[3].getAttribute('name'), 'reputation');
```

Your fourth checkbox should have a `value` of `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) input[type='checkbox']")[3].getAttribute('value'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
