---
id: 66a972137acd1179fa3fe8a0
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

Inside your `select` element, add the following five `option` elements with these corresponding values for the `option` text and `value` attribute:

**Value Attributes:**

- poor
- satisfactory
- good
- very-good
- excellent

**Option Element Text:**

- Poor
- Satisfactory
- Good
- Very Good
- Excellent

# --hints--

You should have an `option` element with the value set to `"poor"`.

```js
assert(document.querySelector('option[value="poor"]'));
```

Your `option` element with the `value` of `"poor"` should have the text `Poor`.

```js
assert.strictEqual(document.querySelector('option[value="poor"]').textContent, 'Poor');
```

You should have an `option` element with the `value` set to `"satisfactory"`.

```js
assert(document.querySelector('option[value="satisfactory"]'));
```

Your `option` element with the `value` of `"satisfactory"` should have the text `Satisfactory`.

```js
assert.strictEqual(document.querySelector('option[value="satisfactory"]').textContent, 'Satisfactory');
```

You should have an `option` element with the `value` set to `"good"`.

```js
assert(document.querySelector('option[value="good"]'));
```

Your `option` element with the `value` of `"good"` should have the text `Good`.

```js

assert.strictEqual(document.querySelector('option[value="good"]').textContent, 'Good');
```

You should have an `option` element with the value set to `"very-good"`.

```js
assert(document.querySelector('option[value="very-good"]'));
```

Your `option` element with the `value` of `"very-good"` should have the text `Very Good`.

```js
assert.strictEqual(document.querySelector('option[value="very-good"]').textContent, 'Very Good');
```

You should have an `option` element with the value set to `"excellent"`.

```js
assert(document.querySelector('option[value="excellent"]'));
```

Your `option` element with the `value` of `"excellent"` should have the text `Excellent`.

```js

assert.strictEqual(document.querySelector('option[value="excellent"]').textContent, 'Excellent');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name*:</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required>

          <label for="email">Email address*:</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />
          <label for="age">Age(optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />

          <label for="location">Location</label>
          <input type="checkbox" id="location" name="location" value="location" />

          <label for="reputation">Reputation</label>
          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />

          <label for="price">Price</label>
          <input type="checkbox" id="price" name="price" value="price" />
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          --fcc-editable-region--
          <select name="service" id="service">
            
          </select>
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
