---
id: bad87fee1348bd9aec908850
title: 给 Bootstrap 按钮添加默认样式
challengeType: 0
forumTopicId: 16657
---

# --description--

Bootstrap 还有另外一种属于按钮的 class 属性叫做 `btn-default`。

为 `button` 元素增加两个 class： `btn` 和 `btn-default`。

# --hints--

将 `btn` class 添加到所有的 `button` 元素中。

```js
assert($('.btn').length > 5);
```

将 `btn-default` class 添加到每一个 `button` 元素中。

```js
assert($('.btn-default').length > 5);
```

# --solutions--

