---
id: bad87fee1348bd9aec908856
title: 给 Bootstrap 按钮贴标签
challengeType: 0
forumTopicId: 18222
---

# --description--

正如我们标注了每个 wells 一样，我们同样想要标注每一个按钮。

为每一个 `button` 元素设置与其 `id` 选择器相同的文本。

# --hints--

给 id 为 `target1` 的 `button` 元素设置文本 `#target1`。

```js
assert(new RegExp('#target1', 'gi').test($('#target1').text()));
```

给 id 为 `target2` 的 `button` 元素设置文本 `#target2`。

```js
assert(new RegExp('#target2', 'gi').test($('#target2').text()));
```

给 id 为 `target3` 的 `button` 元素设置文本 `#target3`。

```js
assert(new RegExp('#target3', 'gi').test($('#target3').text()));
```

给 id 为 `target4` 的 `button` 元素设置文本 `#target4`。

```js
assert(new RegExp('#target4', 'gi').test($('#target4').text()));
```

给 id 为 `target5` 的 `button` 元素设置文本 `#target5`。

```js
assert(new RegExp('#target5', 'gi').test($('#target5').text()));
```

给 id 为 `target6` 的 `button` 元素设置文本 `#target6`。

```js
assert(new RegExp('#target6', 'gi').test($('#target6').text()));
```

# --solutions--

