---
id: 6579cbab9825b8170974c69a
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

You can remove a key-value pair from a dictionary by using the `del` keyword. The syntax is the following:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack`
}

del my_dict['occupation']
```

Just before your `print()` call, use the `del` keyword to delete the `age` key and its value from `copper`.

# --hints--

You should use the `del` keyword to delete `copper['age']` before the `print()` call.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^print\s*\(\s*copper\s*\)/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
