---
id: 61437d575fb98f57fa8f7f36
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

Comience con su plantilla HTML estándar. Agregue una declaración `DOCTYPE`, un elemento `html` que especifique que esta página está en inglés, un elemento `head` y un elemento `body`.

Agregue una etiqueta`<meta>` con el `charset` apropiado y una etiqueta `<meta>` para la capacidad de respuesta móvil dentro del elemento `head`.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio despues de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debe cerrar la declaración `DOCTYPE` con `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Debe tener una etiqueta de apertura `<html>`, y debe tener `lang` como `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debe tener una etiqueta de apertura `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

Debe tener una etiqueta de cierre `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

Debe tener una etiqueta de apertura `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

Debe tener una etiqueta de cierre `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

Los elementos `head` y `body` deben ser hermanos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

El elemento `head` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

El elemento `body` debe estar dentro del elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Un elemento `meta` debe tener un `name` establecido en `viewport`, y `content` establecido en `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

El otro elemento `meta` debe tener el atributo `charset` establecido en `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
