---
id: 614387cbefeeba5f3654a291
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Después de su elemento `header`, cree un `div` con la `class` establecida en `author`.

Dentro de ese `div`, crea un elemento `p` con la `class` configurada en `author-name` y dale el texto `By freeCodeCamp`. Envuelva la porción `freeCodeCamp` en un elemento `a` con `href` establecido en `https://freecodecamp.org`, y el `target` establecido en `_blank`.


Debajo de eso, agregue un segundo elemento `p` con la clase `publish-date` y el texto `March 7, 2019`.

# --hints--

Debe crear un nuevo elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Su elemento `div` debe ir después de su elemento `header`.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

Tu elemento `div` debe tener `class` establecido en `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

Debes crear dos nuevos elementos `p`.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

Sus dos nuevos elementos `p` deben estar dentro de su elemento `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Su primer elemento nuevo `p` debe tener una `class` establecida en `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

Su primer elemento nuevo `p` debe tener el texto `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

Su segundo elemento nuevo `p` debe tener una `class` establecida en `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

Su segundo elemento nuevo `p` debe tener el texto `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

Debe crear un nuevo elemento `a`.

```js
assert.exists(document.querySelector('a'));
```

Su elemento `a` debe estar dentro de su primer elemento `p` nuevo.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

Su elemento `a` debe tener `href` establecido en `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

Su elemento `a` debe tener el `target` establecido en `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

Su elemento `a` debe rodear el texto `freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.textContent === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
