---
id: 614389f601bb4f611db98563
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Debajo de su elemento `.author`, cree un nuevo elemento `div` con la clase `social-icons`.

Agregue cinco elementos `a` dentro de ese nuevo `div` y asígneles los siguientes atributos `href`.

- El primer elemento `a` debe tener un `href` establecido en `https://www.facebook.com/freecodecamp`.
- El segundo elemento `a` debe tener un `href` establecido en `https://twitter.com/freecodecamp`.
- El tercer elemento `a` debe tener un `href` establecido en `https://instagram.com/freecodecamp`.
- El cuarto elemento `a` debe tener un `href` establecido en `https://www.linkedin.com/school/free-code-camp`.
- El quinto elemento `a` debe tener un `href` establecido en `https://www.youtube.com/freecodecamp`.

# --hints--

Debe crear un nuevo elemento `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Tu nuevo elemento `div` debe ir después de tu elemento `.author`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.localName === 'div');
```

Su nuevo elemento `div` debe tener la clase `social-icons`.

```js
assert(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

Su elemento `.social-icons` debe tener cinco elementos `a`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.length === 5);
```

Tu primer elemento `a` debe tener un `href` establecido en `https://www.facebook.com/freecodecamp`.

```js
assert(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href')?.includes('https://www.facebook.com/freecodecamp'));
```

Su segundo elemento `a` debe tener un `href` establecido en `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

Su tercer elemento `a` debe tener un `href` establecido en `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

Su cuarto elemento `a` debe tener un `href` establecido en `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

Tu quinto elemento `a` debe tener un `href` establecido en `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
