---
id: 6153908a366afb4d57185c8d
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

O Flexbox tem um eixo principal e um eixo cruzado. O eixo principal é definido pela propriedade `flex-direction`, que tem quatro valores possíveis:

- `row` (padrão): eixo horizontal com os itens flex items da esquerda para a direita
- `row-reverse`: eixo horizontal com os itens flex da direita para a esquerda
- `column`: eixo vertical com os itens flex de cima para baixo
- `column-reverse`: eixo vertical com os itens flex de baixo para cima

**Observação**: os eixos e direções serão diferentes, dependendo da direção do texto. Os valores mostrados são para uma direção de texto da esquerda para a direita.

Experimente os diferentes valores para ver como eles afetam o layout.

Quando estiver pronto, defina explicitamente uma `flex-direction` de `row` no elemento `.gallery`.

# --hints--

O seletor `.gallery` deve ter uma propridade `flex-direction` definida como `row`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection === 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
