---
id: 61537485c4f2a624f18d7794
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Comece com o boilerplate padrão do HTML. Adicione a declaração `DOCTYPE` e os elementos `html`, `head` e `body`.

# --hints--

O código deve conter a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Você deve incluir um espaço após a referência `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Você deve definir o tipo de documento para `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Você deve fechar a declaração `DOCTYPE` com um `>` após o tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

O elemento `html` deve ter uma tag de abertura.

```js
assert(code.match(/<html\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

A declaração `DOCTYPE` deve estar no início do arquivo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Você deve ter uma tag de abertura para `head`.

```js
assert(code.match(/<head\s*>/i));
```

Você deve ter uma tag de fechamento para `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Você deve ter uma tag de abertura para o elemento `body`.

```js
assert(code.match(/<body\s*>/i));
```

Você deve ter uma tag de fechamento para o elemento `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Os elementos `head` e `body` devem ser irmãos.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

O elemento `head` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

O elemento `body` deve ser filho do elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
