---
id: 61537c5f81f0cf325b4a854c
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Dentro do elemento `body`, crie um elemento `div` com o atributo `class` definido como `header`.

Dentro do elemento `.header`, coloque um elemento `h1` com o texto `CSS FLEXBOX PHOTO GALLERY`.

# --hints--

Você precisa de um elemento `div` dentro do elemento `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('div'))
```

O elemento `div` deve ter para o atributo `class` definido `header`.

```js
assert(document?.querySelector('body')?.querySelector('div')?.classList?.contains('header'))
```

O elemento `.header` deve ter um elemento `h1`.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

O elemento `h1` deve ter o texto `CSS FLEXBOX PHOTO GALLERY`. Lembre-se que a distinção entre maiúsculas e minúsculas é importante.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'CSS FLEXBOX PHOTO GALLERY')
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
