---
id: 615392916d83fa4f02f7e2cf
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Talvez você tenha percebido que as imagens foram todas transferidas para a mesma linha.

A propriedade `flex-wrap` determina como seus itens devem se comportar quando o contêiner flex é muito pequeno. Definir esta propriedade como `wrap` permitirá que seus itens sejam envolvidos na próxima linha/coluna (dependendo do seu eixo principal), enquanto `nowrap` impedirá esse movimento. Quando esse valor estiver definido como `nowrap`, os itens poderão encolher para se encaixarem ou ultrapassar o limite.

Dê ao seletor `#gallery` a propriedade `flex-wrap` com o valor `wrap`. Você deve ver as imagens assumirem um layout de quatro colunas. Isso porque você definiu a `width` como `25%` em um passo anterior.

# --hints--

O seletor `#gallery` deve ter a propriedade `flex-wrap` definida como `wrap`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery')?.flexWrap === 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

.header {
  text-align: center;
  padding: 32px;
  background: #E0DDDD;
}

--fcc-editable-region--
#gallery {
  display: flex;
  flex-direction: row;
}
--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
