---
id: 61537a8054753e2f1f2a1574
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de seu elemento `head`, adicione uma tag `meta` com `name` definido como `viewport` e `content` definido como `width=device-width, initial-scale=1`.

Adicione também uma tag `meta` com `charset` definido como `UTF-8`.

# --hints--

Você deve ter dois elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Um elemento `meta` deve ter o atributo `name` definido como `viewport` e o atributo `content` definido como `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

O outro elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
