---
id: 615389bd81347947ea7ba896
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Flexbox é uma abordagem de layout de CSS unidimensional que se concentra no fluxo de conteúdo. Ele oferece a capacidade de controlar a maneira como os itens são espaçados e alinhados dentro de um contêiner.

Para definir um elemento para que use o Flexbox, você dá a ele uma propriedade `display` definida como `flex`. Crie um seletor de `#gallery` e dê a ele essa propriedade.

# --hints--

Você deve ter um seletor `#gallery`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery'));
```

O seletor `#gallery` deve ter a propriedade `display` definida como `flex`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery')?.display === 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

.header {
  text-align: center;
  padding: 32px;
  background: #E0DDDD;
}

--fcc-editable-region--

--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
