---
id: 6153897c27f6334716ee5abe
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Alinhe o texto de `.header` no centro. Torne o texto maiúsculo usando a propriedade `text-transform` com `uppercase` como valor.

Dê a ele um espaçamento de `32px` em todos os lados. Defina os valores de cor do fundo como `#0a0a23` e do texto como `#fff`.

Adicione uma `border-bottom` com o valor de `4px solid #fdb347`.


# --hints--

Você deve ter um seletor `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

O seletor `.header` deve definir a propriedade `text-align` como `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign === 'center');
```

O seletor `.header` deve definir a propriedade `text-transform` como `uppercase`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform === 'uppercase');
```

O seletor `.header` deve ter a propriedade `padding` definida como `32px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.padding === '32px');
```

O seletor `.header` deve ter a propriedade `background-color` definida como `#0a0a23`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor === 'rgb(10, 10, 35)');
```

O seletor `.header` deve ter a propriedade `color` definida como `#fff`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.color === 'rgb(255, 255, 255)');
```

O seletor `.header` deve ter a propriedade `border-bottom` definida como `4px solid #fdb347`.

```js
console.log(new __helpers.CSSHelp(document).getStyle('.header'))
assert(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom === '4px solid rgb(253, 179, 71)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
