---
id: 61537c9eecea6a335db6da79
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Abaixo do elemento `.header`, crie outro elemento `div` e atribua a ele uma `class` de `gallery`. Essa `div` atuará como um contêiner para as imagens da galeria.

Dentro de `.gallery`, crie nove elementos `img`.

# --hints--

Você deve criar um elemento `div` no elemento `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

O novo elemento `div` deve ter a `class` definida como `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

O novo elemento `div` deve estar depois do elemento `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

O elemento `.gallery` deve ter nove elementos `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
