---
id: 61537c5f81f0cf325b4a854c
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Adicione um elemento `header` dentro do elemento `body` e atribua a ele uma classe `header`.

Dentro de `header`, crie um `h1` com o texto `css flexbox photo gallery`.

# --hints--

Você precisa de um elemento `header` dentro do elemento `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

O elemento `header` deve ter a `class` definida como `header`.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

O elemento `header` deve ter um elemento `h1` dentro de si.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

O elemento `h1` deve conter o texto `css flexbox photo gallery` dentro de si.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
