---
id: 61537c9eecea6a335db6da79
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Abaixo do elemento `.header`, crie um elemento `div` com o `id` definido para `gallery`.

Nesse elemento `#gallery`, crie dez elementos `img`.

# --hints--

Você deve criar um segundo elemento `div` no elemento `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length === 2);
```

O novo elemento `div` deve estar depois do elemento `.header`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[0]?.classList?.contains('header'));
```

O novo elemento `div` deve ter para o atributo `id` o valor de `gallery`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[1]?.id === 'gallery');
```

O elemento `#gallery` deve ter dez elementos `img`.

```js
assert(document.querySelector('#gallery')?.querySelectorAll('img')?.length === 10);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
  </body>
</html>
--fcc-editable-region--
```

```css

```
