---
id: 615f171d05def3218035dc85
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

As imagens são muito grandes. Crie um seletor `.gallery img` para localizar suas imagens. Dê a elas uma `width` de `100%` e uma `max-width` de `350px` para que encolham conforme necessário, mas não fiquem muito grandes.

Também defina a propriedade `height` para `300px` para manter suas imagens com um tamanho uniforme.

# --hints--

Você deve ter um seletor `.gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

O seletor `.gallery img` deve ter uma propridade `width` definida como `100%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width === '100%');
```

O seletor `.gallery img` deve ter uma propridade `max-width` definida como `350px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth === '350px');
```

O seletor `.gallery img` deve ter uma propridade `height` definida como `300px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
