---
id: 615f171d05def3218035dc85
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Atualmente, suas imagens são muito grandes e você não será capaz de ver suas alterações no CSS.

Crie um seletor `#gallery img` para localizar suas imagens. Dê ao seletor uma propriedade de `width` definida para `25%`.

Também defina a propriedade `height` para `300px` para manter suas imagens com um tamanho uniforme.

# --hints--

Você deve ter um seletor `#gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#gallery img'))
```

O seletor `#gallery img` deve ter uma propridade de `width` definida para `25%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery img')?.width === '25%')
```

O seletor `#gallery img` deve ter uma propriedade de `height` definida para `300px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
