---
id: 6494da0daf5df5197963671d
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Observe como a borda azul da imagem se estende para além da borda vermelha da galeria. Isto é devido ao modo como os navegadores calculam o tamanho dos elementos do contêiner.

A propriedade `box-sizing` é usada para definir esse comportamento. Por padrão, o modelo `content-box` é usado. Com esse modelo, quando um elemento tem uma largura específica, essa largura é calculada apenas com base no conteúdo do elemento. Os valores de preenchimento e das bordas são adicionados à largura total. Assim, o elemento cresce para acomodar esses valores.

Tente definir o `box-sizing` para `content-box` explicitamente, com o seletor global `*`. Nesse ponto, você não verá nenhuma mudança, porque você está usando o valor padrão.

# --hints--

Você deve ter um seletor `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

O seletor `*` deve ter a propriedade `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

O seletor `*` deve ter a propriedade `box-sizing` definida como `content-box`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
