---
id: 64dd076a1b33c86d84b4232f
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Para visualizar melhor o tamanho dos elementos, pode ser útil adicionar uma borda.

Dê ao elemento `.gallery` uma `width` de `50%` e uma `border` de `5px solid red`.

Em seguida, dê ao elemento `img` uma `width` de `100%`, um `padding` de `5px` e uma `border` de `5px solid blue`.

# --hints--

Você deve ter um seletor `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

O seletor `.gallery` deve ter uma propriedade `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

O seletor `.gallery` deve ter a propriedade `width` definida como `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

O seletor `.gallery` deve ter uma propriedade `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

O seletor `.gallery` deve ter a propriedade `border` definida como `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Você deve ter um seletor `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

O seletor `img` deve ter uma propriedade `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

O seletor `img` deve ter a propriedade `width` definida como `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

O seletor `img` deve ter uma propriedade `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

O seletor `img` deve ter a propriedade `padding` definida como `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

O seletor `img` deve ter uma propriedade `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

O seletor `img` deve ter a propriedade `border` definida como `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
