---
id: 6690e10ebe2181212abc9652
title: Design a Business Card
challengeType: 14
dashedName: design-a-business-card
---

# --description--

**Objective:** Design a Business card that looks similar to <a href="https://codepen.io/rancid_beet/full/ExBaJdw" target="_blank" rel="noopener noreferrer nofollow">https://business-card.freecodecamp.rocks</a>. Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**
1. Your page background should be `rosybrown`. The overall font should be `Arial`, with a fallback of `sans-serif`.
   
2. All of the `p` elements should have a top and bottom margin of `5px` and a left and right margin of `0px`.

3. All links on the page should have a default of no underline. When you hover over the links, there should be an underline.

4. You should have a `div` with the class `business-card`. 
   
5. The `business-card` class should be `300px` wide and it should have a background color of `lavenderblush`. There should be a padding of `20px` all around and a top margin of `100px`. The text should be center aligned and the font size should be `16px`. The left and right margins should be set to `auto`.

6. Inside the `buisness-card` div, there should be an `img` element with the class `profile-image`. The image source should be set to `flower.jpg` and there should be an `alt` attribute should be set to a meaningful description. The image size should be `150px by 150px`. The bottom margin should be `10px` and there should be a `2px` solid `lightgray` border around the image.

7. Inside the `business-card` `div`, after the `img` element, you should have three `p` elements with the class `full-name`, `designation`, and `company`, respectively.

8. The first `p` element should contain your name. Its font size should be `24px` and the `color` should be `darkslategray`.

9. The second `p` element should contain your designation. Its `font-size` property should be `18px` and the `color` should be set to `gray`.

10. The third `p` element should have your company name. Its font size should be `16px` and the `color` should be set to `darkgray`.

11. There should be an `hr` element below the last `p` element. The `hr` element in general should have a `1px` dashed top border in `slategray` color. There should be a top and bottom `margin` of `20px` and a left and right `margin` of `0px`.

12. After the `hr` element, there should be three `p` elements. In the first `p` element, you should have an email address as the text. In the second `p` element, the text should be a phone number. The third `p` element should have the text `Portfolio: ` and an `a` element with any valid link. Add an `hr` element after the third `p` element.

13. You should have another `div` with the class `social-media`. This div should have a top margin of `10px` and a font size of `14px`. The links inside this div should be `maroon` and should have a margin of `10px` all around. Within this class, there should be a `p` element with the text `Connect with me`. 

14. Inside the `social-media` div, there should be three `a` elements. The `a` elements should have the text `Twitter`, `LinkedIn` and `GitHub` respectively with links to valid sites.

# --hints--

Your page background colour should be `rosybrown`.

```js
assert($('body').css('background-color') === 'rosybrown');
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class = "company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <p>Portfolio: <a href="http://example.com">example.com</a></p>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

a:hover {
    text-decoration: underline;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    border-top: 1px dashed slategray;
    margin: 20px 0;
}
```
