---
id: 699d20623e57e8ffee88545b
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Let's verify the search works. Stephen King has six books in the catalog.

Call `findByAuthor(catalog, "king")` and assign the result to a variable named `kingBooks`. Log `kingBooks.length` to confirm the count.

Then use a `for` loop to log each result. Use a template literal to include the title and year, like `` `${kingBooks[i].title} (${kingBooks[i].year})` ``.

# --hints--

`kingBooks` should have a length of `6`.

```js
assert.strictEqual(kingBooks.length, 6);
```

You should log `kingBooks.length` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*kingBooks\.length\s*\)/);
```

You should use a `for` loop to log each book's title and year.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\([^)]*kingBooks\.length[^)]*\)/);
assert.match(__helpers.removeJSComments(code), /kingBooks\[.+\]\.title\b/);
assert.match(__helpers.removeJSComments(code), /kingBooks\[.+\]\.year\b/);
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

--fcc-editable-region--

--fcc-editable-region--
```
