---
id: 699d20623e57e8ffee885452
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now extract the individual fields from the trimmed array. Declare four variables: `title` assigned to `trimmedParts[0]`, `author` assigned to `trimmedParts[1]`, `year` assigned to `trimmedParts[2]`, and `location` assigned to `trimmedParts[3]`.

# --hints--

You should declare a variable named `title` with the value `trimmedParts[0]`.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+title\s*=\s*trimmedParts\s*\[\s*0\s*\]/);
```

You should declare a variable named `author` with the value `trimmedParts[1]`.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+author\s*=\s*trimmedParts\s*\[\s*1\s*\]/);
```

You should declare a variable named `year` with the value `trimmedParts[2]`.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+year\s*=\s*trimmedParts\s*\[\s*2\s*\]/);
```

You should declare a variable named `location` with the value `trimmedParts[3]`.

```js
assert.match(parseCard.toString(), /(const|let|var)\s+location\s*=\s*trimmedParts\s*\[\s*3\s*\]/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
--fcc-editable-region--

  return trimmedParts;
--fcc-editable-region--
}
```
