---
id: 699d20623e57e8ffee885453
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Now that you have extracted the individual fields, replace the `return trimmedParts;` statement with a `return` that gives back an object with four properties: `title`, `author`, `year`, and `location`.

Some catalog cards have missing data, meaning an empty string for a field. The logical OR operator (`||`) returns the right-hand side when the left side is falsy. An empty string is falsy, so `"" || "Unknown"` evaluates to `"Unknown"`.

Set the `title` property to `title || "Unknown"` and the `author` property to `author || "Unknown"`. For now, set `year` and `location` to their raw variable values. You will add their fallbacks in the next step.

# --hints--

`parseCard` should return an object.

```js
assert.isObject(parseCard("A | B | C | D"));
```

The returned object should have a `title` property that falls back to `"Unknown"` for an empty title.

```js
assert.strictEqual(parseCard(" |Author|2000|Shelf").title, "Unknown");
```

The returned object should have an `author` property that falls back to `"Unknown"` for an empty author.

```js
assert.strictEqual(parseCard("Title| |2000|Shelf").author, "Unknown");
```

The returned object should have `year` and `location` properties.

```js
const result = parseCard("Title|Author|2000|Shelf");
assert.property(result, "year");
assert.property(result, "location");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
--fcc-editable-region--

  return trimmedParts;
--fcc-editable-region--
}
```
