---
id: 699d1f8e51fcdc8230f779fb
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will digitize historical catalog cards for a heritage library. You will parse raw text data into objects, search and group entries, render formatted output, validate data, and export to different formats.

Each catalog card is stored as a single string in the format `"Title | Author | Year | Location"`, with fields separated by pipe (`|`) characters.

Create an array named `rawCatalogCards` containing the following four strings:

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1",
"The Stand | King, Stephen | 1978 | Shelf K2",
"It | King, Stephen | 1986 | Shelf K3",
```

# --hints--

You should declare a variable named `rawCatalogCards`.

```js
assert.isDefined(rawCatalogCards);
```

`rawCatalogCards` should be an array.

```js
assert.isArray(rawCatalogCards);
```

`rawCatalogCards` should have 4 elements.

```js
assert.lengthOf(rawCatalogCards, 4);
```

The first element should be `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"`.

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

The last element should be `"It | King, Stephen | 1986 | Shelf K3"`.

```js
assert.strictEqual(rawCatalogCards[3], "It | King, Stephen | 1986 | Shelf K3");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
