---
id: 699d20623e57e8ffee885450
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

You need a function to turn each raw string into a structured object. Create a function named `parseCard` that takes a `rawString` parameter.

Inside the function, declare a variable named `parts` and assign to it the result of calling `.split("|")` on `rawString`. The `.split()` method breaks a string into an array of substrings at every occurrence of the separator you provide.

Return `parts` from the function.

# --hints--

You should have a function named `parseCard`.

```js
assert.isFunction(parseCard);
```

`parseCard` should take one parameter.

```js
assert.lengthOf(parseCard, 1);
```

`parseCard("A | B | C | D")` should return an array.

```js
assert.isArray(parseCard("A | B | C | D"));
```

`parseCard("A | B | C | D")` should return an array with 4 elements.

```js
assert.lengthOf(parseCard("A | B | C | D"), 4);
```

`parseCard("Test|Author|2000|Shelf A1")` should return an array whose first element is `"Test"`.

```js
const result = parseCard("Test|Author|2000|Shelf A1");
assert.strictEqual(result[0], "Test");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

--fcc-editable-region--

--fcc-editable-region--
```
