---
id: 699d20623e57e8ffee885450
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

The rest of the catalog cards have been added for you, including three entries near the end with missing data: blank fields where an author, a year, or a location should be. Real-world data often has gaps like these. You will handle those later.

Now you need a function to turn each raw string into a structured object. Create a function named `parseCard` that takes a `rawString` parameter.

Inside the function, declare a variable named `parts` and assign to it the result of calling `.split("|")` on `rawString`. The `.split()` method breaks a string into an array of substrings at every occurrence of the separator you provide.

Each part will still have extra whitespace from the pipe-separated format. For example, splitting `"The Shining | King, Stephen"` produces `["The Shining ", " King, Stephen"]`.

Declare an empty array named `trimmedParts`. Use a `for` loop to iterate over `parts` and push `parts[i].trim()` into `trimmedParts` on each iteration.

Return `trimmedParts` from the function.

# --hints--

You should have a function named `parseCard`.

```js
assert.isFunction(parseCard);
```

Your `parseCard` function should have a `rawString` parameter.

```js
assert.match(parseCard.toString(), /rawString/);
```

You should use a `for` loop inside `parseCard`.

```js
assert.match(parseCard.toString(), /for\s*\(/);
```

You should declare an empty array named `trimmedParts` inside `parseCard`.

```js
assert.match(parseCard.toString(), /trimmedParts\s*=\s*\[\s*\]/);
```

You should push trimmed parts into `trimmedParts`.

```js
assert.match(parseCard.toString(), /trimmedParts\.push/);
```

`parseCard` should return a value.

```js
assert.exists(parseCard("A | B | C | D"));
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

--fcc-editable-region--

--fcc-editable-region--
```
