---
id: 699d20623e57e8ffee885451
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Each part still has extra whitespace from the pipe-separated format. For example, splitting `"The Shining | King, Stephen"` produces `["The Shining ", " King, Stephen"]`.

Declare an empty array named `trimmedParts`. Then use a `for` loop to iterate over `parts` and push `parts[i].trim()` into `trimmedParts` on each iteration.

A `for` loop repeats a block of code a set number of times. The syntax is:

```js
for (let i = 0; i < array.length; i++) {
  // code to repeat
}
```

The variable `i` starts at `0`, the loop continues while `i < array.length`, and `i++` increments `i` by one after each iteration.

# --hints--

You should declare an empty array named `trimmedParts`.

```js
assert.match(__helpers.removeJSComments(code), /trimmedParts\s*=\s*\[\s*\]/);
```

You should use a `for` loop to iterate over `parts`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(/);
```

You should push `parts[i].trim()` into `trimmedParts` inside the loop.

```js
assert.match(__helpers.removeJSComments(code), /trimmedParts\.push\s*\(\s*parts\s*\[\s*i\s*\]\s*\.trim\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

  return parts;
--fcc-editable-region--
}
```
