---
id: 60a3e3396c7b40068ad69973
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Ändere die `background-color` der Canvas auf `#4d0f00`.

# --hints--

Du solltest die `background-color`-Eigenschaft auf `#4d0f00` setzen.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'rgb(77, 15, 0)');
assert(hasBackground);
```

Dein `.canvas`-Element sollte eine `background-color` von `#4d0f00` enthalten.

```js
const canvasBackground = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('background-color');
assert(canvasBackground === 'rgb(77, 15, 0)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
