---
id: 60a3e3396c7b40068ad69979
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Füge innerhalb deines `.canvas`-Elements ein neues `div`-Element hinzu.

Weise dem neuen `div` das `class`-Attribut mit einem Wert von `one` zu. Dies wird dein erstes Rechteck sein.

# --hints--

Du solltest ein neues `div`-Element erstellen.

```js
assert(document.querySelectorAll('div').length === 3);
```

Du solltest das neue `div`-Element in dein `.canvas`-Element einbetten.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

Dein neues `div` sollte ein `class`-Attribut mit einem Wert von `one` enthalten.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
