---
id: 60a3e3396c7b40068ad6996c
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In dem CSS-Box-Modell wird jedes HTML-Element wie eine Box mit vier Bereichen behandelt.

Stelle dir vor, dass du eine Box von deinem Lieblings-Onlinehändler erhältst - der Inhalt ist der Artikel in der Box, oder in unserem Fall, ein Header, Paragraph oder Bildelement.

Ändere das `src`-Attribut im `<img>` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` zu `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

Das `img`-Element sollte eine `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png` enthalten.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

