---
id: 60a3e3396c7b40068ad6997f
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

Durch das Hinzufügen von 1 Pixel Padding oben, unten, links und rechts der Canavs änderten sich die Dimensionen zu 502 Pixel x 602 Pixel.

Replace the `padding` property with `overflow` set to `hidden` - changing the canvas back to its original dimensions.

# --hints--

Du solltest die `padding`-Eigenschaft aus dem `.canvas`-Selektor entfernen.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

Du solltest die `overflow`-Eigenschaft auf `hidden` setzen.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

Dein `.canvas`-Element sollte einen `overflow`-Wert von `hidden` enthalten.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
