---
id: 60a3e3396c7b40068ad6996b
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge innerhalb des `head`-Elements ein `meta`-Tag hinzu, das das `charset` auf `UTF-8` setzt und ein `title`-Element mit dem Wert `Rothko Painting` einfügt.

Within the `body` element, add an `img` element with a `src` of `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

Dein Code sollte ein `meta`-Tag enthalten.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Das `meta`-Tag sollte das `charset` auf `UTF-8` einstellen.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Dein Code sollte ein `title`-Element enthalten.

```js
assert(document.querySelectorAll('title').length === 1);
```

Der `title` sollte `Rothko Painting` lauten.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Dein Code sollte ein `img`-Element enthalten.

```js
assert(document.querySelectorAll('img').length === 1);
```

Das `img`-Element sollte ein `src`-Attribut mit dem Wert `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` enthalten.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
