---
id: 60a3e3396c7b40068ad69980
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Füge direkt unter deinem `one`-Element ein weiteres `div` mit einem `class`-Wert von `two` hinzu. Dies wird dein zweites Rechteck sein.

# --hints--

Du solltest das vorhandene `.one`-Element nicht ändern.

```js
assert(document.querySelectorAll('.one').length === 1);
```

Du solltest ein zweites `div`-Element innerhalb deines `.canvas`-Elements haben.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

Dein zweites `div`-Element sollte einen `class`-Wert von `two` enthalten.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

Dein `.two`-Element sollte nach deinem `.one`-Element stehen.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
