---
id: 60a3e3396c7b40068ad69978
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Verwende Margins, um den Abstand außerhab eines Elements anzupassen.

Gib dem `.frame`-Element unter der Verwendung der `margin`-Eigenschaft eine vertikale Margin in Höhe von `20px` und eine horizontale Margin mit dem Wert `auto`. Dies wird den Rahmen um 20 Pixel nach unten verschieben und horizontal auf der Seite zentrieren.

# --hints--

Du solltest die `margin`-Eigenschaft auf `20px auto` setzen.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

Dein `.frame`-Element sollte einen `margin`-Wert von `20px auto` haben.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
