---
id: 60a3e3396c7b40068ad6998c
title: Schritt 34
challengeType: 0
dashedName: step-34
---

# --description--

Die Farben und Formen deines Bildes sind zu scharf, um als Rothko durchzugehen.

Verwende die `filter`-Eigenschaft, um das Bild um `2px` im `.canvas`-Element zu `blur`.

Hier ist ein Beispiel einer Regel, die ein 3px `blur` hinzufügt:

```css
p {
  filter: blur(3px);
}
```

# --hints--

Du solltest die `filter`-Eigenschaft auf `blur(2px)` gesetzt haben.

```js
const hasFilter = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.filter === 'blur(2px)');
assert(hasFilter);
```

Dein `.canvas`-Element sollte einen `filter`-Wert von `blur(2px)` enthalten.

```js
const canvasFilter = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('filter');
assert(canvasFilter === 'blur(2px)');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;

}
--fcc-editable-region--

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
