---
id: 60a3e3396c7b40068ad6996e
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Die Margin ist der Bereich außerhalb der Box und kann genutzt werden, um den Abstand zwischen anderen Boxen oder Elementen zu steuern.

Hier hat das untere Element eine größere obere Margin, wodurch es auf der Seite weiter nach unten verschoben wird.

Lass uns nun, da du das CSS-Boxmodell verstehst, mit dem Rothko-Bild beginnen.

Entferne das `<img>`-Element.

# --hints--

Du solltest in deinem Code kein `img`-Element haben.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
