---
id: 60a3e3396c7b40068ad69972
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Add the `height` property with the value `600px` to your `.canvas` rule.

# --hints--

Du solltest die `height`-Eigenschaft auf `600px` setzen.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

Dein `.canvas`-Element sollte eine `height` von `600px` haben.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
