---
id: 60a3e3396c7b40068ad69974
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Every painting needs a frame.

Bette das `.canvas`-Element in ein anderes `div` ein. Weise diesem `div` die `frame`-Klasse zu.
# --hints--

Du solltest ein neues `div`-Element hinzufügen.

```js
assert(document.querySelectorAll('div').length === 2)
```

Dein `.canvas` Element sollte in dem neuen `div` Element eingebettet sein.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Dein neues `div` sollte eine `class` mit dem Wert `frame` enthalten.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Dein neues `div` sollte sich innerhalb deines `body`-Elements befinden.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
