---
id: 60a3e3396c7b40068ad69975
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Schreibe eine neue Regel, indem du den `.frame`-Klassenselektor verwendest.

Nutze die `border`-Kurzbezeichnung, um dem `.frame`-Element einen soliden, schwarzen Rand mit einer Breite von `50px` zu geben.

# --hints--

Dein Code sollte einen `.frame`-Selektor haben.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Du solltest die `border`-Eigenschaft auf `50px solid black` setzen.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Your `.frame` element should have a `50px solid black` `border`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
