---
id: 60a3e3396c7b40068ad69987
title: Schritt 29
challengeType: 0
dashedName: step-29
---

# --description--

Du musst nicht immer Pixel verwenden, um die Größe eines Element anzupassen.

Erstelle eine neue Regel, `.three`, und setze ihre `width` auf `91%`.

# --hints--

Du solltest den `.three`-Selektor verwenden.

```js
const hasThree = new __helpers.CSSHelp(document).getStyle('.three');
assert(hasThree);
```

Du solltest die `width`-Eigenschaft auf `91%` setzen.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '91%');
assert(hasWidth);
```

Dein `.three`-Element sollte einen `width`-Wert von `91%` haben.

```js
const threeWidth = new __helpers.CSSHelp(document).getStyle('.three')?.getPropertyValue('width');
assert(threeWidth === '91%');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}

--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
