---
id: 60a3e3396c7b40068ad6996d
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Der Inhalt ist von einem Zwischenraum namens Padding umgeben, der ähnlich wie ein Blasenumschlag ein Element von der ihn umgebenden Box trennt.

Stelle dir den Rand wie den Karton vor, in dem dein Artikel versandt wurde.

Ändere das `src`-Attribut zu `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

Das `img`-Element sollte eine `src` von `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png` enthalten.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
