---
id: 66ed41f912d0bb1dc62da5dd
title: Build a JavaScript Trivia Bot
challengeType: 14
dashedName: lab-javascript-trivia-bot
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should log `"Hello! I'm your coding fun fact guide!"` to the console as a greeting message to the user.
1. You should create three variables: `botName`, `botLocation`, and `favoriteLanguage`, that store the bot's name, where it's from, and its favorite coding language, respectively.
1. You should use string concatenation to log `"My name is (botName) and I live on (botLocation)."` to the console.
1. You should use string concatenation to log `"My favorite programming language is (favoriteLanguage)."` to the console.
1. You should use `let` to create a `codingFact` variable and assign it a string that is a fun fact about your bot's favorite coding language and uses string concatenation to include the use of the `favoriteLanguage` variable.
1. You should log the `codingFact` to the console.
1. You should reassign the `codingFact` variable to a new fact about the bot's favorite language using the `favoriteLanguage` variable again.
1. You should log the `codingFact` to the console again.
1. You should reassign the `codingFact` variable again to another new fact about the bot's favorite language using the `favoriteLanguage` variable.
1. You should log the `codingFact` to the console a third time.
1. You should log `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` to the console as a farewell statement from the bot.

# --hints--

You should log `"Hello! I'm your coding fun fact guide!"` to the console.

```js
assert.equal(output[0], "Hello! I'm your coding fun fact guide!")
```

You should declare a `botName` variable and give it a value.

```js
assert.exists(botName);
assert.isString(botName);
```

You should declare a `botLocation` variable and give it a value.

```js
assert.exists(botLocation);
assert.isString(botLocation);
```

You should declare a `favoriteLanguage` variable and give it a value.

```js
assert.exists(favoriteLanguage);
assert.isString(favoriteLanguage);
```

You should log to the console `"My name is (botName) and I live on (botLocation)."` using concatenation to add the variables to the string.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(output[1], `My name is ${botName} and I live on ${botLocation}.`)
assert.match(codeWithoutComments, /is ("|')\s*\+\s*botName\s*\+\s*("|') and I live on \2\s*\+\s*botLocation\s*\+\s*('|")\./)
```

You should log to the console `"My favorite programming language is (favoriteLanguage)."` using concatenation to add the variable to the string.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(output[2], `My favorite programming language is ${favoriteLanguage}.`)
assert.match(codeWithoutComments, /language is ('|")\s*\+\s*favoriteLanguage\s*\+\s*('|")\./);
```

You should use `let` to declare a new variable `codingFact`.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*("|').+?\2\s*\+\s*favoriteLanguage/g)
assert.match(code, /let\s+codingFact/)
assert.include(first, 'let');
assert.exists(first);
```

You should give `codingFact` a value  that includes `favoriteLanguage` using concatenation.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*("|').+?\2\s*\+\s*favoriteLanguage/g)
assert.match(code, /let\scodingFact\s*=\s*("|').+?\1\s*\+\s*favoriteLanguage/)
assert.include(first, 'let');
assert.exists(first);
```

You should log `codingFact` to the console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(output[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

You should assign a new value to `codingFact` that also contains `favoriteLanguage`, and log it to the console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*("|').+?\2\s*\+\s*favoriteLanguage/g)
assert.include(output[4], favoriteLanguage);
assert.notEqual(output[4], output[3]);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.exists(second);
```

You should assign a value to `codingFact` for the third time that also contains `favoriteLanguage`, and log it to the console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
const [first, second, third] = codeWithoutComments.match(/(let )?\s*codingFact\s*=\s*("|').+?\2\s*\+\s*favoriteLanguage/g)
assert.include(output[5], favoriteLanguage);
assert.notEqual(output[5], output[4]);
assert.equal(output[5], codingFact);
assert.lengthOf(loggingCodingFacts, 3);
assert.exists(third);
```

You should log to the console `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` using concatenation to add the values of the variables.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
assert.equal(output[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
assert.match(codeWithoutComments, /\. Goodbye! - ("|')\s*\+\s*botName\s*\+\s*('|") from \2\s*\+\s*botLocation\s*\+\s*("|')\./)
```

# --seed--

## --before-user-code--

```js
const temp = console.log
const output = []
console.log = function (...args) {
temp(...args)
output.push(...args)
}
```

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
