---
id: 6842d5b537c64266420faf62
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

Add a key `age` with the value of the integer `34` and a key `gender` with the value of the string `Female` to your dictionary. Don't forget the comma between the key-value pairs.

# --hints--

Your `medical_records` list should contain a single dictionary.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

You should have a `age` key with the value of the integer `34` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

You should have a `gender` key with the value of the string `Female` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

You should have a `patient_id` key with the value of the string `P1001` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [
    {
        'patient_id': 'P1001'
    }    
]
--fcc-editable-region--
```
