---
id: 6842d659a4ae7973dabab3b0
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Complete the dictionary by adding the following three key-value pairs:

- A key `diagnosis` with the value of `Hypertension`
- A key `medications` with the value of the list `['Lisinopril']`
- A key `last_visit_id` with the value of `V2301`.

# --hints--

Your `medical_records` list should contain a single dictionary.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

You should have a `diagnosis` key with the value of the string `Hypertension` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

You should have a `medications` key with the value of the list `['Lisinopril']` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

You should have a `last_visit` key with the value of the string `V2301` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit", None) == "V2301"
`) })
```

You should have a `patient_id` key with the value of the string `P1001` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

You should have a `age` key with the value of the integer `34` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

You should have a `gender` key with the value of the string `Female` in your dictionary.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female'
    }    
]
--fcc-editable-region--
```
