---
id: 66faa20000000000000000da
title: Debug a Broken Sprite Sheet
challengeType: 26
dashedName: debug-a-broken-sprite-sheet
---

# --description--

**Objective:** A previous teammate shipped this scene and the sprite sheets render wrong, the hero won't animate, and the coins won't disappear. The bugs are all in `script.js`. Fix them so the scene satisfies every user story below.

The Phaser library is loaded for you on the global `Phaser` object. The hero, coin, and slime sprite sheets are hosted under `/curriculum-assets/phaser/spritesheets/`.

**User Stories:**

1. The hero sprite sheet should load from `/curriculum-assets/phaser/spritesheets/hero-walk.png`.
1. The hero sprite sheet should be loaded with a `frameWidth` of `16` and a `frameHeight` of `16`.
1. The coin sprite sheet should be loaded from `/curriculum-assets/phaser/spritesheets/coin.png` with `frameWidth: 16` and `frameHeight: 16`.
1. `this.player.setCollideWorldBounds(true)` should be called so the hero stays inside the canvas.
1. `this.player.setScale(2)` should be called so the hero renders at 32×32 instead of 16×16.
1. The `walk-left` animation should use frames 8 through 11 of `hero`.
1. The `walk-right` animation should use frames 12 through 15 of `hero`.
1. The `coin-spin` animation should loop forever with `repeat: -1`.
1. Each coin sprite should play the `coin-spin` animation inside the `forEach`.
1. `this.physics.add.overlap(this.player, this.coins, this.collectCoin, null, this)` should be registered.
1. `this.player.setVelocity(0)` should run at the top of `update` so input doesn't accumulate.
1. Pressing the left arrow should call `this.player.setFlipX(false)`, not `true`.
1. Pressing the right arrow should call `this.player.anims.play("walk-right", true)`.
1. When no arrow key is pressed, an `else` branch should play the `idle` animation with `this.player.anims.play("idle", true)`.

# --hints--

The hero sprite sheet should load from `/curriculum-assets/phaser/spritesheets/hero-walk.png`.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])hero\1\s*,\s*(['"`])\/curriculum-assets\/phaser\/spritesheets\/hero-walk\.png\2/);
```

The hero sprite sheet should be loaded with `frameWidth: 16` and `frameHeight: 16`.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])hero\1[\s\S]*?frameWidth\s*:\s*16[\s\S]*?frameHeight\s*:\s*16/);
```

The coin sprite sheet should be loaded with `frameWidth: 16` and `frameHeight: 16`.

```js
assert.match(code, /this\.load\.spritesheet\(\s*(['"`])coin\1[\s\S]*?frameWidth\s*:\s*16[\s\S]*?frameHeight\s*:\s*16/);
```

The player should call `setCollideWorldBounds(true)`.

```js
assert.match(code, /this\.player\.setCollideWorldBounds\(\s*true\s*\)/);
```

The player should call `setScale(2)`.

```js
assert.match(code, /this\.player\.setScale\(\s*2\s*\)/);
```

The `walk-left` animation should use frames `start: 8, end: 11` of `hero`.

```js
assert.match(code, /key\s*:\s*(['"`])walk-left\1[\s\S]*?frames\s*:\s*this\.anims\.generateFrameNumbers\(\s*(['"`])hero\2\s*,\s*\{\s*start\s*:\s*8\s*,\s*end\s*:\s*11\s*\}\s*\)/);
```

The `walk-right` animation should use frames `start: 12, end: 15` of `hero`.

```js
assert.match(code, /key\s*:\s*(['"`])walk-right\1[\s\S]*?frames\s*:\s*this\.anims\.generateFrameNumbers\(\s*(['"`])hero\2\s*,\s*\{\s*start\s*:\s*12\s*,\s*end\s*:\s*15\s*\}\s*\)/);
```

The `coin-spin` animation should loop with `repeat: -1`.

```js
assert.match(code, /key\s*:\s*(['"`])coin-spin\1[\s\S]*?repeat\s*:\s*-1/);
```

Each coin sprite should play `coin-spin` inside the `forEach`.

```js
assert.match(code, /coin\.anims\.play\(\s*(['"`])coin-spin\1/);
```

You should register an overlap between `this.player` and `this.coins` that calls `this.collectCoin`.

```js
assert.match(code, /this\.physics\.add\.overlap\(\s*this\.player\s*,\s*this\.coins\s*,\s*this\.collectCoin\s*,\s*null\s*,\s*this\s*\)/);
```

`update` should reset the player's velocity to `0` at the top.

```js
assert.match(code, /update\s*\(\s*\)\s*\{[\s\S]*?this\.player\.setVelocity\(\s*0\s*\)/);
```

Pressing the left arrow should call `this.player.setFlipX(false)`.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*?this\.player\.setFlipX\(\s*false\s*\)/);
```

Pressing the right arrow should play the `walk-right` animation.

```js
assert.match(code, /this\.cursors\.right\.isDown[\s\S]*?this\.player\.anims\.play\(\s*(['"`])walk-right\1\s*,\s*true\s*\)/);
```

When no arrow is held, the `else` branch should play the `idle` animation.

```js
assert.match(code, /\}\s*else\s*\{[\s\S]*?this\.player\.anims\.play\(\s*(['"`])idle\1\s*,\s*true\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Sprite Sheet Debugger</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}
```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 32, frameHeight: 32 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 32, frameHeight: 32 }
    );
  }

  create() {
    this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 13 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12
    });
    this.coins = this.physics.add.group();
    [
      [80, 80],
      [400, 80],
      [80, 280],
      [400, 280],
      [240, 80]
    ].forEach(([x, y]) => {
      const coin = this.coins.create(x, y, "coin");
      coin.setScale(2);
    });
  }

  update() {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    }
  }

  collectCoin(player, coin) {
    coin.destroy();
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Sprite Sheet Debugger</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}
```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12,
      repeat: -1
    });
    this.coins = this.physics.add.group();
    [
      [80, 80],
      [400, 80],
      [80, 280],
      [400, 280],
      [240, 80]
    ].forEach(([x, y]) => {
      const coin = this.coins.create(x, y, "coin");
      coin.setScale(2);
      coin.anims.play("coin-spin");
    });
    this.physics.add.overlap(this.player, this.coins, this.collectCoin, null, this);
  }

  update() {
    this.player.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(false);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-right", true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    } else {
      this.player.anims.play("idle", true);
    }
  }

  collectCoin(player, coin) {
    coin.destroy();
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
