---
id: 66ad8294a0ad902f1b31b612
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop you will learn how to work with JavaScript fundamentals by building a greeting bot. 

In this first step, you will want to output a message to the console from the greeting bot.

Remember that you learned about `console.log()` and strings in the previous lecture videos. 

Here is a reminder of how to use `console.log()` with strings:

```js
console.log("Hello, World!");
```

Add a `console.log()` statement that outputs the string `"Hi there!"` to the console. Don't forget your quotes around the message! 


# --hints--

You should have a `console.log()` statement in your code.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

Your `console` statement should have the message `"Hi there!"`. Don't forget the quotes.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
