---
id: 66ad984e92ce213a66f7e51c
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

When you need to declare variables with multiple words, you can use the <dfn>camelCase</dfn> naming convention.

When using `camelCase`, the first word is all lowercase and the first letter of each subsequent word is capitalized. 

Here is an example:

```js
let thisIsCamelCase;
```

Use `let` to declare a variable named `botLocation`.

# --hints--

You should use `let` to declare the variable.

```js
assert.lengthOf(code.match(/let/g), 2);
```

Your variable should be named `botLocation`.

```js
assert.isNotNull(botLocation);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
