---
id: 66ad984e92ce213a66f7e51c
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

Now, it is time to assign a value to your `bot` variable.

In previous lessons, you learned how to assign values to variables like this:

```js
variableName = "Here is the value";
```

Remember that what is on the right side of the assignment operator `=` is the value that you are assigning to the variable on the left side.

Assign the string `"teacherBot"` to the `bot` variable.

# --hints--

You should assign a string to your `bot` variable.

```js
assert.isString(bot);
```

Your `bot` variable should be assigned the string `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
