---
id: 66ad8294a0ad902f1b31b612
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop you will learn how to work with JavaScript fundamentals by building a greeting bot. 

In this first step, you will want to output a message to the console from the greeting bot.

Remember that you learned about `console.log()` and strings in previous lessons. 

Here is a reminder of how to use `console.log()` with strings:

```js
console.log("Hello, World!");
```

Add a `console.log()` statement that outputs the string `"Hi there!"` to the console. Don't forget your quotes around the message! 

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

You should have a `console.log()` statement in your code.

```js
assert.isNotEmpty(logSpy.calls);
```

Your `console` statement should have the message `"Hi there!"`. Don't forget the quotes.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
