---
id: 66ad9c633fa26d3c1475eae3
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now, it is time to add another message from the bot. 

Add another `console` statement to the code that logs the message `"Allow me to introduce myself."`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

You should have a third `console.log()` statement in your code.

```js
assert.lengthOf(logSpy.calls, 3);
```

Your `console` statement should have the message `"Allow me to introduce myself."`. Don't forget the quotes.

```js
assert.deepInclude(logSpy.calls, ["Allow me to introduce myself."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
