---
id: 62a24068d60b671847d1d4e2
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Now we need some quick styling. Start by giving the `body` a `background-color` set to `darkblue`.

# --hints--

You should have a `body` selector.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.exists(body);
```

Your `body` selector should have a `background-color` property set to `darkblue`.

```js
const background = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert.equal(background, 'darkblue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
</head>
<body>
    <div id="game">
        <div id="stats">
            <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
            <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
            <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
        </div>
        <div id="controls">
            <button id="button1">Go to store</button>
            <button id="button2">Go to cave</button>
            <button id="button3">Fight dragon</button>
        </div>
        <div id="monsterStats">
            <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
            <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
        </div>
        <div id="text">
            Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
        </div>
    </div>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
