---
id: 62a3a488b24fb32b91155d56
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Variables can be <dfn>assigned</dfn> a value. When you do this while you declare it, this is called <dfn>initialization</dfn>. For example:

```js
let age = 32;
```

This would initialize the `age` variable with a value of `32`, a number.

Initialize your `xp` variable to have a value of `0`, a number.

# --hints--

`xp` should have a value of `0`.

```js
assert.equal(xp, 0);
```

You should initialize the `xp` variable to `0`. Don't forget the semi-colon at the end of the line.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp;
--fcc-editable-region--
```
