---
id: 614387cbefeeba5f3654a291
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

After your `header` element, create an `address` element with the `class` set to `author`.

Within that `address`, create a `p` element with the `class` set to `author-name` and give it the text `By freeCodeCamp`. Wrap the `freeCodeCamp` portion in an `a` element with the `href` set to `https://freecodecamp.org`, the `target` set to `_blank`, and the `rel` set to `author`.

Below that, add a `time` element with the `class` set to `publish-date`, the `datetime` attribute set to `2019-03-07`, and the text `March 7, 2019`.

# --hints--

You should create a new `address` element.

```js
assert.exists(document.querySelector('address'));
```

Your `address` element should come after your `header` element.

```js
assert(document.querySelector('address')?.previousElementSibling?.localName === 'header');
```

Your `address` element should have the `class` set to `author`.

```js
assert(document.querySelector('address')?.className === 'author');
```

You should create a new `p` element.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Your new `p` element should be within your `address` element.

```js
assert(document.querySelector('address')?.querySelectorAll('p')?.length === 1);
```

Your `p` element should have a `class` set to `author-name`.

```js
assert(document.querySelector('address')?.querySelector('p')?.className === 'author-name');
```

Your `p` element should have the text `By freeCodeCamp`.

```js
assert(document.querySelector('address')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

Your `a` element should have a `rel` set to `author`.

```js
assert(document.querySelector('address')?.querySelector('a')?.getAttribute('rel') === 'author');
```

You should create a new `time` element.

```js
assert.exists(document.querySelector('time'));
```

Your `time` element should be within your `address` element.

```js
assert(document.querySelector('address')?.querySelector('time')?.localName === 'time');
```

Your `time` element should have a `class` set to `publish-date`.

```js
assert(document.querySelector('address')?.querySelector('time')?.className === 'publish-date');
```

Your `time` element should have the `datetime` attribute set to `2019-03-07`.

```js
assert(document.querySelector('address')?.querySelector('time')?.getAttribute('datetime') === '2019-03-07');
```

Your `time` element should have the text `March 7, 2019`.

```js
assert(document.querySelector('address')?.querySelector('time')?.innerText === 'March 7, 2019');
```

You should create a new `a` element.

```js
assert.exists(document.querySelector('a'));
```

Your `a` element should be within your first new `p` element.

```js
assert(document.querySelector('address')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

Your `a` element should have the `href` set to `https://freecodecamp.org`.

```js
assert(document.querySelector('address')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

Your `a` element should have the `target` set to `_blank`.

```js
assert(document.querySelector('address')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

Your `a` element should surround the text `freeCodeCamp`.

```js
assert(document.querySelector('address')?.querySelector('p')?.firstElementChild?.innerText === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
