---
id: 6169cd8a558aa8434e0ad7f6
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

The `Referer` HTTP header contains information about the address or URL of a page that a user might be visiting from. This information can be used in analytics to track how many users from your page visit freecodecamp.org, for example. Setting the `rel` attribute to `noreferrer` omits this information from the HTTP request. Update the existing `a` element so its `rel` attribute includes both `author` and `noreferrer`.

# --hints--

Your `a` element should have the `rel` attribute include `noreferrer`.

```js
assert(document.querySelector('address')?.querySelector('a')?.getAttribute('rel')?.split(' ').includes('noreferrer'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        <address class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="author"
              >freeCodeCamp</a
            >
          </p>
          <time class="publish-date" datetime="2019-03-07">March 7, 2019</time>
        </address>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
