---
id: 66faab00000000000000000a
title: Build a Brick Breaker
challengeType: 26
dashedName: lab-brick-breaker
---

# --description--

**Objective:** Build a complete brick breaker using Phaser 3 that fulfills the user stories below.

The Phaser library is loaded for you on the global `Phaser` object.

**User Stories:**

1. You should declare a class named `MainScene` that extends `Phaser.Scene`.
1. `MainScene` should have a constructor that calls `super("MainScene")`.
1. `MainScene` should have `create` and `update` methods.
1. You should declare a class named `GameOverScene` that extends `Phaser.Scene` and whose constructor calls `super("GameOverScene")`.
1. You should declare a class named `ControlsScene` that extends `Phaser.Scene` and whose constructor calls `super("ControlsScene")`.
1. You should declare an `ACTIONS` object that maps the action names `left`, `right`, `launch`, and `pause` to keyboard key names.
1. You should declare a `config` object with a `type` of `Phaser.AUTO`, a `width` of `480`, a `height` of `640`, a `parent` of `"game-container"`, and arcade physics with `gravity: { y: 0 }`.
1. `config` should register all three scenes via `scene: [MainScene, GameOverScene, ControlsScene]`.
1. Inside `MainScene.create`, you should declare a `this.paddle` (rectangle, sprite, or image) with a physics body whose body is set to immovable.
1. The paddle should respond to both keyboard cursor keys and pointer events (mouse/touch).
1. You should declare a `this.ball` (circle, sprite, or image) with a physics body, `setBounce(1, 1)`, and `setCollideWorldBounds(true)`.
1. You should disable the bottom-edge collision so the ball can fall off the bottom: `this.physics.world.checkCollision.down = false`.
1. You should declare a static group `this.bricks` via `this.physics.add.staticGroup` and lay out the bricks from a `LEVELS` array.
1. Bricks should track HP — destructible bricks should decrement HP on hit and indestructible bricks should be set to `Infinity`.
1. The ball should collide with `this.bricks` via `this.physics.add.collider`, and a destructible-brick hit should add to `this.score`.
1. You should render a HUD with `this.add.text` showing the score and at least three starting lives.
1. The high score should persist across page reloads using `localStorage.getItem` and `localStorage.setItem`.
1. `MainScene.create` should run `ControlsScene` in parallel with `this.scene.launch("ControlsScene")`.

# --hints--

You should declare a class named `MainScene` that extends `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{/);
```

`MainScene` should call `super("MainScene")` in its constructor.

```js
assert.match(code, /super\(\s*(['"`])MainScene\1\s*\)/);
```

`MainScene` should define `create` and `update` methods.

```js
assert.match(code, /create\s*\(\s*\)\s*\{/);
assert.match(code, /update\s*\(\s*\)\s*\{/);
```

You should declare a class named `GameOverScene` that extends `Phaser.Scene` and calls `super("GameOverScene")`.

```js
assert.match(code, /class\s+GameOverScene\s+extends\s+Phaser\.Scene\s*\{/);
assert.match(code, /super\(\s*(['"`])GameOverScene\1\s*\)/);
```

You should declare a class named `ControlsScene` that extends `Phaser.Scene` and calls `super("ControlsScene")`.

```js
assert.match(code, /class\s+ControlsScene\s+extends\s+Phaser\.Scene\s*\{/);
assert.match(code, /super\(\s*(['"`])ControlsScene\1\s*\)/);
```

You should declare an `ACTIONS` object containing `left`, `right`, `launch`, and `pause` keys.

```js
assert.match(code, /ACTIONS\s*=\s*\{[\s\S]*?left[\s\S]*?right[\s\S]*?launch[\s\S]*?pause[\s\S]*?\}/);
```

You should declare a `config` object with the required Phaser settings.

```js
assert.match(code, /const\s+config\s*=\s*\{/);
assert.match(code, /type\s*:\s*Phaser\.AUTO/);
assert.match(code, /width\s*:\s*480\b/);
assert.match(code, /height\s*:\s*640\b/);
assert.match(code, /parent\s*:\s*(['"`])game-container\1/);
assert.match(code, /physics\s*:\s*\{[\s\S]*?default\s*:\s*(['"`])arcade\1[\s\S]*?arcade\s*:\s*\{\s*gravity\s*:\s*\{\s*y\s*:\s*0\s*\}\s*\}/);
```

`config` should register all three scenes as an array.

```js
assert.match(code, /scene\s*:\s*\[\s*MainScene\s*,\s*GameOverScene\s*,\s*ControlsScene\s*\]/);
```

The paddle should be added to the scene with a physics body and set to immovable.

```js
assert.match(code, /this\.paddle\s*=\s*this\.(?:add|physics\.add)\.(?:rectangle|sprite|image)\(/);
assert.match(code, /this\.paddle\.body\.setImmovable\(\s*true\s*\)/);
```

The paddle should respond to both keyboard cursors and pointer events.

```js
assert.match(code, /this\.input\.keyboard\.createCursorKeys\(\s*\)/);
assert.match(code, /this\.input\.on\(\s*(['"`])pointer(?:move|down)\1/);
```

The ball should be created with bounce, world-bounds collision, and an attached physics body.

```js
assert.match(code, /this\.ball\s*=\s*this\.(?:add|physics\.add)\.(?:circle|sprite|image)\(/);
assert.match(code, /\.setBounce\(\s*1\s*,\s*1\s*\)/);
assert.match(code, /\.setCollideWorldBounds\(\s*true\s*\)/);
```

You should disable the bottom-edge collision so the ball can fall off the bottom.

```js
assert.match(code, /this\.physics\.world\.checkCollision\.down\s*=\s*false/);
```

You should declare a static brick group fed from a `LEVELS` array.

```js
assert.match(code, /this\.bricks\s*=\s*this\.physics\.add\.staticGroup\(/);
assert.match(code, /(?:const|let|var)\s+LEVELS\s*=\s*\[/);
```

Bricks should track HP, with indestructible bricks set to `Infinity` and destructible bricks decrementing on hit.

```js
assert.match(code, /\.hp\b[\s\S]{0,80}?\bInfinity\b/);
assert.match(code, /\.hp\s*-=\s*1/);
```

The ball should collide with `this.bricks` and a brick hit should add to `this.score`.

```js
assert.match(code, /this\.physics\.add\.collider\(\s*this\.ball\s*,\s*this\.bricks/);
assert.match(code, /this\.score\s*\+=/);
```

You should render the score and at least three starting lives via `this.add.text`.

```js
assert.match(code, /this\.add\.text\(/);
assert.match(code, /this\.score\s*=\s*(?:data\.score\s*\|\|\s*)?\d+/);
assert.match(code, /this\.lives\s*=\s*(?:data\.lives\s*\|\|\s*)?[3-9]/);
```

The high score should persist across page reloads using `localStorage`.

```js
assert.match(code, /localStorage\.getItem\(/);
assert.match(code, /localStorage\.setItem\(/);
```

`MainScene.create` should launch `ControlsScene` in parallel.

```js
assert.match(code, /this\.scene\.launch\(\s*(['"`])ControlsScene\1\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js

const ACTIONS = {
  left: ["LEFT", "A"],
  right: ["RIGHT", "D"],
  launch: ["SPACE"],
  pause: ["P"]
};

const SAVE_KEY = "phaser-bb-high-score";

const LEVELS = [
  [
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
  ],
  [
    [2, 2, 2, 2, 2, 2, 2, 2, 2, 2],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [2, 2, 2, 2, 2, 2, 2, 2, 2, 2]
  ],
  [
    [9, 1, 1, 1, 1, 1, 1, 1, 1, 9],
    [1, 1, 1, 1, 1, 1, 1, 1, 1, 1],
    [9, 1, 1, 1, 1, 1, 1, 1, 1, 9]
  ],
  [
    [3, 3, 2, 2, 1, 1, 2, 2, 3, 3],
    [9, 1, 1, 1, 1, 1, 1, 1, 1, 9],
    [1, 1, 1, 9, 9, 9, 9, 1, 1, 1]
  ],
  [
    [3, 3, 3, 3, 3, 3, 3, 3, 3, 3],
    [9, 2, 2, 2, 2, 2, 2, 2, 2, 9],
    [9, 1, 1, 9, 9, 9, 9, 1, 1, 9]
  ]
];

const COLORS = [0x666666, 0x44aaff, 0x44ee88, 0xffee44, 0xffaa44, 0xff6644];

class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  init(data) {
    data = data || {};
    this.level = data.level || 1;
    this.score = data.score || 0;
    this.lives = data.lives || 3;
  }

  create() {
    this.add.rectangle(240, 320, 480, 640, 0x111122);
    this.paddle = this.add.rectangle(240, 580, 80, 16, 0xeeeeee);
    this.physics.add.existing(this.paddle);
    this.paddle.body.setImmovable(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.keys = this.input.keyboard.addKeys("A,D,SPACE,P");
    this.targetX = this.paddle.x;
    this.input.on("pointermove", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
    });
    this.input.on("pointerdown", pointer => {
      if (pointer.y > 540) this.targetX = pointer.x;
      if (this.ballAttached) this.launchBall();
    });
    this.ball = this.add.circle(240, 540, 8, 0xffaa44);
    this.physics.add.existing(this.ball);
    this.ball.body.setBounce(1, 1);
    this.ball.body.setCollideWorldBounds(true);
    this.physics.world.checkCollision.down = false;
    this.ball.body.setVelocity(0, 0);
    this.ballAttached = true;
    this.physics.add.collider(this.ball, this.paddle, this.onPaddleBounce, null, this);
    this.bricks = this.physics.add.staticGroup();
    this.spawnBricks();
    this.physics.add.collider(this.ball, this.bricks, this.onBrickHit, null, this);
    this.scoreText = this.add.text(10, 10, "SCORE " + this.score, { fontSize: "16px", color: "#ffffff" });
    this.livesText = this.add.text(10, 32, "LIVES " + this.lives, { fontSize: "16px", color: "#ffffff" });
    this.highText = this.add.text(470, 10, "", { fontSize: "16px", color: "#aaaaaa" }).setOrigin(1, 0);
    this.updateHud();
    this.paused = false;
    this.scene.launch("ControlsScene");
    this.events.on("action:left", () => { this.targetX -= 30; });
    this.events.on("action:right", () => { this.targetX += 30; });
    this.events.on("action:pause", () => this.togglePause());
  }

  update() {
    if (Phaser.Input.Keyboard.JustDown(this.keys.P)) this.togglePause();
    if (this.paused) return;
    if (this.isDown("left")) this.targetX -= 6;
    else if (this.isDown("right")) this.targetX += 6;
    this.targetX = Phaser.Math.Clamp(this.targetX, 40, 440);
    this.paddle.x = Phaser.Math.Linear(this.paddle.x, this.targetX, 0.3);
    this.paddle.x = Phaser.Math.Clamp(this.paddle.x, 40, 440);
    if (this.ballAttached) {
      this.ball.x = this.paddle.x;
      this.ball.y = this.paddle.y - 16;
      if (this.isDown("launch")) this.launchBall();
      return;
    }
    if (this.ball.y > 660) this.loseLife();
  }

  isDown(action) {
    return ACTIONS[action].some(name => {
      const k = this.keys[name] || this.cursors[name.toLowerCase()];
      return k && k.isDown;
    });
  }

  onPaddleBounce(ball, paddle) {
    const offset = (ball.x - paddle.x) / (paddle.width / 2);
    const speed = Math.sqrt(ball.body.velocity.x ** 2 + ball.body.velocity.y ** 2) || 240;
    const angle = -Math.PI / 2 + offset * (Math.PI / 3);
    ball.body.setVelocity(speed * Math.cos(angle), speed * Math.sin(angle));
  }

  spawnBricks() {
    const layout = LEVELS[(this.level - 1) % LEVELS.length];
    for (let row = 0; row < layout.length; row++) {
      for (let col = 0; col < layout[row].length; col++) {
        const cell = layout[row][col];
        if (cell === 0) continue;
        const x = 40 + col * 40 + 20;
        const y = 80 + row * 22;
        const isIndestructible = cell === 9;
        const color = isIndestructible ? COLORS[0] : COLORS[cell];
        const brick = this.add.rectangle(x, y, 36, 18, color);
        this.bricks.add(brick);
        brick.hp = isIndestructible ? Infinity : cell;
        brick.points = (3 - row + cell) * 10;
      }
    }
  }

  onBrickHit(ball, brick) {
    if (brick.hp === Infinity) return;
    brick.hp -= 1;
    if (brick.hp <= 0) {
      this.score += brick.points;
      brick.destroy();
      this.updateHud();
      const remaining = this.bricks.getChildren().filter(b => b.hp !== Infinity).length;
      if (remaining === 0) this.advanceLevel();
    } else {
      brick.setFillStyle(COLORS[Phaser.Math.Clamp(brick.hp, 0, 5)]);
    }
  }

  launchBall() {
    this.ballAttached = false;
    const baseSpeed = 220 + (this.level - 1) * 12;
    this.ball.body.setVelocity(baseSpeed * 0.4, -baseSpeed);
  }

  loseLife() {
    this.lives -= 1;
    this.updateHud();
    if (this.lives <= 0) {
      this.gameOver();
      return;
    }
    this.ball.body.setVelocity(0, 0);
    this.ballAttached = true;
  }

  gameOver() {
    this.persistHighScore();
    this.scene.stop("ControlsScene");
    this.scene.start("GameOverScene", { score: this.score, win: false });
  }

  persistHighScore() {
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    if (this.score > high) localStorage.setItem(SAVE_KEY, String(this.score));
  }

  advanceLevel() {
    this.persistHighScore();
    this.level += 1;
    if (this.level > LEVELS.length) {
      this.scene.stop("ControlsScene");
      this.scene.start("GameOverScene", { score: this.score, win: true });
      return;
    }
    this.scene.restart({ level: this.level, score: this.score, lives: this.lives });
  }

  updateHud() {
    this.scoreText.setText("SCORE " + this.score);
    this.livesText.setText("LIVES " + this.lives);
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    this.highText.setText("HIGH " + Math.max(high, this.score));
  }

  togglePause() {
    this.paused = !this.paused;
    if (this.paused) this.physics.world.pause();
    else this.physics.world.resume();
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create(data) {
    this.add.rectangle(240, 320, 480, 640, 0x000000, 0.85);
    this.add.text(240, 240, data.win ? "YOU WIN" : "GAME OVER", { fontSize: "32px", color: data.win ? "#44ee88" : "#ff6644" }).setOrigin(0.5);
    this.add.text(240, 300, "Score: " + data.score, { fontSize: "20px", color: "#ffffff" }).setOrigin(0.5);
    const high = parseInt(localStorage.getItem(SAVE_KEY) || "0", 10);
    this.add.text(240, 340, "High: " + high, { fontSize: "16px", color: "#aaaaaa" }).setOrigin(0.5);
    this.add.text(240, 420, "Press R to restart", { fontSize: "14px", color: "#ffee44" }).setOrigin(0.5);
    this.input.keyboard.once("keydown-R", () => {
      this.scene.start("MainScene", { level: 1, score: 0, lives: 3 });
    });
  }
}

class ControlsScene extends Phaser.Scene {
  constructor() {
    super("ControlsScene");
  }

  create() {
    const isMobile = Phaser.Device.os.android || Phaser.Device.os.iOS;
    if (!isMobile) return;
    let leftHeld = false;
    let rightHeld = false;
    const leftBtn = this.add.rectangle(60, 600, 80, 60, 0x556677, 0.5).setInteractive();
    leftBtn.on("pointerdown", () => { leftHeld = true; });
    leftBtn.on("pointerup", () => { leftHeld = false; });
    leftBtn.on("pointerout", () => { leftHeld = false; });
    const rightBtn = this.add.rectangle(420, 600, 80, 60, 0x556677, 0.5).setInteractive();
    rightBtn.on("pointerdown", () => { rightHeld = true; });
    rightBtn.on("pointerup", () => { rightHeld = false; });
    rightBtn.on("pointerout", () => { rightHeld = false; });
    const pauseBtn = this.add.rectangle(240, 600, 60, 40, 0x556677, 0.5).setInteractive();
    pauseBtn.on("pointerdown", () => {
      this.scene.get("MainScene").events.emit("action:pause");
    });
    this.events.on("update", () => {
      const main = this.scene.get("MainScene").events;
      if (leftHeld) main.emit("action:left");
      if (rightHeld) main.emit("action:right");
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#111122",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene, GameOverScene, ControlsScene]
};

const game = new Phaser.Game(config);
```
