---
id: 6551fe3b1df7c9740f13f270
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Key aspects of variable naming in Python are:

- Algunas palabras son palabras clave reservadas (por ejemplo, `for`, `while`, `True`). Tienen un significado especial en Python, así que no puedes usarlos para nombres de variables.
- Los nombres de variables no pueden comenzar con un número, y sólo pueden contener caracteres alfanuméricos o guiones bajos.
- Los nombres de las variables son sensibles a mayúsculas y minúsculas, es decir, `my_var` es diferente de `my_Var` y de `MY_VAR`.
- Finalmente, es una convención común escribir nombres de variables usando `snake_case`, donde cada espacio se sustituye por un carácter de guión bajo y las palabras se escriben en minúscula.

Remove both calls to `print()` and declare another variable called `alphabet`. Assign the string `abcdefghijklmnopqrstuvwxyz` to this variable.

# --hints--

You should declare a variable called `alphabet`.

```js
({ test: () => assert(__userGlobals.has("alphabet")) })
```

You should assign the string `abcdefghijklmnopqrstuvwxyz` to your `alphabet` variable.

```js
({ test: () => assert.equal(__userGlobals.get("alphabet"), "abcdefghijklmnopqrstuvwxyz") })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
