---
id: 6551ef504d91936d2d4e54f8
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

You can use the built-in function `print()` to print the output of your code on the terminal. Functions are reusable code blocks that you can call (run) when you need them. To call a function, you just need to write a pair of parentheses next to its name. You will learn more about functions very soon. For now, print your `text` variable to the screen by including the variable name between the opening and closing parentheses of the `print()` function.


# --hints--

You should call the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(.*\)/) })

```

You should pass `text` to your `print()` function by including the name of this variable within parentheses.

```js
({ test: () => assert.match(code, /print\s*\(\s*text\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
