---
id: 6551f22653e6ce6e4f1d1370
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

Call the `type()` function and print the data type of your `shift` variable.

# --hints--

You should pass `shift` to the `type()` function.

```js
({ test: () => assert.match(code, /type\s*\(\s*shift\s*\)/) })

```

You should pass `type(shift)` to your `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/) })

```

You should keep your existing `print()` call.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
