---
id: 655208d59b131e7816f18c96
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Now assign the value returned by `alphabet.find(text[0])` to a variable named `index`. Then, print its value.

# --hints--

You should declare a variable named `index`.

```js
({ test: () => assert(__userGlobals.has("index")) })

```

You should assign the value returned by `alphabet.find(text[0])` to your `index` variable.

```js
({ test: () => assert.match(code, /index\s*=\s*alphabet\.find\s*\(\s*text\[\s*\s*0\s*\]\s*\)/) })

```

You should print your `index` variable.

```js
({ test: () => assert.match(code, /print\s*\(\s*index\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
