---
id: 6555e3f387381a19d5e00333
title: Paso 88
challengeType: 20
dashedName: step-88
---

# --description--

Next, modify the other two `print()` calls to use f-strings.

# --hints--

You should modify the `print('Key: ' + custom_key)` call so that it prints the same result using an f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Key:\s\{\s*custom_key\s*\}\1\s*\)/) })
```

You should modify the `print('Decrypted text: ' + decryption)` call so that it prints the same result using an f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Decrypted\stext:\s\{\s*decryption\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message

def encrypt(message, key):
    return vigenere(message, key)

def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print('Key: ' + custom_key)
decryption = decrypt(text, custom_key)
print('Decrypted text: ' + decryption)
--fcc-editable-region--
```
