---
id: 61329b210dac0b08047fd6ab
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Kuendelea na vipengele vya `meta`, ufafanuzi wa `viewport` huambia kivinjari jinsi ya kutoa ukurasa. Ikijumuisha ufikivu bora wa kuona kwenye simu ya mkononi, na inaboresha _SEO_ (search engine optimization).

Ongeza ufafanuzi wa `viewport` na sifa ya `content` inayoelezea `width` na `initial-scale` ya ukurasa.

# --hints--

Unapaswa kuunda kipengele kingine cha `meta` katika sehemu ya `head`.

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

Unapaswa kuipa `meta` sifa ya `name` ya `viewport`.

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

Unapaswa kuipa `meta` sifa ya `content` ya `width=device-width, initial-scale=1`.

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
