---
id: 613297a923965e0703b64796
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Huenda unafahamu kipengele cha `meta` tayari; kinatumika kubainisha habari kuhusu ukurasa, kama vile kichwa, maelezo, manenomsingi na mwandishi.

Upe ukurasa wako kipengele cha `meta` chenye thamani inayofaa ya `charset`.

Sifa ya `charset` inabainisha usimbaji wa herufi ya ukurasa, na, siku hizi, `UTF-8` ndio usimbaji pekee unaoauniwa na vivinjari vingi.

# --hints--

Unapaswa kuunda kipengele cha `meta` ndani ya kipengele cha `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Unapaswa kuipa tagi ya `meta` `charset` ya `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
