---
id: 6133acc353338c0bba9cb553
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Hatimaye katika sehemu ya `head`, kipengele cha `title` ni muhimu kwa visomaji skrini kuelewa maudhui ya ukurasa. Zaidi ya hayo, ni sehemu muhimu ya _SEO_.

Upe ukurasa wako `title` yenye maelezo mafupi.

# --hints--

Unapaswa kuongeza kipengele cha `title` kwenye sehemu ya `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Hupaswi kufanya `title` kuwa ndefu zaidi ya herufi 60.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Jaribu kuwa na maelezo zaidi ukitumia kipengele chako cha `title`. _Kidokezo: Angalau herufi 15_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
