---
id: 6133d11ef548f51f876149e3
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Uongozaji ni sehemu kuu ya ufikivu, na visoma skrini vinakutegemea wewe kutoa muundo wa ukurasa wako. Hii inakamilishwa na vipengele vya HTML vya semantic.

Ongeza `header` na kipengele kikuu cha `main` kwenye ukurasa wako.

Kipengele cha `header` kitatumika kutambulisha ukurasa, na pia kutoa menyu ya mwongozo.

Kipengele cha `main` kitakuwa na maudhui ya msingi ya ukurasa wako.

# --hints--

Unapaswa kuongeza kipengele cha `header` kwenye `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Unapaswa kuongeza kipengele cha `main` kwenye `body`.

```js
assert.exists(document.querySelector('body > main'));
```

Kipengele cha `header` kinafaa kuja kabla ya kipengele cha `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
