---
id: 613e2546d0594208229ada50
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Ndani ya `header`, toa muktadha kuhusu ukurasa kwa kuweka kipengee kimoja cha `img`, `h1` na `nav`.

The `img` should point to `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`, have an `id` of `logo`, and have an `alt` text of `freeCodeCamp`.

`h1` inapaswa kuwa na maandishi `HTML/CSS Quiz`.

# --hints--

Unapaswa kuongeza kipengele cha `img` kwenye kipengele cha `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Unapaswa kuongeza kipengele cha `h1` kwenye kipengele cha `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Unapaswa kuongeza kipengele cha `nav` kwenye kipengele cha `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Unapaswa kuweka vipengele vya `img`, `h1`, na `nav` kwa mpangilio huo.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Unapaswa kukipa kipengele cha `img` `src` sifa ya `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Unapaswa kukipa kipengele cha `img` na `id` sifa ya `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

You should give the `img` element an `alt` attribute of `freeCodeCamp`.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

Unapaswa kukipa kipengele cha `h1` maandishi `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
