---
id: 6140827cff96e906bd38fc2b
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Kama ilivyoelezwa katika <a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">Mwongozo wa Mtindo wa freeCodeCamp</a>, logo inapaswa kuhifadhi uwiano wa `35 / 4`, na iwe na pedi kwenye maandishi.

Kwanza, badilisha `background-color` hadi `#0a0a23` ili uweze kuona logo. Kisha, tumia sifa ya `aspect-ratio` ili kuweka uwiano unaohitajika kuwa `35 / 4`. Hatimaye, ongeza `padding` ya `0.4rem` pande zote.

# --hints--

Unapaswa kuipa `#logo` `background-color` ya `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

Unapaswa kutumia sifa ya `aspect-ratio`.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

Hupaswi kutumia sifa ya `height`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

Haupaswi kubadilisha sifa ya `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

Unapaswa kuipa `img` `aspect-ratio` ya `35 / 4`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

Unapaswa kuipa `img` `padding` ya `0.4rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--
#logo {
  width: max(100px, 18vw);

}
--fcc-editable-region--

```
