---
id: 6140883f74224508174794c0
title: Hatua ya 10
challengeType: 0
dashedName: step-10
---

# --description--

Fanya `header` ichukue upana kamili wa kontena kuu, weka `height` kuwa `50px`, na uweke `background-color` kuwa `#1b1b32`. Kisha, weka `display` kutumia _Flexbox_.

# --hints--

Unapaswa kutumia kipengele kichaguzi cha `header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

Unapaswa kuipa `header` `width` ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

Unapaswa kuipa `header` `height` ya `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

Unapaswa kuipa `header` `background-color` ya `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

Unapaswa kuipa `header` `display` ya `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
