---
id: 61408f155e798909b6908712
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Ili kuwezesha ufikiaji kwenye ukurasa, ongeza orodha ambayo haijapangwa na vitu vitatu vifuatavyo:

- `INFO`
- `HTML`
- `CSS`

Maandishi ya vipengee vya orodha yanapaswa kufungwa kwa tagi za nanga (anchor tags).

# --hints--

Unapaswa kuweka kipengee kimoja cha `ul` ndani ya `nav`.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

Unapaswa kuweka vipengele vitatu vya `li` ndani ya kipengele cha `ul`.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

Unapaswa kuweka kipengele kimoja cha `a` ndani ya kila kipengele cha `li`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

Unapaswa kukipa kipengele cha kwanza cha `a` maandishi `INFO`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

Unapaswa kukipa kipengele cha pili cha `a` maandishi `HTML`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

Unapaswa kukipa kipengele cha tatu cha `a` maandishi `CSS`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
