---
id: 6141fabd6f75610664e908fd
title: Hatua ya 14
challengeType: 0
dashedName: step-14
---

# --description--

Kwa vile hili ni swali, utahitaji fomu kwa watumiaji kuwasilisha majibu. Unaweza kutenganisha maudhui ndani ya fomu kimantiki kwa kutumia vipengele vya `section`.

Ndani ya kipengele cha `main`, unda fomu yenye vipengele vitatu vya `section`. Kisha, fanya fomu iwasilishwe kwa `https://freecodecamp.org/practice-project/accessibility-quiz`, kwa kutumia mbinu sahihi.

# --hints--

Unapaswa kuweka kipengele cha `form` ndani ya kipengele chako cha `main`.

```js
assert.exists(document.querySelector('main > form'));
```

Unapaswa kuweka vipengele vitatu vya `section` ndani ya kipengele chako cha `form`.

```js
assert.equal(document.querySelectorAll('main > form > section')?.length, 3);
```

Unapaswa kukipa kipengele cha `form` sifa ya `action`.

```js
assert.notEmpty(document.querySelector('main > form')?.action);
```

Unapaswa kuipa sifa ya `action` thamani ya `https://freecodecamp.org/practice-project/accessibility-quiz`.

```js
assert.equal(document.querySelector('main > form')?.action, 'https://freecodecamp.org/practice-project/accessibility-quiz');
```

Unapaswa kukipa kipengele cha `form` sifa ya `method`.

```js
assert.notEmpty(document.querySelector('main > form')?.method);
```

Unapaswa kukipa kipengele cha `form` sifa ya `method` yenye thamani ya `post`.

```js
assert.equal(document.querySelector('main > form')?.method?.toLowerCase(), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
