---
id: 6143639d5eddc7094161648c
title: Hatua ya 20
challengeType: 0
dashedName: step-20
---

# --description--

Ni muhimu kuunganisha kila `input` kwa `label` inayolingana. Hii huwapa watumiaji wa teknolojia ya usaidizi rejeleo la kuona kwa ingizo.

Hii inafanywa kwa kuipa `label` sifa ya `for`, ambayo ina `id` ya `input`.

Sehemu hii itachukua jina la mwanafunzi, anwani ya barua pepe na tarehe ya kuzaliwa. Vipe vipengele vya `label` sifa inayofaa ya `for`, pamoja na maudhui ya maandishi. Kisha, unganisha vipengele vya `input` kwa vipengele vinavyolingana vya `label`.

# --hints--

Unapaswa kukipa kipengele cha `label` sifa ya `for` inayofaa.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

Unapaswa kukipa kipengele cha pili cha `label` sifa inayofaa ya `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

Unapaswa kukipa kipengele cha tatu cha `label` sifa inayofaa ya `for`.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

Unapaswa kukipa kipengele cha kwanza cha `label` maudhui ya maandishi yanayofaa.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.textContent?.length, 1);
```

Unapaswa kukipa kipengele cha pili cha `label` maudhui ya maandishi yanayofaa.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.textContent?.length, 1);
```

Unapaswa kukipa kipengele cha tatu cha `label` maudhui ya maandishi yanayofaa.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.textContent?.length, 1);
```

Unapaswa kukipa kipengele cha kwanza cha `input` sifa ya `id` inayolingana na sifa ya `for` ya `label` ya kwanza.

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

Unapaswa kukipa kipengele cha pili cha `input` sifa ya `id` inayolingana na sifa ya `for` ya `label` ya pili.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

Unapaswa kukipa kipengele cha tatu cha `input` sifa ya `id` inayolingana na sifa ya `for` ya `label` ya tatu.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

Hupaswi kutumia sifa ya `id` kwa zaidi ya kipengele kimoja cha `input`.

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

Hufai kutumia sifa sawa ya `for` kwa zaidi ya kipengele kimoja cha `label`.

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```
