---
id: 614ccc21ea91ef1736b9b578
title: Hatua ya 1
challengeType: 0
dashedName: step-1
---

# --description--

Karibu kwenye sehemu ya kwanza ya Maswali ya Ufikiaji (Accessibility Quiz). Unapoendelea kuwa msanidi wa HTML na CSS aliyeboreshwa, tumekuanzisha na mpangilio wa msingi.

Anza safari hii ya ufikivu kwa kutoa sifa ya `lang` kwa kipengele chako cha `html`. Hii itasaidia visoma skrini katika kutambua lugha ya ukurasa.

# --hints--

Unapaswa kukipa kipengele cha `html` sifa ya `lang`. _Kidokezo: Unaweza kutumia thamani `en` kwa Kiingereza._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
