---
id: 66d9639a8db48f247aa643d2
title: Build a Product Landing Page
challengeType: 14
dashedName: build-a-product-landing-page
demoType: onClick
---

# --description--

Project description coming soon!

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Product Landing Page</title>
  </head>

<body>

</body>

</html>
```

```css

```

```js

```

# --solutions--

```html

```

```css

```
