---
id: 646ceb843412c74edee27a79
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Utafanya kila sikio lionekane kama pembetatu.

Kwa kutumia kichaguzi cha darasa, kipe kipengele cha `.cat-left-ear` mpaka wa kushoto na kulia wa `35px solid transparent` kila moja. Pia, weka mpaka wa chini kuwa `70px solid #5e5e5e`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.cat-left-ear`.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

Kichaguzi chako cha `.cat-left-ear` kinafaa kuwa na sifa ya `border-left` iliyowekwa kuwa `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

Kichaguzi chako cha `.cat-left-ear` kinafaa kuwa na sifa ya `border-right` iliyowekwa kuwa `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

Kichaguzi chako cha `.cat-left-ear` kinapaswa kuwa na sifa ya `border-bottom` iliyowekwa kuwa `70px solid #5e5e5e`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
