---
id: 646c48df8674cf2b91020ecb
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Ndani ya kipengele chako cha `head`, ongeza tagi ya `meta` yenye sifa ya `charset` iliyowekwa kuwa `utf-8`. Pia, ongeza kipengele cha `title` chenye maandishi `fCC Cat Painting`.

# --hints--

Unapaswa kuongeza kipengele kimoja cha `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya `charset` iliyowekwa kuwa `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Unapaswa kuongeza kipengele kimoja cha `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Kipengele chako cha `title` kinapaswa kuwa na maandishi `fCC Cat Painting`. Kumbuka kwamba tahajia na herufi ni muhimu.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
