---
id: 646c5e727e56e743c9aed4a1
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Mpangilio wa CSS hukuruhusu kuweka jinsi unavyotaka kipengele kiwekwe kwenye kivinjari. Ina sifa ya `position` unayoweza kuweka kuwa `static`, `absolute`, `relative`, `sticky` au `fixed`.

Mara tu unapoweka sifa ya `position` ya kipengele, unaweza kusogeza kipengele kwa kuweka pikseli au thamani ya asilimia kwa moja au zaidi ya sifa za `top`, `right`, `left`, au `bottom`.

`static` ni nafasi chaguomsingi ya vipengele vyote. Ukiikabidhi kwa kipengele, hutaweza kuisogeza na `top`, `right`, `left` au `bottom`.

Ipe `.cat-head` `position` ya `static`, kisha uweke sifa za `top` na `left` kuwa `100px` kila moja.


# --hints--

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `static`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'static')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `100px`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa `100px`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {



  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
