---
id: 646c5ffef5598d449b52ec12
title: Hatua ya 19
challengeType: 0
dashedName: step-19
---

# --description--

Sasa unapaswa kufanya kazi kwenye masikio ya paka. Kutakuwa na sikio la kulia na la kushoto, na ndani ya kila moja, kutakuwa na sikio la ndani.

Ndani ya kipengele chako cha `.cat-head`, unda kipengele cha `div` ukitumia darasa la `cat-ears`.

# --hints--

Hupaswi kubadilisha kipengele kilichopo cha `div` kilicho na class ya `cat-head`.

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

Unapaswa kuunda kipengele kimoja cha `div` ndani ya kipengele chako cha `.cat-head`.

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la `cat-ears`.

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">

      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
