---
id: 646ce8bb4b35544d501c7184
title: Hatua ya 20
challengeType: 0
dashedName: step-20
---

# --description--

Ndani ya kipengele chako cha `.cat-ears`, unda vipengele viwili vya `div` ukitumia madarasa `cat-left-ear` na `cat-right-ear` mtawalia.

# --hints--

You should not change the existing `div` element with the class `cat-ears`.

```js
assert(document.querySelectorAll('div.cat-ears').length === 1);
```

Unapaswa kuunda vipengee viwili vya `div` ndani ya kipengele chako cha `.cat-ears`.

```js
assert(document.querySelectorAll('.cat-ears div').length === 2);
```

Kipengele chako cha kwanza cha `div` kinapaswa kuwa na darasa la `cat-left-ear`.

```js
assert(document.querySelectorAll('.cat-ears div')[0]?.classList.contains('cat-left-ear'));
```

Kipengele chako cha pili cha `div` kinapaswa kuwa na darasa la `cat-right-ear`.

```js
assert(document.querySelectorAll('.cat-ears div')[1]?.classList.contains('cat-right-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        --fcc-editable-region--
        <div class="cat-ears">

        </div>
        --fcc-editable-region--
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
