---
id: 646c586be7180e40ddf74ff6
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Kipe kipengele cha `body` `background-color` ya `#c9d2fc`.

# --hints--

Unapaswa kuwa na kichaguzi cha `body`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Kichaguzi chako cha `body` kinafaa kuwa na sifa ya `background-color` iliyowekwa kuwa `#c9d2fc`.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(201, 210, 252)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```

