---
id: 646cf48d8f8e1f535a1821d3
title: Step 27
challengeType: 0
dashedName: step-27
---

# --description--

Sogeza sikio la kulia hadi kwenye nafasi yenye sifa ya `position` iliyowekwa kuwa `absolute`, `top` ya `-26px`, na `left` ya `163px`.

# --hints--

Kichaguzi chako cha `.cat-right-ear` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`. Usisahau kuongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.position === 'absolute')
```

Kichaguzi chako cha `.cat-right-ear` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `-26px`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.top === '-26px')
```

Kichaguzi chako cha `.cat-right-ear` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa `163px`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.left === '163px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  transform: rotate(-45deg);
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}

.cat-right-ear {
  --fcc-editable-region--

  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
