---
id: 646c59e942f35541923104bf
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Katika HTML yako, unda kipengele cha `main`. Ndani ya kipengele hicho cha `main`, ongeza kipengele cha `div` chenye darasa la `cat-head`.

# --hints--

Unapaswa kuwa na kipengele cha `main`.

```js
assert(document.querySelectorAll('main').length === 1);
```

Unapaswa kuwa na kipengele cha `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la `cat-head`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya tagi ya `main`.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
