---
id: 646c48df8674cf2b91020ecc
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Ongeza kipengele cha `link` ndani ya kipengele chako cha `head`. Kwa kipengele hicho cha `link`, weka sifa ya `rel` kuwa `stylesheet` na `href` kuwa `./styles.css`.

# --hints--

Unapaswa kuwa na kipengele cha `link`.

```js
assert.match(code, /<link/)
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya `rel` iliyowekwa kuwa `stylesheet`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1/)
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `./styles.css`.

```js
assert.match(code, /<link\s+rel=("|'|`)stylesheet\1\s+href=\1\.\/styles\.css\1\s*\/?>/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
