---
id: 646c5ace05e4be4211407935
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--

To see the `.cat-head` element, give it a linear gradient background with `#5e5e5e` at `85%` and `#45454f` at `100%`.

Huenda usione tofauti kati ya rangi hizi mbili, lakini zipo.

# --hints--

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `background`.

```js
assert.match(code, /background:/)
```

Sifa yako ya `background` inapaswa kutumia kitendakazi cha `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

Kitendakazi chako cha `linear-gradient` kinapaswa kuweka rangi ya kwanza kuwa `#5e5e5e` kwa `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

Kitendakazi chako cha `linear-gradient` kinapaswa kuweka rangi ya pili kuwa `#45454f` kwa `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
