---
id: 646cf1206cac5f51804f49cf
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

Ili kuweka sikio la kushoto vizuri, unaweza kutumia CSS transform ili kulizungusha kwa kiwango fulani.

Sifa ya `transform` hukuruhusu kurekebisha umbo, nafasi, na ukubwa wa kipengele bila kubadilisha mpangilio au kuathiri vipengele vinavyozunguka. Ina vitendaji kama vile `translate()`, `rotate()`, `scale()`, `skew()` na `matrix()`.

Weka sifa ya `transform` kuwa `rotate(-45deg)` na uone kitakachotokea.

# --hints--

Unapaswa kuweka sifa ya `transform` ya kipengele cha `.cat-left-ear` kuwa `rotate(-45deg)`. Usisahau kuongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.transform === 'rotate(-45deg)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.cat-left-ear {
  position: absolute;
  top: -26px;
  left: -31px;
  border-top-left-radius: 90px;
  border-top-right-radius: 10px;
  --fcc-editable-region--

  --fcc-editable-region--
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-bottom: 70px solid #5e5e5e;
}
```
