---
id: 6476fc5cf14b276e6d04e82a
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Tumia kichaguzi cha darasa ili kukipa kipengele cha `.box` upana wa `200px`, urefu wa `600px` na rangi ya mandharinyuma ya `#000`. Pia, ipe `position` ya `absolute`, `top` ya `800px` na `left` ya `650px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.box`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

Kichaguzi chako cha `.box` kinafaa kuwa na sifa ya `width` iliyowekwa kuwa `200px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

Kichaguzi chako cha `.box` kinafaa kuwa na sifa ya `height` iliyowekwa kuwa `600px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

Kichaguzi chako cha `.box` kinafaa kuwa na sifa ya `background-color` iliyowekwa kuwa `#000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

Kichaguzi chako cha `.box` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

Kichaguzi chako cha `.box` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa `650px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

Kichaguzi chako cha `.box` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `800px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

