---
id: 6476f7a4827bcc61682f2347
title: Hatua ya 11
challengeType: 0
dashedName: step-11
---

# --description--

Sifa ya nafasi inayofuata ni `absolute`. Unapotumia thamani ya `absolute` kwa sifa ya `position` yako, kipengele kinatolewa nje ya mtiririko wa kawaida wa hati, na kisha nafasi yake kuamuliwa na sifa za `top`, `right`, `bottom`, na `left`.

Weka sifa ya nafasi ya kipengele chako cha `.cat-head` kuwa `absolute`, kisha weka sifa za `top` na `left` kuwa thamani yoyote.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `top` iliyowekwa kwa nambari yoyote katika `px`. Hakikisha unaongeza nusu koloni.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa nambari yoyote katika `px`. Hakikisha unaongeza nusu koloni.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
