---
id: 6476f5c17f99146071ee884c
title: Hatua ya 10
challengeType: 0
dashedName: step-10
---

# --description--

Unaweza kuona kwamba hakuna kinachotokea. Kipengele cha `.cat-head` hakikusogezwa licha ya kuweka `top` na `left` ya `100px` kila moja. Lakini sivyo ilivyo kwa nafasi ya `relative`.

Unapotumia thamani ya `relative`, kipengele bado kimewekwa kulingana na mtiririko wa kawaida wa hati, lakini `top`, `left`, `bottom`, na `right` thamani zinaanza kutumika.

Instead of `static`, give your `.cat-head` a position of `relative`, and leave both `top` and `left` properties as they are.

# --hints--

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `relative`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `100px`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa `100px`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
