---
id: 6477062778c85972eb648030
title: Hatua ya 16
challengeType: 0
dashedName: step-16
---

# --description--

Unapaswa sasa kuweka kichwa cha paka katikati.

Give the `.cat-head` element a `position` property set to `absolute`. Weka thamani ya `0` kwa sifa za `right`, `left`, `top`, `bottom`, kisha weka sifa `margin` pande zote kuwa `auto`. Hiyo ni njia moja ya kuweka kipengee katikati kiwima na mlalo kwa kutumia uwekaji wa CSS.

# --hints--

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `absolute`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `0`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `left` iliyowekwa kuwa `0`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '0px')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `right` iliyowekwa kuwa `0`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.right === '0px')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `bottom` iliyowekwa kuwa `0`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.bottom === '0px')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `margin` iliyowekwa kuwa `auto`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.margin === 'auto')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: sticky;
  top: 0;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 400px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
