---
id: 64770351e8586671ec0911f0
title: Hatua ya 15
challengeType: 0
dashedName: step-15
---

# --description--

Thamani ya nafasi ya mwisho ni `sticky`. Nafasi ya `sticky` ni mseto wa nafasi za `relative` na `fixed`. Huruhusu kipengee **kushikamana** na nafasi maalum ndani ya kipengele kilicho na kipengee chake au kituo cha kutazama, kulingana na nafasi ya kusogeza.

Badilisha thamani ya sifa ya `position` ya `.cat-head` hadi `sticky`, weka `top` hadi `0`, kisha uondoe `left` na thamani yake.

**Kumbuka**: Ili kuona jinsi `sticky` inavyofanya kazi, inabidi uweke maandishi kadhaa kabla na baada ya kipengele chako cha `.cat-head` `div`. If you scroll down after that, you'll see that the `.cat-head` gets stuck to the top and remains there.

# --hints--

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `position` iliyowekwa kuwa `sticky`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

Kichaguzi chako cha `.cat-head` kinafaa kuwa na sifa ya `top` iliyowekwa kuwa `0`. Hakikisha unaongeza nusu koloni.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

Hupaswi kuwa na kipengele cha `left` na thamani yake katika msimbo wako.

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
