---
id: 6676a8a8d9861319ee901a90
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

Now you can begin working on your cat's ears.

Clean up your review code by removing the `.cat-right-ear` selector and all of its properties.

Using a class selector, give the `.cat-left-ear` element a left and right border of `35px solid transparent` each. Also, set the bottom border to `70px solid #5e5e5e`.

# --hints--


Your `.cat-right-ear` selector should not have a `height` property.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

Your `.cat-right-ear` selector should not have a `width` property.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


Your `.cat-right-ear` selector should not have a `border-left` property.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

Your `.cat-right-ear` selector should not have a `border-right` property.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

Your `.cat-right-ear` selector should not have a `border-top` property.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

Your `.cat-right-ear` selector should not have a `border-bottom` property.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

You should not have a `.cat-right-ear` selector.

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

You should have a `.cat-left-ear` selector.

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

Your `.cat-left-ear` selector should have a `border-left` property set to `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

Your `.cat-left-ear` selector should have a `border-right` property set to `35px solid transparent`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

Your `.cat-left-ear` selector should have a `border-bottom` property set to `70px solid #5e5e5e`

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
