---
id: 658252f6b1526d57b103d48a
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

В межах інструкції `if` потрібно попросити користувача ввести суму витрати та зберегти її в змінній.

Передайте рядок `'Enter amount: '` до порожнього виклику `input()`, щоб мати змогу зберігати витрати.

# --hints--

Інструкція `if` повинна містити `amount = input('Enter amount: ')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+if\s+choice\s*==\s*("|')1\1\s*:\s*amount\s*=\s*input\s*\(\s*("|')Enter amount:\s\2\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')

        if choice == '1':
            amount = input()
--fcc-editable-region--
```
