---
id: 65823040bdb32949d4c1a96f
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Вам знадобиться функція для підрахунку загальної суми витрат. Визначте функцію під назвою `total_expenses`, яка приймає один параметр `expenses`. Наразі заповніть тіло функції інструкцією `pass`.

# --hints--

Визначте функцію під назвою `total_expenses`.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Функція `total_expenses` має приймати параметр `expenses`.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
