---
id: 6582601b2987045e8e7da994
title: Крок 39
challengeType: 20
dashedName: step-39
---

# --description--

Після виклику `print()` потрібно відфільтрувати витрати та надрукувати відфільтрований список. Оголосіть змінну `expenses_from_category` та призначте до неї виклик `filter_expenses_by_category`, передавши `expenses` та `category` як аргументи.

# --hints--

Інструкція `elif` повинна містити `expenses_from_category = filter_expenses_by_category(expenses, category)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')4\1\s*:\s*category\s*=\s*input\s*\(\s*("|')Enter\scategory\sto\sfilter:\s\2\s*\)\s*print\s*\(\s*f("|')\\nExpenses for \{\s*category\s*\}:\3\s*\)\s*expenses_from_category\s*=\s*filter_expenses_by_category\s*\(\s*expenses\s*,\s*category\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)

        elif choice == '3':
            print('\nTotal Expenses: ', total_expenses(expenses))
--fcc-editable-region--
        elif choice == '4':
            category = input('Enter category to filter: ')
            print(f'\nExpenses for {category}:')

--fcc-editable-region--
```
