---
id: 65822308ded1ba4632f66aa3
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Словник — ще один вбудований тип даних у Python. Словник — це набір даних у вигляді пари *ключ*-*значення*. Словники визначають за допомогою фігурних дужок `{}`, які містять пари ключ-значення, розділені комами. Після кожного ключа стоїть двокрапка `:` та значення:

```py
{'amount': 50.0, 'category': 'Food'}
```

У наведеному вище прикладі `'amount'` та `'category'` є ключами, а `50.0` та `'Food'` є їхніми відповідними значеннями.

Створіть словник з ключем `'amount'` та значенням параметра `amount` і передайте новий словник до виклику `.append()`.

# --hints--

Функція повинна містити `expenses.append({'amount': amount})`.

```js
assert.match(code, /^\s+expenses\.append\s*\(\s*\{\s*("|')amount\1\s*:\s*amount\s*\}\s*\)/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
