---
id: 658212ba98182f3e855e85f9
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви створите простий, але функціональний трекер витрат на Python.

Почніть з визначення функції під назвою `add_expense`, яка приймає три параметри: `expenses`, `amount` та `category`. Використайте ключове слово `pass`, щоб заповнити тіло функції.

# --hints--

You should use the `def` keyword to declare a function named `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Функція `add_expense` має приймати три параметри: `expenses`, `amount` та `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
