---
id: 658246c28575d653d1f89b59
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Цикл `while` — ще один вид циклу, який виконує частину коду, допоки вказана умова оцінюється як `True`:

```py
while condition:
    <code>
```

Створіть цикл `while` під списком `expenses`. Використайте `True` для умови та надрукуйте рядок `'\nExpense Tracker'` в межах тіла циклу, щоб відтворити назву програми.

# --hints--

Функція `main()` повинна містити `while True:`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+while\s+True\s*:/m));
  }
})
```

Цикл `while` повинен містити `print('\nExpense Tracker')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+print\s*\(\s*("|')\\nExpense\s*Tracker\1\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []

--fcc-editable-region--
```
