---
id: 6582352cbdbcde4d0c4f7b0b
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Щоб викликати лямбда-функцію, можна використати синтаксис звичайної функції з парою круглих дужок після назви змінної.

Викличте лямбда-функцію `test` та передайте `3` як аргумент. Потім надрукуйте результат.

# --hints--

Код повинен містити `print(test(3))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
