---
id: 65822934c8a0454803e3ef12
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Далі ви відтворите деталі всіх витрат.

Замініть `pass` на виклик `print()` в межах циклу `for` та передайте до нього наступний f-рядок: `f'Amount: {expense}, Category: {expense}'`. Наразі залиште заповнювачі порожніми.

# --hints--

Цикл `for` повинен містити `print(f'Amount: {expense}, Category: {expense}')`.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\}, Category: \{\s*expense\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
