---
id: 65824dfdb6815d563b2d3256
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Функція `input()` приймає введені користувачем дані та повертає їх у вигляді рядка.

Викличте функцію `input()` в межах циклу `while`, передавши рядок `'Enter your choice: '` як аргумент, та призначте результат до змінної під назвою `choice`.

# --hints--

Оголосіть змінну під назвою `choice` в межах циклу `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].has_variable("choice")
`)) })
```

Призначте результат функції `input()`, аргументом якої є рядок `'Enter your choice: '`, до змінної `choice`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_variable("choice").is_equivalent("choice = input('Enter your choice: ')")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

--fcc-editable-region--
```
