---
id: 658238f7604f154ea9a23e1e
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Функція `sum()` повертає суму елементів ітерованого об’єкта, який передається як аргумент. Ви використовуватимете `sum()` разом із функціями `map()` та `lambda`, щоб отримати загальну суму витрат.

Наразі створіть невеликий текст та змініть наявний виклик `print()`, передавши `sum(map(test, [2, 3, 5, 8]))` як аргумент.

# --hints--

Надрукуйте `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
