---
id: 65822934c8a0454803e3ef12
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Далі ви відтворите деталі всіх витрат.

Замініть `pass` на виклик `print()` в межах циклу `for` та передайте до нього наступний f-рядок: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Надрукуйте `f'Amount: {expense}, Category: {expense}'` у циклі `for`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

Цикл не повинен містити `pass`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
