---
id: 658225d560369446ece5518b
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Додайте ще одну пару ключ-значення до словника, який додаєте до списку `expense`. Використайте рядок `'category'` як ключ, а параметр `category` як значення.

# --hints--

Код повинен містити `expenses.append({'amount': amount, 'category': category})`.

```js
assert.match(code, /^\s+expenses\.append\s*\(\s*\{\s*("|')amount\1\s*:\s*amount\s*,\s*("|')category\2\s*:\s*category\s*\}\s*\)/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
