---
id: 65821c1c40ccac44adf38e6b
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Список у Python — це вбудований тип даних, який дозволяє зберігати багато елементів в одній структурі даних. In Python, you create a list by putting items inside square brackets (`[]`), with each item separated from the following one by a comma.

```py
numbers = [1, 2, 3, 4]
```

Use a pair of square brackets to create an empty list named `expenses`. Ви використовуватимете його, щоб зберігати витрати.

# --hints--

Оголосіть нову змінну під назвою `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Призначте порожній список до змінної `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass

--fcc-editable-region--

--fcc-editable-region--
```
