---
id: 65823cfc74aa564ffc460489
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Тепер викличте `map()`, передавши функцію `lambda` як перший аргумент та список `expenses` як другий аргумент.

# --hints--

Код повинен містити `map(lambda expense: expense['amount'], expenses)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "total_expenses");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+map\s*\(\s*lambda\s+expense\s*:\s*expense\s*\[\s*("|')amount\1\s*\]\s*,\s*expenses\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
