---
id: 6582324b5d37fe4b622bc3c0
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Лямбда-функції — це лаконічні анонімні функції в Python, ідеальні для простих одноразових завдань. Їх визначають за допомогою ключового слова `lambda`, використовуючи такий синтаксис:

```py
lambda x: expr
```

`x` в прикладі вище є параметром, який потрібно використати у виразі `expr`. Створіть змінну `test` та призначте до неї лямбда-функцію, яка приймає параметр `x` та повертає `x * 2`.

# --hints--

Код повинен містити `test = lambda x: x * 2`.

```js
({ test: () => assert.match(code,/^test\s*=\s*lambda\s+x\s*:\s*x\s*\*\s*2/m)})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
