---
id: 65822bd82d708c4895080c35
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Ви можете отримати доступ до значень в словнику через ключі. Для цього використайте дужкову нотацію та помістіть ключ між квадратними дужками:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Зараз ви інтерполюєте словник `expense` в f-рядок. Змініть вираз f-рядка, щоб отримати значення ключів `'amount'` та `'category'` в словнику `expense`.

# --hints--

Передайте `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` до виклику `print()`. Не забудьте, що в f-рядку з одинарними лапками потрібно використовувати подвійні лапки та навпаки.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
