---
id: 65823e87c9741750a22085a7
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Потім визначте функцію під назвою `filter_expenses_by_category`, яка приймає два параметри: `expenses` та `category`. Використайте `pass`, щоб заповнити тіло функції.

# --hints--

Визначте функцію під назвою `filter_expenses_by_category`.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

Функція `filter_expenses_by_category` має приймати `expenses` та `category` як параметри.

```js
({ test: () => assert(__pyodide.runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
