---
id: 65823b0167fd894f4a7ea60a
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Далі ви будете реалізовувати таку ж логіку до функції `total_expenses`.

Наразі видаліть функцію `test` та виклик `print()`.

# --hints--

Видаліть змінну `test`.

```js
({ test: () => assert.notMatch(code, /test\s*=/) })
```

Видаліть виклик `print(sum(map(test, [2, 3, 5, 8])))`.

```js
({ test: () => assert.notMatch(code, /print\s*\(\s*sum/) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
