---
id: 65823634a308cb4d6bcc5fa6
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Лямбда-функції можна поєднати з функцією `map()`, яка виконує певну функцію для кожного елемента в колекції об’єктів, як-от список:

```py
map(lambda x: x * 2, [1, 2, 3])
```

Результатом прикладу вище буде `[2, 4, 6]`, де кожен елемент списку, переданого до `map()`, було подвоєно лямбда-функцією.

Змініть виклик `print()`, щоб надрукувати результат виклику `map()`, де першим аргументом є `test`, а другим — `[2, 3, 5, 8]`.

# --hints--

Код повинен містити `print(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
