---
id: 65823ff0d4b991510fade1a8
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Замініть `pass` на функцію `lambda` в межах функції `filter_expenses_by_category`. Використайте `expense` як параметр і поверніть порівняння між значенням ключа `'category'` словника `expense` та `category`.

# --hints--

Функція `filter_expenses_by_category` повинна містити `lambda expense: expense['category'] == category`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "filter_expenses_by_category");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+lambda\s+expense\s*:\s*expense\s*\[\s*("|')category\1\s*\]\s*==\s*category/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
