---
id: 658246c28575d653d1f89b59
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--

Цикл `while` — ще один вид циклу, який виконує частину коду, допоки вказана умова оцінюється як `True`:

```py
while condition:
    <code>
```

Створіть цикл `while` під списком `expenses`. Використайте `True` для умови та надрукуйте рядок `'\nExpense Tracker'` в межах тіла циклу, щоб відтворити назву програми.

# --hints--

You should create a `while` loop using `True` as the condition within your `main()` function.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

You should print `'\nExpense Tracker'` in your `while` loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []

--fcc-editable-region--
```
