---
id: 65824872894f59546e3084e2
title: Крок 25
challengeType: 20
dashedName: step-25
---

# --description--

Цикл `while`, який ви створили в попередньому кроці, є нескінченним циклом, що дозволить програмі постійно показувати параметри меню, доки користувач не вирішить вийти.

Додайте ще один виклик після виклику `print()`, щоб надрукувати рядок `'1. Add an expense'`.

# --hints--

Цикл `while` повинен містити `print('1. Add an expense')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+print\s*\(\s*("|')1\.\sAdd\san\sexpense\1\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')

--fcc-editable-region--
```
