---
id: 658255d5f955175b270f251d
title: Step 32
challengeType: 20
dashedName: step-32
---

# --description--

Inside your `if` statement, create a variable `category` to store the expense category. Assign it `input('Enter category: ')`.

# --hints--

You should have `category = input('Enter category: ')` in your `if` statement.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+if\s+choice\s*==\s*("|')1\1\s*:\s*amount\s*=\s*float\s*\(\s*input\s*\(\s*("|')Enter\samount:\s\2\s*\)\s*\)\s*category\s*=\s*input\s*\(\s*("|')Enter\scategory:\s\3\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')
--fcc-editable-region--       
        if choice == '1':
            amount = float(input('Enter amount: '))

--fcc-editable-region--
```
