---
id: 6582324b5d37fe4b622bc3c0
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Лямбда-функції — це лаконічні анонімні функції в Python, ідеальні для простих одноразових завдань. Їх визначають за допомогою ключового слова `lambda`, використовуючи такий синтаксис:

```py
lambda x: expr
```

`x` у прикладі вище представляє параметр, який буде використано у виразі `expr` та який поводиться як будь-який інший параметр традиційної функції. `expr` — це вираз, який оцінюється та повертається, якщо викликано лямбда-функцію.

Створіть змінну під назвою `test` та призначте до неї лямбда-функцію, яка приймає параметр `x` та повертає `x * 2`.

# --hints--

Оголосіть змінну під назвою `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Призначте `lambda x: x * 2` до змінної `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
