---
id: 658258c0e5fbe85c14c060cf
title: Крок 34
challengeType: 20
dashedName: step-34
---

# --description--

Щоб перерахувати всі витрати, можна використати інструкцію `elif` після інструкції `if`. Інструкція `elif` перевіряє додаткові умови та працює лише тоді, коли розміщена після інструкції `if`.

Створіть інструкцію `elif`, щоб перевірити, чи вибір користувача дорівнює рядку `2`. Надрукуйте рядок `'\nAll Expenses:'` в межах інструкції `elif`.

# --hints--

Інструкція `elif choice == '2':` має розташовуватись після інструкції `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')2\1\s*:/m));
  }
})
```

Інструкція `elif` повинна містити `print('\nAll Expenses:')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')2\1\s*:\s*print\s*\(\s*("|')\\nAll\sExpenses:\2\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

--fcc-editable-region--
```
