---
id: 65821fcc010c3245718f2a06
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Параметром `expenses` функції `add_expense` буде список витрат. Важливо мати можливість додавати елементи в кінець списку. For that you'll use the `.append()` list method:

```py
my_list = [2, 4, 7]
my_list.append(3)
```

In the example above, after appending `3`, `my_list` would be `[2, 4, 7, 3]`.

Replace `pass` with a call to the `.append()` method on the `expenses` list. Don't pass any arguments to `.append()` for now.

# --hints--

You should have `expenses.append()` in your function.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

You should not have `pass` in your function.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py    
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
