---
id: 65822773f4c60a479ec89f66
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

В межах функції `print_expenses` створіть цикл `for`, який перебирає кожен елемент в списку `expenses`. Використайте `expense` як змінну циклу та перемістіть `pass` до тіла циклу.

# --hints--

Функція `print_expenses` повинна містити `for expense in expenses:`. Не забудьте використати `pass` в межах тіла циклу.

```js
({ test: () => assert.match(code, /for\s+expense\s+in\s+expenses\s*:/ ) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
