---
id: 65822308ded1ba4632f66aa3
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Словник — ще один вбудований тип даних у Python. Словник — це набір даних у вигляді пари *ключ*-*значення*. Dictionaries are defined with curly braces (`{}`) and they contain key-value pairs separated by commas. Each key is followed by a colon (`:`) and the value:

```py
{'amount': 50.0, 'category': 'Food'}
```

У наведеному вище прикладі `'amount'` та `'category'` є ключами, а `50.0` та `'Food'` є їхніми відповідними значеннями.

Створіть словник з ключем `'amount'` та значенням параметра `amount` і передайте новий словник до виклику `.append()`.

# --hints--

You should pass a dictionary as the argument to the `.append()` method.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

You should pass `{'amount': amount}` as the argument to the `.append()` method.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
