---
id: 65825b758fe85a5cebabc8c5
title: Крок 36
challengeType: 20
dashedName: step-36
---

# --description--

Щоб показати загальні витрати, створіть інструкцію `elif`, яка перевіряє, чи `choice == '3'`.

Якщо так, то користувач хоче бачити загальні витрати. Тож викличте `print()` та передайте рядок `'\nTotal Expenses: '` як перший аргумент і `total_expenses(expenses)` як другий аргумент.

# --hints--

Код повинен містити `elif choice == '3':`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert.match(function_body, /\s+elif\s+choice\s*==\s*("|')3\1\s*:/);
  }
})
```

Інструкція `elif` повинна містити `print('\nTotal Expenses: ', total_expenses(expenses))`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+elif\s+choice\s*==\s*("|')3\1\s*:\s*print\s*\(\s*("|')\\nTotal\sExpenses:\s\2\s*,\s*total_expenses\s*\(\s*expenses\s*\)\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')

        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)
--fcc-editable-region--
        elif choice == '2':
            print('\nAll Expenses:')
            print_expenses(expenses)

--fcc-editable-region--
```
