---
id: 65822773f4c60a479ec89f66
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Створіть цикл `for` в межах функції `print_expenses`, який ітерує над кожним елементом в списку `expenses`. Використайте `expense` як змінну циклу та перемістіть `pass` до тіла циклу.

# --hints--

You should create a `for` loop that iterates over `expenses` in your `print_expenses` function. Не забудьте використати `pass` в межах тіла циклу.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

You should use `expense` as the loop variable.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

You should have `pass` only inside your loop body.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
