---
id: 6582507654b3ed5712341382
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Ви використовуватимете умовні інструкції, щоб перевірити вибір користувача. Якщо користувач вибрав `'1'`, то він хоче додати витрати.

Напишіть інструкцію `if` в циклі `while` під змінною `choice`, щоб перевірити, чи `choice` дорівнює рядку `'1'`. Якщо так, то це буде відправною точкою для додавання нових витрат.

Оголосіть змінну `amount` в межах тіла інструкції `if` та передайте до неї порожній виклик `input()`.

# --hints--

Цикл `while` повинен містити `if choice == '1':`:

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+if\s+choice\s*==\s*("|')1\1\s*:/m));
  }
})
```

Тіло інструкції `if` повинне містити `amount = input()`:

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+if\s+choice\s*==\s*("|')1\1\s*:\s*amount\s*=\s*input\s*\(\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')

--fcc-editable-region--
```
