---
id: 65821fcc010c3245718f2a06
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Параметром `expenses` функції `add_expense` буде список витрат. Важливо мати можливість додавати елементи в кінець списку. Для цього можна використати метод списку `.append()`:

```py
my_list = [2, 4, 7]
my_list.append(3)
```

Після додавання `3` у прикладі вище, `my_list` зміниться на `[2, 4, 7, 3]`. Замініть `pass` на виклик методу `.append()` на списку `expenses`. Наразі не передавайте жодних аргументів до `.append()`.

# --hints--

Функція повинна містити `expenses.append()`.

```js
assert.match(code, /^\s+expenses\.append\s*\(\s*\)/m)
```

# --seed--

## --seed-contents--

```py    
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
