---
id: 658258c0e5fbe85c14c060cf
title: Крок 35
challengeType: 20
dashedName: step-35
---

# --description--

Щоб перерахувати всі витрати, можна використати умову `elif` після інструкції `if`. Інструкція `elif` перевіряє додаткові умови та працює лише тоді, коли розміщена після інструкції `if`.

Створіть умову `elif`, щоб перевірити, чи вибір користувача дорівнює рядку `'2'`. Надрукуйте рядок `'\nAll Expenses:'` в межах умови `elif`.

# --hints--

Після інструкції `if` створіть умову `elif`, яка перевіряє, чи `choice` дорівнює `'2'`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1]
cond.is_equivalent("choice == '2'") or cond.is_equivalent("'2' == choice")
`)) })
```

Надрукуйте `'\nAll Expenses:'` в межах умови `elif`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("print('\\\\nAll Expenses:')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)


def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')

        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            add_expense(expenses, amount, category)

--fcc-editable-region--
```
