---
id: 65823dde36392f505a39f7c7
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Зрештою, передайте виклик `map()` до функції `sum()`, щоб отримати загальні витрати та поверніть результат.

# --hints--

Викличте функцію `sum()`, передавши до неї поточний виклик `map()` як перший аргумент, та поверніть результат з функції `total_expenses`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
