---
id: 65824111a09164518320088d
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Функція `filter()` дозволяє вибрати елементи з ітерованого об’єкта (наприклад, списку) на основі результату функції:

```py
filter(my_function, my_list)
```

`filter()` приймає функцію як перший аргумент та ітерований об’єкт як другий аргумент. Вона повертає ітератор, який є спеціальним об’єктом, що дозволяє ітерувати над елементами колекції (наприклад, списку).

Результатом прикладу вище є ітератор, який містить елементи з `my_list`, для якого `my_function` повертає `True`.

Викличте `filter()` в межах функції `filter_expenses_by_category`, передавши функцію `lambda`, яку ви написали в попередньому кроці, як перший аргумент та список `expenses` як другий аргумент.

# --hints--

Викличте `filter()` в межах функції `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Передайте `lambda expense: expense['category'] == category` як перший аргумент до виклику `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Передайте `expenses` як другий аргумент до виклику `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
