---
id: 6582267ef39f5047411c1743
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Next, define a function named `print_expenses` that takes one parameter `expenses`. Пізніше ця функція використовуватиметься для відтворення всіх витрат зі списку.

Заповніть тіло нової функції інструкцією `pass`.

# --hints--

Визначте нову функцію під назвою `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Функція `print_expenses` має приймати параметр `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
