---
id: 65824872894f59546e3084e2
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Цикл `while`, який ви створили в попередньому кроці, є нескінченним циклом, що дозволить програмі постійно показувати параметри меню, доки користувач не вирішить вийти.

Додайте ще один виклик `print()` після виклику `print()`, щоб надрукувати рядок `'1. Add an expense'`.

# --hints--

Надрукуйте `'1. Add an expense'` в межах циклу `while` після наявного виклику `print()`.

```js
({ test: () => assert(runPython(`
call_lst = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")
test_lst = ["print('\\\\nExpense Tracker')", "print('1. Add an expense')"]
all(call_lst[n].is_equivalent(i) for n, i in enumerate(test_lst))
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')

--fcc-editable-region--
```
