---
id: 65824c7b4e2da85597693dcf
title: Крок 27
challengeType: 20
dashedName: step-27
---

# --description--

Надайте інші параметри меню, надрукувавши наступні три рядки в межах циклу `while`: `'3. Show total expenses'`, `'4. Filter expenses by category'` та `'5. Exit'`.

# --hints--

Цикл `while` повинен містити `print('3. Show total expenses')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+print\s*\(\s*("|')3\.\sShow\stotal\sexpenses\1\s*\)/m));
  }
})
```

Цикл `while` повинен містити `print('4. Filter expenses by category')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+print\s*\(\s*("|')4\.\sFilter\sexpenses\sby\scategory\1\s*\)/m));
  }
})
```

Цикл `while` повинен містити `print('5. Exit')`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const foo = __helpers.python.getDef("\n"+transformedCode, "main");
    const {function_body} = foo;    
    assert(function_body.match(/^\s+print\s*\(\s*("|')5\.\sExit\1\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))

def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)

def main():
    expenses = []
--fcc-editable-region--
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')

--fcc-editable-region--
```
