---
id: 658238f7604f154ea9a23e1e
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Функція `sum()` повертає суму елементів ітерованого об’єкта, який передається як аргумент. Ви використовуватимете `sum()` разом із функціями `map()` та `lambda`, щоб отримати загальну суму витрат.

Поки зробіть невеличкий тест та змініть наявний виклик `print()`, замінивши виклик `list()` на виклик до функції `sum()`, передавши тест як аргумент до поточного виклику `map()`.

# --hints--

Надрукуйте `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
