---
id: 658225d560369446ece5518b
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Додайте ще одну пару ключ-значення до словника, який додаєте до списку `expense`. Використайте рядок `'category'` як ключ, а параметр `category` як значення.

# --hints--

Додайте другу пару ключ-значення до словника, який передаєте як аргумент до методу `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

Передайте `{'amount': amount, 'category': category}` як аргумент до методу `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
