---
id: 66314f2add0a8b57da5d6d91
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Ви маєте побачити щось схоже до `<map object at 0xd273a8>` на консолі, що є рядковим представленням map-об’єкта, повернутого функцією `map()`.

Щоб отримати читабельний вивід, потрібно перетворити map-об’єкт на список. Для цього передайте виклик `map()` як аргумент до функції `list()`.

# --hints--

Передайте виклик `map()` як аргумент до функції `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})

def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass

--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
