---
id: 69a5f35669099ed52f8563b1
title: Build a Smart Pantry Restocker
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

In this lab, you will build a small pantry management program using basic JavaScript concepts like arrays, objects, loops, and conditionals.

You will simulate receiving a shipment of pantry items, deciding what to do with each item, and organizing the results for storage. 

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should implement a `parseShipment(rawData)` function that takes an array of strings and returns an array of `{ sku, name, qty, expires }` objects.

2. You should implement a `planRestock(pantry, shipment)` function that compares the current pantry with the incoming shipment and returns an array of actions in the form `{ type: "restock" | "discard" | "donate", item }`.

3. You should implement a `groupByZone(actions)` function that groups the actions into storage zones based on each item’s zone property.

4. You should implement a `clonePantry(pantry)` function that returns a deep copy of the pantry so planning changes do not affect the original list.
A deep copy means creating a new array with new objects, so modifying the copy does not change the original pantry.

5. You should use all of the functions together to process a shipment and log the final grouped result to the console.

# --hints--

You should define a function named `parseShipment` that accepts one
parameter called `rawData`.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

Your `parseShipment` function should convert shipment strings into objects with the
properties `sku`, `name`, `qty`, and `expires`.

``` js
const data = ["A12|Tomatoes|5|2025-01-10"];
const result = parseShipment(data);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A12",
  name: "Tomatoes",
  qty: 5,
  expires: "2025-01-10"
});
```

The `qty` value should be converted into a number.

``` js
const result = parseShipment(["A1|Beans|5|2025-01-01"]);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Duplicate SKUs in the shipment should be ignored.

``` js
const data = [
  "C77|Rice|3|2026-01-01",
  "C77|Rice|3|2026-01-01"
];

const result = parseShipment(data);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C77");
```

You should define a function named `planRestock` that accepts two
parameters: `pantry` and `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

Your `planRestock` function should return an array of actions.

``` js
const pantry = [{ sku: "A12", name: "Tomatoes", qty: 4 }];
const shipment = [{ sku: "A12", name: "Tomatoes", qty: 5 }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Items with a quantity of `0` or less should return an action with type
`"discard"`.

``` js
const pantry = [];
const shipment = [{ sku: "Z1", name: "Milk", qty: 0 }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
```

If a shipment item already exists in the pantry (same `sku`), the action
should be `"restock"`.

``` js
const pantry = [{ sku: "A12", name: "Tomatoes", qty: 4 }];
const shipment = [{ sku: "A12", name: "Tomatoes", qty: 3 }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

If a shipment item does not exist in the pantry, the action should be
`"donate"`.

``` js
const pantry = [{ sku: "A12", name: "Tomatoes", qty: 4 }];
const shipment = [{ sku: "B99", name: "Beans", qty: 2 }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

You should define a function named `groupByZone` that accepts one
parameter called `actions`.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

`groupByZone` should group actions by the `zone` property of each item's
object.

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

If an item does not have a `zone`, it should be grouped under
`"general"`.

``` js
const actions = [
  { type: "donate", item: { sku: "X1" } }
];

const result = groupByZone(actions);

assert.isArray(result.general);
```

You should define a function named `clonePantry` that accepts one
parameter called `pantry`.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` should return a new array instead of the original pantry
array.

``` js
const pantry = [{ sku: "A1", name: "Rice", qty: 3 }];
const copy = clonePantry(pantry);

assert.isArray(copy);
assert.notStrictEqual(copy, pantry);
```

The objects inside the cloned pantry should also be new objects (deep
copy).

``` js
const pantry = [{ sku: "A1", name: "Rice", qty: 3 }];
const copy = clonePantry(pantry);

assert.notStrictEqual(copy[0], pantry[0]);
assert.deepEqual(copy[0], pantry[0]);

copy[0].qty = 100;
assert.strictEqual(pantry[0].qty, 3);
```

The functions should work together to process a shipment and group the
resulting actions.

``` js
const pantry = [{ sku: "A12", name: "Tomatoes", qty: 4, zone: "fridge" }];
const raw = ["A12|Tomatoes|3|2025-01-10"];

const shipment = parseShipment(raw);
const actions = planRestock(pantry, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.isArray(grouped.fridge);
```

# --seed--

## --seed-contents--

```js
function parseShipment(rawData) {

}

function planRestock(pantry, shipment) {

}

function groupByZone(actions) {

}

function clonePantry(pantry) {

}

const pantry = [
  { sku: "A12", name: "Tomatoes", qty: 4, expires: "2025-01-10", zone: "fridge" },
  { sku: "B44", name: "Pasta", qty: 2, expires: "2026-03-02", zone: "pantry" }
];

const rawShipment = [
  "A12|Tomatoes|5|2025-01-10",
  "C77|Rice|3|2026-01-01",
  "C77|Rice|3|2026-01-01"
];

const shipment = parseShipment(rawShipment);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

console.log(grouped);

```

# --solutions--

```js

function parseShipment(rawData) {
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let i = 0; i < rawData.length; i++) {
    const parts = rawData[i].split("|");

    const sku = parts[0];
    const name = parts[1];
    const qty = Number(parts[2]);
    const expires = parts[3];

    let duplicate = false;
    for (let j = 0; j < seenSKUs.length; j++) {
      if (seenSKUs[j] === sku) {
        duplicate = true;
        break;
      }
    }
    if (duplicate) continue;

    seenSKUs.push(sku);

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires
    });
  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let i = 0; i < shipment.length; i++) {
    const item = shipment[i];

    if (item.qty <= 0) {
      actions.push({ type: "discard", item: item });
      continue;
    }

    let found = false;
    for (let j = 0; j < pantry.length; j++) {
      if (pantry[j].sku === item.sku) {
        found = true;
        break;
      }
    }

    if (found) {
      actions.push({ type: "restock", item: item });
    } else {
      actions.push({ type: "donate", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {
  const zones = {};

  for (let i = 0; i < actions.length; i++) {
    const action = actions[i];
    const zone = action.item.zone || "general";

    if (!zones[zone]) {
      zones[zone] = [];
    }

    zones[zone].push(action);
  }

  return zones;
}

function clonePantry(pantry) {
  const copy = [];

  for (let i = 0; i < pantry.length; i++) {
    const item = pantry[i];
    copy.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return copy;
}

const pantry = [
  { sku: "A12", name: "Tomatoes", qty: 4, expires: "2025-01-10", zone: "fridge" },
  { sku: "B44", name: "Pasta", qty: 2, expires: "2026-03-02", zone: "pantry" }
];

const rawShipment = [
  "A12|Tomatoes|5|2025-01-10",
  "C77|Rice|3|2026-01-01",
  "C77|Rice|3|2026-01-01"
];

const shipment = parseShipment(rawShipment);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

console.log(grouped);

```
