---
id: 68e97fe79367ad7b5dd6c9cd
title: Step 3
challengeType: 0
dashedName: step-3
demoType: onLoad
---

# --description--

In this step add an `h1` element with the text `XYZ Bookstore`.

# --hints--

You should have an `h1` element.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` element's text should be `XYZ Bookstore`. Double check for spelling or capitalization errors.

```js
assert.equal(document.querySelector('h1')?.innerText.trim(), 'XYZ Bookstore');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>
--fcc-editable-region--
<body>

</body>
--fcc-editable-region--
</html>
```
