---
id: 68ebdbacdd3fa474132cc975
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Below the `<h1>` element, add a `<p>` element with this text: `Browse our collection of amazing books!`

# --hints--

Your `p` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('p'));
```

Your `p` element should have a closing tag. Closing tags have a `/` and the same element name: `</p>`.

```js
assert.match(code, /<\/p\s*>/);
```

Your `p` element's text should be: `Browse our collection of amazing books!`. You have either omitted the text, have a typo, or it is not between the `p` element's opening and closing tags.

```js
assert.equal(document.querySelector('p')?.innerText?.trim().toLowerCase(), 'browse our collection of amazing books!');
```

Your `p` element should be below your `h1` element.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
--fcc-editable-region--
```
