---
id: 68ebdbacdd3fa474132cc975
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Below the `h1` element, add a `p` element with this text: `Browse our collection of amazing books!`.

# --hints--

You should have a `p` element.

```js
assert.exists(document.querySelector('p'));
```

Your `p` element's text should be: `Browse our collection of amazing books!`. 

```js
assert.equal(document.querySelector('p')?.innerText.trim(), 'Browse our collection of amazing books!');
```

Your `p` element should be below your `h1` element.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
