---
id: 68ec6d9a315221aa31e54816
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

The `class` attribute is used to identify one or more elements for styling. Multiple elements can share the same class. Unlike the `id` attribute, class names do not need to be unique.

Add a `class` attribute to your `div` from Step 3 and set its value to `card-container`.

# --hints--

Your `div` element should have a `class` attribute.

```js
assert.exists(document.querySelector('div')?.getAttribute('class'));
```

Your `div` element should have a `class` with value `card-container`. 

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div></div>
--fcc-editable-region--
```
