---
id: 68ec99e478211dc578699944
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Below the `<h2>` element in the first card, add a `<p>` element with this text:
`This is an epic story of Sally and her dog Rex as they navigate through other worlds.`


# --hints--

You should have a `p` element nested inside the `.card` div.

```js
assert.exists(document.querySelector('.card p'));
```

Your `p` element should have both opening and closing tags. Opening tags have this syntax: `<elementName>`. Closing tags have a `/` and the same element name: `</elementName>`.

```js
const cardRegex = /<div\s+class\s*=\s*['"]card['"]\s+id\s*=\s*['"]sally-adventure-book['"][^>]*>([\s\S]*?)<\/div\s*>/i;
const cardMatch = code.match(cardRegex);
assert.exists(cardMatch, 'Could not find the card div');
const cardContent = cardMatch[1];
const openPCount = (cardContent.match(/<p/gi) || []).length;
const closePCount = (cardContent.match(/<\/p/gi) || []).length;
assert.equal(openPCount, closePCount, 'Each <p> tag inside the card should have a closing </p> tag');
```

Your `p` element's text should be `This is an epic story of Sally and her dog Rex as they navigate through other worlds.` . You have either omitted the text, have a typo, or it is not between the p element's opening and closing tags.

```js
assert.equal(document.querySelector('.card p')?.innerText?.trim(), "This is an epic story of Sally and her dog Rex as they navigate through other worlds.");
```

Your `p` element should be below your `h2` element.

```js
assert.exists(document.querySelector('.card h2 + p'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
  </div>
</div>
--fcc-editable-region--
```
