---
id: 68eca6e5a6759ed4ea0034dc
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Inside the element with a class of `btn-container`, add two `button` elements:

First button:

- Id: `view-cart-btn`
- Class: `btn`
- Text: `View Cart`

Second button:

- Id: `checkout-btn`
- Class: `btn`
- Text: `Checkout`


Congratulations! You have successfully built the structure of a bookstore page using divs, classes, and ids to organize your content.

# --hints--

You should have two `button` elements nested inside the `.btn-container` element.

```js
assert.lengthOf(document.querySelectorAll('.btn-container button'), 2);
```

Your first button inside element with the class of `btn-container` should have an `id` attribute with the value of `view-cart-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.id, 'view-cart-btn');
```

Your first button inside the element having a class of `btn-container` should have a `class` attribute with the value of `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.className, 'btn');
```

Your first button's text should be `View Cart`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[0]?.innerText.trim(), 'View Cart');
```

Your second button inside the element with a class of `btn-container` should have an `id` attribute with the value of `checkout-btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.id, 'checkout-btn');
```

Your second button inside the element with `class="btn-container"` should have a `class` attribute with the value of `btn`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.className, 'btn');
```

Your second button's text should be `Checkout`.

```js
const buttons = document.querySelectorAll('.btn-container button');
assert.equal(buttons[1]?.innerText.trim(), 'Checkout');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
    <p>This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.</p>
    <button class="btn">Buy Now</button>
  </div>
</div>
--fcc-editable-region--
<p>Review your selections and continue to checkout.</p>
<div class="btn-container">

</div>
--fcc-editable-region--
```

# --solutions--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>

<div class="card-container">
     <div class="card" id="sally-adventure-book">
     <h2>Sally's SciFi Adventure</h2>
     <p>
          This is an epic story of Sally and her dog Rex as they navigate
          through other worlds.
     </p>
     <button class="btn">Buy Now</button>
     </div>

     <div class="card" id="dave-cooking-book">
     <h2>Dave's Cooking Adventure</h2>
     <p>
          This is the story of Dave as he learns to cook everything from
          pancakes to pasta, one recipe at a time.
     </p>
     <button class="btn">Buy Now</button>
     </div>
</div>

<p>Review your selections and continue to checkout.</p>

<div class="btn-container">
     <button id="view-cart-btn" class="btn">View Cart</button>
     <button id="checkout-btn" class="btn">Checkout</button>
</div>
```
