---
id: 68eca5c0065c82d256d29ca3
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Add a `<button>` element inside the second card with the `class` attribute set to `btn` and the text:
`Buy Now`

Both buttons now share the same class, which means they can be styled consistently together.

# --hints--

You should have a `button` element nested inside the second `.card` div.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button'));
```

Your second card's `button` element should have both opening and closing tags. Opening tags have this syntax: `<elementName>`. Closing tags have a `/` and the same element name: `</elementName>`.

```js
const secondCardRegex = /<div\s+class\s*=\s*['"]card['"]\s+id\s*=\s*['"]dave-cooking-book['"][^>]*>([\s\S]*?)<\/div\s*>/i;
const secondCardMatch = code.match(secondCardRegex);
assert.exists(secondCardMatch, 'Could not find the second card div');
const cardContent = secondCardMatch[1];
const openButtonCount = (cardContent.match(/<button/gi) || []).length;
const closeButtonCount = (cardContent.match(/<\/button/gi) || []).length;
assert.equal(openButtonCount, closeButtonCount, 'Each  tag inside the second card should have a closing  tag');
```

Your second card's `button` element should have a `class` attribute.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('button')?.getAttribute('class'));
```

Your second card's `button` element should have a `class` with value `btn`.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.className, 'btn');
```

Your second card's `button` element's text should be `Buy Now` . You have either omitted the text, have a typo, or it is not between the button element's opening and closing tags.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('button')?.innerText?.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
     <div class="card" id="sally-adventure-book">
          <h2>Sally's SciFi Adventure</h2>
          <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
           <button class="btn">Buy Now</button>
     </div>
     <div class="card" id="dave-cooking-book">
          <h2>Dave's Cooking Adventure</h2>
          <p>
               This is the story of Dave as he learns to cook everything from
               pancakes to pasta, one recipe at a time.
          </p>
     </div>
</div>
--fcc-editable-region--
```
