---
id: 68eca11e99e3c5c894ca9d69
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

The `button` element is used to create clickable buttons on a webpage. Buttons are interactive elements that users can click to perform actions.

You can add multiple elements inside a `div` element to group related content. Add a `button` element inside the first `.card` with the `class` attribute set to `btn` and the text`Buy Now`

# --hints--

You should have a `button` element nested inside the `.card` element.

```js
assert.exists(document.querySelector('.card button'));
```

Your `button` element should have both opening and closing tags. Opening tags have this syntax: `<elementName>`. Closing tags have a `/` and the same element name: `</elementName>`.

```js
const cardRegex = /<div\s+class\s*=\s*['"]card['"]\s+id\s*=\s*['"]sally-adventure-book['"][^>]*>([\s\S]*?)<\/div\s*>/i;
const cardMatch = code.match(cardRegex);
assert.exists(cardMatch, 'Could not find the card div');
const cardContent = cardMatch[1];
const openButtonCount = (cardContent.match(/<button/gi) || []).length;
const closeButtonCount = (cardContent.match(/<\/button/gi) || []).length;
assert.equal(openButtonCount, closeButtonCount, 'Each  tag inside the card should have a closing  tag');
```

Your `button` element should have a `class` attribute.

```js
assert.exists(document.querySelector('.card button')?.getAttribute('class'));
```

Your `button` element should have a `class` with the value of `btn`.

```js
assert.equal(document.querySelector('.card button')?.className, 'btn');
```

Your `button` element's text should be `Buy Now`. 

```js
assert.equal(document.querySelector('.card button')?.innerText?.trim(), 'Buy Now');
```

# --seed--

## --seed-contents--

```html
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class="card-container">
  --fcc-editable-region--
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>

  </div>
  --fcc-editable-region--
</div>
```
