---
id: 68ec9332a9b5b2b32487bd00
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

The `id` attribute adds a unique identifier to an HTML element. Each `id` should be unique within a page and should only be used once. `id` values cannot contain spaces and should only contain letters, digits, underscores, and dashes.

Add an `id` attribute to your first card `div` and set its value to `sally-adventure-book`.

# --hints--

Your `.card` div should have an `id` attribute.

```js
assert.exists(document.querySelector('.card')?.getAttribute('id'));
```

Your `.card` div should have an `id` with value `sally-adventure-book`.

```js
assert.equal(document.querySelector('.card')?.id, 'sally-adventure-book');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
     <div class="card"></div>
</div>
--fcc-editable-region--
```
