---
id: 68ec6d9a315221aa31e54816
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

The `class` attribute is used to identify one or more elements for styling. Unlike the `id` attribute, class names do not need to be unique: multiple elements can share the same class.

Here is an example:

```html
<p class="example">example paragraph</p>
```

Add a `class` attribute to your `div` element and set its value to `card-container`.

# --hints--

Your `div` element should have a `class` attribute.

```js
assert.isTrue(document.querySelector('div')?.hasAttribute('class'));
```

Your `div` element should have a `class` attribute with the value of `card-container`.

```js
assert.equal(document.querySelector('div')?.className, 'card-container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title>XYZ Bookstore Page</title>
</head>

<body>
  <h1>XYZ Bookstore</h1>
  <p>Browse our collection of amazing books!</p>
  
--fcc-editable-region--
  <div>
--fcc-editable-region--
    
  </div>
</body>

</html>
```
