---
id: 68eca2a795b333ca5fee30a8
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Now create a second book card. Add another `div` with the `class` attribute set to `card` . Notice how you can reuse the same class name for multiple elements to apply consistent styling.

# --hints--

You should have two `div` elements with the `class` of the value `card`.

```js
assert.equal(document.querySelectorAll('.card').length, 2);
```

Your second `.card` element should be inside the `.card-container` element.

```js
assert.equal(document.querySelectorAll('.card-container .card').length, 2);
```

Your second `.card` element should be a sibling of the first `.card` element.

```js
assert.exists(document.querySelector('.card + .card'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
</div>
--fcc-editable-region--
```
