---
id: 68eca5412b5dd9d06b3d6404
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Below the `<h2>` element in the second card, add a `<p>` element with this text:
`This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`

# --hints--

You should have a `p` element nested inside the second `.card` div.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('p'));
```

Your second card's `p` element should have both opening and closing tags. Opening tags have this syntax: `<elementName>`. Closing tags have a `/` and the same element name: `</elementName>`.

```js
const secondCardRegex = /<div\s+class\s*=\s*['"]card['"]\s+id\s*=\s*['"]dave-cooking-book['"][^>]*>([\s\S]*?)<\/div\s*>/i;
const secondCardMatch = code.match(secondCardRegex);
assert.exists(secondCardMatch, 'Could not find the second card div');
const cardContent = secondCardMatch[1];
const openPCount = (cardContent.match(/<p/gi) || []).length;
const closePCount = (cardContent.match(/<\/p/gi) || []).length;
assert.equal(openPCount, closePCount, 'Each  tag inside the second card should have a closing  tag');
```

Your second card's `p` element's text should be `This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.`. You have either omitted the text, have a typo, or it is not between the p element's opening and closing tags.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('p')?.innerText?.trim(), "This is the story of Dave as he learns to cook everything from pancakes to pasta, one recipe at a time.");
```

Your second card's `p` element should be below your `h2` element.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2 + p'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  <div class="card" id="dave-cooking-book">
    <h2>Dave's Cooking Adventure</h2>
  </div>
</div>
--fcc-editable-region--
```
