---
id: 68ec6cd5b7e8f5a8f7319e32
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

The `div` element is used as a container to group other HTML elements. You will mainly use the `div` element when you want to group HTML elements that will share a set of CSS styles.

Below the `p` element, add a `div` element. This `div` will be a container for your book cards.

# --hints--

You should have a `div` element.

```js
assert.exists(document.querySelector('div'));
```

Your `div` element should be below your `p` element.

```js
assert.exists(document.querySelector('p + div'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>

--fcc-editable-region--
```
