---
id: 68eca4c4b0f952cf09fabe09
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Inside the second card div, add an `<h2>` element with the text:
`Dave's Cooking Adventure`

# --hints--

You should have an `h2` element nested inside the second `.card` div.

```js
const cards = document.querySelectorAll('.card');
assert.exists(cards[1]?.querySelector('h2'));
```

Your second card's `h2` element should have both opening and closing tags. Opening tags have this syntax: `<elementName>`. Closing tags have a `/` and the same element name: `</elementName>`.

```js
const secondCardRegex = /<div\s+class\s*=\s*['"]card['"]\s+id\s*=\s*['"]dave-cooking-book['"][^>]*>([\s\S]*?)<\/div\s*>/i;
const secondCardMatch = code.match(secondCardRegex);
assert.exists(secondCardMatch, 'Could not find the second card div');
const cardContent = secondCardMatch[1];
const openH2Count = (cardContent.match(/<h2/gi) || []).length;
const closeH2Count = (cardContent.match(/<\/h2/gi) || []).length;
assert.equal(openH2Count, closeH2Count, 'Each  tag inside the second card should have a closing  tag');
```

Your second card's `h2` element's text should be `Dave's Cooking Adventure` . You have either omitted the text, have a typo, or it is not between the h2 element's opening and closing tags.

```js
const cards = document.querySelectorAll('.card');
assert.equal(cards[1]?.querySelector('h2')?.innerText?.trim(), "Dave's Cooking Adventure");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
  <div class="card" id="sally-adventure-book">
    <h2>Sally's SciFi Adventure</h2>
    <p>This is an epic story of Sally and her dog Rex as they navigate through other worlds.</p>
    <button class="btn">Buy Now</button>
  </div>
  <div class="card" id="dave-cooking-book"></div>
</div>
--fcc-editable-region--
```
