---
id: 68ec98b38d83a3c28dd30efe
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Inside the first card div, add an `<h2>` element with the text:
`Sally's SciFi Adventure`

# --hints--

You should have an `h2` element nested inside the `.card` div.

```js
assert.exists(document.querySelector('.card h2'));
```

Your `h2` element should have a closing tag. Closing tags have a `/` and the same element name: `</h2>`.

```js
assert.match(code, /<\/h2\s*>/);
```

Your `h2` element's text should be `Sally's SciFi Adventure`. You have either omitted the text, have a typo, or it is not between the `h2` element's opening and closing tags.

```js
assert.equal(document.querySelector('.card h2')?.innerText?.trim(), "Sally's SciFi Adventure");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<h1>XYZ Bookstore</h1>
<p>Browse our collection of amazing books!</p>
<div class='card-container'>
  <div class="card" id="sally-adventure-book"></div>
</div>
--fcc-editable-region--
```
