---
id: 69134b6cb43aa5f254950a10
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will build a bookstore page by creating book cards that display information about different books. You'll practice organizing content using `div` elements, classes, and IDs.

Start your bookstore page by creating the HTML boilerplate.

Add the `<!DOCTYPE html>` declaration and `html` and `head` elements.

Add a `lang` attribute to the `html` element and set it to `"en"`.

# --hints--

Your code should include a `<!DOCTYPE html>` declaration.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Your code should include both opening and closing `<html>` tags.

```js
assert.match(code,/<html\b[^>]*>[\s\S]*<\/html>/i);
```

Your `html` element should have a `lang` attribute set to `"en"`.

```js
assert.match(code, /<html[^>]*\blang\s*=\s*["']en["'][^>]*>/i);
```

Your code should include both opening and closing `<head>` tags.

```js
assert.match(code,/<head>[\s\S]*<\/head>/i);
```

Your `head` element should be nested inside `html` element.

```js
assert.match(code, /<\/html>\s*$/);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
