---
id: 69134d14150d51f8066a3246
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now, improve the structure of your HTML document to ensure your page is encoded correctly and looks good on all screen sizes.

Add a language attribute to the `html` element and set it to `en`.Then,inside the `head` element add:

- a `<meta charset="UTF-8">` tag  
- a `<meta name="viewport" content="width=device-width, initial-scale=1.0">` tag  

Now add the `body` element below the head section. This is where all of your visible page content will go.

# --hints--

Your `html` element should have a `lang` attribute set to `"en"`.

```js
assert.match(code, /<html[^>]*\blang=["']en["'][^>]*>/i);
```

Your code should include a `<meta charset="UTF-8">` element.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Your code should include a viewport `meta` element.

```js
assert.exists(document.querySelector('head > meta[name="viewport"][content="width=device-width, initial-scale=1.0"]'));
```

Your code should include opening and closing `<body>` tags.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Your `body` element should be below `head` element.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
--fcc-editable-region--
   <title>XYZ Bookstore Page</title>
</head>

--fcc-editable-region-- 
</html>
```
