---
id: 69134d14150d51f8066a3246
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now, improve the structure of your HTML document to ensure your page is encoded correctly.

Inside the `head` element, add the `<meta charset="UTF-8">` element.

Lastly, add a `body` element below the `head` section. This is where all of your visible page content will go.

# --hints--

Your code should include a `<meta charset="UTF-8">` element.

```js
assert.exists(document.querySelector('head > meta[charset="utf-8"]'));
```

Your code should include opening and closing `<body>` tags.

```js
assert.match(code,/<body>[\s\S]*<\/body>/i);
```

Your `body` element should be below `head` element.

```js
assert.match(code, /<\/head>[\s\S]*<body>[\s\S]*<\/body>[\s\S]*<\/html>/i);
```

# --seed--
## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
<head>
  <title>XYZ Bookstore Page</title>
</head>

</html>
--fcc-editable-region--
```
