---
id: 69134d14150d51f8066a3246
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Improve the structure of your HTML document by:

1. Adding a language attribute to the `html` element and set it to `en`.
2. Inside the `head`element add:

   - a `<meta charset="UTF-8">` tag  
   - a `<meta name="viewport" content="width=device-width, initial-scale=1.0">` tag  

These help ensure your page is encoded correctly and looks good on all screen sizes.

Now add the `body` element below the head section. This is where all of your visible page content will go.

# --hints--

Your `html` element should have a `lang` attribute set to `"en"`.

```js
assert.match(code, /<html[^>]*\blang=["']en["'][^>]*>/i);
```

Your code should include a `<meta charset="UTF-8">` element.

```js
assert.match(code, /<meta\s+charset=["']?utf-8["']?\s*\/?>/i);
```

Your code should include a viewport `meta` element.

```js
assert.match(code,/<meta\s+name=["']viewport["']\s+content=["']width=device-width,\s*initial-scale=1\.0["']\s*\/?>/i);
```

Your code should include opening and closing `<body>` tags.

```js
assert.exists(document.querySelector('body'));
```

# --seed--
## --seed-contents--

```html
  --fcc-editable-region--
<!DOCTYPE html>
<html>
<head>
   <title>XYZ Bookstore Page</title>
</head>

  --fcc-editable-region-- 
</html>
```
