---
id: 69134b6cb43aa5f254950a10
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--


In this workshop, you will build a bookstore page by creating book cards that display information about different books. You'll practice organizing content using `div` elements, classes, and IDs.

Start your bookstore page by creating the HTML boilerplate.

Add the `<!DOCTYPE html>` declaration and `<html>`, `<head>`, and `<title>` tags.

Set the page title to `XYZ Bookstore Page`.

# --hints--

Your code should include a `<!DOCTYPE html>` declaration.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Your code should include `<html>` and `<head>` tags.

```js
assert.match(code,/<html/i);
assert.match(code,/<head/i);
```

Your title should be `XYZ Bookstore Page`.

```js
assert.match(code,/<title>XYZ Bookstore Page<\/title>/);
```

# --seed--
## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
