---
id: 69134d14150d51f8066a3246
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now add the opening `<body>` tag. This is where all of your visible page content will go.

# --hints--


Your code should include an opening `<body>` tag.

```js
assert.match(code, /<body>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>XYZ Bookstore Page</title>
</head>
--fcc-editable-region--
```
