---
id: 691f85085787a8c4ca589fbb
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Add the `title` element inside the `head` element. 

Set the page title to `XYZ Bookstore Page`.

# --hints--

Your code should include both opening and closing `<title>` tags. 

```js
assert.match(code,/<title>[\s\S]*<\/title>/i);
```

Your title should be `XYZ Bookstore Page`.

```js
assert.equal(document.querySelector('title')?.innerText.trim(),'XYZ Bookstore Page');
```

Your `title` element should be nested within `head` element.

```js
assert.match(code, /<head>[\s\S]*<title>[\s\S]*<\/title>[\s\S]*<\/head>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
--fcc-editable-region--
  
--fcc-editable-region-- 
</head>

</html>
```
