---
id: 5a9036ee38fddaf9a66b5d36
title: Erstelle einen Zeilenabstands mit grid-row-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cPbJ2Cv'
forumTopicId: 301125
dashedName: create-a-row-gap-using-grid-row-gap
---

# --description--

Du kannst einen Abstand zwischen den Zeilen eines Rasters mit `grid-row-gap` auf die gleiche Weise einf체gen, wie du in der vorherigen Aufgabe einen Abstand zwischen den Spalten eingef체gt hast.

# --instructions--

Erstelle einen Abstand f체r die Zeilen, die `5px` hoch ist.

# --hints--

Die Klasse `container` sollte eine Eigenschaft `grid-row-gap` besitzen, die den Wert `5px` enth채lt.

```js
assert(
  code.match(/.container\s*?{[\s\S]*grid-row-gap\s*?:\s*?5px\s*?;[\s\S]*}/gi)
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1{background:LightSkyBlue;}
  .d2{background:LightSalmon;}
  .d3{background:PaleTurquoise;}
  .d4{background:LightPink;}
  .d5{background:PaleGreen;}

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */


    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>.container {grid-row-gap: 5px;}</style>
```
