---
id: 5a94fe5469fb03452672e461
title: Flexible Layouts mit Auto-Fill erstellen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzdycW'
forumTopicId: 301126
dashedName: create-flexible-layouts-using-auto-fill
---

# --description--

Die repeat-Funktion verfügt über eine Option namens <dfn>auto-fill</dfn>. Auf diese Weise kannst du je nach Größe des Containers automatisch so viele Zeilen oder Spalten deiner gewünschten Größe einfügen wie möglich. Du kannst flexible Layouts erstellen, wenn du `auto-fill` mit `minmax` kombinierst, so wie hier:

```css
repeat(auto-fill, minmax(60px, 1fr));
```

Wenn der Container seine Größe ändert, fügt dieses Setup immer wieder 60px-Spalten ein und streckt sie, bis es eine weitere einfügen kann. **Hinweis:** Wenn dein Container nicht alle Elemente in einer Reihe unterbringen kann, verschiebt er sie nach unten in eine neue Reihe.

# --instructions--

Im ersten Raster (Grid) verwendest du `auto-fill` mit `repeat`, um das Raster mit Spalten zu füllen, die eine Mindestbreite von `60px` und eine Maximalbreite von `1fr` haben. Ändere dann die Größe der Vorschau, um das automatische Ausfüllen in Aktion zu sehen.

# --hints--

Die Klasse `container` sollte eine `grid-template-columns`-Eigenschaft mit `repeat` und `auto-fill` besitzen, die das Raster mit Spalten füllt, die eine minimale Breite von `60px` und maximal `1fr` haben.

```js
assert(
  code.match(
    /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fill\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1{background:LightSkyBlue;}
  .item2{background:LightSalmon;}
  .item3{background:PaleTurquoise;}
  .item4{background:LightPink;}
  .item5{background:PaleGreen;}

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1{background:LightSkyBlue;}
  .item2{background:LightSalmon;}
  .item3{background:PaleTurquoise;}
  .item4{background:LightPink;}
  .item5{background:PaleGreen;}

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
