---
id: 66c057041df6394ca796bf33
title: Build a Sentence Maker
challengeType: 14
dashedName: build-a-sentence-maker
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should declare the following variables using `let`:
   
- `adjective`
- `noun`
- `verb`
- `place`
- `adjective2`
- `noun2`

1. You should assign the above variables some string values of your choice.

2. You should declare a `firstStory` variable using `const` to store the first story.
   
3. You should use the following story template to create the first story and assign it to the `firstStory` variable:

`"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. You should output your first story to the console using the message `"First story:\n" + firstStory`.

1. You should assign new values to your `adjective`, `noun`, `verb`, `place`, `adjective2`, and `noun2` variables.

1. Create another story using the same template and assign it to the `secondStory` variable.

1. You should output your second story to the console using the message `"Second story:\n" + secondStory`.


# --hints--

Your `link` element should be within your `head` element.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. 
The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story:\n" + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. 
The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story:\n" + secondStory);
```

