---
id: 66e2edb23013ca596f8d9568
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Create an `else if` block that checks if the index of the current consonant is not `-1`. In the `else if` block, increment the `consonantCount` variable the same way you did to `vowelsCount`.

Create another `else if` block that checks if the index of the current punctuation is not `-1`. In the block of that `else if` increment `punctuationCount`.

# --hints--

You should create an `else if` block with the condition `consonants.indexOf(char) !== -1`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(consonants\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*/)
```

You should use postfix incrementer on the `consonantCount` variable.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(consonants\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*consonantCount\+\+;?/)
```

You should create another `else if` block with the condition `punctuations.indexOf(char) !== -1`.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(punctuations\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*/)
```

You should use postfix incrementer on the `punctuationCount` variable.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(punctuations\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*punctuationCount\+\+;?/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiou";
const consonants = "bcdfghjklmnpqrstvwxyz";
const punctuations = '.,!?;:-()[]{}"\'–';

for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
--fcc-editable-region--
  } 

--fcc-editable-region--
}
```
