---
id: 66e2eab8a5638f57b637b7cc
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

It's time to count the number of vowels, consonants, and punctuation marks. For this, you'll use an `if/else` statement and the `indexOf()` method.

The `indexOf()` method lets you find the index (position) of a specific value within a string or an array.

Create an `if` statement with the condition `vowels.indexOf(char) !== -1`. This checks if `char` is in the list of `vowels` variable. If `char` is found in the `vowels`, the index will not be `-1`.

Inside the `if` statement, use the postfix increment operator on the `vowelsCount` variable.

# --hints--

You should create an `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/)
```

Your `if` statement should have `vowels.indexOf(char) !== -1` as its condition.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(vowels\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*/)
```

You should use the postfix incrementer on `vowelCount` inside your `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(vowels\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*vowelCount\+\+;?\s*\}/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];
--fcc-editable-region--

--fcc-editable-region--
}
```
