---
id: 66e2e5d890ac8a54c9d703bd
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

It's time to count the consonants. Create a `getConsonantCount` function with a sentence parameter.

Inside the function, create a loop to count the number of consonants in the `sentence` that will be passed into the function when it is called.

Your `getConsonantCount` function must return a number.

# --hints--

You should create a `getConsonantCount` function.

```js
assert.isFunction(getConsonantCount);
```

You `getConsonantCount` function should have a `sentence` parameter.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

Your `getConsonantCount` function should return a number.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
