---
id: 66e2edb23013ca596f8d9568
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

To start testing your `getPunctuationCount` function, create a `punctuationCount` variable and set it to the calling of `getPunctuationCount` function with a sentence, for example, `"WHAT?!?!?!?!?"`.

After that, log the `punctuationCount` variable to the console with the text `"Punctuation Count:"` and the interpolation of the `punctuationCount` variable. You can also use plus (`+`) concatenator if you want.

# --hints--

You should create a `punctuationCount` variable.

```js
assert.isNotNull(punctuationCount)
```

Your `punctuationCount` variable should be set to the result of `getPunctuationCount("WHAT?!?!?!?!?")`.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

You should log the `punctuationCount` variable to the console with the parameter `"Punctuation Count: ${punctuationCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

