---
id: 66e2f06b191e305c00574e4d
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

You should count the number of words too. For this, the `while` loop would suffice.

Here's the basic syntax of the `while` loop:

```js
while (condition) {
  // code to run
}
```

And here's an example:

```js
let i = 0;

while (i < 5) {
  console.log(i);  // Logs 0, 1, 2, 3, 4 in their respective lines
  i++;  // Increment i by 1 each time
}
```

Use `let` to create an `i` variable set to `0`. Just under that, create a `while` loop with the condition `i < sentence.length`. In the body of the `while` loop, increment `i` to prevent an infinite loop.

# --hints--

You should use `let` to create an `i` variable set to `0`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0;?/)
```

You should create a `while` loop with the condition `i < sentence.length`.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(i\s*<\s*sentence\.length\s*\)\s*\{\s*/)
```

You should increment `i` with postfix incrementer inside your `while` loop.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(i\s*<\s*sentence\.length\s*\)\s*\{\s*i\+\+;?\s*\}/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiou";
const consonants = "bcdfghjklmnpqrstvwxyz";
const punctuations = '.,!?;:-()[]{}"\'–';

for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
