---
id: 66e2ffa564f06b6baafff948
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Comment out the `while` loop too.

# --hints--

You should comment out the `while` loop.

```js
assert.match(code, /\/\*\s*while\s*\(i\s*<\s*sentence\.length\)\s*\{\s*if\s*\(sentence\[i\]\s*===\s*'\s*'\s*\|\|\s*i\s*===\s*sentence\.length\s*-\s*1\)\s*\{\s*wordCount\+\+;?\s*\}\s*i\+\+\s*\}\s*\*\//)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

/*
for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}
*/

let i = 0;

--fcc-editable-region--
while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}
--fcc-editable-region--

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total Words: ${wordCount}`);
console.log(`Vowel Count: ${vowelCount}`);
console.log(`Consonant Count: ${consonantCount}`);
console.log(`Punctuation Count: ${punctuationCount}`);
```
