---
id: 66e2e5d890ac8a54c9d703bd
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

It's time to write the logic for analyzing the `sentence`. For this, you'll use a loop.

Loops let you automate repetitive tasks by executing the same block of code multiple times. There are `for`, `while`, `do...while`, `for...in`, `for...of` loops, and others.

A `for` loop is good enough to analyze the `sentence`. Here's the basic syntax of the `for` loop:

```js
for (initialization; condition; increment) {
  // code to run
}
```

And here's an example that runs 5 times based on the `condition` specified:

```js
for (let i = 0; i < 5; i++) {
  console.log(i);  // Logs 0, 1, 2, 3, 4 in their respective lines
}
```

Create a for loop with `i` as the initializer, `i < sentence.length` as the condition, and `i++` as the incrementer. Leave the block empty for now.

# --hints--

You should create a `for` loop.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(/)
```

Your `for` loop should initialize `i` to `0`.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(((let|var)\s+)?i\s*=\s*0;/)
```

Your `for` loop should have `i < sentence.length` as its condition.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(((let|var)\s+)?i\s*=\s*0;\s*i\s*<\s*sentence\.length;/)
```

Your `for` loop should increment `i` with the postfix incrementer.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(((let|var)\s+)?i\s*=\s*0;\s*i\s*<\s*sentence\.length;\s*i\+\+\s*\)\s*\{\s*/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

--fcc-editable-region--

--fcc-editable-region--
```
