---
id: 66e2fc6af2688f6972b0495f
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Good job so far! You can take things further by using the cleaner `for...of` loop instead of the good old `for` loop and `while` loop.

Use the multi-line comment syntax to comment off the `for` loop.

# --hints--

You should comment out the `for` loop using multi-line comment syntax.

```js
assert.match(code, /\/\*\s*for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*sentence\.length;\s*i\+\+\)\s*\{\s*let\s+char\s*=\s*sentence\[i\];\s*if\s*\(\s*vowels\.indexOf\(char\)\s*!==\s*-1\)\s*\{\s*vowelCount\+\+;\s*\}\s*else\s*if\s*\(\s*consonants\.indexOf\(char\)\s*!==\s*-1\)\s*\{\s*consonantCount\+\+;\s*\}\s*else\s*if\s*\(\s*punctuations\.indexOf\(char\)\s*!==\s*-1\)\s*\{\s*punctuationCount\+\+;\s*\}\s*\}\s*\*\//)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

--fcc-editable-region--
for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}
--fcc-editable-region--

let i = 0;

while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total Words: ${wordCount}`);
console.log(`Vowel Count: ${vowelCount}`);
console.log(`Consonant Count: ${consonantCount}`);
console.log(`Punctuation Count: ${punctuationCount}`);
```
