---
id: 66e2df8c5676fe51cead8ec8
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Also, create the following variables and values with any keyword you like:

| Variable Name | Value |
| ----------- | ------- |
| `vowels`       | `"aeiouAEIOU"`|
| `consonants`   | `"bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ"`|
| `punctuations`   | `".,!?;:-()[]{}"\'–"`|

# --hints--

You should create a `vowels` variable set to `"aeiou"`.

```js
assert.equal(vowels, "aeiouAEIOU");
```

You should create a `consonants` variable set to `"bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ"`.

```js
assert.equal(consonants, "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ");
```

You should create a `punctuations` variable set to `".,!?;:-()[]{}\"\'–"`.

```js
assert.equal(punctuations, ".,!?;:-()[]{}\"\'–")
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

--fcc-editable-region--

--fcc-editable-region--
```
