---
id: 66e2fac586cc48683839f5b0
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

It's time to test out your sentence analyzer.

Log the following to the console:

| Text to log| 
| ----------- |
| `"Sentence Analysis Results"` | 
| `Original Sentence: ${sentence}` |
| `Total Words: ${wordCount}`   | 
| `Vowel Count: ${vowelCount}`   | 
| `Consonant Count: ${consonantCount}`   | 
| `Punctuation Count: ${punctuationCount}`   | 

Feel free to modify the `sentence` variable so you can see different results!

# --hints--

You should log the text `"Sentence Analysis Results"` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console.log\(('|"|`)Sentence\s+Analysis\s+Results("|'|`)\);?/)
```

You should log `Original Sentence: ${sentence}` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((?:('|"|`)Original\s+Sentence:\s+('|"|`)\s+\+\s+sentence|`Original\s+Sentence: \${sentence}`)\);?/)
```

You should log `Total Words: ${wordCount}` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((?:('|"|`)Total\s+Words:\s+('|"|`)\s+\+\s+wordCount|`Total\s+Words:\s+\${wordCount}`)\);?/)
```

You should log `Vowel Count: ${vowelCount}` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

You should log `Consonant Count: ${consonantCount}` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

You should log `Punctuation Count: ${punctuationCount}` to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}

let i = 0;

while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}

--fcc-editable-region--
  
--fcc-editable-region--
```
