---
id: 66e2d680e129e1423116a541
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this project, you'll learn JavaScript loops by building a sentence analyzer.

To begin, use `let` to create the following variables and values:

| Variable Name | Value   |
| ----------- | ------- |
| `sentence`       | `"Welcome to freeCodeCamp!"` |
| `wordCount`       | `0` |
| `vowelCount`   | `0`   |
| `consonantCount`   | `0`   |
| `punctuationCount`   | `0`   |

You have to use the `let` keyword because all the variables will be reassigned at some point.

# --hints--

You should use let to create a `sentence` variable set to `"Welcome to freeCodeCamp!"`.

```js
assert.match(code, /let\s+sentence\s*=\s*("|'|`)Welcome\s+to\s+freeCodeCamp!("|'|`);?/);
```

You should use let to create a `wordCount` variable set to `0`.

```js
assert.match(code, /let\s+wordCount\s*=\s*0;?/);
```

You should use let to create a `vowelCount` variable set to `0`.

```js
assert.match(code, /let\s+vowelCount\s*=\s*0;?/);
```

You should use let to create a `consonantCount` variable set to `0`.

```js
assert.match(code, /let\s+consonantCount\s*=\s*0;?/);
```

You should use let to create a `punctuationCount` variable set to `0`.

```js
assert.match(code, /let\s+punctuationCount\s*=\s*0;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
