---
id: 66e2f376df6f315ee81de81a
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Inside your `while` loop, create an `if` statement with the condition `sentence[i] === ' ' || i === sentence.length - 1`.

`sentence[i] === ' '` checks if the current character is a space and `i === sentence.length - 1` checks if `i` is at the last character of the string. These two helps identify the boundary between words and the last word is counted, even if it doesn’t end with a space.

Inside the `if` statement, increment the `wordCount` variable.

# --hints--

You should create an if statement with the condition `sentence[i] === ' ' || i === sentence.length - 1`.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(sentence\[i\]\s===\s*('|"|`) ('|"|`)\s*\|\|\s*i\s*===\s*sentence\.length\s*-\s*1\)\s*\{\s*/)
```

You should increment `wordCount` with a postfix incrementer inside your `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(sentence\[i\]\s===\s*('|"|`) ('|"|`)\s*\|\|\s*i\s*===\s*sentence\.length\s*-\s*1\)\s*\{\s*wordCount\+\+;?\s*\}/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}

let i = 0;

while (i < sentence.length) {
--fcc-editable-region--
  
--fcc-editable-region--
  i++
}
```
