---
id: 66e2e99caa0cb3570f4822b4
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

In the block of the `for` loop, use `let` to create a `char` variable and set it to `sentence[i]`. 

`i` represents the current character starting from the initialzer variable. For example, at `i = 0`, the character is `W`, and at `i = 1`, the character is `e`, and so on. Remember the string you're testing against is `"Welcome to freeCodeCamp!"`.

# --hints--

You should use let to create a `char` variable set to `sentence[i]`.

```js
assert.match(__helpers.removeJSComments(code), /let\s+char\s*=\s*sentence\[i\];?/);
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

for (let i = 0; i < sentence.length; i++) {
--fcc-editable-region--

--fcc-editable-region--
}
```
