---
id: 66e2df8c5676fe51cead8ec8
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Time to test your `getVowelCount` function.

Create a `vowelCount` variable and set it to the calling of `getVowelCount` function with a sentence, for example, `"Apples are tasty fruits"`.

After that, log the `vowelCount` variable to the console with the text `"Vowel Count:"` and the interpolation of the `vowelCount` variable. You can also use plus (`+`) concatenator if you want.

# --hints--

You should create a `vowelCount` variable.

```js
assert.isNotNull(vowelCount)
```

Your `vowelCount` variable should be set to the result of `getVowelCount("Apples are tasty fruits")`.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

You should log the `vowelCount` variable to the console with the parameter `"Vowel Count: ${vowelCount}"`.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
