---
id: 66e2d680e129e1423116a541
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will build a sentence analyzer that will take a sentence and get the count for the number of words, vowels, consonants, and punctuation marks.

To begin, create a `getVowelCount` function with a parameter called `sentence`. Your function should return the total number of vowels in a sentence.

# --hints--

You should create a `getVowelCount` function.

```js
assert.isFunction(getVowelCount);
```

You `getVowelCount` function should have a `sentence` parameter.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

Your `getVowelCount` function should return a number.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
