---
id: 66e3006a27baf76c9f833d02
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Create a `for...of` loop with `char` as the variable name and sentence as the iterable.

Here's the basic syntax of the `for...of` loop:

```js
for (const variableName of iterable) {
  // code to run
}
```

Leave the `for...of` loop empty for now.

# --hints--

You should create a `for..of` loop with the condition `char` as the varable and `sentence` as the iterable.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(const\s+char\s+of\s+sentence\)\s*\{\s*\}/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

/*
for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}
*/

let i = 0;

/*
while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}
*/

--fcc-editable-region--

--fcc-editable-region--

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total Words: ${wordCount}`);
console.log(`Vowel Count: ${vowelCount}`);
console.log(`Consonant Count: ${consonantCount}`);
console.log(`Punctuation Count: ${punctuationCount}`);
```
