---
id: 66e4102a3e6356ade9f073ef
title: Step 14
challengeType: 1
dashedName: step-14
---

# --description--

Create an `else if` block that checks if the index of the current consonant is not `-1`, then increment the `consonantCount` variable in its block.

# --hints--

You should create an `else if` block that checks for the index of the cirrent consonant.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(consonants\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*/)
```

You should increment the `consonantCount` variable inside your `else if` block.

```js
assert.match(__helpers.removeJSComments(code), /else\s*if\s*\(consonants\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*consonantCount\+\+;?/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

/*
for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}
*/

let i = 0;

/*
while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}
*/

for (const char of sentence) {
  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
--fcc-editable-region--
  } 

--fcc-editable-region--
}

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total words: ${wordCount}`);
console.log(`Vowel count: ${vowelCount}`);
console.log(`Consonant count: ${consonantCount}`);
console.log(`Punctuation count: ${punctuationCount}`);
```
