---
id: 66e40f6267ffbfad00331fcb
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

Inside the `for...of` loop, create an `if` statement to check if the index of the current vowel is not `-1`, then increment the `vowelCount` variable in the block of the `if` statement.

Once you do that, you should see the vowel count update to `9` on the right, as long as you have everything implemented rightly.

# --hints--

You should create an `if` statement that checks for the index of the current vowel.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(vowels\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*/)
```

You should increment the `vowelCount` variable inside your `if` statement.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(vowels\.indexOf\(char\)\s*!==\s*-1\s*\)\s*\{\s*vowelCount\+\+;?\s*\}/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

/*
for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}
*/

let i = 0;

/*
while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}
*/

for (const char of sentence) {
--fcc-editable-region--

--fcc-editable-region--
}

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total Words: ${wordCount}`);
console.log(`Vowel Count: ${vowelCount}`);
console.log(`Consonant Count: ${consonantCount}`);
console.log(`Punctuation Count: ${punctuationCount}`);
```
