---
id: 66e41132ff8a26af06ed1963
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

What's left is to count the number of words. This will use a separate `if` statement.

Create an `if` statement with the appropriate condition to count the number of words. In its block, increment the `wordCount` variable.

After that, outside the `if` statement but within the `for...of` loop, increment the `i` variable with a postfix incrementer so the loop can move to the next character index.

Play around with the sentence variable one more time to see different results.

With that, your sentence analyzer project is complete!

# --hints--

You should create an `if` block with the condition `sentence[i] === ' ' || i === sentence.length -1`.

```js
assert.match(__helpers.removeJSComments(code), /if\(sentence\[i\]\s*===\s*("|'|`)\s{1}("|'|`)\s*\|\|\s*i\s*===\s*sentence\.length\s*-1\)\s*\{\s*/)
```

You should increment the `wordCount` variable inside your `if` block.

```js
assert.match(__helpers.removeJSComments(code), /if\(sentence\[i\]\s*===\s*("|'|`)\s{1}("|'|`)\s*\|\|\s*i\s*===\s*sentence\.length\s*-1\)\s*\{\s*wordCount\+\+;?\s*\}\s*/)
```

You should increment the `i` variable within your `for...of` loop and outside the last `if` block.

```js
assert.match(__helpers.removeJSComments(code), /if\(sentence\[i\]\s*===\s*("|'|`)\s{1}("|'|`)\s*\|\|\s*i\s*===\s*sentence\.length\s*-1\)\s*\{\s*wordCount\+\+;?\s*\}\s*i\+\+;?/)
```

# --seed--

## --seed-contents--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiouAEIOU";
const consonants = "bcdfghjklmnpqrstvwxyzBCDFGHJKLMNPQRSTVWXYZ";
const punctuations = '.,!?;:-()[]{}"\'–';

/*
for (let i = 0; i < sentence.length; i++) {
  let char = sentence[i];

  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
}
*/

let i = 0;

/*
while (i < sentence.length) {
  if (sentence[i] === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
  i++
}
*/

for (const char of sentence) {
  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  }
--fcc-editable-region--


--fcc-editable-region--
}

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total words: ${wordCount}`);
console.log(`Vowel count: ${vowelCount}`);
console.log(`Consonant count: ${consonantCount}`);
console.log(`Punctuation count: ${punctuationCount}`);
```

## --solutions--

```js
let sentence = "Welcome to freeCodeCamp!";

let wordCount = 0;
let vowelCount = 0;
let consonantCount = 0;
let punctuationCount = 0;

const vowels = "aeiou";
const consonants = "bcdfghjklmnpqrstvwxyz";
const punctuations = '.,!?;:-()[]{}"\'–';

let i = 0;

for (const char of sentence) {
  if (vowels.indexOf(char) !== -1) {
    vowelCount++;
  } else if (consonants.indexOf(char) !== -1) {
    consonantCount++;
  } else if (punctuations.indexOf(char) !== -1) {
    punctuationCount++;
  } else if (char === ' ' || i === sentence.length - 1) {
    wordCount++;
  }
}

console.log('Sentence Analysis Results');
console.log(`Original Sentence: ${sentence}`);
console.log(`Total words: ${wordCount}`);
console.log(`Vowel count: ${vowelCount}`);
console.log(`Consonant count: ${consonantCount}`);
console.log(`Punctuation count: ${punctuationCount}`);
```
