---
id: 5dc174fcf86c76b9248c6eb2
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will continue working with basic HTML elements like headings, paragraphs, and lists by building a cat photo app. 

Begin the workshop by adding an `h1` element with the text of `CatPhotoApp`.

# --hints--

The text `CatPhotoApp` should be present in the code. You may want to check your spelling.

```js
assert.match(code, /catphotoapp/i);
```

Your `h1` element should have an opening tag. Opening tags have this syntax: `<elementName>`.

```js
assert.exists(document.querySelector('h1'));
```

Your `h1` element should have a closing tag. Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

Your `h1` element's text should be `CatPhotoApp`. You have either omitted the text, have a typo, or it is not between the `h1` element's opening and closing tags.

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
   
--fcc-editable-region--
  </body>
</html>
```
