---
id: 67cd57e96257ed1fabbe4879
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now would be a good time to see the updated `isVisible` value.

Below the `setIsVisible(!isVisible);`, log the `isVisible` state variable to the console. 

A `handleToggleVisibility` function call has been added for you so you can see the results. Open up the console and you will notice that the value is not what you expected it to be. 

# --hints--

You should use `console.log` to log the `isVisible` state variable to the console. 

```js
assert.match(code, /console\.log\s*\(\s*isVisible\s*)\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
  <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>
  <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
import { useState } from "react";

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(true);

  --fcc-editable-region--
  const handleToggleVisibility = () => {
    setIsVisible(!isVisible);

  }
  handleToggleVisibility();
  --fcc-editable-region--

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
