---
id: 660c1d301c2a1c2a1c2a1c2c
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now it's time to use D3.js to fetch the revenue data and set up our SVG container. Use `d3.json()` to load the data from the provided URL. Inside the `.then()` callback, select the `.visHolder` and append an `svg` element with a width of `800` and height of `400`.

# --hints--

You should use `d3.json()` to fetch the data.

```js
assert.match(code, /d3\.json\(/);
```

You should append an `svg` element to the `.visHolder`.

```js
assert.exists(document.querySelector('.visHolder svg'));
```

The `svg` element should have a `width` of `800` and a `height` of `400`.

```js
const svg = document.querySelector('svg');
assert.equal(svg.getAttribute('width'), '800');
assert.equal(svg.getAttribute('height'), '400');
```

# --seed--

## --seed-contents--

```html
<div class="main">
  <div class="container">
    <div id="title">Global Music Industry Revenue</div>
    <div class="visHolder"></div>
  </div>
</div>
```

```js
--fcc-editable-region--
const url = 'https://raw.githubusercontent.com/freeCodeCamp/ProjectReferenceData/master/global-music-revenue.json';

--fcc-editable-region--
```

# --solutions--

```html
<div class="main">
  <div class="container">
    <div id="title">Global Music Industry Revenue</div>
    <div class="visHolder"></div>
  </div>
</div>
```

```js
const url = 'https://raw.githubusercontent.com/freeCodeCamp/ProjectReferenceData/master/global-music-revenue.json';

d3.json(url).then(data => {
  const width = 800;
  const height = 400;

  const svg = d3.select('.visHolder')
    .append('svg')
    .attr('width', width)
    .attr('height', height);
});
```
