---
id: 660c1d2e1c2a1c2a1c2a1c2a
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will build a bar chart showing the Global Music Industry Revenue from 1999 to 2023.

Start by setting up the HTML structure. Inside the `body` tag, add a `div` with a class of `main`. Inside that, add another `div` with a class of `container`. Finally, add a `div` with an ID of `title` containing the text `Global Music Industry Revenue`, and a `div` with a class of `visHolder` where the chart will be rendered.

# --hints--

You should have a `div` with a class of `main`.

```js
assert.exists(document.querySelector('div.main'));
```

You should have a `div` with a class of `container` inside the `main` div.

```js
assert.exists(document.querySelector('div.main div.container'));
```

You should have a `div` with an ID of `title` containing the text `Global Music Industry Revenue`.

```js
const title = document.querySelector('#title');
assert.exists(title);
assert.equal(title.innerText.trim(), 'Global Music Industry Revenue');
```

You should have a `div` with a class of `visHolder`.

```js
assert.exists(document.querySelector('.visHolder'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Music Industry Revenue</title>
    <link rel="stylesheet" href="styles.css">
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
</head>
<body>
--fcc-editable-region--

--fcc-editable-region--
    <script src="script.js"></script>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Music Industry Revenue</title>
    <link rel="stylesheet" href="styles.css">
    <script src="https://cdnjs.cloudflare.com/ajax/libs/d3/7.9.0/d3.min.js"></script>
</head>
<body>
    <div class="main">
      <div class="container">
        <div id="title">Global Music Industry Revenue</div>
        <div class="visHolder"></div>
      </div>
    </div>
    <script src="script.js"></script>
</body>
</html>
```
