---
id: 660c1d2f1c2a1c2a1c2a1c2b
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Now, let's add some basic styling to center the chart and make it look premium. Use CSS to style the `.main` container with a minimum height of `100vh` and center its content using Flexbox. Set a background color of `#f0f2f5`.

# --hints--

The `*` selector should have `margin`, `padding` set to `0` and `box-sizing` set to `border-box`.

```js
const style = __helpers.getStyleAny('*');
assert.equal(style?.margin, '0px');
assert.equal(style?.padding, '0px');
assert.equal(style?.boxSizing, 'border-box');
```

The `.main` class should use Flexbox to center its content.

```js
const style = __helpers.getStyleAny('.main');
assert.equal(style?.display, 'flex');
assert.equal(style?.alignItems, 'center');
assert.equal(style?.justifyContent, 'center');
```

The `.main` class should have a `min-height` of `100vh`.

```js
const style = __helpers.getStyleAny('.main');
assert.equal(style?.minHeight, '100vh');
```

# --seed--

## --seed-contents--

```html
<div class="main">
  <div class="container">
    <div id="title">Global Music Industry Revenue</div>
    <div class="visHolder"></div>
  </div>
</div>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<div class="main">
  <div class="container">
    <div id="title">Global Music Industry Revenue</div>
    <div class="visHolder"></div>
  </div>
</div>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  font-family: Arial, sans-serif;
}

.main {
  min-height: 100vh;
  width: 100%;
  display: flex;
  align-items: center;
  justify-content: center;
  background: #f0f2f5;
}

.container {
  background: white;
  padding: 2rem;
  border-radius: 8px;
  box-shadow: 0 4px 6px rgba(0,0,0,0.1);
}

#title {
  font-size: 2rem;
  text-align: center;
  margin-bottom: 2rem;
  color: #1a1a1a;
}
```
