---
id: 660c1d311c2a1c2a1c2a1c2d
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

With the SVG set up, we now need to create our scales and axes. Use `d3.scaleTime()` for the x-axis (years) and `d3.scaleLinear()` for the y-axis (revenue). Don't forget to add padding so the axes are visible within the SVG.

# --hints--

You should define an `xScale` using `d3.scaleTime()`.

```js
assert.match(code, /d3\.scaleTime\(\)/);
```

You should define a `yScale` using `d3.scaleLinear()`.

```js
assert.match(code, /d3\.scaleLinear\(\)/);
```

You should append two `g` elements to the `svg` for the axes.

```js
assert.equal(document.querySelectorAll('svg g').length, 2);
```

# --seed--

## --seed-contents--

```js
const url = 'https://raw.githubusercontent.com/freeCodeCamp/ProjectReferenceData/master/global-music-revenue.json';

d3.json(url).then(data => {
  const width = 800;
  const height = 400;
  const padding = 60;

  const svg = d3.select('.visHolder')
    .append('svg')
    .attr('width', width)
    .attr('height', height);

--fcc-editable-region--

--fcc-editable-region--
});
```

# --solutions--

```js
const url = 'https://raw.githubusercontent.com/freeCodeCamp/ProjectReferenceData/master/global-music-revenue.json';

d3.json(url).then(data => {
  const width = 800;
  const height = 400;
  const padding = 60;

  const svg = d3.select('.visHolder')
    .append('svg')
    .attr('width', width)
    .attr('height', height);

  const years = data.map(d => new Date(d.year, 0));
  const revenues = data.map(d => d.revenue);

  const xScale = d3.scaleTime()
    .domain([d3.min(years), d3.max(years)])
    .range([padding, width - padding]);

  const yScale = d3.scaleLinear()
    .domain([0, d3.max(revenues)])
    .range([height - padding, padding]);

  const xAxis = d3.axisBottom(xScale);
  const yAxis = d3.axisLeft(yScale);

  svg.append('g')
    .attr('transform', `translate(0, ${height - padding})`)
    .call(xAxis);

  svg.append('g')
    .attr('transform', `translate(${padding}, 0)`)
    .call(yAxis);
});
```
