---
id: 660c1d321c2a1c2a1c2a1c2e
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Finally, let's render the bars and add a tooltip to display the revenue for each year. Bind the data to `rect` elements, set their attributes, and use `.on('mouseover')` and `.on('mouseout')` to show and hide the tooltip.

# --hints--

You should have one `rect` element for each data point.

```js
assert.equal(document.querySelectorAll('rect.bar').length, 25);
```

The `rect` elements should have a class of `bar`.

```js
assert.exists(document.querySelector('rect.bar'));
```

The tooltip should be displayed on `mouseover`.

```js
// This is a simplified test for the existence of the event listener
assert.match(code, /\.on\(['"]mouseover['"]/);
```

# --seed--

## --seed-contents--

```js
const url = 'https://raw.githubusercontent.com/freeCodeCamp/ProjectReferenceData/master/global-music-revenue.json';

d3.json(url).then(data => {
  const width = 800;
  const height = 400;
  const padding = 60;

  const svg = d3.select('.visHolder')
    .append('svg')
    .attr('width', width)
    .attr('height', height);

  const years = data.map(d => new Date(d.year, 0));
  const revenues = data.map(d => d.revenue);

  const xScale = d3.scaleTime()
    .domain([d3.min(years), d3.max(years)])
    .range([padding, width - padding]);

  const yScale = d3.scaleLinear()
    .domain([0, d3.max(revenues)])
    .range([height - padding, padding]);

  const xAxis = d3.axisBottom(xScale);
  const yAxis = d3.axisLeft(yScale);

  svg.append('g')
    .attr('transform', `translate(0, ${height - padding})`)
    .call(xAxis);

  svg.append('g')
    .attr('transform', `translate(${padding}, 0)`)
    .call(yAxis);

--fcc-editable-region--

--fcc-editable-region--
});
```

# --solutions--

```js
const url = 'https://raw.githubusercontent.com/freeCodeCamp/ProjectReferenceData/master/global-music-revenue.json';

d3.json(url).then(data => {
  const width = 800;
  const height = 400;
  const padding = 60;

  const svg = d3.select('.visHolder')
    .append('svg')
    .attr('width', width)
    .attr('height', height);

  const years = data.map(d => new Date(d.year, 0));
  const revenues = data.map(d => d.revenue);

  const xScale = d3.scaleTime()
    .domain([d3.min(years), d3.max(years)])
    .range([padding, width - padding]);

  const yScale = d3.scaleLinear()
    .domain([0, d3.max(revenues)])
    .range([height - padding, padding]);

  const xAxis = d3.axisBottom(xScale);
  const yAxis = d3.axisLeft(yScale);

  svg.append('g')
    .attr('transform', `translate(0, ${height - padding})`)
    .call(xAxis);

  svg.append('g')
    .attr('transform', `translate(${padding}, 0)`)
    .call(yAxis);

  svg.selectAll('rect')
    .data(data)
    .enter()
    .append('rect')
    .attr('class', 'bar')
    .attr('x', (d, i) => xScale(years[i]))
    .attr('y', d => yScale(d.revenue))
    .attr('width', (width - 2 * padding) / data.length)
    .attr('height', d => height - padding - yScale(d.revenue))
    .attr('fill', '#3498db')
    .on('mouseover', (event, d) => {
      d3.select('.visHolder')
        .append('div')
        .attr('id', 'tooltip')
        .style('left', event.pageX + 10 + 'px')
        .style('top', event.pageY - 20 + 'px')
        .style('position', 'absolute')
        .style('background', 'rgba(0,0,0,0.8)')
        .style('color', 'white')
        .style('padding', '5px')
        .style('border-radius', '4px')
        .html(`${d.year}: $${d.revenue}B`);
    })
    .on('mouseout', () => {
      d3.select('#tooltip').remove();
    });
});
```
