---
id: 66b102ee0bdbad34a9f85ba0
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you are going to continue learning about strings by building a Teacher Chatbot. 

To begin, add a `console` statement, with the message of `"Hi there!"`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

You should have a `console` statement.

```js
assert.isAbove(logSpy.calls.length, 0);
```

Your console statement should output the message `"Hi there!"`.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
