---
id: 63508c898d753754757bd5e3
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Die `.sort()`-Methode wandelt die Elemente eines Arrays standardmäßig in Strings um und sortiert sie dann in alphabetischer Reihenfolge. Dies funktioniert gut für Strings, aber nicht so gut für Zahlen. Zum Beispiel steht `10` vor `2`, wenn sie als Strings sortiert werden, aber `2` steht vor `10`, wenn sie als Zahlen sortiert werden.

Um dies zu beheben, kannst du der `.sort()`-Methode eine Callback-Funktion übergeben. Diese Funktion verwendet zwei Argumente, die die beiden zu vergleichenden Elemente darstellen. Die Funktion sollte einen Wert kleiner als `0` zurückgeben, wenn das erste Element vor dem zweiten Element kommen soll, einen Wert größer als `0`, wenn das erste Element nach dem zweiten Element kommen soll, und `0`, wenn die beiden Elemente an ihren derzeitigen Positionen verbleiben sollen.

Um deine Zahlen von der kleinsten bis zur größten Zahl zu sortieren, übergib eine Callback-Funktion mit den Parametern `a` und `b`, die das Ergebnis der Subtraktion von `b` von `a` zurückgibt.

# --hints--

Your `sort` method should have a callback function which takes an `a` and a `b` parameters.

```js
assert.match(getMedian.toString(), /(array\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>|array\.sort\(\s*function\s*\(\s*a\s*,\s*b\s*\)\s*\{)/)
```

Deine `sort`-Methode sollte das Ergebnis der Subtraktion von `b` und `a` zurückgeben.

```js
assert.match(getMedian.toString(), /(array\.sort\(\s*\(\s*a\s*,\s*b\s*\)\s*=>|array\.sort\(\s*function\s*\(\s*a\s*,\s*b\s*\)\s*\{)\s*(return\s+)?\s*a\s*\-\s*b/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort();
}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));

  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
