---
id: 63507d810f1a2e38f1908fd8
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Ein Nutzer kann einen beliebigen Text in das Eingabefeld eingeben. Du solltest sicherstellen, dass du nur mit Zahlen arbeitest. Der `Number()`-Konstruktor wird `NaN` zurückgeben (was für "keine Zahl" steht), wenn der an ihn übergebene Wert nicht in eine Zahl umgewandelt werden kann.

Du musst diese Werte herausfiltern - zum Glück haben Arrays dafür eine Methode. Mit der `.filter()`-Methode kannst du Elemente aus einem Array herausfiltern und dabei ein neues Array erstellen.

Deklariere eine `filtered`-Variable und weise ihr `numbers.filter()` zu.

# --hints--

Deine `calculate`-Funktion sollte über eine `filtered`-Variable verfügen.

```js
assert.match(calculate.toString(), /filtered/);
```

Deine `calculate`-Funktion sollte die `.filter()`-Methode für dein `numbers`-Array verwenden.

```js
assert.match(calculate.toString(), /numbers\.filter\(\s*\)/)
```

Du solltest der `filtered`-Variable das Ergebnis von `numbers.filter()` zuweisen.

```js
assert.match(calculate.toString(), /filtered\s*=\s*numbers\.filter\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el));
}
--fcc-editable-region--
```
