---
id: 635090f47eb6d9563a6fed05
title: Schritt 22
challengeType: 0
dashedName: step-22
---

# --description--

Im nächsten Schritt ist die Suche nach der Zahl in der Mitte der Liste. Wenn die Liste eine ungerade Anzahl von Zahlen enthält, ist die mittlere Zahl der Median. Wenn die Liste eine gerade Anzahl von Zahlen enthält, ist der Median der Durchschnitt der beiden mittleren Zahlen.

You can check if a number is even or odd with the <dfn>modulus operator</dfn>, which is represented by the `%` symbol. Dieser Operator gibt den Rest der Division von zwei Zahlen zurück. Wenn der Rest `0` ist, ist die Zahl gerade. Wenn der Rest `1` ist, ist die Zahl ungerade:

```js
array.length % 2 === 0;
```

Deklariere eine `median`-Variable. Überprüfe mit Hilfe des ternären Operators, ob die Länge von `array` eine gerade Zahl ist. Wenn die Länge von `array` eine gerade Zahl ist, finde die beiden mittleren Zahlen und berechne den Mittelwert dieser Zahlen. Wenn die Länge von `array` eine ungerade Zahl ist, finde die mittlere Zahl und weise sie der `median`-Variablen zu.

# --hints--

Deine `getMedian`-Funktion sollte über eine `median`-Variable verfügen.

```js
assert.match(getMedian.toString(), /median\s*=/);
```

Deine `median`-Variable sollte einen ternären Operator verwenden, um zu überprüfen, ob die `array.length` eine gerade Zahl ist.

```js
assert.match(getMedian.toString(), /median\s*=\s*array\.length\s*%\s*2\s*===\s*0\s*\?/);
```

If `array.length` is even, pass an array with the two middle numbers from the `sorted` array to the `getMean` function. The first item in the array argument should be the value at index `array.length / 2` from `sorted` and the second should use the value at the index `array.length / 2 - 1` from `sorted`.

```js
assert.match(getMedian.toString(), /median\s*=\s*array\.length\s*%\s*2\s*===\s*0\s*\?\s*getMean\(\s*\[\s*sorted\s*\[\s*array\.length\s*\/\s*2\s*\]\s*,\s*sorted\s*\[\s*array\.length\s*\/\s*2\s*-\s*1\s*\]\s*\]\s*\)\s*\:/);
```

Wenn die `array.length` eine ungerade Zahl ist, sollte deine `median`-Variable der `median`-Variablen die mittlere Zahl zuweisen. The middle number is the value of `sorted` at `Math.floor(array.length / 2)`.

```js
assert.match(getMedian.toString(), /median\s*=\s*array\.length\s*%\s*2\s*===\s*0\s*\?\s*getMean\(\s*\[\s*sorted\s*\[\s*array\.length\s*\/\s*2\s*\]\s*,\s*sorted\s*\[\s*array\.length\s*\/\s*2\s*-\s*1\s*\]\s*\]\s*\)\s*\:\s*sorted\s*\[\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*\]\s*;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort((a, b) => a - b);
}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));

  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
