---
id: 6350805fe0fe283dd347b0dc
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Der <dfn>mean</dfn> ist der Durchschnittswert aller Zahlen in einer Liste. Der erste Schritt zur Berechnung des Mittelwerts besteht darin, die Summe aller Zahlen in der Liste zu berechnen. Arrays haben eine weitere Methode namens `.reduce()`, die für diese Situation perfekt ist. Die `.reduce()`-Methode verwendet ein Array und wendet eine Callback-Funktion an, um das Array zu einem einzigen Wert zu verdichten.

Deklariere eine `sum`-Variable und weise ihr `array.reduce()` zu.

# --hints--

Deine `getMean`-Funktion sollte eine `sum`-Variable enthalten.

```js
assert.match(getMean.toString(), /sum/);
```

Your `getMean` function should use the `.reduce()` method of the `array` parameter.

```js
assert.match(getMean.toString(), /array\.reduce\(\s*\)/);
```

Du solltest der `sum`-Variable das Ergebnis von `array.reduce()` zuweisen.

```js
assert.match(getMean.toString(), /sum\s*=\s*array\.reduce\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const getMean = (array) => {

}
--fcc-editable-region--

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
}
```
