---
id: 60eebd07ea685b0e777b5583
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Willkommen beim Registrierungsformular-Projekt! Beginne, indem du die `!DOCTYPE html`-Deklaration oben im Dokument hinzufügst, damit der Browser weiß, welche Art von Dokument er ließt.

# --hints--

Dein Code sollte die `DOCTYPE`-Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Du solltest ein Leerzeichen nach der `DOCTYPE`-Referenz einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Du solltest den Dokumententyp als `html` definieren.

```js
assert(code.match(/html/gi));
```

Du solltest die `DOCTYPE`-Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/html\s*>/gi));
```

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
