---
id: 60f027099a15b00485563dd2
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Füge unter dem `DOCTYPE` ein `html`-Element mit einem `lang`-Attribut von `en` hinzu, so dass du einen Ort hast, an dem du mit der Eingabe vom Code beginnen kannst.

# --hints--

Deine `DOCTYPE`-Deklaration sollte am Anfang deines HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dein `html`-Element sollte ein öffnendes Tag enthalten sowie ein `lang`-Attribut mit dem Wert `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein abschließendes Tag enthalten.

```js
assert(code.match(/<\/html\s*>/));
```

Deine `html`-Tags sollten in der richtigen Reihenfolge stehen.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Du solltest nur ein `html`-Element haben.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
