---
id: 60f0286404aefb0562a4fdf9
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Füge ein `title`- und ein `meta`-Element zum `head` hinzu. Weise deinem Projekt den Titel `Registration Form` zu und füge deinem `meta`-Element ein `charset`-Attribut mit dem Wert `UTF-8` hinzu.

# --hints--

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Das `title`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > title'));
```

Dein Projekt sollte den Titel `Registration Form` haben.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
