---
id: 60f030d388cb74067cf291c3
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Below the heading, use the following text within a paragraph element to encourage users to register:

```md
Please fill out this form with the required information
```

# --hints--

Du solltest ein `p`-Element innerhalb des `body` hinzufügen.

```js
assert.exists(document.querySelector('body > p'));
```

Du solltest das `p`-Element unter dem `h1` hinzufügen.

```js
assert.exists(document.querySelector('h1 + p'));
```

Du solltest dem `p`-Element einen Text von `Please fill out this form with the required information` zuweisen.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
