---
id: 60f027c87bc98f050395c139
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Füge als Nächstes öffnende und abschließende `head`- und `body`-Tags innerhalb des `html`-Elements hinzu.

# --hints--

Du solltest ein öffnendes `head`-Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Du solltest ein abschließendes `head`-Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Du solltest ein öffnendes `body`-Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Du solltest ein abschließendes `body`-Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head`- und `body`-Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

Das `head`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body`-Element sollte sich innerhalb des `html`-Elements befinden.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
