---
id: 60f1a9cbd23023082e149fee
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Das ist besser. Gestalte jetzt den Hintergrund optisch ansprechend, indem du die `body` `background-color` auf `#1b1b32` setzt. Ändere dann, um den Text zu sehen, die `color` auf `#f5f6f7`.

# --hints--

Du solltest innerhalb des `body`-Elementselektors die `background-color` hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Du solltest der `background-color` den Wert `#1b1b32` zuweisen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Du solltest die `color` innerhalb des `body`-Elementselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Du solltest der `color` den Wert `#f5f6f7` zuweisen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
