---
id: 60f5d2776c854e069560fbe6
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Das erste `fieldset` wird Name-, E-Mail- und Passwort-Felder enthalten. Beginne, indem du vier `label`-Elemente zum ersten `fieldset` hinzufügst.

# --hints--

Es sollte insgesamt drei `fieldset`-Elemente geben.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

Die `fieldset`-Elemente sollten alle direkte untergeordnete Elemente des `form`-Elements sein.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

Die vier `label`-Elemente sollten alle direkte untergeordnete Elemente des ersten `fieldset`-Elements sein. Stelle sicher, dass du die `label`-Elemente schließt.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Die letzten zwei `fieldset`-Elemente sollten leer sein.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Stelle sicher, dass du die `label`-Elemente schließt.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
