---
id: 60f5dc35c07ac1078f140916
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Füge den folgenden Text zu den `label`-Elementen hinzu:

- `Enter Your First Name:`
- `Enter Your Last Name:`
- `Enter Your Email:`
- `Create a New Password:`

# --hints--

Das erste `label`-Element sollte den Text `Enter Your First Name:` enthalten.

```js
assert.match(document.querySelector('label')?.innerHTML, /Enter Your First Name:/i);
```

Das zweite `label`-Element sollte den Text `Enter Your Last Name:` enthalten.

```js
assert.match(document.querySelector('fieldset > label:nth-child(2)')?.innerHTML, /Enter Your Last Name:/i);
```

Das dritte `label`-Element sollte den Text `Enter Your Email:` enthalten.

```js
assert.match(document.querySelector('fieldset > label:nth-child(3)')?.innerHTML, /Enter Your Email:/i);
```

Das vierte `label`-Element sollte den Text `Create a New Password:` enthalten.

```js
assert.match(document.querySelector('fieldset > label:nth-child(4)')?.innerHTML, /Create a New Password:/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label></label>
        <label></label>
        <label></label>
        <label></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
