---
id: 60f5c3e399ff1a05629964e4
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Wie der Titel vorschlägt, erstellst du ein Formular. Füge also nach dem `p`-Element eine `form` mit einem `action`-Attribut ein, das `https://register-demo.freecodecamp.org` anvisiert.

# --hints--

Du solltest neben dem `p`-Element ein `form`-Element hinzufügen.

```js
assert.exists(document.querySelector('p + form'));
```

Du solltest der `form` ein `action`-Attribut zuweisen.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Du solltest der `action` einen Wert von `https://register-demo.freecodecamp.org` zuweisen.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
