---
id: 60f81167d0d4910809f88945
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Das erste `input`-Element mit einem `type` von `submit` ist automatisch so eingestellt, dass es das n채chstgelegene 체bergeordnete `form`-Element absendet.

Um den Formular- Versand zu bearbeiten, f체ge nach dem letzten `fieldset`-Element ein `input`-Element mit dem `type`-Attribut von `submit` und dem `value`-Attribut von `Submit` hinzu.

# --hints--

Du solltest das `input`-Element nach dem letzten `fieldset`-Element hinzuf체gen.

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

Du solltest dem `input`-Element ein `type`-Attribut von `submit` zuweisen.

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

Du solltest dem `input`-Element ein `value`-Attribut von `Submit` zuweisen.

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>

--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
