---
id: 60f83e7bfc09900959f41e20
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Bestimmte `type`-Attributwerte besitzen eine eingebaute Formularprüfung. Zum Beispiel erfordert `type="email"`, dass der Wert eine gültige E-Mail-Adresse ist.

Füge eine benutzerdefinierte Überprüfung für das `input`-Passwort-Element hinzu, indem du ein `minlength`-Attribut mit einem Wert in Höhe von `8` hinzufügst. Dadurch wird verhindert, dass Eingaben mit weniger als 8 Zeichen übermittelt werden.

# --hints--

Du solltest dem `input`-Passwort-Element ein `minlength`-Attribut zuweisen.

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

Du solltest dem `minlength`-Attribut einen Wert in Höhe von `8` zuweisen.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
