---
id: 60f805f813eaf2049bc2ceea
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Bette ein `input`-Element in jedes `label` ein. Vergewissere dich, dass du jedes `input` nach dem `label`-Text einfügst und füge ein Leerzeichen nach dem Doppelpunkt hinzu.

# --hints--

Du solltest dem `fieldset`-Element vier `input`-Elemente hinzufügen.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

Du solltest die `input`-Elemente innerhalb der `label`-Elemente einbetten.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

Du solltest den ersten `input` nach dem `label`-Text `Enter Your First Name:` hinzufügen und ein Leerzeichen nach dem Doppelpunkt einfügen.

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

Du solltest den zweiten `input` nach dem `label`-Text `Enter Your Last Name:` hinzufügen und ein Leerzeichen nach dem Doppelpunkt einfügen.

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

Du solltest den dritten `input` nach dem `label`-Text `Enter Your Email:` hinzufügen und ein Leerzeichen nach dem Doppelpunkt einfügen.

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

Du solltest den vierten `input` nach dem `label`-Text `Create a New Password:` hinzufügen und ein Leerzeichen nach dem Doppelpunkt einfügen.

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
  --fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
