---
id: 60f80e0081e0f2052ae5b505
title: Schritt 19
challengeType: 0
dashedName: step-19
---

# --description--

Specifying the `type` attribute of a form element is important for the browser to know what kind of data it should expect. If the `type` is not specified, the browser will default to `text`.

Weise den ersten zwei `input`-Elementen ein `type`-Attribut von `text`, dem dritten ein `type`-Attribut von `email` und dem vierten ein `type`-Attribut von `password` zu.

The `email` type only allows emails with a `@` and a `.` in the domain. Der `password`-Typ verdeckt die Eingabe und warnt, wenn die Seite kein HTTPS verwendet.

# --hints--

Du solltest dem ersten `input`-Element ein `type`-Attribut von `text` zuweisen.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Du solltest dem zweiten `input`-Element ein `type`-Attribut von `text` zuweisen.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Du solltest dem dritten `input`-Element ein `type`-Attribut von `email` zuweisen.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

You should give the fourth `input` element a `type` attribute of `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
