---
id: 60fabf0dd4959805dbae09e6
title: Schritt 32
challengeType: 0
dashedName: step-32
---

# --description--

Füge ein weiteres `label` nach dem ersten mit dem Text `Input your age (years):` hinzu. Bette dann ein `input` mit dem `type` von `number` ein.

Füge als Nächstes ein `min`-Attribut zum `input` mit einem Wert von `13` hinzu, da sich Benutzer unter 13 Jahren nicht registrieren sollten. Außerdem werden die Benutzer wahrscheinlich nicht älter als 120 Jahre sein; füge ein `max`-Attribut mit dem Wert `120` hinzu.

Wenn nun jemand versucht das Formular mit Werten außerhalb des Bereichs abzusenden, wird eine Warnung angzeigt und das Formular wird nicht abgeschickt. Probiere es aus.

# --hints--

Du solltest ein `label` zum dritten `fieldset` nach dem vorhandenen `label` hinzufügen.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

Du solltest dem `label` den Text `Input your age (years):` geben.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

Du solltest dem `label` einen `input` mit `type` von `number` geben.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

Du solltest dem `input` ein `min`-Attribut mit einem Wert von `13` zuweisen.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

Du solltest dem `input` ein `max`-Attribut mit einem Wert von `120` zuweisen.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
