---
id: 60facde2d0dc61085b41063f
title: Schritt 37
challengeType: 0
dashedName: step-37
---

# --description--

Das `textarea`-Element wirkt wie ein `input`-Element des Typs `text`, aber bietet den zusätzlichen Vorteil, dass du mehrzeiligen Text und eine ursprüngliche Anzahl an Textzeilen und Spalten empfangen kannst.

Benutzer werden sich mit einer Bio registrieren können. Füge ein `label`-Element mit dem Text `Provide a bio:` am Ende des `fieldset`-Elements hinzu. Füge ein `textarea`-Element innerhalb des `label`-Elements hinzu. Beachte, dass das `textarea`-Element ein schließendes Tag benötigt.

# --hints--

Du solltest am Ende des dritten `fieldset`-Elements ein `label`-Element nach den vorhandenen `label`-Elementen hinzufügen.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)'));
```

Du solltest dem `label`-Element den Text `Provide a bio:` zuweisen.

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)')?.innerText, /Provide a bio/);
```

Du solltest ein `textarea`-Element innerhalb des `label`-Elements einbetten.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4) > textarea'));
```

Du solltest dem `textarea`-Element ein öffnendes und schließendes Tag geben.

```js
assert.match(code, /<textarea\s*>[\s\S]*<\/textarea\s*>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" /></label>
        <label>How did you hear about us?
          <select>
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
