---
id: 62cc5b1779e4d313466f73c5
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Bette ein selbstabschließendes `link`-Element innerhalb des `head`-Elements ein. Weise ihm ein `rel`-Attribut mit dem Wert von `stylesheet` und ein `href`-Attribut mit einem Wert von `styles.css` zu.

# --hints--

Dein Code sollte ein `link`-Element enthalten.

```js
assert.exists(document.querySelector('link'));
```

Du solltest deine vorhandenen `head`-Tags nicht ändern. Stelle sicher, dass du das abschließende Tag nicht gelöscht hast.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Du solltest ein selbstabschließendes `link`-Element haben.

```js
assert(document.querySelectorAll('link').length === 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.exists(document.querySelector('head > link'));
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
