---
id: 62ff8b9dab5ac88e4d3d43a3
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Verkn체pfe nach den bew채hrten Methoden der Zug채nglichkeit die `input`-Elemente mit den `label`-Elementen unter Verwendung des `for`-Attributs.

Verwende `first-name`, `last-name`, `email` und `new-password` als Werte f체r die jeweiligen `id`-Attribute.

# --hints--

Das erste `input`-Element sollte die `id` `first-name` enthalten.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Das zweite `input`-Element sollte eine `id` von `last-name` haben.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Das dritte `input`-Element sollte die `id` `email` enthalten.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Das vierte `input`-Element sollte die `id` `new-password` enthalten.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Das erste `label`-Element sollte ein `for`-Attribut mit einem Wert von `first-name` enthalten.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Das zweite `label`-Element sollte ein `for`-Attribut mit einem Wert von `last-name` enthalten.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Das dritte `label`-Element sollte ein `for`-Attribut mit einem Wert von `email` enthalten.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Das vierte `label`-Element sollte ein `for`-Attribut mit einem Wert von `new-password` enthalten.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
