---
id: 62ff8e998d3e7eae14d6ae3b
title: Schritt 30
challengeType: 0
dashedName: step-30
---

# --description--

You need to confirm that the user has read the terms and conditions.

Füge ein `label`-Element hinzu. Füge innerhalb des neu erstellten `label`-Elements ein `input`-Element hinzu und setze das `type`-Attribut auf `checkbox`. Stelle dieses `input`-Element auf `required` ein, so dass Nutzer sich nicht registrieren können, ohne die Nutzungsbedingung zu akzeptieren.

Füge ein `id`- und `for`-Attribut mit dem Wert `terms-and-conditions` bei den Elementen für die Zugänglichkeit hinzu.

# --hints--

You should add an `label` after the third `fieldset` element.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

You should add an `input` to the `label` element.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Du solltest dem `input`-Element ein `type`-Attribut mit dem Wert `checkbox` zuweisen.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Du solltest dem `input`-Element ein `required`-Attribut zuweisen.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

Das `input`-Element sollte die `id` `terms-and-conditions` haben.

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

Das `label`-Element sollte ein `for`-Attribut mit einem Wert von `terms-and-conditions` enthalten.

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
