---
id: 63541ef4f96cd82e8e6c788a
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

Das `method`-Attribut legt fest, wie Formulardaten an die im `action`-Attribut angegebene URL gesendet werden sollen. Die Formulardaten können als URL-Parameter (mit `method="get"`) über eine `GET`-Anfrage oder als Daten im Request-Body (mit `method="post"`) über eine `POST`-Anfrage gesendet werden.

Lege das `method`-Attribut fest, um deine Formulardaten über eine `POST`-Anfrage zu senden.

# --hints--

Du solltest kein neues `form`-Element hinzufügen.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Dein `form`-Element sollte über ein `method`-Attribut verfügen.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Dein `method`-Attribut sollte auf `post` gesetzt sein.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
