---
id: 60a3e3396c7b40068ad69977
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Utilizza il padding per regolare la spaziatura all'interno di un elemento.

In `.frame`, utilizza la proprietà shorthand `padding` per aumentare di `50px` lo spazio tra gli elementi `.frame` e `.canvas`. La shorthand aumenterà lo spazio in alto, in basso, a sinistra e a destra del bordo dell'elemento e la tela al suo interno.

# --hints--

Dovresti assegnare alla proprietà `padding` il valore `50px`.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

Il selettore `.frame` dovrebbe avere una proprietà `padding` con il valore `50px`.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
