---
id: 60a3e3396c7b40068ad69971
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

È il momento di un po' di CSS.

Anche se l'elemento `<div>` non contiene testo, viene comunque trattato come una scatola con del contenuto. Scrivi una regola CSS che utilizza il selettore di classe `.canvas` e assegna alla sua proprietà `width` il valore 500 pixel. Ecco una regola CSS che imposta la larghezza della classe `card` a 300 pixel:

```css
.card {
  width: 300px;
}
```

# --hints--

Il codice dovrebbe avere un selettore `.canvas`.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

Dovresti impostare la proprietà `width` su `500px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

Il selettore `.canvas` dovrebbe avere una proprietà `width` con il valore `500px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
