---
id: 60a3e3396c7b40068ad6996c
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Nel box model CSS, ogni elemento HTML viene trattato come una scatola con quattro aree.

Immagina di ricevere una scatola dal tuo rivenditore online preferito -- il contenuto è l'elemento nella scatola, o nel nostro caso, degli elementi di intestazione, paragrafo o immagine.

Cambia l'attributo `src` nell'elemento `<img>` da `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` a `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

L'elemento `img` dovrebbe avere un `src` con il valore `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

