---
id: 60a3e3396c7b40068ad6996d
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Il contenuto è circondato da uno spazio chiamato padding, simile al pluriball che separa un oggetto dalla scatola che lo contiene.

Pensa al bordo come alla scatola di cartone in cui il tuo articolo è stato spedito.

Imposta l'attributo `src` su `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

# --hints--

L'elemento `img` dovrebbe avere un `src` con il valore `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
