---
id: 60a3e3396c7b40068ad69972
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Aggiungi la proprietà `height` con il valore `600px` alla regola `.canvas`.

# --hints--

Dovresti impostare la proprietà `height` su `600px`.

```js
const hasHeight = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.height === '600px');
assert(hasHeight);
```

Il selettore `.canvas` dovrebbe avere una proprietà `height` con il valore `600px`.

```js
const canvasHeight = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('height');
assert(canvasHeight === '600px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
--fcc-editable-region--

--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
