---
id: 60a3e3396c7b40068ad6996b
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

All'interno dell'elemento `head`, aggiungi un tag `meta` con un attributo `charset` del valore di `UTF-8` e un elemento `title` con il valore `Rothko Painting`.

All'interno dell'elemento `body`, aggiungi un elemento `img` con un `src` del valore di `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.
# --hints--

Il codice dovrebbe avere un tag `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Il tag `meta` dovrebbe impostare il `charset` su `UTF-8`.

```js
assert(document.querySelector('meta').getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Il codice dovrebbe avere un elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

L'elemento `title` dovrebbe contenere il testo `Rothko Painting`.

```js
assert(document.querySelector('title').innerText === 'Rothko Painting');
```

Il codice dovrebbe avere un elemento `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

L'elemento `img` dovrebbe avere un `src` con il valore `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
  <body>
  </body>
--fcc-editable-region--
</html>
```
