---
id: 60a3e3396c7b40068ad6996a
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Ormai dovresti avere familiarità con gli elementi di base di una pagina HTML.

Aggiungi la dichiarazione `DOCTYPE` e gli elementi `html` (specificando che la pagina è in inglese), `head` e `body`.

# --hints--

Il codice dovrebbe includere un tag `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Dovresti avere un tag `<html>` di apertura con un attributo `lang` con il valore `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il codice dovrebbe avere un elemento `head` all'interno dell'elemento `html`.

```js
assert(document.querySelectorAll('head').length === 1);
```

Il codice dovrebbe avere un elemento `body` all'interno dell'elemento `html`.

```js
assert(document.querySelectorAll('body').length === 1);
```

L'elemento `head` dovrebbe trovarsi prima dell'elemento `body`.

```js
assert(document.querySelector('body').previousElementSibling.tagName === 'HEAD');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
