---
id: 60a3e3396c7b40068ad69974
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Ogni dipinto ha bisogno di una cornice.

Inserisci l'elemento `.canvas` in un altro `div`. Assegna al nuovo `div` la classe `frame`.
# --hints--

Dovresti aggiungere un nuovo elemento `div`.

```js
assert(document.querySelectorAll('div').length === 2)
```

L'elemento `.canvas` dovrebbe essere annidato all'interno del nuovo elemento `div`.

```js
assert(document.querySelector('.canvas').parentElement.tagName === 'DIV');
```

Il nuovo `div` dovrebbe avere un attributo `class` con il valore `frame`.

```js
assert(document.querySelector('div').className.split(' ').includes('frame'));
```

Il nuovo `div` dovrebbe essere all'interno dell'elemento `body`.

```js
assert(document.querySelector('div').parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--

    <div class="canvas">
    </div>

--fcc-editable-region--
  </body>
  ```
