---
id: 60a3e3396c7b40068ad6996e
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Il margine è l'area esterna alla scatola e può essere utilizzato per controllare lo spazio tra i diversi elementi.

Qui l'elemento inferiore ha un margine superiore più grande, che fa sì che si trovi più in basso nella pagina.

Ora che hai capito il modello a scatola CSS, iniziamo il progetto Rothko painting.

Rimuovi l'elemento `<img>`.

# --hints--

Non dovresti avere un elemento `img` nel codice.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
