---
id: 6a6b0a80d50f2685dcaad2e3
title: Step 1
challengeType: 1
dashedName: step-1
---

# --description--

In this workshop, you will work with a smart campus directory and create lookup tables from nested data. 
You'll use these lookup tables to perform fast and efficient searches. The `rawData` has been provided for you.

Take a moment to look at its shape: each key is a department name, and its value holds a `Dept_code` along with
a `programs` array. Each program then holds its own array of `instructors`.

To begin, log `rawData` to the console so you can inspect this nested structure before working with it.

# --hints--

You should log the `rawData` into the console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*rawData\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
