---
id: 6a70ad7439aa38c7878b4323
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

Inside the loop, declare a variable called `departmentId` and assign it the value of `departmentData.Dept_code`.

You should also remove the console statements from your code.

# --hints--

You should delete `console.log(departmentName)`.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*departmentName\s*\)/
);
```

You should delete `console.log(departmentData)`.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*departmentData\s*\)/
);
```

You should declare a variable called `departmentId`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+departmentId\s*=/
);
```

`departmentId` should be set to `departmentData.Dept_code`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+departmentId\s*=\s*departmentData\s*\.\s*Dept_code/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@edu.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@edu.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@edu.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];
    console.log(departmentName);
    console.log(departmentData);

    --fcc-editable-region--

    --fcc-editable-region--
  }
}

normalizeDirectory(rawData);
```
