---
id: 6a7061b30dc6edf4528c6a07
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Inside `normalizeDirectory`, declare four empty objects: `departmentsById`, `programsById`, `instructorsById`, and `instructorsByEmail`.

These objects are lookup tables: they use an ID or another unique value, such as an email, as a key. This lets you access a record directly instead of searching through an array.

# --hints--

You should declare `departmentsById` as an empty object.

```js
const explorer = await __helpers.Explorer(code);
const fn = explorer.functions.normalizeDirectory;
const { departmentsById } = fn.variables;

assert.exists(departmentsById);
assert.isTrue(departmentsById.value.matches("{}"));
```

You should declare `programsById` as an empty object.

```js
const explorer = await __helpers.Explorer(code);
const fn = explorer.functions.normalizeDirectory;
const { programsById } = fn.variables;

assert.exists(programsById);
assert.isTrue(programsById.value.matches("{}"));
```

You should declare `instructorsById` as an empty object.

```js
const explorer = await __helpers.Explorer(code);
const fn = explorer.functions.normalizeDirectory;
const { instructorsById } = fn.variables;

assert.exists(instructorsById);
assert.isTrue(instructorsById.value.matches("{}"));
```

You should declare `instructorsByEmail` as an empty object.

```js
const explorer = await __helpers.Explorer(code);
const fn = explorer.functions.normalizeDirectory;
const { instructorsByEmail } = fn.variables;

assert.exists(instructorsByEmail);
assert.isTrue(instructorsByEmail.value.matches("{}"));
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@edu.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@edu.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@edu.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
--fcc-editable-region--

--fcc-editable-region--
}

normalizeDirectory(rawData);

```
