---
id: 6a70b10d9a1a0ff3cb4775b8
title: Step 18
challengeType: 1
dashedName: step-18
---

# --description--

Now that you have a `departmentId`, use it to store the department's data inside `departmentsById`.

Assign `departmentsById[departmentId]` an object containing:

`id`: set to `departmentId`,  
`code`: set to `departmentData.Dept_code`,  
`name`: set to `departmentName`  
`programIds`: an empty array (you'll fill this in later steps)

Lastly, log the `departmentsById[departmentId]` to the console.

# --hints--

You should assign an object to `departmentsById[departmentId]`.

```js
assert.match(code, /departmentsById\s*\[\s*departmentId\s*\]\s*=\s*{/);
```

The object should have an `id` property set to `departmentId`.

```js
assert.match(code, /id\s*:\s*departmentId/);
```

The object should have a `code` property set to `departmentData.Dept_code`.

```js
assert.match(code, /code\s*:\s*departmentData\s*\.\s*Dept_code/);
```

The object should have a `name` property set to `departmentName`.

```js
assert.match(code, /name\s*:\s*departmentName/);
```

The object should have a `programIds` property set to an empty array.

```js
assert.match(code, /programIds\s*:\s*\[\s*\]/);
```

You should call `console.log()` with `departmentsById[departmentId]` as the argument.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*departmentsById\s*\[\s*departmentId\s*\]\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;
      
    --fcc-editable-region--

    --fcc-editable-region--
  }
}

normalizeDirectory(rawData);
```
