---
id: 6a7061b30dc6edf4528c6a07
title: Step 12
challengeType: 1
dashedName: step-12
---

# --description--

Inside `normalizeDirectory`, declare four empty objects: `departmentsById`, `programsById`, `instructorsById`, and `instructorsByEmail`.

These objects are lookup tables: they use an ID or another unique value, such as an email, as a key. This lets you access a record directly instead of searching through an array.

# --hints--

You should declare `departmentsById` as an empty object.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+departmentsById\s*=\s*{\s*}/
);
```

You should declare `programsById` as an empty object.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+programsById\s*=\s*{\s*}/
);
```

You should declare `instructorsById` as an empty object.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+instructorsById\s*=\s*{\s*}/
);
```

You should declare `instructorsByEmail` as an empty object.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+instructorsByEmail\s*=\s*{\s*}/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
--fcc-editable-region--

--fcc-editable-region--
}

normalizeDirectory(rawData);

```
