---
id: 6a70655455cb163db5e1b482
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

To process the department names and values as array entries, use `Object.entries(rawData)`. It returns an array of `[key, value]` pairs for the object's own enumerable string-keyed properties.

```js
const person = {
  name: "Alice",
  age: 30
};

console.log(Object.entries(person));
/*
Output: 
[ ["name", "Alice"], ["age", 30] ]
*/
```

Declare a variable called `departmentEntries` and set it to `Object.entries(rawData)`. 
Then log `departmentEntries` to the console to see the `[key, value]` pairs it produced.

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, "log");
const getLogs = () => globalThis.spy.calls.map(call => call?.[0]);
```

# --hints--

You should declare a variable called `departmentEntries`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+departmentEntries\s*=/
);
```

`departmentEntries` should be set to `Object.entries(rawData)`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /const\s+departmentEntries\s*=\s*Object\s*\.\s*entries\s*\(\s*rawData\s*\)/
);
```

You should call `console.log()` with `departmentEntries` as the argument.

```js
const expectedEntries = Object.entries(rawData);
const logs = getLogs();

assert.isTrue(
  logs.some(log => JSON.stringify(log) === JSON.stringify(expectedEntries))
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};
  
  --fcc-editable-region--

  --fcc-editable-region--
}

normalizeDirectory(rawData);
```
