---
id: 6a7050a3eea4b3ab9ebfb889
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

At this point, `rawData` is deeply nested: `departments → programs → instructors`. To find an instructor by email, you would have to loop through each department, then each program, and finally each instructor until you find a match.

While that approach works, it becomes slow and repetitive as the data grows. 

Later on in the workshop, you will learn how to normalize the data which is a more efficient approach. 

But for the next few steps, you will create a function that takes an email as its argument, uses loops to search through the nested `rawData` structure, and returns the instructor's name when it finds a matching email address.

Declare a function called `getInstructorByEmail` and use `email` for the parameter name. Then return `null` from the function.

# --hints--

You should declare a function `getInstructorByEmail`.

```js
assert.isFunction(getInstructorByEmail);
```

Your `getInstructorByEmail` function should return `null` for any email.

```js
const result = getInstructorByEmail("test@example.com");
assert.isNull(result);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

console.log(rawData);

--fcc-editable-region--

--fcc-editable-region--
```
