---
id: 6a70ad7439aa38c7878b4323
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

In this step, for each entry of `departmentEntries`, you will assign a vairable `departmentID` equals to `Dept_code`.  

Now create variable `departmentId` equals to the `departmentData.Dept_code`.

# --hints--

You should declare a variable called `departmentId`.

```js
assert.match(code, /const\s+departmentId\s*=/);
```

`departmentId` should be set to `departmentData.Dept_code`.

```js
assert.match(code, /const\s+departmentId\s*=\s*departmentData\s*\.\s*Dept_code/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];
    console.log(departmentName);
    console.log(departmentData);

    --fcc-editable-region--

    --fcc-editable-region--
  }
}

normalizeDirectory(rawData);
```
