---
id: 6a705f3f4d36006ee0d90f87
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Declare a function called `normalizeDirectory` that takes `rawData` as a parameter. 

Then below that function declaration, call `normalizeDirectory` and pass in `rawData` for the argument. 

# --hints--

You should declare a function named `normalizeDirectory`.

```js
assert.isFunction(normalizeDirectory);
```

`normalizeDirectory` function should take rawData` as a parameter.

```js
assert.strictEqual(normalizeDirectory.length, 1);
```

You should call `normalizeDirectory` function.

```js
assert.match(code, /normalizeDirectory\s*\(/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
