---
id: 6a783e4917bab7b18657f0be
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Inside the instructor loop, check if the current instructor's email matches the `email` parameter. If it matches, return the instructor's `name` immediately.

# --hints--

You should create an `if` statement that checks if `instructor.email` is equal to `email`.

```js
assert.match(code, /if\s*\(\s*instructor\s*\.\s*email\s*={2,3}\s*email\s*\)\s*\{[\s\S]*\}/);
```

You should return `instructor.name` inside the `if` statement.

```js
assert.match(code, /return\s+instructor\s*\.\s*name\s*;?\s*\}/);
```

You should check the email and return the instructor's name if it matches.

```js
assert.match(code, /if\s*\(\s*instructor\s*\.\s*email\s*={2,3}\s*email\s*\)\s*\{\s*return\s+instructor\s*\.\s*name\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      for (const instructor of program['instructors']) {
        console.log(instructor);

--fcc-editable-region--

--fcc-editable-region--
      }
    }
  }
  return null;
}

console.log(getInstructorByEmail("fatima@edu.com"));
```
