---
id: 6a783d71dfd45563ae72d93d
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

Now add a third loop inside the program loop to iterate over each instructor. 
Use `for...of` to iterate over the `program["instructors"]` array. Use `instructor` as the loop variable and log each instructor object.

# --hints--

You should create a `for...of` loop that iterates over `program['instructors']`.

```js
assert.match(code, /for\s*\(\s*const\s+instructor\s+of\s+program\s*\[\s*['"]instructors['"]\s*\]\s*\)\s*\{[\s\S]*\}/);
```

You should log the `instructor` to the console inside the loop.

```js
assert.match(code, /for\s*\(\s*const\s+instructor\s+of\s+program\s*\[\s*['"]instructors['"]\s*\]\s*\)\s*\{\s*console\.log\s*\(\s*instructor\s*\);?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      console.log(program);

--fcc-editable-region--
      
--fcc-editable-region--
    }

  }
  return null;
}

console.log(getInstructorByEmail("fatima@edu.com"));
```
