---
id: 6a783e4917bab7b18657f0be
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Inside the instructor loop, check if the current instructor's email matches the `email` parameter. If it matches, return the instructor's `name` immediately.

You should also remove the `console.log(instructor);` line from your code.

# --hints--

You should remove the `console.log(instructor);` line from your code.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*instructor\s*\)/
);
```

Your `getInstructorByEmail` function should return `"Alejandro"` when called with `"alejandro@example.com"`.

```js
assert.strictEqual(getInstructorByEmail('alejandro@example.com'), 'Alejandro');
```

Your `getInstructorByEmail` function should return `"Fatima"` when called with `"fatima@example.com"`.

```js
assert.strictEqual(getInstructorByEmail('fatima@example.com'), 'Fatima');
```

Your `getInstructorByEmail` function should return `"Sofia"` when called with `"sofia@example.com"`.

```js
assert.strictEqual(getInstructorByEmail('sofia@example.com'), 'Sofia');
```

Your `getInstructorByEmail` function should return `null` when called with `"missing@example.com"`.

```js
assert.isNull(getInstructorByEmail('missing@example.com'));
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      for (const instructor of program['instructors']) {
        
--fcc-editable-region--
        console.log(instructor);
        
--fcc-editable-region--
      }
    }
  }
  return null;
}

console.log(getInstructorByEmail("fatima@example.com"));
```
