---
id: 6a783159ed2ac315ea575013
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Inside the department loop, you should access the `programs` array for each department using bracket 
notation. Store it in a variable called `programs` and log it to the console.

You should also delete the `console.log(dept)`.

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, "log");
const getLogs = () => globalThis.spy.calls.map(call => call?.[0]);
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should delete `console.log(dept)` of previous step.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*dept\s*\)/
);
```

You should declare a variable `programs` inside `for...in` loop as `rawData[dept].programs`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /(?:const|let)\s+programs\s*=\s*rawData\s*\[\s*dept\s*\]\s*\.\s*programs\s*;?/
);
```

You should log the `programs` inside `for...in` loop.

```js
const expectedPrograms = Object.values(rawData).map(dept => dept.programs);
const logs = getLogs();

assert.isTrue(
  expectedPrograms.every(programs =>
    logs.some(log => JSON.stringify(log) === JSON.stringify(programs))
  )
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    console.log(dept);

    --fcc-editable-region--

    --fcc-editable-region--
  }
  return null;
}

console.log(getInstructorByEmail("fatima@edu.com"));
```
