---
id: 6a782d8f7681a0851167900a
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Call the `getInstructorByEmail` function with the argument `'fatima@edu.com'`. Place that function call inside of a `console.log()`.

# --hints--

You should write a `console.log()`.

```js
assert.match(code, /console\.log\s*\(.*\)/);
```

Call `getInstructorByEmail` function with `'fatima@edu.com'` argument.

```js
assert.match(code, /getInstructorByEmail\s*\(\s*["']fatima@edu\.com["']\s*\)/);
```

Call `getInstructorByEmail` function inside the `console.log()`.

```js
assert.match(code, /console\.log\s*\(\s*getInstructorByEmail\s*\(\s*["']fatima@edu\.com["']\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  return null;
}

--fcc-editable-region--

--fcc-editable-region--
```
