---
id: 6a7857f3a0819ed988102785
title: Step 26
challengeType: 1
dashedName: step-26
---

# --description--

In an earlier step, you had called the `normalizeDirectory` function with `rawData` for the argument. Now update that line by assigning that function call to a variable called `normalizedData`.

Log `normalizedData` to the console, then check it against the original `normalizedData`:


# --hints--

You should declare a variable called `normalizedData`.

```js
assert.match(code, /const\s+normalizedData\s*=/);
```

`normalizedData` should be set to the result of calling `normalizeDirectory(rawData)`.

```js
assert.match(code, /const\s+normalizedData\s*=\s*normalizeDirectory\s*\(\s*rawData\s*\)/);
```

You should call `console.log()` with `normalizedData` as the argument.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*normalizedData\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

  return {
    departmentsById,
    programsById,
    instructorsById,
    instructorsByEmail,
  };

}

--fcc-editable-region--
normalizeDirectory(rawData);

--fcc-editable-region--
```
