---
id: 6a783e4917bab7b18657f0be
title: Step 8
challengeType: 1
dashedName: step-8
---

# --description--

Inside the instructor loop, check if the current instructor's email matches the `email` parameter. If it matches, return the instructor's `name` immediately.

Before starting the loop delete `console.log(instructor);` of previous step.

# --hints--

Before starting, delete `console.log(instructor)` from the previous step.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*instructor\s*\)/
);
```

You should use an `if` statement to check whether `instructor.email` is equal to `email` and return `instructor.name` when it matches.

```js
assert.match(
  __helpers.removeJSComments(code),
  /if\s*\(\s*instructor\s*\.\s*email\s*={2,3}\s*email\s*\)\s*\{\s*return\s+instructor\s*\.\s*name\s*;?\s*\}/
);
```

Your `getInstructorByEmail` function should return `"Alejandro"` when called with `"alejandro@edu.com"`.

```js
assert.strictEqual(getInstructorByEmail('alejandro@edu.com'), 'Alejandro');
```

Your `getInstructorByEmail` function should return `"Fatima"` when called with `"fatima@edu.com"`.

```js
assert.strictEqual(getInstructorByEmail('fatima@edu.com'), 'Fatima');
```

Your `getInstructorByEmail` function should return `"Sofia"` when called with `"sofia@edu.com"`.

```js
assert.strictEqual(getInstructorByEmail('sofia@edu.com'), 'Sofia');
```

Your `getInstructorByEmail` function should return `null` when called with `"missing@edu.com"`.

```js
assert.isNull(getInstructorByEmail('missing@edu.com'));
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      for (const instructor of program['instructors']) {
        console.log(instructor);

--fcc-editable-region--
        
--fcc-editable-region--
      }
    }
  }
  return null;
}

console.log(getInstructorByEmail("fatima@edu.com"));
```
