---
id: 6a78549572797d79eedc16e5
title: Step 25
challengeType: 1
dashedName: step-25
---

# --description--

Your four lookup tables are now fully built. Have `normalizeDirectory` return an object containing 
all four objects: `departmentsById`, `programsById`, `instructorsById`, and `instructorsByEmail`.

Since the property names match the variable names exactly, you can use object 
shorthand — writing `{ departmentsById }` instead of `{ departmentsById: departmentsById }`.

# --hints--

You should use a `return` statement.

```js
assert.match(code, /return\s*{/);
```

You should return an object containing `departmentsById`.

```js
assert.match(code, /return\s*{[^}]*\bdepartmentsById\b[^}]*}/);
```

You should return an object containing `programsById`.

```js
assert.match(code, /return\s*{[^}]*\bprogramsById\b[^}]*}/);
```

You should return an object containing `instructorsById`.

```js
assert.match(code, /return\s*{[^}]*\binstructorsById\b[^}]*}/);
```

You should return an object containing `instructorsByEmail`.

```js
assert.match(code, /return\s*{[^}]*\binstructorsByEmail\b[^}]*}/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];
      console.log(program);

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

--fcc-editable-region--

--fcc-editable-region--
}

normalizeDirectory(rawData);
```
