---
id: 6a782d8f7681a0851167900a
title: Step 3
challengeType: 1
dashedName: step-3
---

# --description--

Call the `getInstructorByEmail` function with the argument `'fatima@edu.com'`. Place that function call inside of a `console.log()`.

# --hints--

You should have a `console.log()`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\.log\s*\(.*\)/
);
```

You should call the `getInstructorByEmail` function with `'fatima@edu.com'` for the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /getInstructorByEmail\s*\(\s*['"]fatima@edu\.com['"]\s*\)/
);
```

Your `getInstructorByEmail` function call should be inside the `console.log()`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*getInstructorByEmail\s*\(\s*(['"])fatima@edu\.com\1\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@edu.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@edu.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@edu.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  return null;
}

--fcc-editable-region--

--fcc-editable-region--
```
