---
id: 6a78523edb374913e58d71d9
title: Step 24
challengeType: 1
dashedName: step-24
---

# --description--

Not every `instructor` has an `email` property — some may only have a phone number or no contact 
info at all. Writing `instructorsByEmail[instructor.email]` without checking first would create 
an entry with an `undefined` key for those `instructors`.

Now, you should add a conditional `if (instructor.email)`, then only get inside that block, 
assign `instructorsByEmail[instructor.email]` it to `instructor.id`.

By `instructorsByEmail` lookup table, you will find the ID of the `instructor`, using ID you can 
find all the information of `instructor`.

# --hints--

You should use an `if` statement.

```js
assert.match(code, /if\s*\(/);
```

Your `if` statement should check `instructor.email`.

```js
assert.match(code, /if\s*\(\s*instructor\.email\s*\)/);
```

Inside the `if` block, you should assign to `instructorsByEmail[instructor.email]`.

```js
assert.match(code, /if\s*\(\s*instructor\.email\s*\)\s*{\s*instructorsByEmail\s*\[\s*instructor\.email\s*\]\s*=/);
```

`instructorsByEmail[instructor.email]` should be set to `instructor.id`.

```js
assert.match(code, /if\s*\(\s*instructor\.email\s*\)\s*{\s*instructorsByEmail\s*\[\s*instructor\.email\s*\]\s*=\s*instructor\.id/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];
      console.log(program);

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        console.log(instructorsById[instructor.id]);

--fcc-editable-region--

--fcc-editable-region--
      }

    }
  }
}

normalizeDirectory(rawData);
```
