---
id: 6a78513949298537bfa222b6
title: Step 23
challengeType: 1
dashedName: step-23
---

# --description--

In this step, you will store each `instructor` in `instructorsById`, keyed by its own id. 
Just like you did for `programs`, assign `instructorsById[instructor.id]` an object that 
spreads `instructor` and adds `programId` (set it to `program.id`) and `departmentId`.

Log `instructorsById[instructor.id]` in the console to see stored object.

# --hints--

You should assign an object to `instructorsById[instructor.id]`.

```js
assert.match(code, /instructorsById\s*\[\s*instructor\.id\s*\]\s*=\s*{/);
```

The object should spread `instructor`.

```js
assert.match(code, /instructorsById\s*\[\s*instructor\.id\s*\]\s*=\s*{\s*\.\.\.\s*instructor/);
```

The object should include `programId` set to `program.id`.

```js
assert.match(code, /programId\s*:\s*program\.id/);
```

The object should include `departmentId`.

```js
assert.match(code, /instructorsById\s*\[\s*instructor\.id\s*\]\s*=\s*{\s*\.\.\.\s*instructor\s*,\s*programId\s*:\s*program\.id\s*,\s*departmentId/);
```

You should call `console.log()` with `instructorsById[instructor.id]` as the argument.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*instructorsById\s*\[\s*instructor\.id\s*\]\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];
        console.log(instructor);

--fcc-editable-region--

--fcc-editable-region--
      }

    }
  }
}

normalizeDirectory(rawData);
```
