---
id: 6a788bcb73a3d8d17a6240f0
title: Step 28
challengeType: 1
dashedName: step-28
---

# --description--

Now, call the function with `"fatima@edu.com"` and `normalizedData`.
Log the output from the function call.

# --hints--

You should call `getInstructorByEmail` with `"fatima@edu.com"` and `normalizedData` as arguments.

```js
assert.match(code, /getInstructorByEmail\s*\(\s*["']fatima@edu\.com["']\s*,\s*normalizedData\s*\)/);
```

You should call `console.log()` with the result of that call as the argument.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getInstructorByEmail\s*\(\s*["']fatima@edu\.com["']\s*,\s*normalizedData\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

  return {
    departmentsById,
    programsById,
    instructorsById,
    instructorsByEmail,
  };

}

const normalizedData = normalizeDirectory(rawData);

function getInstructorByEmail(email, normalizedData) {
  return null;
}
--fcc-editable-region--

--fcc-editable-region--

```
