---
id: 6a7838b74249924e0cd9012b
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Now add a second loop inside the department loop to iterate over each `program` and use `program` as the loop variable.   
Use a `for...of` loop to loop through the `programs` array and log each `program` object.

You should also remove the `console.log(programs)` line from your code.

Here is a reminder on how to work with a `for...of` loop:

```js
const dishes = [
  { id: "FOOD101", name: "Margherita Pizza", price: 14 },
  { id: "FOOD102", name: "Chicken Tacos", price: 12 }
];

for (const dish of dishes) {
  console.log(dish);
}

// { id: "FOOD101", name: "Margherita Pizza", price: 14 }
// { id: "FOOD102", name: "Chicken Tacos", price: 12 }
```

# --before-each--

```js
globalThis.spy = __helpers.spyOn(console, "log");
const getLogs = () => globalThis.spy.calls.map(call => call?.[0]);
```

# --after-each--

```js
globalThis.spy.restore();
delete globalThis.spy;
```

# --hints--

You should remove the `console.log(programs)` line from your code.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*programs\s*\)/
);
```

You should create a `for...of` loop that iterates over `programs`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(?:const|let)\s+program\s+of\s+programs\s*\)\s*\{[\s\S]*\}/
);
```

You should log the `program` to the console inside the loop.

```js
const expectedPrograms = Object.values(rawData).flatMap(dept => dept.programs);
const logs = getLogs();

assert.isTrue(
  expectedPrograms.every(program =>
    logs.some(log => JSON.stringify(log) === JSON.stringify(program))
  )
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;
    console.log(programs);

    --fcc-editable-region--
    
    --fcc-editable-region--
  }
  return null;
}

console.log(getInstructorByEmail("fatima@example.com"));
```
