---
id: 6a784bd0d86389ad5528535c
title: Step 22
challengeType: 1
dashedName: step-22
---

# --description--

Each program has its own `instructors` array. Add a third `for` loop, nested inside the programs loop, iterating over `program.instructors` from `k = 0` up to `program.instructors.length`.

Before adding the instructor loop, delete `console.log(departmentsById[departmentId])` from the previous step.

Inside this loop, declare a variable called `instructor` set to `program.instructors[k]` and log `instructor` in the console.

# --hints--

You should delete `console.log(departmentsById[departmentId]);`.

```js
assert.notMatch(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*departmentsById\s*\[\s*departmentId\s*\]\s*\)/
);
```

You should use a `for` loop.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*let\s+k\s*=\s*0\s*/
  
);
```

Your `for` loop should run while `k` is less than `program.instructors.length`.

```js
assert.match(
  __helpers.removeJSComments(code), 
  /for\s*\(\s*let\s+k\s*=\s*0\s*;\s*k\s*<\s*program\.instructors\.length\s*/
);
```

You should declare a variable called `instructor` set to `program.instructors[k]`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /(?:const|let)\s+instructor\s*=\s*program\s*\.\s*instructors\s*\[\s*k\s*\]/
);
```

You should call `console.log()` with `instructor` as the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*instructor\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "405",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@example.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@example.com", officeHours: "Mon-Wed 10AM-2PM" },
        ],
      },
      {
        id: "CSE103",
        name: "Operating Systems",
        room: "401",
        instructors: [
          { id: "INS005", name: "Wei", email: "wei@example.com", officeHours: "Thu 1PM-3PM" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@example.com" }],
      },
      {
        id: "BBA102",
        name: "Marketing Management",
        room: "204",
        instructors: [
          { id: "INS006", name: "Robin", email: "robin@example.com", officeHours: "Tue-Thu 9AM-11AM" },
        ],
      },
      {
        id: "BBA103",
        name: "Business Law",
        room: "203",
        instructors: [
          { id: "INS007", name: "Priya", email: "priya@example.com", officeHours: "Mon 2PM-4PM" },
        ],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);

--fcc-editable-region--

--fcc-editable-region--
    }

    console.log(departmentsById[departmentId]);
  }
}

normalizeDirectory(rawData);
```
