---
id: 6a7848aad918b10afae7d1f0
title: Step 20
challengeType: 1
dashedName: step-20
---

# --description--

Now store each program in `programsById`, keyed by its own id.

The stored program should have its original properties plus a new `departmentId` property linking it to its department. Object spread syntax (`...`) copies an object's own enumerable properties into a new object, where you can add extra properties.

Assign `programsById[program.id]` an object that spreads `program` and adds `departmentId`.

Log `programsById[program.id]` to the console to inspect the stored object.

# --hints--

You should assign an object to `programsById[program.id]`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programsById\s*\[\s*program\.id\s*\]\s*=\s*{/
);
```

The object should spread `program`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programsById\s*\[\s*program\.id\s*\]\s*=\s*{\s*\.\.\.\s*program/
);
```

The object should include `departmentId`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /programsById\s*\[\s*program\.id\s*\]\s*=\s*{\s*\.\.\.\s*program\s*,\s*departmentId/
);
```

You should call `console.log()` with `programsById[program.id]` as the argument.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*programsById\s*\[\s*program\.id\s*\]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com", officeHours: "Sun-Tue 10AM-12PM" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: null,
        instructors: [
          {id: "INS003", name: "Fatima", email: "fatima@edu.com", "officeHours": "Mon-Wed 10AM-2PM"},
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];
      console.log(program);

--fcc-editable-region--

--fcc-editable-region--
    }
  }
}

normalizeDirectory(rawData);
```
