---
id: 6a788eb3b69df0b1009bc649
title: Step 30
challengeType: 1
dashedName: step-30
---

# --description--

If `email` doesn't exist in `instructorsByEmail`, `instructorId` will be undefined. Before doing
anything else, add a guard clause: an early check that stops the function and returns immediately 
if something's wrong, so you don't have to write nested logic for the "invalid" case.

Add an `if (!instructorId)` check that returns the string `Instructor not found`.

# --hints--

You should use an `if` statement.

```js
assert.match(code, /if\s*\(/);
```

Your `if` statement should check `!instructorId`.

```js
assert.match(code, /if\s*\(\s*!\s*instructorId\s*\)/);
```

Inside the `if` block, you should return the string `"Instructor not found"`.

```js
assert.match(code, /if\s*\(\s*!\s*instructorId\s*\)\s*{\s*return\s*["']Instructor not found["']/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

  return {
    departmentsById,
    programsById,
    instructorsById,
    instructorsByEmail,
  };

}

const normalizedData = normalizeDirectory(rawData);

function getInstructorByEmail(email, normalizedData) {
  const instructorId = normalizedData.instructorsByEmail[email];
  console.log(instructorId);

--fcc-editable-region--

--fcc-editable-region--

  return null;
}

console.log( getInstructorByEmail("fatima@edu.com", normalizedData));
```
