---
id: 6a7840732007006f073005d8
title: Step 9
challengeType: 1
dashedName: step-9
---

# --description--

Now, test the function with email `"alejandro@edu.com"` and `"sofia@edu.com"`, log the output in console like you
have done for `"fatima@edu.com"`.


# --hints--

You should write a `console.log()`.

```js
assert.match(code, /console\.log\s*\(.*\)/);
```

Call `getInstructorByEmail` function with `'alejandro@edu.com'` argument.

```js
assert.match(code, /getInstructorByEmail\s*\(\s*["']alejandro@edu\.com["']\s*\)/);
```

Call `getInstructorByEmail` function inside the `console.log()`.

```js
assert.match(code, /console\.log\s*\(\s*getInstructorByEmail\s*\(\s*["']alejandro@edu\.com["']\s*\)\s*\)/);
```

You should write another `console.log()`.

```js
assert.match(code, /console\.log\s*\(.*\)/g);
```

Call `getInstructorByEmail` function with `'sofia@edu.com'` argument.

```js
assert.match(code, /getInstructorByEmail\s*\(\s*["']sofia@edu\.com["']\s*\)/);
```

Call `getInstructorByEmail` function inside the second `console.log()`.

```js
assert.match(code, /console\.log\s*\(\s*getInstructorByEmail\s*\(\s*["']sofia@edu\.com["']\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      for (const instructor of program['instructors']) {
        if (instructor.email === email) {
          return instructor.name; 
        }
      }
    }
  }
  return null;
}

console.log(getInstructorByEmail("fatima@edu.com"));
--fcc-editable-region--

--fcc-editable-region--
```
