---
id: 6a86f98046a13f6752229620
title: Step 51
challengeType: 1
dashedName: step-51
---

# --description--

Finally, return a template literal that formats all five counts into a readable report:

```js
`
Campus Directory Summary
------------------------
Departments: ${departments}
Programs: ${programs}
Instructors: ${instructors}
Missing Office Hours: ${missingOfficeHours}
Programs Without Rooms: ${programsWithoutRooms}
`
```

# --hints--

You should use a `return` statement.

```js
assert.match(code, /return\s*`/);
```

The returned template literal should include `${departments}`.

```js
assert.match(code, /return\s*`[^`]*\$\{\s*departments\s*\}/);
```

The returned template literal should include `${programs}`.

```js
assert.match(code, /return\s*`[^`]*\$\{\s*programs\s*\}/);
```

The returned template literal should include `${instructors}`.

```js
assert.match(code, /return\s*`[^`]*\$\{\s*instructors\s*\}/);
```

The returned template literal should include `${missingOfficeHours}`.

```js
assert.match(code, /return\s*`[^`]*\$\{\s*missingOfficeHours\s*\}/);
```

The returned template literal should include `${programsWithoutRooms}`.

```js
assert.match(code, /return\s*`[^`]*\$\{\s*programsWithoutRooms\s*\}/);
```

`renderDirectorySummary(normalizedData)` should return the correct summary report.

```js
const result = renderDirectorySummary(normalizedData);

assert.match(
  result,
  /Campus Directory Summary\s*-+\s*Departments:\s*\d+\s*Programs:\s*\d+\s*Instructors:\s*\d+\s*Missing Office Hours:\s*\d+\s*Programs Without Rooms:\s*\d+/
);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

  return {
    departmentsById,
    programsById,
    instructorsById,
    instructorsByEmail,
  };

}

const normalizedData = normalizeDirectory(rawData);


function getInstructorByEmail(email, normalizedData) {
  const instructorId = normalizedData.instructorsByEmail[email];

  if (!instructorId) {
    return "Instructor not found";
  }

  const instructor = normalizedData.instructorsById[instructorId];
  const department = normalizedData.departmentsById[instructor.departmentId];


  return {
    name: instructor.name ?? "Unknown Instructor",
    officeHours: instructor.officeHours ?? "Not Available",
    department: department.name ?? "Unknown Department",
  };
}


function listRoomsByBuilding(buildingCode, normalizedData) {

  const programValues = Object.values(normalizedData.programsById);
  const rooms = [];

  for (let i = 0; i < programValues.length; i++) {
    const program = programValues[i];
    const department = normalizedData.departmentsById[program.departmentId];

    if (department.code === buildingCode && program.room != null) {
      rooms.push(program.room);
    }
    
  }

  rooms.sort();
  return rooms;
}

function renderDirectorySummary(normalizedData) {
  const departments = Object.keys(normalizedData.departmentsById).length;
  const programs = Object.keys(normalizedData.programsById).length;
  const instructors = Object.keys(normalizedData.instructorsById).length;

  const instructorValues = Object.values(normalizedData.instructorsById);
  let missingOfficeHours = 0;

  for (let i = 0; i < instructorValues.length; i++) {
    if (instructorValues[i].officeHours == null) {
      missingOfficeHours++;
    }
  }

  const programValues = Object.values(normalizedData.programsById);
  let programsWithoutRooms = 0;

  for (let i = 0; i < programValues.length; i++) {
    if (programValues[i].room == null) {
      programsWithoutRooms++;
    }
  }
  
  --fcc-editable-region--
          
  --fcc-editable-region--
}
```
