---
id: 6a86f2c4902af075bdfc58cf
title: Step 44
challengeType: 1
dashedName: step-44
---

# --description--

The total number of `departments`, `programs`, and `instructors` can each be found by counting the keys in their 
respective lookup tables.

Declare three variables:

`departments` set to `Object.keys(normalizedData.departmentsById).length`  

`programs` set to `Object.keys(normalizedData.programsById).length`  

`instructors`, set to `Object.keys(normalizedData.instructorsById).length`  

# --hints--

You should declare a variable called `departments`.

```js
assert.match(code, /const\s+departments\s*=/);
```

`departments` should be set to `Object.keys(normalizedData.departmentsById).length`.

```js
assert.match(code, /const\s+departments\s*=\s*Object\s*\.\s*keys\s*\(\s*normalizedData\.departmentsById\s*\)\s*\.\s*length/);
```

You should declare a variable called `programs`.

```js
assert.match(code, /const\s+programs\s*=/);
```

`programs` should be set to `Object.keys(normalizedData.programsById).length`.

```js
assert.match(code, /const\s+programs\s*=\s*Object\s*\.\s*keys\s*\(\s*normalizedData\.programsById\s*\)\s*\.\s*length/);
```

You should declare a variable called `instructors`.

```js
assert.match(code, /const\s+instructors\s*=/);
```

`instructors` should be set to `Object.keys(normalizedData.instructorsById).length`.

```js
assert.match(code, /const\s+instructors\s*=\s*Object\s*\.\s*keys\s*\(\s*normalizedData\.instructorsById\s*\)\s*\.\s*length/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

    const departmentId = departmentData.Dept_code || `dept-${i + 1}`;

    departmentsById[departmentId] = {
      id: departmentId,
      code: departmentData.Dept_code,
      name: departmentName,
      programIds: [],
    };

    for (let j = 0; j < departmentData.programs.length; j++) {
      const program = departmentData.programs[j];

      programsById[program.id] = {
        ...program,
        departmentId,
      };

      departmentsById[departmentId].programIds.push(program.id);
      for (let k = 0; k < program.instructors.length; k++) {
        const instructor = program.instructors[k];

        instructorsById[instructor.id] = {
          ...instructor,
          programId: program.id,
          departmentId,
        };

        if (instructor.email) {
          instructorsByEmail[instructor.email] = instructor.id;
        }
      }
    }
  }

  return {
    departmentsById,
    programsById,
    instructorsById,
    instructorsByEmail,
  };

}

const normalizedData = normalizeDirectory(rawData);


function getInstructorByEmail(email, normalizedData) {
  const instructorId = normalizedData.instructorsByEmail[email];

  if (!instructorId) {
    return "Instructor not found";
  }

  const instructor = normalizedData.instructorsById[instructorId];
  const department = normalizedData.departmentsById[instructor.departmentId];


  return {
    name: instructor.name ?? "Unknown Instructor",
    officeHours: instructor.officeHours ?? "Not Available",
    department: department.name ?? "Unknown Department",
  };
}


function listRoomsByBuilding(buildingCode, normalizedData) {

  const programValues = Object.values(normalizedData.programsById);
  const rooms = [];

  for (let i = 0; i < programValues.length; i++) {
    const program = programValues[i];
    const department = normalizedData.departmentsById[program.departmentId];

    if (department.code === buildingCode && program.room != null) {
      rooms.push(program.room);
    }
    
  }

  rooms.sort();
  return rooms;
}

function renderDirectorySummary(normalizedData) {

  --fcc-editable-region--
      
  --fcc-editable-region--
}

```
