---
id: 6a86afe3266c32e402611261
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

While your current code works, the nested loops make the code harder to read and maintain. Every lookup has to go through the departments, programs, and instructors one by one.

A better approach is to organize the data into lookup tables. This will make it easier to find what you are looking for without going through all of that nested data.

Delete the `getInstructorByEmail` function with all `console.log()` printing output of `getInstructorByEmail`.

# --hints--

You should remove the `getInstructorByEmail` function definition entirely.

```js
assert.notMatch(code, /function\s+getInstructorByEmail\s*\(/);
```

You should not have any `console.log()` calls that print `getInstructorByEmail` output.

```js
assert.notMatch(code, /console\.log\s*\(\s*getInstructorByEmail\s*\(/);
```

You should not have any `getInstructorByEmail` function calls remaining in the code.

```js
assert.notMatch(code, /getInstructorByEmail\s*\(/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

--fcc-editable-region--

function getInstructorByEmail(email) {
  for (const dept in rawData) {
    const programs = rawData[dept].programs;

    for (const program of programs) {
      for (const instructor of program['instructors']) {
        if (instructor.email === email) {
          return instructor.name; 
        }
      }
    }
  }
  return null;
}

console.log(getInstructorByEmail("fatima@edu.com"));
console.log(getInstructorByEmail("alejandro@edu.com"));
console.log(getInstructorByEmail("sofia@edu.com"));


--fcc-editable-region--
```
