---
id: 6a8bec19ff25dc6373dc087f
title: Step 17
challengeType: 1
dashedName: step-17
---

# --description--

But, not every department is guaranteed to have a `Dept_code`. You need a fallback for when it's missing, so the department still gets a usable ID.

The `||` operator returns its left-hand value if that value is `truthy`, otherwise it returns the right-hand value.  

For example: `a || b` evaluates to `a` when `a` is `truthy`, and falls back to `b` when `a` is `falsy` (like `undefined`, `null`, or `""`).

```js
const code = undefined;
const fallback = code || "default-code";
console.log(fallback); // "default-code"
```

Now add a fallback to `departmentId` equals to the `departmentData.Dept_code`, if `departmentData.Dept_code` is `undefined` take `dept-${i + 1}` as fallback.



# --hints--

You should declare a variable called `departmentId`.

```js
assert.match(code, /const\s+departmentId\s*=/);
```

`departmentId` should use `departmentData.Dept_code` as its primary value.

```js
assert.match(code, /const\s+departmentId\s*=\s*departmentData\s*\.\s*Dept_code\s*\|\|/);
```

`departmentId` should fall back to `` `dept-${i + 1}` `` when `Dept_code` is missing.

```js
assert.match(code, /const\s+departmentId\s*=\s*departmentData\s*\.\s*Dept_code\s*\|\|\s*`dept-\$\{\s*i\s*\+\s*1\s*\}`/);
```

# --seed--

## --seed-contents--

```js
const rawData = {
  "Computer Science": {
    Dept_code: "CSE",
    programs: [
      {
        id: "CSE101",
        name: "Database Management System",
        room: "402",
        instructors: [
          { id: "INS001", name: "Alejandro", email: "alejandro@edu.com" },
          { id: "INS002", name: "Kenji", officeHours: null },
        ],
      },
      {
        id: "CSE102",
        name: "Data Structures",
        room: "205",
        instructors: [
          { id: "INS003", name: "Fatima", email: "fatima@edu.com" },
        ],
      },
    ],
  },
  "Business Administration": {
    Dept_code: "BBA",
    programs: [
      {
        id: "BBA101",
        name: "Financial Accounting",
        room: "201",
        instructors: [{ id: "INS004", name: "Sofia", email: "sofia@edu.com" }],
      },
    ],
  },
};

function normalizeDirectory(rawData) {
  const departmentsById = {};
  const programsById = {};
  const instructorsById = {};
  const instructorsByEmail = {};

  const departmentEntries = Object.entries(rawData);

  for (let i = 0; i < departmentEntries.length; i++) {

    const [departmentName, departmentData] = departmentEntries[i];

--fcc-editable-region--

    const departmentId = departmentData.Dept_code ;

--fcc-editable-region--



  }
}

normalizeDirectory(rawData);
```
