---
id: 66a37f37ef5823a313de8c26
title: Learn CSS by Building a Moon Orbit
challengeType: 14
dashedName: build-a-moon-orbit
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `div` with the class `space`.

2. Inside the `.space` `div`, there should be two more `div` elements with the class `earth` and `orbit`, respectively.

3. Inside the `.orbit` `div`, there should be another `div` with the class `moon`.

4. The `space` `div` should be centered on the page and have a width and height of `200px`.

5. The `earth` `div` should be centered inside the `space` `div` and have a width and height of `100px`.

6. The `orbit` `div` should be centered with respect to the earth and have a width and height of `200px`.

7. The orbit path for the moon around the Earth should be a circle. 

8. The moon `div` should be centered inside the `orbit` `div` and have a width and height of `30px`.

9. Your `.earth` and `.moon` `div` elements should have a background color and `border-radius` to make them look like planets.

10. You should define a `@keyframes orbit` animation that rotates the `.orbit` element 360 degrees around its center. You should apply this animation to the `.orbit` element with a duration of `5` seconds, a linear timing function, and infinite iterations.

# --hints--

You should have a main `div` with the class `space`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'space');
```

The `space` `div` should have a width and height of `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.space')?.width === '200px' && new __helpers.CSSHelp(document).getStyle('.space')?.height === '200px');
```

You should have a `div` with the class `earth` inside the `.space` `div`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

The `earth` `div` should have a width and height of `100px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.earth')?.width === '100px' && new __helpers.CSSHelp(document).getStyle('.earth')?.height === '100px');
```

You should have a `div` with the class `orbit` inside the `.space` `div`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

The `orbit` `div` should have a width and height of `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.orbit')?.width === '200px' && new __helpers.CSSHelp(document).getStyle('.orbit')?.height === '200px');
```

You should have a `div` with the class `moon` inside the `.orbit` `div`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

The `moon` `div` should have a width and height of `30px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.moon')?.width === '30px' && new __helpers.CSSHelp(document).getStyle('.moon')?.height === '30px');
```

Your `.earth` `div` element should have a background color.

```js
assert(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

Your `.earth` element should have `border-radius` to make it look like a circle.

```js
assert(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius);
```

Your `.moon` `div` element should have a background color.

```js
assert(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

Your `.moon` element should have `border-radius` to make it look like a circle.

```js
assert(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius);
```

You should have a `@keyframes` rule.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.length === 1);
```

Your new `@keyframes` rule should be named `orbit`.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name === 'orbit');
```

Your `@keyframes orbit` rule should have a `0%` selector.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert(rules?.[0]?.keyText === '0%' || rules?.[0]?.keyText === '0%');
```

Your `@keyframes orbit` rule should have a `100%` selector.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Your `0%` selector should have a `transform` property set to `rotate(0deg) translate(-50%, -50%)`.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform === 'rotate(0deg) translate(-50%, -50%)');
```

Your `100%` selector should come after your `0%` selector.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText === '100%')
```

Your `100%` selector should have a `transform` property set to `rotate(360deg) translate(-50%, -50%)`.

```js
assert(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform === 'rotate(360deg) translate(-50%, -50%)')
```

Your `.orbit` selector should have an `animation` property set to `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert(
  orbitStyles?.animationName === "orbit" &&
    orbitStyles?.animationDuration === "5s" &&
    orbitStyles?.animationTimingFunction === "linear" &&
    orbitStyles?.animationIterationCount === "infinite"
);

```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

