---
id: 616968c2c94c8071b349c146
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Sasa uko tayari kuanza kuongeza maudhui kwenye ukurasa.

Ndani ya kipengele cha `body`, weka kipengele cha `h1` chenye maandishi `CSS Color Markers`.

# --hints--

Code yako inapaswa kuwa na kipengele cha `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Unapaswa kuwa na tagi ya kufungua ya `h1`.

```js
assert(code.match(/<h1\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `h1`.

```js
assert(code.match(/<\/h1\s*>/i));
```

Kipengele cha `h1` kinafaa kuwa ndani ya kipengele cha `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Kumbuka, ukubwa wa herufi na tahajia ni jambo la muhimu kwa maandishi ya `h1`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
