---
id: 616965351e74d4689eb6de30
title: Hatua ya 5
challengeType: 0
dashedName: step-5
---

# --description--

Unaweza kuwa na vipengele vingi vya kujifunga vya `meta` kwenye ukurasa wa wavuti. Kila kipengele cha `meta` huongeza maelezo kuhusu ukurasa ambayo hayawezi kuonyeshwa na vipengele vingine vya HTML.

Ongeza kipengele kingine cha kujifunga cha `meta` ndani ya `head`. Ipe sifa ya `name` iliyowekwa kuwa `viewport` na sifa ya `content` iliyowekwa kuwa `width=device-width, initial-scale=1.0` ili ukurasa wako uonekane sawa kwenye vifaa vyote.

# --hints--

Unapaswa kuwa na vipengele viwili vya `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Kipengele chako kipya cha `meta` kinapaswa kuwa kipengele cha kujifunga.

```js
assert(code.match(/<\/meta>/i) === null);
```

Kipengele chako kipya cha `meta` kinapaswa kuwa na sifa ya `name` iliyowekwa kuwa `viewport`, na sifa ya `content` iliyowekwa kuwa `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
