---
id: 61696ef7ac756c829f9e4048
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Katika mradi huu, utafanya kazi na faili ya nje ya CSS ili kuunda ukurasa. Tayari tumekuundia faili ya `styles.css`. Lakini kabla ya kuitumia, utahitaji kuiunganisha kwenye ukurasa.

Ongeza kipengele cha `link` ndani ya kipengele chako cha `head`. Ipe sifa ya `rel` iliyowekwa kuwa `stylesheet` na sifa ya `href` iliyowekwa kuwa `styles.css`.

# --hints--

Unapaswa kuwa na kipengele kimoja cha kujifunga cha `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya `rel` yenye thamani ya `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
