---
id: 61695c4aad56f95497c19583
title: Hatua ya 3
challengeType: 0
dashedName: step-3
---

# --description--

Kumbuka kwamba kipengele cha `title` kinaipa injini za utafutaji maelezo ya ziada kuhusu ukurasa. Pia huonyesha maudhui ya kipengele hicho cha `title` kwa njia mbili zaidi:

* kwenye upau wa kichwa wakati ukurasa umefunguliwa
* kwenye kichupo cha kivinjari cha ukurasa unapoelea juu yake. Hata kama kichupo hicho hakitumiki, mara tu unapoelea kwenye kichupo, maandishi ya `title` yanaonyeshwa.

Ndani ya kipengele cha `head`, weka kipengele cha `title` chenye maandishi `Colored Markers`.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `title`.

```js
assert(code.match(/<title\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Mradi wako unapaswa kuwa na kichwa `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Kumbuka, ukubwa wa herufi na tahajia ni jambo la muhimu katika kichwa.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
