---
id: 61695d1fbc003856628bf561
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Ili kuviambia vivinjari jinsi ya kusimba herufi kwenye ukurasa wako, weka `charset` kuwa `utf-8`. `utf-8` ni seti ya vibambo zima inayojumuisha takriban kila herufi kutoka lugha zote za binadamu.

Ndani ya kipengele cha `head`, weka kipengele cha `meta` chenye sifa ya `charset` iliyowekwa `utf-8`. Kumbuka kwamba vipengele vya `meta` vinajifunga, na havihitaji tagi ya kufunga.

# --hints--

Unapaswa kuwa na kipengele kimoja cha `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

Kipengele chako cha `meta` kinapaswa kuwa kipengele cha kujifunga.

```js
assert(code.match(/<\/meta>/i) === null);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya `charset` iliyowekwa kuwa `utf-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
