---
id: 616d4a84b756d9c4b8255093
title: Hatua ya 11
challengeType: 0
dashedName: step-11
---

# --description--

Ni wakati wa kuongeza rangi kwenye alama. Kumbuka kwamba njia moja ya kuongeza rangi kwenye kipengele ni kutumia <dfn>color keyword</dfn> kama `black`, `cyan` au `yellow`.

Kama kumbusho, hii ndio jinsi ya kulenga class ya `freecodecamp`:

```css
.freecodecamp {

}
```

Unda sheria mpya ya CSS inayolenga class ya `marker` na uweke sifa ya `background-color` kuwa `red`.

**Kumbuka:** Hutaona mabadiliko yoyote baada ya kuongeza CSS.

# --hints--

Unapaswa kuunda kiteuzi cha class ili kulenga class ya `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
