---
id: 616d595270d902f0e7086e18
title: Hatua ya 15
challengeType: 0
dashedName: step-15
---

# --description--

Ingawa una vipengele vitatu tofauti vya alama ya `div`, vinaonekana kama mstatili mmoja mkubwa. Unapaswa kuongeza nafasi kati yao ili iwe rahisi kuona kila kipengele.

Wakati kipengele cha mkato cha `margin` kina thamani mbili, huweka `margin-top` na `margin-bottom` kwa thamani ya kwanza, na `margin-left` na `margin-right` kwa thamani ya pili.

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `margin` iliyowekwa kuwa `10px auto`.

# --hints--

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `margin` iliyowekwa kuwa `10px auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
