---
id: 616d525007b8c5d8b3308b1c
title: Hatua ya 13
challengeType: 0
dashedName: step-13
---

# --description--

Alama yako ingeonekana bora ikiwa ingewekwa katikati kwenye ukurasa. Njia rahisi ya kufanya hivyo ni kwa kutumia `margin` <dfn>shorthand property</dfn>.

Katika mradi wa mwisho, uliweka eneo la ukingo wa vipengee kando na sifa kama `margin-top` na `margin-left`. Sifa ya mkato ya `margin` hurahisisha kuweka maeneo mengi ya ukingo kwa wakati mmoja.

Ili kuweka alama yako katikati kwenye ukurasa, weka sifa ya `margin` kuwa `auto`. Hii inaweka `margin-top`, `margin-right`, `margin-bottom`, na `margin-left` zote kuwa `auto`.

# --hints--

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `margin` iliyowekwa kuwa `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
