---
id: 616d50b93ba424d6282c99cf
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Rangi ya mandharinyuma ilitumika, lakini kwa kuwa kipengee cha alama cha `div` hakina maudhui ndani yake, hakina urefu wowote kwa chaguo-msingi.

Katika sheria yako ya CSS ya `.marker`, weka sifa za `height` kuwa `25px` na `width` kuwa `200px`

# --hints--

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `width` iliyowekwa kuwa `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `height` iliyowekwa kuwa `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
