---
id: 61764c602bee6974e7790f35
title: Hatua ya 16
challengeType: 0
dashedName: step-16
---

# --description--

Ili kutoa alama za rangi tofauti, utahitaji kuongeza class ya kipekee kwa kila moja. Classes yingi zinaweza kuongezwa kwa kipengele kwa kuorodhesha katika sifa ya `class` na kuzitenganisha na nafasi. Kwa mfano, ifuatayo huongeza classes za `animal` na `dog` kwenye kipengele cha `div`.

```html
<div class="animal dog">
```

Ukiongeza classes nyingi kwenye kipengele cha HTML, mitindo ya class za kwanza inayoorodhesha inaweza kubatilishwa na classes za baadaye.

Ili kuanza, ongeza class `one` kwa alama ya kwanza ya kipengele cha `div`.


# --hints--

Unapaswa kuongeza class `one` kwa alama ya kwanza ya kipengele cha `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Alama yako ya kwanza ya `div` inapaswa kuwa na classes za `marker` na `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
