---
id: 6176519636a76b810ab1219a
title: Hatua ya 20
challengeType: 0
dashedName: step-20
---

# --description--

Unda sheria ya CSS inayolenga class ya `two` na uweke kipengele chake cha `background-color` kuwa `green`.

Pia, unda sheria tofauti ya CSS inayolenga class ya `three` na uweke `background-color` kuwa `blue`.

# --hints--

Unapaswa kutumia kichaguzi cha class kulenga class ya `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

Unapaswa kutumia kichaguzi cha class kulenga class ya `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

Sheria yako ya CSS ya `.three` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
