---
id: 61764fdda535587e1fefb3aa
title: Hatua ya 19
challengeType: 0
dashedName: step-19
---

# --description--

Ongeza class ya `two` kwa alama ya pili `div`, na uongeze class ya `three` kwa alama ya tatu `div`.

# --hints--

Unapaswa kuongeza class ya `two` kwa alama ya pili ya kipengele cha pili cha `div` katika `container` `div`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('two'));
```

`div` ya alama yako ya pili inapaswa kuwa na classes za `marker` na `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('marker') && containerSecondChild.classList?.contains('two'));
```

Unapaswa kuongeza class ya `three` kwa alama ya tatu ya kipengele cha pili cha `div` katika `container` `div`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('three'));
```

`div` ya alama yako ya tatu inapaswa kuwa na classes za `marker` na `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('marker') && containerThirdChild?.classList?.contains('three'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

```
