---
id: 617bd6ec666b1da2587e4e37
title: Hatua ya 34
challengeType: 0
dashedName: step-34
---

# --description--

Kuna rangi tatu zaidi za tertiary: chartreuse green (kijani + njano), azure (bluu + cyan), na rose (nyekundu + magenta).

Ili kuunda chartreuse green, sasisha kitendakazi cha `rgb` katika sheria ya CSS ya `.one` ili nyekundu iwe `127`, na uweke kijani kuwa thamani ya juu zaidi.

Kwa azure, sasisha kitendakazi cha `rgb` katika sheria ya CSS ya `.two` ili kijani kiwe `127` na bluu iwe katika thamani ya juu zaidi.

Na kwa rose, ambayo wakati mwingine huitwa bright pink, sasisha kitendakazi cha `rgb` katika sheria ya CSS ya `.three` ili bluu iwe `127` na nyekundu. iko kwenye thamani ya juu.

# --hints--

Sheria yako ya CSS ya `.one` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(127, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 127, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

Sheria yako ya CSS ya `.three` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(255, 0, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
