---
id: 617b97abd9f3f61d1590b815
title: Hatua ya 26
challengeType: 0
dashedName: step-26
---

# --description--

Katika modeli ya rangi ya ziada ya RGB, <dfn>primary colors</dfn> ni rangi ambazo, zikiunganishwa, huunda nyeupe safi. Lakini kwa hili kutokea, kila rangi inahitaji kuwa katika kiwango chake cha juu.

Kabla ya kuchanganya rangi, rudisha alama yako ya kijani kuwa kijani kibichi. Kwa kitendakazi cha `rgb` katika sheria ya CSS ya `.two`, rudisha thamani ya green kuwa juu zaidi `255`.

# --hints--

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
