---
id: 617b954d9f4f6217a749380e
title: Hatua ya 25
challengeType: 0
dashedName: step-25
---

# --description--

Sasa ongeza nafasi wima zaidi kati ya vialamisho vyako na ukingo wa kipengele cha `container` walichomo.

Katika sheria ya CSS ya `.container`, tumia sifa ya mkato ya `padding` kuongeza `10px` ya padding ya juu na chini, na kuweka padding ya kushoto na kulia kuwa `0`. Hii inafanya kazi sawa na sifa ya mkato ya `margin` ulichotumia awali.

# --hints--

Hupaswi kuondoa sifa ya `background-color` na thamani yake kutoka kwa sheria ya CSS ya `.container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

Sheria yako ya CSS ya `.container` inapaswa kuwa na sifa ya `padding` iliyowekwa kuwa `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
