---
id: 617b92b9de349513466f6156
title: Hatua ya 24
challengeType: 0
dashedName: step-24
---

# --description--

Wakati alama nyekundu na bluu zinaonekana sawa, moja ya kijani ni nyepesi zaidi kuliko ilivyokuwa hapo awali. Hii ni kwa sababu neno kuu la rangi ya `green` ni kivuli cheusi, na iko karibu nusu kati ya nyeusi na thamani ya juu zaidi ya kijani.

Katika sheria ya CSS ya `.two`, weka thamani ya kijani katika kitendakazi cha `rgb` kuwa `127` ili kupunguza ukubwa wake.

# --hints--

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 127, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
