---
id: 617bc3386dc7d07d6469bf20
title: Hatua ya 32
challengeType: 0
dashedName: step-32
---

# --description--

Fahamu kwamba, ili kuunda rangi ya chungwa, ilibidi uongeze nguvu ya nyekundu na kupunguza ukubwa wa thamani za kijani katika `rgb`. Hii ni kwa sababu machungwa ni mchanganyiko wa nyekundu na njano.

Ili kuunda rangi ya tertiary ya chemchemi ya kijani, changanya cyan na kijani. Sasisha kitendakazi cha `rgb` katika sheria ya CSS ya `.two` ili kijani kifikie thamani ya juu zaidi, na uweke bluu kuwa `127`.

# --hints--

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 255, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
