---
id: 618a00ed1ca871a2b3aca0cb
title: Hatua ya 35
challengeType: 0
dashedName: step-35
---

# --description--

Kwa kuwa sasa umepitia rangi zote za primary, secondary, na tertiary kwenye gurudumu la rangi, itakuwa rahisi kuelewa dhana nyingine za nadharia ya rangi na jinsi zinavyoathiri muundo.

Kwanza, katika sheria za CSS za `.one`, `.two`, na `.three`, rekebisha thamani katika kitendakazi cha `rgb` ili `background-color` ya kila kipengele iwekwe nyeusi kabisa. Kumbuka kwamba kitendakazi cha `rgb` hutumia modeli ya rangi ya addictive, ambapo rangi huanza kuwa nyeusi na kubadilika kadri thamani za nyekundu, kijani na buluu zinavyoongezeka.

# --hints--

Sheria yako ya CSS ya `.one` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

Sheria yako ya CSS ya `.three` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
