---
id: 619b761916dac02643940022
title: Hatua ya 46
challengeType: 0
dashedName: step-46
---

# --description--

Njia ya kawaida sana ya kuweka rangi kwenye kipengele kilicho na CSS ni kwa <dfn>hexadecimal</dfn> au thamani za heksi. Ingawa thamani za hex zinakaa kuwa ngumu, kwa kweli ni aina nyingine ya thamani za RGB.

Thamani za rangi za Hex huanza na alama ya `#` na kuchukua nambari sita kutoka 0-9 na A-F. Jozi ya kwanza ya wahusika inawakilisha nyekundu, jozi ya pili inawakilisha kijani, na jozi ya tatu inawakilisha bluu. Kwa mfano, `#4B5320`.

Katika kichaguzi cha class ya `.green`, weka `background-color` kuwa msimbo wa rangi wa hex wenye thamani `00` kwa nyekundu, `FF` kwa kijani, na `00` bluu.

# --hints--

Sheria yako ya CSS ya `.green` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `#00FF00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
