---
id: 61a493ead935c148d2b76312
title: Hatua ya 50
challengeType: 0
dashedName: step-50
---

# --description--

Kitendaji cha `linear-gradient` kinaweza kunyumbulika sana -- hapa kuna sintaksia ya msingi utakayotumia katika somo hili:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` ni mwelekeo wa njia inayotumika kwa mpito. `color1` na `color2` ni hoja za rangi, ambazo ndizo rangi zitakazotumika katika mpito wenyewe. Hizi zinaweza kuwa aina yoyote ya rangi, ikijumuisha manenomsingi ya rangi, heksi, `rgb` au `hsl`.

Sasa utaweka upinde rangi nyekundu hadi kijani kibichi kando ya mstari wa digrii 90 kwenye alama ya kwanza.

Kwanza, katika sheria ya CSS ya `.red`, weka sifa ya `background` kuwa `linear-gradient()`, na uipitishe thamani `90deg` kama `gradientDirection`.

# --hints--

Sheria yako ya CSS ya `.red` inapaswa kuwa na sifa ya `background` yenye thamani `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
