---
id: 61a489b8579e87364b2d2cdb
title: Hatua ya 49
challengeType: 0
dashedName: step-49
---

# --description--

Umejifunza njia chache za kuweka rangi bapa katika CSS, lakini pia unaweza kutumia mpito wa rangi, au <dfn>gradient</dfn>, kwenye kipengele.

Gradient ni wakati rangi moja inabadilika kuwa nyingine. Kitendakazi cha CSS cha `linear-gradient` hukuwezesha kudhibiti mwelekeo wa mpito kando ya mstari, na rangi zipi zinatumika.

Jambo moja la kukumbuka ni kwamba kitendakazi cha `linear-gradient` huunda kipengele cha `image`, na kwa kawaida huoanishwa na `background` ambayo inaweza kukubali picha kama thamani.

Katika sheria ya CSS ya `.red`, badilisha sifa ya `background-color` kuwa `background`.

# --hints--

Sheria yako ya CSS ya `.red` inapaswa kuwa na sifa ya `background` yenye thamani `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
