---
id: 61b093329e7fc020b43b1bb5
title: Hatua ya 75
challengeType: 0
dashedName: step-75
---

# --description--

Tayari una uzoefu wa kutumia kitendakazi cha `rgb` kuweka rangi. Ili kuongeza kituo cha alpha kwenye rangi ya `rgb`, tumia kitendakazi cha `rgba` badala yake.

Kitendakazi cha `rgba` hufanya kazi kama kitendakazi cha `rgb`, lakini huchukua nambari moja zaidi kutoka `0` hadi `1.0` kwa kituo cha alpha:

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

Unaweza pia kutumia kituo cha alpha chenye rangi za `hsl` na `hex`. Utaona jinsi ya kufanya hivyo hivi karibuni.

Katika sheria ya `.sleeve`, tumia kitendakazi cha `rgba` ili kuweka `background-color` kuwa nyeupe kabisa yenye opacity ya 50%.

# --hints--

Sheria yako ya CSS ya `.sleeve` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgba(255, 255, 255, 0.5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor === 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
