---
id: 61b30a286c228bd0c493c09a
title: Hatua ya 89
challengeType: 0
dashedName: step-89
---

# --description--

Lakini ni nini ikiwa ungependa kuongeza uwazi wa kivuli zaidi? Unaweza kufanya hivyo kwa kutumia thamani ya `spreadRadius`:

```css
box-shadow: offsetX offsetY blurRadius spreadRadius color;
```

Kama vile `blurRadius`, `spreadRadius` hujipa chaguomsingi ya `0` ikiwa haijajumuishwa.

Fanya mazoezi kwa kuongeza kivuli cha pikseli 5 moja kwa moja karibu na alama ya bluu.

Katika sheria ya CSS ya `.blue`, ongeza sifa ya `box-shadow` yenye thamani `0` kwa `offsetX`, `0` kwa `offsetY`, na `0` kwa `blurRadius`, `5px` kwa `spreadRadius`, na `blue` kwa `color`.

# --hints--

Sheria yako ya CSS ya `.blue` inapaswa kuwa na sifa ya mkato ya `box-shadow` yenye thamani ya `0 0 0 5px blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.boxShadow === 'blue 0px 0px 0px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
  box-shadow: -5px -5px red;
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
  box-shadow: 5px 5px 5px green;
}

--fcc-editable-region--
.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}
--fcc-editable-region--

```
