---
id: cf1111c1c11feddfaeb3bdef
title: Додайте два числа з JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` - це такий тип в JavaScript, який передає числові дані.

Тепер спробуємо додати два числа за допомогою JavaScript.

JavaScript використовує символ `+` в ролі оператора додавання, і розміщується між двома числами.

**Наприклад:**

```js
const myVar = 5 + 10;
```

`myVar` тепер матиме значення `15`.

# --instructions--

Змініть `0` так, щоб сума дорівнювала `20`.

# --hints--

`sum` повинен дорівнювати `20`.

```js
assert(sum === 20);
```

Вам слід використовувати оператор `+`.

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
