---
id: bd7123c9c441eddfaeb4bdef
title: Поясніть ваш JavaScript код
challengeType: 1
removeComments: false
videoUrl: 'https://scrimba.com/c/c7ynnTp'
forumTopicId: 16783
dashedName: comment-your-javascript-code
---

# --description--

Пояснення - це рядки коду, які JavaScript буде навмисно ігнорувати. Коментарі - це чудовий спосіб залишити нотатки для себе та інших людей, які згодом мають дізнатися, що робить цей код.

Є два шляхи написати пояснення у JavaScript:

Використання `//` дозволить JavaScript ігнорувати решту тексту в поточному рядку. Це коментар всередині рядка:

```js
// This is an in-line comment.
```

Ви можете зробити багаторядковий коментар починаючи з `/*` і закінчуватися `*/`. Це багаторядковий коментар:

```js
/* This is a
multi-line comment */
```

**ПРИМІТКА:**, написавши код, вам слід постійно додавати коментарі, щоб уточнити функцію частин вашого коду. Хороші і якісні коментарі можуть допомогти передати мету вашого коду – як для інших людей, *так і* для себе в майбутньому.

# --instructions--

Спробуйте створити один коментар кожного типу.

# --hints--

Ви повинні створити стиль `//`, що містить щонайменше п'ять літер.

```js
assert(code.match(/(\/\/)...../g));
```

Ви повинні створити коментар стилю `/* */`, що містить щонайменше п'ять літер.

```js
assert(code.match(/(\/\*)([^\/]{5,})(?=\*\/)/gm));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
// Fake Comment
/* Another Comment */
```
