---
id: 56533eb9ac21ba0edf2244a9
title: Ініціалізація змінних за допомогою оператора присвоєння
challengeType: 1
videoUrl: 'https://scrimba.com/c/cWJ4Bfb'
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

Часто змінній <dfn>ініціалізовують</dfn> початкове значення в тому ж рядку, що оголошення.

```js
var myVar = 0;
```

Створює нову змінну під назвою `myVar` та призначає їй початкове значення `0`.

# --instructions--

Визначте змінну `a` за допомогою `var` та ініціалізуйте їй значення `9`.

# --hints--

Ви повинні ініціалізувати `a` зі значенням `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(code));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
