---
id: cf1111c1c11feddfaeb5bdef
title: Ітерація з циклами for у JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/c9yNVCe'
forumTopicId: 18219
dashedName: iterate-with-javascript-for-loops
---

# --description--

Ви можете виконати один і той самий код декілька разів за допомогою циклу.

Найпоширенішим типом циклу в JavaScript є `for`, оскільки він виконується певну кількість разів.

Цикли for оголошуються з трьома додатковими виразами, які розділені крапкою з комою:

`for (a; b; c)`, де `a` є інструкцією ініціалізації, `b` є інструкцією умови, а `c` є кінцевим виразом.

Інструкція ініціалізації виконується лише один раз (перед початком циклу). Її переважно використовують, щоб визначити та налаштувати змінну циклу.

Інструкція умови обчислюється на початку кожної ітерації циклу і продовжуватиметься до тих пір, поки не обчислиться до `true`. Коли на початку ітерації умова становить `false`, цикл припинить виконуватись. Це означає те, що якщо умова починається як «false», ваш цикл ніколи не виконається.

Кінцевий вираз виконується в кінці кожної ітерації циклу, перед перевіркою наступної умови і, зазвичай, використовується для того, щоб збільшити або зменшити лічильник.

У наступному прикладі ми ініціалізуємо з `i = 0` та ітеруємо, доки наша умова `i < 5` не буде істинною. Ми збільшимо `i` на `1` у кожній ітерації циклу з кінцевим виразом `i++`.

```js
const ourArray = [];

for (let i = 0; i < 5; i++) {
  ourArray.push(i);
}
```

Тепер `ourArray` матиме значення `[0, 1, 2, 3, 4]`.

# --instructions--

Використайте цикл `for`, щоб передати значення від 1 до 5 до `myArray`.

# --hints--

Ви повинні використати цикл `for`.

```js
assert(/for\s*\([^)]+?\)/.test(code));
```

`myArray` має дорівнювати `[1, 2, 3, 4, 5]`.

```js
assert.deepEqual(myArray, [1, 2, 3, 4, 5]);
```

# --seed--

## --after-user-code--

```js
if (typeof myArray !== "undefined"){(function(){return myArray;})();}
```

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 6; i++) {
  myArray.push(i);
}
```
