---
id: 6552008c0d9d9075cbec9772
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

The first kind of cipher you are going to build is called a *Caesar* cipher. Specifically, you will take each letter in your message, find its position in the alphabet, take the letter located after 3 positions, and replace the original letter with the new letter.

Start by finding the position of the first letter in the string. One way is to call the `.find()` string method on the string you want to parse:

```py
text.find('W')
```

Above, uppercase `'W'` is the character you want to locate inside the string stored in the `text` variable. The method will return `6`, which is the index of the `'W'` character inside the string stored in the `text` variable.

At the end of your code, call `.find()` on your `alphabet` string and pass `text[0]` to the method. Note that a method is just a function that belongs to an object (you will learn more about that in another project).

# --hints--

You should call the `.find()` method.

```js
assert.match(code, /find\s*\(.*\)/)
```

You should call the `.find()` method on the `alphabet` variable.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```

You should call the `.find()` method on the `alphabet` variable and pass `text[0]` to the method. Pay attention to place the method call at the beginning of the line.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
