---
id: 5f33071498eb2472b87ddee4
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

In diesem Projekt lernst du die Grundlagen von CSS (Cascading Style Sheets) indem du eine Speisekarte für ein Café erstellst. CSS ist die Sprache, die verwendet wird, um ein HTML-Dokument zu gestalten. Es beschreibt, wie HTML-Elemente auf dem Bildschirm angezeigt werden sollen.

Wie du in den letzten Schritten der Katzenfoto-App bereits gelernt hast, ist eine grundlegende Struktur erforderlich, um deine Webseite zu erstellen. Jedes HTML-Dokument sollte eine `DOCTYPE` Deklaration und ein `HTML-` Element enthalten. Der `DOCTYPE` teilt dem Browser mit, in welcher HTML-Version das Dokument verfasst ist. Und das `HTML-` Element stellt das Wurzelelement dar, das alle anderen Elemente enthält.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Füge das `<!DOCTYPE html>`-Tag und ein `html`-Element mit einem `lang`-Element von `en` hinzu.

# --hints--

Du solltest eine `<!DOCTYPE html>`-Deklaration haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Du solltest ein öffnendes `<html lang="en">` Element haben, das die Sprache auf Englisch stellt.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Du solltest ein abschließendes `</html>`-Tag haben.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
