---
id: 5f33294a6af5e9188dbdb8f3
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Es ist an der Zeit, ein paar Menü-Inhalte hinzuzufügen. Füge ein `main`-Element innerhalb des vorhandenen `body`-Elements ein. Es wird irgendwann Preisinformationen über die vom Café angebotenen Kaffees und Desserts enthalten.

# --hints--

Dein Code sollte über ein öffnendes `<main>`-Tag verfügen.

```js
assert(code.match(/<main>/i));
```

Dein Code sollte über ein abschließendes `</main>`-Tag verfügen.

```js
assert(code.match(/<\/main>/i));
```

Du solltest dein `body`-Element nicht ändern. Make sure you don't accidentally delete your closing tag.

```js
assert($('body').length === 1);
```

Dein `main`-Tag sollte sich innerhalb deines `body`-Tags befinden.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
