---
id: 5f33071498eb2472b87ddee4
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Wie du in den letzten Schritten der Katzenfoto-App bereits gelernt hast, ist eine grundlegende Struktur erforderlich, um deine Webseite zu erstellen.

Add the `<!DOCTYPE html>` tag, and an `html` element with a `lang` attribute of `en`.

# --hints--

Du solltest eine `DOCTYPE`-Deklaration haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

You should have an opening `<html>` tag with a `lang` attribute of `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Du solltest ein schließendes `<html>`-Tag haben. Remember that closing tags have a `/` following the opening `<` bracket.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
