---
id: 5f3313e74582ad9d063e3a38
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. Der Text des `title`-Elements sollte `Cafe Menu` sein.

# --hints--

You should have an opening `<head>` tag.

```js
assert(code.match(/<head>/i));
```

Du solltest ein schließendes `</head>`-Tag haben.

```js
assert(code.match(/<head>/i));
```

You should have an opening `<title>` tag.

```js
assert(code.match(/<title>/i));
```

Du solltest ein schließendes `</title>`-Tag haben.

```js
assert(code.match(/<\/title>/i));
```

Your `title` element should be nested in your `head` element.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Your `title` element should have the text `Cafe Menu`. Möglicherweise musst du deine Rechtschreibung überprüfen.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
