---
id: 5f33071498eb2472b87ddee4
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Wie du in den letzten Schritten der Katzenfoto-App bereits gelernt hast, ist eine grundlegende Struktur erforderlich, um deine Webseite zu erstellen.

Füge das `<!DOCTYPE html>`-Tag und ein `html`-Element mit einem `lang`-Element von `en` hinzu.

# --hints--

Du solltest eine `DOCTYPE`-Deklaration haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Du solltest ein öffnendes `<html>`-Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Du solltest ein schließendes `<html>`-Tag haben. Vergiss nicht, dass schließende Tags ein `/` nach der Öffnungsklammer (`<`) haben.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
