---
id: 5f332a88dc25a0fd25c7687a
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Der Name des Cafés ist `CAMPER CAFE`. Add an `h1` element within your `main` element. Gib den Namen des Cafés in Großbuchstaben an, um ihn hervorzuheben.

# --hints--

You should have an opening `<h1>` tag.

```js
assert(code.match(/<h1>/i));
```

Du solltest ein schließendes `</h1>`-Tag haben.

```js
assert(code.match(/<\/h1>/i));
```

Du solltest dein bestehendes `main`-Element nicht ändern.

```js
assert($('main').length === 1);
```

Dein `h1`-Element sollte in deinem `main`-Element verschachtelt sein.

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

Dein `h1`-Element sollte den Text `CAMPER CAFE` enthalten.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

