---
id: 5f332b23c2045fb843337579
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Füge ein `p`-Element unterhalb des `h1`-Elements mit dem Text `Est. 2020` ein, um die Besucher wissen zu lassen, dass das Café im Jahr 2020 gegründet wurde.

# --hints--

Du solltest ein öffnendes `<p>`-Tag haben.

```js
assert(code.match(/<p>/i));
```

Du solltest ein schließendes `</p>` Tag haben.

```js
assert(code.match(/<\/p>/i));
```

Du solltest dein vorhandenes `h1`-Element nicht ändern. Stelle sicher, dass du das schließende Tag nicht gelöscht hast.

```js
assert($('h1').length === 1);
```

Dein `p`-Element sollte unter deinem `h1`-Element stehen.

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

Dein `p`-Element sollte den Text `Est. 2020` haben.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

