---
id: 5f3326b143638ee1a09ff1e3
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare to create some actual content, add a `body` element below the `head` element.

# --hints--

Du solltest ein öffnendes `<body>`-Tag haben.

```js
assert(code.match(/<body>/i));
```

Du solltest ein schließendes `</body>`-Tag haben.

```js
assert(code.match(/<\/body>/i));
```

Du solltest dein `head`-Element nicht ändern. Stelle sicher, dass du dein schließendes Tag nicht gelöscht hast.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Dein `body`-Element sollte nach deinem `head`-Element stehen.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

