---
id: 5f331e55dfab7a896e53c3a1
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Der `title` ist eines von mehreren Elementen, die zusätzliche Informationen enthalten, die auf der Webseite nicht sichtbar aber nützlich für Suchmaschinen sind oder die Art und Weise ändern, wie die Seite dargestellt wird.

Im `head`-Element, ist ein `meta`-Element mit einem Attribut namens `charset` verschachtelt, das auf den Wert `utf-8` gesetzt ist, um dem Browser zu erklären, wie Zeichen für die Seite kodiert werden sollen. Beachte, dass `meta`-Elemente selbstschließend sind.

# --hints--

Du solltest ein `meta`-Tag haben.

```js
assert(code.match(/<meta.*?\/?>/is));
```

Dein `meta`-Tag sollte ein `charset`-Attribut haben.

```js
assert(code.match(/<meta charset=/i));
```

Dein `charset` sollte einen Wert von `utf-8` haben.

```js
assert(code.match(/charset=('|")utf-8\1/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
