---
id: 5f3313e74582ad9d063e3a38
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstelle ein `head`-Element innerhalb des `html`-Elements, so dass du ein `title`-Element hinzufügen kannst. Der Text des `title`-Elements sollte `Cafe Menu` sein.

# --hints--

Du solltest ein öffnendes `<head>`-Tag haben.

```js
assert(code.match(/<head>/i));
```

Du solltest ein schließendes `</head>`-Tag haben.

```js
assert(code.match(/<head>/i));
```

Du solltest ein öffnendes `<title>` Tag haben.

```js
assert(code.match(/<title>/i));
```

Du solltest ein schließendes `</title>`-Tag haben.

```js
assert(code.match(/<\/title>/i));
```

Dein `title`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Dein `title`-Element sollte den Text `Cafe Menu` haben. Möglicherweise musst du deine Rechtschreibung überprüfen.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
