---
id: 5f331e55dfab7a896e53c3a1
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Der `title` ist eines von mehreren Elementen, die zusätzliche Informationen enthalten, die auf der Webseite nicht sichtbar aber nützlich für Suchmaschinen sind oder die Art und Weise ändern, wie die Seite dargestellt wird.

Verschachtele innerhalb des `head` Elements ein `meta-` -Element mit einem Attribut mit dem Namen `charset`, das auf den Wert `utf-8` gesetzt ist, um dem Browser mitzuteilen, wie Zeichen für die Seite entschlüsselt werden sollen.

# --hints--

Du solltest ein `meta`-Tag haben.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Das `meta` Element ist ein leeres Element, es sollte kein schließendes Tag `</meta>`haben.

```js
assert.notMatch(code, /<\/meta>/i);
```

Dein `meta`-Tag sollte ein `charset`-Attribut haben.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Dein `charset` sollte einen Wert von `utf-8` haben.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
