---
id: 5f3477cbcb6ba47918c1da92
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Damit das Design der Seite auf dem Handy ähnlich aussieht wie auf einem Desktop oder Laptop, musst du ein `meta`-Element mit einem speziellen `content`-Attribut hinzufügen.

Füge das Folgende innerhalb des `head`-Elements hinzu:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Dein Code sollte zwei `meta`-Elemente haben.

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

Dein `meta`-Element sollte ein `name`-Attribut mit dem Wert `viewport` haben.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

Dein `meta`-Element sollte ein `content`-Attribut mit einem Wert von `width=device-width, initial-scale=1.0` haben.

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
