---
id: 5f3477cb2e27333b1ab2b955
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Jetzt musst du die `styles.css`-Datei verknüpfen, damit die Formate wieder angewendet werden. Füge ein sich selbst schließendes `link`-Element in das `head`-Element ein. Gib ihm einen `rel`-Attributwert von `stylesheet` und einen `href`-Attributwert von `styles.css`.

# --hints--

Dein Code sollte ein `link`-Element haben.

```js
const link = document.querySelector('link');
assert(link);
```

Du solltest nicht das vorhandene `head`-Element ändern. Stelle sicher, dass du das schließende Tag nicht gelöscht hast.

```js
assert($('head').length === 1);
```

Du solltest ein selbstschließendes `link`-Element haben.

```js
const link = document.querySelectorAll('link');
assert(link.length === 1);
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
const link = document.querySelector('head > link');
assert(link);
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link = document.querySelector('link')
const rel = link.getAttribute('rel')
assert(rel == `stylesheet`)
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link')
assert(link.dataset.href == 'styles.css')

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
