---
id: 5f344fc1520b6719f2e35605
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Auf der Speisekarte gibt es zwei Abschnitte, einen für Kaffee und einen für die Desserts. Füge ein `section`-Element innerhalb des `main`-Elements hinzu, so dass du einen Platz hast, um alle verfügbaren Kaffeesorten zu zeigen.

# --hints--

Du solltest ein einleitendes `<section>`-Tag haben.

```js
assert(code.match(/<section\s*>/i));
```

Du solltest ein abschließendes `</section>`-Tag haben.

```js
assert(code.match(/<\/section\s*>/i));
```

Du solltest das vorhandene `main`-Element nicht ändern. Stelle sicher, dass du den abschließenden Tag nicht gelöscht hast.

```js
assert($('main').length === 1);
```

Dein `section`-Element sollte sich innerhalb deines `main`-Elements befinden.

```js
const main = document.querySelectorAll('main')?.[0];
assert(main.children[0].tagName === 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
