---
id: 5f344fbc22624a2976425065
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Erstelle ein `h2`-Element innerhalb des `section`-Elements und gib ihm den Text `Coffee`.

# --hints--

You should have an opening `<h2>` tag.

```js
assert(code.match(/<h2\s*>/i));
```

Du solltest ein schließendes `</h2>`-Tag haben.

```js
assert(code.match(/<\/h2\s*>/i));
```

Du solltest das vorhandene `section`-Element nicht verändern. Stelle sicher, dass du das schließende Tag nicht gelöscht hast.

```js
assert($('section').length === 1);
```

Dein `h2`-Element sollte sich innerhalb deines `section`-Elements befinden.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Dein `h2`-Element sollte den Text `Coffee` haben.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
