---
id: 5f3477ae34c1239cafe128be
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Du hast jetzt drei Typselektoren mit der gleichen Formatierung. Du kannst die gleiche Gruppe von Stilen zu vielen Elementen hinzufügen, indem du eine Selektorenliste erstellst. Jeder Selektor wird wie folgt mit Kommas getrennt:

```css
selector1, selector2 {
  property: value;
}
```

Delete the three existing type selectors and replace them with one selector list that centers the text for the `h1`, `h2`, and `p` elements.

# --hints--

Du solltest für alle drei Elemente einen einzigen Typselektor wählen: `h1, h2, p`. Stelle sicher, dass du diese Reihenfolge benutzt.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Du solltest nur einen Selektor in deinem `style`-Element haben.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

