---
id: 5f3477ae8466a9a3d2cc953c
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Jetzt, wo du das CSS in der `styles.css`-Datei hast, kannst du das `style`-Element und seinen gesamten Inhalt entfernen. Sobald es entfernt wird, wird der Text, der zentriert war, wieder nach links verschoben.

# --hints--

Du solltest keine `style`-Tags in deinem Code haben.

```js
assert(!code.match(/style/i));
```

Du solltest keine CSS-Selektoren in deiner HTML-Datei haben.

```js
(getUserInput) => {
  const html = getUserInput('editableContents');
  assert(!html.includes('style'));
  assert(!html.includes('text-align'));
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

