---
id: 5f35e5c44359872a137bd98f
title: Schritt 28
challengeType: 0
dashedName: step-28
---

# --description--

Da das Hauptprodukt des Cafés der Kaffee ist, kannst du ein Bild von Kaffeebohnen für den Hintergrund der Seite verwenden.

Lösche den Kommentar und dessen Inhalt innerhalb des `body`-Typselektors. Füge jetzt eine `background-image`-Eigenschaft hinzu und setze deren Wert auf `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Du solltest die auskommentierte `background-color`-Eigenschaft entfernen.

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

Dein `body`-Selektor sollte keine Kommentare enthalten.

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

Du solltest die `background-image`-Eigenschaft auf `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)` setzen.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

Your `body` element should have the coffee beans background image.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

