---
id: 5f356ed69db0a491745e2bb6
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Füge, um den Stil der Klasse auf das `div`-Element anzuwenden, ein `class`-Attribut zum Öffnungstag des `div`-Elements hinzu und setze seinen Wert auf `menu`.

# --hints--

Dein `div` sollte weiterhin wiedergegeben werden. Stelle sicher, dass du das `<div>`-Tag nicht falsch gebildet hast.

```js
assert($('div').length === 1);
```

Dein `div`-Element sollte die `menu`-Klasse besitzen.

```js
assert($('div').attr('class').includes('menu'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div>
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

