---
id: 5f356ed6199b0cdef1d2be8f
title: Schritt 26
challengeType: 0
dashedName: step-26
---

# --description--

Bisher hast du Typ- und ID-Selektoren für verschiedene Style-Elemente verwendet. Allerdings ist es üblicher einen anderen Selektor zu verwenden, um dein Element zu gestalten.

Ein <dfn>Klassenselektor</dfn> ist durch einen Namen mit einem Punkt definiert, der sich direkt vor ihm befindet, wie folgendes Beispiel:

```css
.class-name {
  styles
}
```

Ändere den vorhandenen `#menu`-Selektor in einen Klassenselektor, indem `#menu` durch eine Klasse namens `.menu` ersetzt wird.

# --hints--

Du solltest einen `.menu`-Klassenselektor haben.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

Du solltest keinen `#menu`-Selektor haben.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

