---
id: 5f3c866daec9a49519871816
title: Schritt 30
challengeType: 0
dashedName: step-30
---

# --description--

`article`-Elemente enthalten in der Regel mehrere Elemente mit verwandten Informationen. In diesem Fall wird es eine Kaffeesorte und einen Preis für dieses Aroma enthalten. Füge zwei `p`-Elemente deinem `article`-Element hinzu. Der erste Text sollte `French Vanilla` sein, der zweite `3.00`.

# --hints--

Du solltest das existierende Element `article` nicht verändern.

```js
assert($('article').length === 1);
```

Dein `article`-Element sollte zwei `p`-Elemente enthalten.

```js
assert($('article').children('p').length === 2);
```

Dein erstes `p`-Element sollte den Text `French Vanilla` enthalten.

```js
const firstP = $('article').children('p')[0];
assert(firstP.innerText.match(/French Vanilla/i));
```

Dein zweites `p`-Element sollte den Text `3.00` haben.

```js
const secondP = $('article').children('p')[1];
assert(secondP.innerText.match(/3\.00/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
