---
id: 5f3c866dbf362f99b9a0c6d0
title: Schritt 37
challengeType: 0
dashedName: step-37
---

# --description--

Die `p`-Elemente sind innerhalb eines `article`-Elements mit dem Klassenattribut `item` verschachtelt. Du kannst all die `p`-Elemente formatieren, die irgendwo in ein Element mit einer Klasse namens `item` verschachtelt sind, und zwar wie folgt:

```css
.item p { }
```

Verwende den obigen Selektor, füge eine `display`-Eigenschaft mit dem Wert `inline-block` hinzu, so dass die `p`-Elemente sich mehr wie Inline-Elemente verhalten.

# --hints--

Du solltest den `.item p` Selektor verwenden.

```js
const hasItemP = new __helpers.CSSHelp(document).getStyle('.item p');
assert(hasItemP);
```

Du solltest die `display`-Eigenschaft auf `inline-block` setzen.

```js
const hasDisplay = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.display === 'inline-block');
assert(hasDisplay);
```

Dein `.item p`-Selektor sollte die `display`-Eigenschaft auf `inline-block` setzen.

```js
const itemPDisplay = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('display');
assert(itemPDisplay === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

--fcc-editable-region--

--fcc-editable-region--

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
