---
id: 5f3c866d0fc037f7311b4ac8
title: Schritt 38
challengeType: 0
dashedName: step-38
---

# --description--

Das ist näher dran, aber der Preis ist nicht auf der rechten Seite geblieben. Das liegt daran, dass `inline-block`-Elemente nur die Breite ihres Inhalts annehmen. Füge, um sie etwas breiter zu machen, eine `width`-Eigenschaft den Klassenselektoren `flavor` und `price` hinzu, die jeweils einen Wert von `50%` haben.

# --hints--

Du solltest die `width`-Eigenschaft auf `50%` in deinem `.flavor`-Selektor setzen.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert(flavorWidth === '50%');
```

Du solltest die `width`-Eigenschaft auf `50%` in deinem `.price`-Selektor setzen.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert(priceWidth === '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
--fcc-editable-region--
```

