---
id: 5f3c866de7a5b784048f94b1
title: Schritt 36
challengeType: 0
dashedName: step-36
---

# --description--

So in etwa ist es das, was du möchtest, aber jetzt wäre es noch schön, wenn die Sorte und der Preis sich in der gleichen Zeile befinden würden. `p`-Elemente sind <dfn>block-level</dfn>-Elemente, sie nehmen also die gesamte Breite des übergeordneten Elements in Anspruch.

Um sie auf die gleiche Zeile zu bringen, musst du die `p`-Elemente ein wenig gestalten, damit sie sich mehr wie <dfn>inline</dfn>-Elemente verhalten. Füge für diesen Zweck ein `class`-Attribut mit dem Wert `item` zum ersten `article`-Element unter der `Coffee`-Überschrift hinzu.

# --hints--

Du solltest die `item`-Klasse auf dein `article`-Element übertragen.

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

Du solltest nur ein `item`-Klassen-Element haben.

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

Dein erstes `article`-Element sollte die `item`-Klasse enthalten.

```js
const el = document.querySelectorAll('article')[0];

assert.equal(el.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
