---
id: 5f3c866d5414453fc2d7b480
title: Schritt 32
challengeType: 0
dashedName: step-32
---

# --description--

Füge unterhalb des existierenden Paares Kaffee/Preis folgende Kaffee- und Preisangaben mithilfe eines `article`-Elements mit zwei verschachtelten `p`-Elementen hinzu. Wie zuvor sollte der Text des ersten `p`-Elements die Kaffeesorte und des zweiten `p`-Elements die Preise enthalten.

```bash
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

# --hints--

Du solltest fünf `article`-Elemente haben.

```js
assert($('article').length === 5);
```

Jedes `article`-Element sollte zwei `p`-Elemente haben.

```js
const articles = $('article');
assert(articles[0].children.length === 2);
assert(articles[1].children.length === 2);
assert(articles[2].children.length === 2);
assert(articles[3].children.length === 2);
assert(articles[4].children.length === 2);
```

Dein erstes `article`-Element sollte `p`-Elemente mit den Texten `French Vanilla` und `3.00` enthalten.

```js
const children = $('article')[0].children;
assert(children[0].innerText.match(/French Vanilla/i));
assert(children[1].innerText.match(/3\.00/i));
```

Dein zweites `article`-Element sollte `p`-Elemente mit den Texten `Caramel Macchiato` und `3.75` enthalten.

```js
const children = $('article')[1].children;
assert(children[0].innerText.match(/Caramel Macchiato/i));
assert(children[1].innerText.match(/3\.75/i));
```

Dein drittes `article`-Element sollte `p`-Elemente mit den Texten `Pumpkin Spice` und `3.50` enthalten.

```js
const children = $('article')[2].children;
assert(children[0].innerText.match(/Pumpkin Spice/i));
assert(children[1].innerText.match(/3\.50/i));
```

Dein viertes `article`-Element sollte `p`-Elemente mit den Texten `Hazelnut` und `4.00` enthalten.

```js
const children = $('article')[3].children;
assert(children[0].innerText.match(/Hazelnut/i));
assert(children[1].innerText.match(/4\.00/i));
```

Dein fünftes `article`-Element sollte `p`-Elemente mit den Texten `Mocha` und `4.50` enthalten.

```js
const children = $('article')[4].children;
assert(children[0].innerText.match(/Mocha/i));
assert(children[1].innerText.match(/4\.50/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <header>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
      </header>
      <main>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
