---
id: 5f716bee5838c354c728a7c5
title: Schritt 52
challengeType: 0
dashedName: step-52
---

# --description--

Unterhalb des Desserts, das du gerade hinzugefügt hast, fügst du die restlichen Desserts und Preise mit drei weiteren `article`-Elementen hinzu, jedes mit zwei verschachtelten `p`-Elementen. Jedes Element sollte den richtigen Dessert- und Preistext haben und alle Elemente sollten die richtigen Klassen haben.

```bash
Cherry Pie 2.75
Cheesecake 3.00
Cinnamon Roll 2.50
```

# --hints--

Du solltest vier `article` Elemente in deiner zweiten `section` haben.

```js
const secondSection = document.querySelectorAll('section')[1]
const articles = secondSection.querySelectorAll('article')
assert(articles.length === 4)
```

Du solltest vier `.dessert` Elemente haben.

```js
assert($('.dessert').length === 4);
```

Du solltest vier neue `.price` Elemente haben.

```js
assert($('section').last().find('.price').length === 4);
```

Dein `section` Element sollte acht `p` Elemente haben.

```js
assert($('section').last().find('p').length === 8);
```

Deine `.dessert`-Elemente sollten den Text `Donut`, `Cherry Pie`, `Cheesecake` und `Cinnamon Roll` haben.

```js
const dessert = $('.dessert');
assert(dessert[0].innerText.match(/donut/i));
assert(dessert[1].innerText.match(/cherry pie/i));
assert(dessert[2].innerText.match(/cheesecake/i));
assert(dessert[3].innerText.match(/cinnamon roll/i));
```

Deine neuen `.price`-Elemente sollten den Text `1.50`, `2.75`, `3.00` und `2.50` haben.

```js
const prices = $('section').last().find('.price');
assert(prices[0].innerText.match(/1\.50/));
assert(prices[1].innerText.match(/2\.75/));
assert(prices[2].innerText.match(/3\.00/));
assert(prices[3].innerText.match(/2\.50/));
```

Du solltest keine Leerzeichen zwischen deinen `p` Elementen haben.

```js
assert(!code.match(/<\/p>\s+<p/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
--fcc-editable-region--
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```
