---
id: 66a9c14e3b34c719e34bae20
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

To add student data to the table, you will need to use the `table row` and <dfn>table data</dfn> elements.

The `table data` element, `<td>`, is used to create a cell in the table.

Here is an example of using the `td` element for a sports players table:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Inside your `table body` element, add a `table row` element, `<tr>`, with three `table data` elements, `<td>`. 

The first `table data` element should contain the last name of `Davis`. 

The second `table data` element should contain the first name of `Alex`. 

The third `table data` element should contain the grade of `54`.

# --hints--

You should have an opening `tr` tag.

```js
assert.match(code, /<tr>/i);
```

You should have a closing `tr` tag.

```js
assert.match(code, /<\/tr>/i);
```

Your `tr` element should be inside your `tbody` element.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

You should have a `td` element with the text `Davis` inside the `tr` element.

```js
assert.strictEqual(document.querySelector('tbody tr td').textContent, 'Davis');
```

You should have a `td` element with the text `Alex` inside the `tr` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1].textContent, 'Alex');
```

You should have a `td` element with the text `54` inside the `tr` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2].textContent, '54');
```

You should have all three of the `td` elements inside the `tr` element.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        --fcc-editable-region--

        --fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
