---
id: 66a9bcf00f13a418368a272e
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

The `table head` element consists of a `table row` element, `<tr>`, which contains the `table header` cell elements, `<th>`.

Here is an example using the `tr` and `th` elements for a sports table:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Inside your `thead` element, add a `tr` element.

Inside your `tr` element, add three `th` elements. 

The first `th` element should contain the text `Last Name`. The second `th` element should contain the text `First Name`. The third `th` element should contain the text `Grade`.

# --hints--

You should have a `tr` element inside your `thead` element.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl.parentElement.tagName, 'THEAD');
```

You should have a `tr` element with the text of `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl.textContent, 'Last Name');
```

You should have a `tr` element with the text of `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl.textContent, 'First Name');
```

You should have a `tr` element with the text of `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl.textContent, 'Grade');
```

You should have a total of three `th` elements inside your `tr` element.

```js
const thEls = document.querySelectorAll('th');
assert.strictEqual(thEls.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      --fcc-editable-region--
      <thead>     
        
      </thead>
      --fcc-editable-region--
    </table>
  </body>
</html>
```
