---
id: 66a9c006e4f58819396605b3
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now that you have completed the head section, it is time to add the `table body`. The `table body` will represent all of the student names and their grades.

Add a `table body` element, `<tbody>`, to your table.

# --hints--

Your table should have a `tbody` element.

```js
assert(document.querySelector('tbody'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      --fcc-editable-region--

      --fcc-editable-region--
    </table>
  </body>
</html>
```
