---
id: 6a1c40f36755ecfb7342d18b
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

Welcome to the Library Book Search workshop! In this project, you will build a book title search by handling errors gracefully using `try`, `catch`, `throw`, and `finally`.

You have been provided with an array of available books called `library`, as well as some basic DOM variables to handle user input. 

If you look at the bottom of the `script.js` file, there is an event listener attached to the `searchBtn`. When clicked, it expects to call a function named `findBook`, passing in the user's input.

To get started, declare an empty function named `findBook`. This function should take a single parameter named `title`.

# --hints--

You should declare an empty function named `findBook` that takes `title` as a parameter. You can use a standard function, an arrow function, or a function expression.

```js
const explorer = await __helpers.Explorer(code);
const { findBook } = explorer.allFunctions;
assert.exists(findBook);
assert.strictEqual(findBook.parameters.length, 1);
assert.strictEqual(findBook.parameters[0].toString(), "title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Find book title</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter book title">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
}

input {
  padding: 0.5rem;
  width: 200px;
}

button {
  padding: 0.5rem 1rem;
  cursor: pointer;
  background-color: #000000;
  color: white;
}

#statusMessage {
  margin-top: 1.5rem;
  font-weight: bold;
  height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById('search');
const statusDisplay = document.getElementById('statusMessage');
const bookInput = document.getElementById('bookTitleInput');

--fcc-editable-region--

--fcc-editable-region--
searchBtn.addEventListener('click', () => {
  const result = findBook(bookInput.value.trim());
  statusDisplay.textContent = result;
});
```
