---
id: 6a1c4a5b9ecacf56d79ee405
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

You have successfully handled both the success and error states. But what if there is code you want to run regardless of what happens? That is where the `finally` block comes in.

Directly below your `catch` block, add a `finally` block. Inside it, set the `value` property of `bookInput` to an empty string (`""`) to clear the input field so the user can immediately type a new search.

# --hints--

You should add a `finally` block directly after your `catch` block, and set `bookInput.value` to an empty string inside it.

```js
document.getElementById('bookTitleInput').value = "The Hobbit";
document.getElementById('search').click();
assert.strictEqual(document.getElementById('bookTitleInput').value, "");

document.getElementById('bookTitleInput').value = "Test Search";
document.getElementById('search').click();
assert.strictEqual(document.getElementById('bookTitleInput').value, "");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Library Book Search</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Find book title</h1>
      <input type="text" id="bookTitleInput" placeholder="Enter book title">
      <button id="search">Search</button>
      <p id="statusMessage"></p>
      <p id="scanStatus"></p>
    </div>
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
}

input {
  padding: 0.5rem;
  width: 200px;
}

button {
  padding: 0.5rem 1rem;
  cursor: pointer;
  background-color: #000000;
  color: white;
}

#statusMessage {
  margin-top: 1.5rem;
  font-weight: bold;
  height: 1.5rem;
}
```

```js
const library = [
  "The Hobbit", "1984", "Dune", "Fahrenheit 451", "To Kill a Mockingbird",
  "The Great Gatsby", "Moby Dick", "Pride and Prejudice", "Frankenstein"
];

const searchBtn = document.getElementById("search");
const statusDisplay = document.getElementById("statusMessage");
const bookInput = document.getElementById("bookTitleInput");

const findBook = (title) => {
  if (title === "") {
    throw new Error("Search field cannot be empty");
  }

  if (!library.includes(title)) {
    throw new Error(title + " not found!");
  }

  return title + " is available to borrow!";
};

searchBtn.addEventListener("click", () => {
  try {
    const result = findBook(bookInput.value.trim());
    statusDisplay.textContent = result;
  } catch (err) {
    statusDisplay.textContent = err.message;
    console.error(err);
    debugger;
--fcc-editable-region--
  }
  
--fcc-editable-region--
});
```
