---
id: cf1111c1c11feddfaeb3bdef
title: Suma dos números con JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` (número) es un tipo de datos en JavaScript que representa datos numéricos.

Ahora intentemos sumar dos números usando JavaScript.

JavaScript utiliza el símbolo `+` como un operador de adición cuando se coloca entre dos números.

**Ejemplo:**

```js
const myVar = 5 + 10;
```

`myVar` ahora tiene el valor `15`.

# --instructions--

Cambia el `0` para que la suma sea igual a `20`.

# --hints--

`sum` debe ser igual a `20`.

```js
assert(sum === 20);
```

Debes usar el operador `+`.

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
