---
id: bd7993c9ca9feddfaeb7bdef
title: Divide un decimal entre otro con JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cBZe9AW'
forumTopicId: 18255
dashedName: divide-one-decimal-by-another-with-javascript
---

# --description--

Ahora dividamos un decimal entre otro.

# --instructions--

Cambia el `0.0` para que el cociente (`quotient`) sea igual a `2.2`.

# --hints--

La variable `quotient` debe ser igual a `2.2`

```js
assert(quotient === 2.2);
```

Debes usar el operador `/` para dividir 4.4 entre 2

```js
assert(/4\.40*\s*\/\s*2\.*0*/.test(code));
```

La variable quotient debe ser asignada solo una vez

```js
assert(code.match(/quotient/g).length === 1);
```

# --seed--

## --seed-contents--

```js
const quotient = 0.0 / 2.0; // Change this line
```

# --solutions--

```js
const quotient = 4.4 / 2.0;
```
