---
id: cf1231c1c11feddfaeb5bdef
title: Multiplica dos números con JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

También podemos multiplicar un número por otro.

JavaScript utiliza el símbolo `*` para multiplicar dos números.

**Ejemplo**

```js
const myVar = 13 * 13;
```

`myVar` ahora tendrá el valor `169`.

# --instructions--

Cambia el `0` para que el producto sea igual a `80`.

# --hints--

La variable `product` debe ser igual a 80.

```js
assert(product === 80);
```

Debes usar el operador `*`.

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'product = '+z;})(product);
```

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
