---
id: 61fd6ab779390f49148773bb
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Erstelle unter deinem `h1`-Element ein `div`-Element. Gib ihm ein `id`-Attribut mit dem Wert `years`. Du möchtest, dass dieses bestimmte Element von den Screenreader versteckt wird. Gib ihm also das `aria-hidden`-Attribut mit dem Wert `true`.

# --hints--

Du solltest ein neues `div`-Element nach deinem `h1`-Element erstellen.

```js
assert(document.querySelector('h1')?.nextElementSibling?.localName === 'div');
```

Dein neues `div`-Element sollte ein `id`-Attribut mit dem Wert `years` haben.

```js
assert(document.querySelector('div')?.getAttribute('id') === 'years');
```

Dein neues `div`-Element sollte das `aria-hidden`-Attribut mit dem Wert `true` haben.

```js
assert(document.querySelector('div')?.getAttribute('aria-hidden') === 'true');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
