---
id: 61fd719788899952e67692b9
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Das `thead`- und `tbody`-Element werden verwendet, um anzugeben, welcher Teil deiner Tabelle der Header ist und welcher Teil die Primärdaten oder den Inhalt enthält.

Füge ein `thead` und `tbody` zu deinem ersten `table` unter dem `caption`-Element hinzu.

# --hints--

Dein erstes `table`-Element sollte ein `thead`-Element enthalten.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Dein erstes `table`-Element sollte ein `tbody`-Element enthalten.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

Dein `thead`-Element sollte sich direkt unter deinem `caption`-Element befinden.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

Dein `tbody`-Element sollte sich direkt unter deinem `thead`-Element befinden.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
