---
id: 61fd6343fb42013d99bcd7f3
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Bette innerhalb deines `body`-Elements ein `section`-Element innerhalb eines `main`-Element ein.

# --hints--

Das erste Element innerhalb des `body`-Elements sollte dein `main`-Element sein.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

Du solltest ein `section`-Element innerhalb deines `main`-Elements haben.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
