---
id: 61fd778081276b59d59abad6
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Füge in deinem ersten `tr` ein `th`-Element mit dem Text `Cash This is the cash we currently have on hand.` hinzu. Umbreche den ganzen Text, außer `Cash`, in ein `span`-Element mit der `class` gesetzt auf `description`.

Füge demzufolge drei `td`-Elemente mit dem folgenden Text (in Reihenfolge) hinzu: `$25`, `$30`, `$28`. Gib dem dritten `td`-Element ein `class`-Attribut mit dem Wert `current`.

# --hints--

Dein erstes `tr` sollte ein `th`-Element haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th'));
```

Dein `th`-Element sollte den Text `Cash This is the cash we currently have on hand.` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th')?.innerText === 'Cash This is the cash we currently have on hand.');
```

Du solltest den Text `This is the cash we currently have on hand.` in einem `span`-Element umschließen.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.textContent === 'This is the cash we currently have on hand.');
```

Dein `span`-Element sollte das `class`-Attribut auf `description` gesetzt haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.classList?.contains('description'));
```

Du solltest drei `td`-Elemente haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td').length === 3);
```

Dein erstes `td`-Element sollte den Text `$25` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[0]?.textContent === '$25');
```

Dein zweites `td`-Element sollte den Text `$30` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[1]?.textContent === '$30');
```

Dein drittes `td`-Element sollte den Text `$28` haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.textContent === '$28');
```

Dein drittes `td`-Element sollte die `class` auf `current` gesetzt haben.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
