---
id: 61fd75ea7f663457612dba02
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Bette innerhalb jedes deiner neuen `th`-Elemente ein `span`-Element mit der `class` `sr-only year` ein. Weise ihnen den folgenden Text zu (in Reihenfolge): `2019`, `2020`, und `2021`.

Weise deinem dritten `th`-Element das `class`-Attribut mit dem Wert `current` zu.

Lasse das `td`-Element leer. Dieses Element existiert nur, um sicherzustellen, dass deine Tabelle ein vierspaltiges Layout hat, und um die Kopfzeilen mit den richtigen Spalten zu verknüpfen.

# --hints--

Jedes deiner `th`-Elemente sollte ein `span`-Element enthalten.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

Jedes deiner neuen `span`-Elemente sollte das `class`-Attribut auf `sr-only year` gesetzt haben.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

Dein erstes `span`-Element sollte den Text `2019` enthalten.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

Dein zweites `span`-Element sollte den Text `2020` enthalten.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

Dein drittes `span`-Element sollte den Text `2021` enthalten.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

Dein drittes `th`-Element sollte die `class` auf `current` gesetzt haben.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

Dein `td`-Element sollte leer sein.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
