---
id: 61fd6cc9475a784b7776233e
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

F체ge unter deinem `div`-Element ein neues `div`-Element mit einer `class` von `table-wrap` hinzu. Dies wird der Container f체r deine Tabellen sein.

F체ge in dieses drei `table`-Elemente hinzu. Du wirst CSS verwenden, um diese in einer einzelnen Tabelle zu gestalten, jedoch erleichtert die Verwendung von separaten Tabellen Screenreadern das Verst채ndnis des Dokumentenflusses.

# --hints--

Du solltest ein neues `div`-Element erstellen.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Dein neues `div`-Element sollte die `class` auf `table-wrap` gesetzt haben.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

Dein `.table-wrap`-Element sollte nach deinem vorhandenen `div` stehen.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

Dein `.table-wrap`-Element sollte drei `table`-Elemente enthalten.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
