---
id: 61fd70336ebb3e4f62ee81ba
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML-Tabellen verwenden das `caption`-Element, um zu beschreiben, worum es in der Tabelle geht. Das `caption`-Element sollte immer das erste untergeordnete Element von einem `table` sein, aber es kann mit der `caption-side`-CSS-Eigenschaft positioniert werden.

Füge ein `caption`-Element zu deinem ersten `table` hinzu und gib ihm den Text `Assets`.

# --hints--

Dein erstes `table`-Element sollte ein `caption`-Element haben.

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

Dein `caption`-Element sollte den Text `Assets` haben.

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
