---
id: 61fd5a93fd62bb35968adeab
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Richte deinen HTML-Code mit der `DOCTYPE`-Deklaration, dem `html`-Element, das angibt, dass dieses Dokument eng­lisch­spra­chig ist, mit dem `head`-Element sowie mit dem `body`-Element ein.

Gib deinem `head`-Element die passenden `meta`-Elemente für das `charset` und den `viewport`, ein `title`-Element mit einem passenden Titel und ein `link`-Element für dein Stylesheet.

# --hints--

Dein Code sollte eine `<!DOCTYPE html>`-Deklaration enthalten.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Dein Code sollte ein `html`-Element enthalten.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Dein `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut mit dem Wert `en` enthalten.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dein `html`-Element sollte ein schließendes Tag enthalten.

```js
assert(code.match(/<\/html\s*>/));
```

Dein Code sollte ein `head`-Element innerhalb des `html`-Elements enthalten.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Dein Code sollte ein `body`-Element innerhalb des `html`-Elements enthalten.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Dein `head`-Element sollte vor deinem `body`-Element stehen.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Du solltest zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Ein `meta`-Element sollte `name` auf `viewport` und `content` auf `width=device-width, initial-scale=1.0` eingestellt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset`-Attribut auf `UTF-8` eingestellt haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Dein Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein `title` sollte etwas Text enthalten.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Dein Code sollte ein `link`-Element enthalten.

```js
assert(/<link/.test(code))
```

Dein `link`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

Dein `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` enthalten.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` enthalten.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
