---
id: 6626a060c4006f793e10cb33
title: 步骤 2
challengeType: 1
dashedName: step-2
---

# --description--

现在，老师需要你帮助将学生们的分数转换成一个字母评级。

完善 `getGrade` 函数， 这个函数以一个数字类型的 `score` 作为参数。 你的函数应该根据分数返回一个字符串，代表字母评级。

这里是分数段和它们所对应的字母评级：

| 分数段       | 字母评级    |
| --------- | ------- |
| `100`     | `"A++"` |
| `90 - 99` | `"A"`   |
| `80 - 89` | `"B"`   |
| `70 - 79` | `"C"`   |
| `60 - 69` | `"D"`   |
| `0 - 59`  | `"F"`   |


**提示：**

- 记得你学习过一些条件语句（`if`、`else if`、`else`）。
- 记得你学习过一些比较运算符（`>`、`<`、`>=`、`<=`、`===`）。

# --hints--

如果分数是 `100`，你的 `getGrade` 函数应该返回 `"A++"`。

```js
assert.strictEqual(getGrade(100), "A++");
```

如果分数是 `94`，你的 `getGrade` 函数应该返回 `"A"`。

```js
assert.strictEqual(getGrade(94), "A");
```

如果分数在 `80` 到 `89` 之间，你的 `getGrade` 函数应该返回 `"B"`。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

如果分数是 `78`，你的 `getGrade` 函数应该返回 `"C"`。

```js
assert.strictEqual(getGrade(75), "C");
```

如果分数在 `60` 到 `69` 之间，你的 `getGrade` 函数应该返回 `"D"`。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

如果分数是 `57`，你的 `getGrade` 函数应该返回 `"F"`。

```js
assert.strictEqual(getGrade(57), "F");
```

如果分数在 `0` 到 `59` 之间，你的 `getGrade` 函数应该返回 `"F"`。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
