---
id: 662693f82c91a66be46c881b
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

一位老师已经完成了学生的考试评级，需要你帮忙计算全班的平均分。

完成 `getAverage` 函数，该函数接收考试分数数组并返回平均分数。

平均分是通过累加所有得分并除以得分的数量来计算的。

```js
average = sum of all scores / total number of scores
```

这里有一些写好的函数调用的例子，便于你来测试你自己的代码。

**提示：**

- 你可以使用循环来累加 `scores` 数组以计算总分。
- 你可以使用 `length` 属性来获取得分的数量。

# --hints--

你的 `getAverage` 函数应该返回一个数字。

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` 应该返回 `71.7`。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` 应该返回 `85.4`。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` 应该返回 `92.4`。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

你的 `getAverage` 函数应该返回考试分数的平均分。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
