---
id: 6551f3bfed58796ebb1268dc
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Each string character can be referenced by a numerical index. The index count starts at zero. So the first character of a string has an index of `0`. For example, in the string `"Hello World"`, "H" is at index `0`, "e" is at index `1`, and so on. Each character of a string can be accessed by using bracket notation:

```python
my_string = 'Hello World'
r = my_string[8]
```

Now, instead of printing `text`, print just the character at index `6`.

# --hints--

You should still call the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(.*\)/) })

```

You should pass `text[6]` to your `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
