---
id: 6551f4259d66026ee28cb205
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

You can also access string characters starting from the end of the string. The last character has an index of `-1`, the second to last `-2` and so on. Now modify your existing `print()` call to print the last character in your string.

# --hints--

You should still call the `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(.*\)/) })

```

You should pass `text[-1]` to your `print()` function.

```js
({ test: () => assert.match(code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
