---
id: 6551f1e8737dbe6e0cb08ea4
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

And now print your new variable.

# --hints--

You should call the `print()` function passing in the `shift` variable.

```js
({ test: () => assert.match(code, /print\s*\(\s*shift\s*\)/) })

```

You should not remove your existing `print()` call.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/) })

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
