---
id: 65524b3aa6a1938a069a91ab
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Next, print `new_char` and see the output.

# --hints--

You should print your `new_char` variable.

```js
({ test: () => assert.match(code, /print\s*\(\s*new_char\s*\)/) })
```

You should print your `new_char` variable at the end of your loop body.

```js
({ test: () => assert.match(code, /:(.*?)index.*\1print\s*\(\s*new_char\s*\)/s) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
    new_char = alphabet[new_index]
--fcc-editable-region--
```
