---
id: 655a2a7210094920069b117c
title: Schritt 37
challengeType: 20
dashedName: step-37
---

# --description--

Comparison operators allow you to compare two objects based on their values. You can use a comparison operator by placing it between the objects you want to compare. They return a *Boolean* value — namely `True` or `False` — depending on the truthfulness of the expression.

Python has the following comparison operators:

<table>
  <thead>
    <tr>
      <th>Operator</th>
      <th>Meaning</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Equal</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Not equal</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Greater than</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Less than</td>
    </tr>
    <tr>
      <td>&ge;</td>
      <td>Greater than or equal to</td>
    </tr>
    <tr>
      <td>&le;</td>
      <td>Less than or equal to</td>
    </tr>
  </tbody>
</table>

At the top of your loop, print the result of comparing `char` with an empty space. Use the equality operator `==` for that.

# --hints--

You should compare `char` with an empty space using the equality operator.

```js
({
  test: () => assert.match(code, /char\s*==\s*("|')\s\1/)
})
```

You should print the result of comparing `char` with an empty space.

```js
({
  test: () => assert.match(code, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/)
})
```

You should print the result of comparing `char` with an empty space at the beginning of your loop.

```js
({
  test: () => assert.match(code, /:\n\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
