---
id: 6568c86dc9193000d11ca5e0
title: Schritt 35
challengeType: 20
dashedName: step-35
---

# --description--

Instead of assigning `alphabet[new_index]` to `encrypted_text`, assign the current value of `encrypted_text` plus `alphabet[new_index]` to this variable.

# --hints--

You should assign the value of `encrypted_text + alphabet[new_index]` to your `encrypted_text` variable.

```js
({ test: () => assert.match(code, /encrypted_text\s*=\s*encrypted_text\s*\+\s*alphabet\s*\[\s*new_index\s*\]/s) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
