---
id: 61fd71d596e8f253b9408b39
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

يتم استخدام عنصر `tr` للإشارة إلى صف الجدول (table row). اضف عنصر `tr` ضمن عنصر `thead`. في عنصر `tr` الجديد، أضف عنصر `td` يليه ثلاث عناصر `th`.

عنصر `td` يشير إلى خانة بيانات (data cell)، في حين أن عنصر `th` يشير إلى خانة العناوين (header cell).

# --hints--

يجب أن يحتوي عنصر `thead` الخاص بك على عنصر `tr`.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

عنصر `tr` الخاص بك يجب أن يحتوي على عنصر `td` كالطفل الأول (first child).

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

عنصر `tr` الخاص بك يجب أن يحتوي على ثلاث عناصر `th` بعد عنصر `td`.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
