---
id: 61fd6b7c83dbf54a08cf0498
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

ضمن عنصرك `div`، أضف ثلاث عناصر `span`. اعطي كل منهم سمة `class` تم تعيينها إلى `year`، وأضف النصوص التالية (بالترتيب): `2019`و `2020` و `2021`.

# --hints--

يجب أن يحتوي عنصرك `div` على ثلاث عناصر `span`.

```js
assert(document.querySelector('div')?.children?.length === 3);
```

كل عنصر `span` يجب أن يحتوي على `class` تم تعيينه إلى `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

يجب أن يحتوي أول `span` على النص `2019`.

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

يجب أن يحتوي ثاني `span` على النص `2020`.

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

يجب أن يحتوي ثالث `span` على النص `2021`.

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
