---
id: 61fd6343fb42013d99bcd7f3
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل عنصر `body`، قم بدمج عنصر `section` ضمن عنصر `main`.

# --hints--

العنصر الأول داخل عنصر `body` يجب أن يكون عنصر `main` الخاص بك.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

يجب أن يكون لديك عنصر `section` ضمن عنصر `main` الخاص بك.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
