---
id: 61fd778081276b59d59abad6
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

في أول `tr` خاص بك، أضف عنصر `th` مع النص `Cash This is the cash we currently have on hand.`. غلف كل هذه النصوص باستثناء `Cash` في عنصر `span` مع `class` محدد إلى `description`.

بعد ذلك، أضف ثلاثة عناصر `td` مع النصوص التالية (بالترتيب): `$25` و `$30` و `$28`. اعطي عنصر `td` الثالث السمة `class` التي تم تعيينها إلى `current`.

# --hints--

يجب أن يحتوي أول `tr` على عنصر `th`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th'));
```

عنصر `th` الخاص بك يجب أن يحتوي على النص `Cash This is the cash we currently have on hand.`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th')?.innerText === 'Cash This is the cash we currently have on hand.');
```

يجب أن تغلف النص `This is the cash we currently have on hand.` في عنصر `span`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.textContent === 'This is the cash we currently have on hand.');
```

يجب أن يحتوي عنصر `span` على سمة `class` بقيمة `description`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelector('th > span')?.classList?.contains('description'));
```

يجب أن يكون لديك ثلاثة عناصر `td`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td').length === 3);
```

يجب أن يحتوي العنصر `td` الاول على النص `$25`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[0]?.textContent === '$25');
```

يجب أن يحتوي العنصر `td` الثاني على النص `$30`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[1]?.textContent === '$30');
```

يجب أن يحتوي العنصر `td` الثالث على النص `$28`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.textContent === '$28');
```

عنصر `td` الثالث الخاص بك يجب أن يحتوي على `class` بقيمة `current`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[0]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
