---
id: 61fd70336ebb3e4f62ee81ba
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

تستخدم HTML tables عنصر `caption` لوصف موضوع الـ table. يجب أن يكون عنصر `caption` دائما أول child لـ `table`، ولكن يمكن أن يتم التحكم بموضعه باستخدام `caption-side` وهي خاصية CSS.

أضف عنصر `caption` إلى أول `table` خاص بك، وأعطيه النص `Assets`.

# --hints--

يجب أن يحتوي عنصر `table` الاول الخاص بك على عنصر `caption`.

```js
assert(document.querySelector('table')?.children?.[0]?.localName === 'caption');
```

يجب أن يحتوي عنصر `caption` على النص `Assets`.

```js
assert(document.querySelector('caption')?.textContent === 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
