---
id: 61fd5a93fd62bb35968adeab
title: الخطوة 1
challengeType: 0
dashedName: step-1
---

# --description--

حدد نوع HTML الخاص بك باستخدام العناصر الآتية `DOCTYPE`, و `html` لتحديد لغة الصفحة لإنجليزية, و `head`, و `body`.

اعطي عنصر `head` عناصر `meta` المناسبة مثل `charset`, و `viewport`, و عنصر `title` يحتوي على عنوان مناسب، و `link` للـصفحتك للأسلوب (stylesheet).

# --hints--

يجب أن يحتوي الكود على إعلان `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يحتوي كودك على عنصر `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

يجب أن يحتوي عنصرك `html` على علامة افتتاحية (opening tag) مع خاصية `lang` من `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصرك `html` على وسم الأغلاق (closing tag).

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يحتوي كودك على عنصر `head` داخل عنصر `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

يجب أن يحتوي كودك على عنصر `body` داخل عنصر `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

يجب أن يأتي عنصرك `head` قبل عنصرك `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

يجب أن يحتوي أحد عناصر `meta` على `name` بقيمة `viewport`، و `content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي عنصر `meta` الآخر على سمة `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

يجب أن يحتوي كودك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أحتواء `title` إلى نص.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

يجب أن يحتوي كودك على عنصر `link`.

```js
assert(/<link/.test(code))
```

يجب أن يكون عنصرك `link` داخل عنصرك `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

يجب أن يحتوي عنصرك `link` على سمة `rel` بقيمة `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

يجب أن يحتوي عنصرك `link` على سمة `href` بقيمة `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
