---
id: 61fd75ea7f663457612dba02
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

داخل كل عنصر من عناصر `th` الجديدة، أدخل عنصر `span` مع `class` تم تعينه إلى `sr-only year`. أعطيهم النصوص التالية (بالترتيب): `2019`، و `2020`، و `2021`.

اعطي عنصر `th` الثالث، السمة `class` باسم `current`.

اترك عنصر `td` فارغاً. هذا العنصر موجود فقط للتأكد من أن جدولك لديه تخطيط (layout) من أربع أعمدة وربط العناوين مع الأعمدة الصحيحة.

# --hints--

يجب أن يحتوا كل من عناصرك `th` على عنصر `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

يجب أن يحتوي كل عنصر من عناصر `span` الجديدة على `class` باسم `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

يجب أن يحتوي العنصر `span` الأول على النص `2019`.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

يجب أن يحتوي العنصر `span` الثاني على النص `2020`.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

يجب أن يحتوي العنصر `span` الثالث على النص `2021`.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

يجب أن يحتوي عنصرك `th` على `class` بقيمة `current`.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

يجب أن يكون عنصر `td` فارغاً.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
