---
id: 61fd77f7ad2aeb5ae34d07d6
title: الخطوة 14
challengeType: 0
dashedName: step-14
---

# --description--

في عنصر `tr` الثاني، أضف عنصر `th` مع النص `Checking Our primary transactional account.`. غلف ذلك النص، باستثناء `Checking`، في عنصر `span` مع سمة `class` بقيمة `description`.

بعد ذلك، أضف ثلاث عناصر `td` مع النصوص التالية (بالترتيب): `$54`, و `$56`, و `$53`. اعطي عنصر `td` الثالث سمة `class` تسمى `current`.

# --hints--

يجب أن يحتوي عنصر `tr` الثاني على عنصر `th`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th'));
```

يجب أن يحتوي عنصرك `th` على النص `Checking Our primary transactional account.`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th')?.innerText === 'Checking Our primary transactional account.');
```

يجب أن تغلف النص `Our primary transactional account.` في عنصر `span`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.textContent === 'Our primary transactional account.');
```

يجب أن يحتوي عنصر `span` على سمة `class` بقيمة `description`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelector('th > span')?.classList?.contains('description'));
```

يجب أن يكون لديك ثلاث عناصر `td`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td').length === 3);
```

يجب أن يحتوي العنصر `td` الاول على النص `$54`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[0]?.textContent === '$54');
```

يجب أن يحتوي العنصر `td` الثاني على النص `$56`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[1]?.textContent === '$56');
```

يجب أن يحتوي العنصر `td` الثالث على النص `$53`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.textContent === '$53');
```

يجب أن يحتوي عنصرك `td` الثالث على `class` بقيمة `current`.

```js
assert(document.querySelector('tbody')?.querySelectorAll('tr')?.[1]?.querySelectorAll('td')?.[2]?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
              <tr class="total">
              </tr>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
