---
id: 61fd6cc9475a784b7776233e
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

تحت عنصر `div` الموجود ، أضف عنصر `div` جديد مع `class` تم تعيينه إلى `table-wrap`. سيكون هذا هو الحاوية (container) للـ tables الخاصة بك.

ضمن ذلك، أضف ثلاثة عناصر `table`. سوف تستخدم CSS لتصميمهم في table واحد، ولكن استخدام tables منفصلة سيساعد قارئ الشاشة (screen readers) على فهم تدفق و تسلسل المستند (document).

# --hints--

يجب عليك إنشاء عنصر `div` جديد.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

عنصر `div` الجديد الخاص بك يجب أن يحتوي على `class` يسمي `table-wrap`.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

عنصر `.table-wrap` الخاص بك يجب أن يأتي بعد الـ `div` الحالي الخاص بك.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

عنصر `.table-wrap` الخاص بك يجب أن يحتوي على ثلاثة عناصر `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
