---
id: 61fd7648a7ba2e5882436831
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

ضمن عنصر `tbody` ، أضف أربعة عناصر `tr`. أعطِ الثلاث الأولى `class` بقيمة `data`، و الرابع `class` تم تعيينه إلى `total`.

# --hints--

يجب أن يحتوي عنصر `tbody` الخاص بك على أربعة عناصر `tr`.

```js
const children = [...document.querySelector('tbody')?.children];
assert(children?.length === 4);
children.forEach(child => assert(child?.localName === 'tr'));
```

يجب أن تحتوي العناصر الثلاثة `tr` على سمة `class` بقيمة `data`.

```js
const children = [...document.querySelector('tbody')?.children];
children.forEach((child, index) => {
  if (index < 3) {
    assert(child?.classList?.contains('data'));
  }
});
```

عنصر `tr` الرابع الخاص بك يجب أن يحتوي على سمة `class` بقيمة `total`.

```js
const children = [...document.querySelector('tbody')?.children];
assert(children?.[3]?.classList?.contains('total'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
--fcc-editable-region--
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
