---
id: 66faa20000000000000000f3
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Make the HUD react. After updating `this.scoreText`, tween it: `this.tweens.add({ targets: this.scoreText, scale: 1.4, duration: 100, yoyo: true })`. The HUD bumps up and back every time the player picks a coin.

# --hints--

You should tween the score text's scale to 1.4 with `yoyo: true`.

```js
assert.match(code, /targets:\s*this\.scoreText[\s\S]*?scale:\s*1\.4[\s\S]*?yoyo:\s*true/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.bg = this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12,
      repeat: -1
    });
    this.player.anims.play("idle");

    this.coins = this.physics.add.group();
    [
      [80, 80],
      [400, 80],
      [80, 280],
      [400, 280],
      [240, 80]
    ].forEach(([x, y]) => {
      const coin = this.coins.create(x, y, "coin");
      coin.setScale(2);
      coin.setDepth(10);
      coin.anims.play("coin-spin");
    });

    this.bg.setDepth(0);
    this.player.setDepth(10);
    this.physics.add.overlap(this.player, this.coins, this.collectCoin, null, this);

    this.ring = this.add.graphics();
    this.ring.setDepth(50);
    this.ring.setBlendMode(Phaser.BlendModes.ADD);
    this.ring.lineStyle(2, 0xffffff, 1);
    this.ring.strokeCircle(0, 0, 24);
    this.ring.setPosition(this.player.x, this.player.y);

    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Score: 0", {
      fontSize: "16px",
      fill: "#ffffff"
    });
    this.scoreText.setScrollFactor(0);
    this.scoreText.setDepth(100);

    this.uiCamera = this.cameras.add(0, 0, 480, 360);
    this.cameras.main.ignore(this.scoreText);
    this.uiCamera.ignore([this.bg, this.player, this.coins, this.ring]);
  }

  update() {
    this.player.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(false);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-right", true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    } else {
      this.player.anims.play("idle", true);
    }
    this.ring.setPosition(this.player.x, this.player.y);
  }

  collectCoin(player, coin) {
    coin.setTint(0xff3366);
    this.tweens.add({
      targets: coin,
      scale: 0,
      duration: 200,
      onComplete: () => coin.destroy()
    });
    const floater = this.add.text(coin.x, coin.y, "+10", {
      fontSize: "16px",
      fill: "#ffeeaa"
    });
    floater.setOrigin(0.5);
    floater.setDepth(50);
    this.tweens.add({
      targets: floater,
      y: floater.y - 30,
      alpha: 0,
      duration: 500,
      onComplete: () => floater.destroy()
    });
    this.tweens.add({
      targets: this.ring,
      scale: 1.5,
      duration: 200,
      yoyo: true
    });
    this.score += 10;
    this.scoreText.setText("Score: " + this.score);
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
