---
id: 66faa20000000000000000e4
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Define the hero animations and a coin-spin animation, then play `"idle"` on the player. The hero sheet has 16 frames laid out as four 4-frame rows: row 0 walk-down (`0..3`), row 1 walk-up (`4..7`), row 2 walk-left (`8..11`), row 3 walk-right (`12..15`). Use a single-frame `"idle"` mapped to frame 0. Walk animations play at `frameRate: 8`, `idle` at `6`, and `coin-spin` at `12`. All loop with `repeat: -1`.

The `"coin-spin"` animation iterates `coin` frames `0..7`. Finish the step by calling `this.player.anims.play("idle")`.

# --hints--

You should create animations keyed `"idle"`, `"walk-down"`, `"walk-up"`, `"walk-left"`, and `"walk-right"`.

```js
assert.match(code, /key:\s*["']idle["']/);
assert.match(code, /key:\s*["']walk-down["']/);
assert.match(code, /key:\s*["']walk-up["']/);
assert.match(code, /key:\s*["']walk-left["']/);
assert.match(code, /key:\s*["']walk-right["']/);
```

You should create a `"coin-spin"` animation that iterates `coin` frames `0..7`.

```js
assert.match(code, /key:\s*["']coin-spin["'][\s\S]*?frames:[\s\S]*?["']coin["'][\s\S]*?start:\s*0[\s\S]*?end:\s*7/);
```

You should play the `"idle"` animation on the player.

```js
assert.match(code, /this\.player\.anims\.play\(\s*["']idle["']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.bg = this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
