---
id: 66faa20000000000000000e1
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome! In this workshop you'll build a particle-pickup effect on top of the sprite-sheet hero from the previous workshop. The hero walks, coins recolor and shrink on pickup, a vector "+10" floats up and fades, a glowing ring pulses around the player, and a multi-camera HUD tracks your score. The HTML and CSS are wired — work in `script.js`.

Start by declaring a class named `MainScene` that extends `Phaser.Scene` and gives the scene the key `"MainScene"` in its constructor.

# --hints--

You should declare a class named `MainScene` that extends `Phaser.Scene`.

```js
assert.match(code, /class\s+MainScene\s+extends\s+Phaser\.Scene\s*\{/);
```

The constructor should call `super("MainScene")`.

```js
assert.match(code, /constructor\s*\(\s*\)\s*\{\s*super\(\s*["']MainScene["']\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
--fcc-editable-region--

--fcc-editable-region--
```
