---
id: 66faa20000000000000000f4
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Ship the win overlay. After the HUD pulse, check `if (this.score >= 50) { ... }` for the all-collected condition.

Inside the if-block, build a full-screen Graphics rectangle: `const overlay = this.add.graphics(); overlay.fillStyle(0x000000, 1); overlay.fillRect(0, 0, 480, 360); overlay.setDepth(200); overlay.setAlpha(0); overlay.setScrollFactor(0); this.cameras.main.ignore(overlay);`. Tween it in: `this.tweens.add({ targets: overlay, alpha: 0.6, duration: 500 })`.

Then add restart text: `const restartText = this.add.text(240, 180, "Press SPACE to restart", { fontSize: "18px", fill: "#ffffff" }); restartText.setOrigin(0.5); restartText.setDepth(201); restartText.setScrollFactor(0); this.cameras.main.ignore(restartText);`. Finally, listen for spacebar once and restart: `this.input.keyboard.once("keydown-SPACE", () => this.scene.restart())`.

Now play the game: walk around, collect five coins (50 points total), watch the HUD pulse, the ring glow, the floaters drift up. When the overlay fades in, press space — fresh run.

Congratulations — Chapter 2 ships.

# --hints--

You should gate the overlay on `this.score >= 50`.

```js
assert.match(code, /if\s*\(\s*this\.score\s*>=\s*50\s*\)/);
```

You should draw a full-screen graphics rectangle with depth 200, alpha 0, and `setScrollFactor(0)`.

```js
assert.match(code, /this\.add\.graphics\(\s*\)[\s\S]*?fillRect\(\s*0\s*,\s*0\s*,\s*480\s*,\s*360[\s\S]*?setDepth\(\s*200\s*\)[\s\S]*?setAlpha\(\s*0\s*\)[\s\S]*?setScrollFactor\(\s*0\s*\)/);
```

You should tween the overlay's alpha to 0.6 over 500 ms.

```js
assert.match(code, /alpha:\s*0\.6[\s\S]*?duration:\s*500/);
```

You should add a `Press SPACE to restart` text and listen for `keydown-SPACE` once, restarting the scene.

```js
assert.match(code, /Press\s+SPACE\s+to\s+restart/);
assert.match(code, /this\.input\.keyboard\.once\(\s*["']keydown-SPACE["'][\s\S]*?this\.scene\.restart\(/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.bg = this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12,
      repeat: -1
    });
    this.player.anims.play("idle");

    this.coins = this.physics.add.group();
    [
      [80, 80],
      [400, 80],
      [80, 280],
      [400, 280],
      [240, 80]
    ].forEach(([x, y]) => {
      const coin = this.coins.create(x, y, "coin");
      coin.setScale(2);
      coin.setDepth(10);
      coin.anims.play("coin-spin");
    });

    this.bg.setDepth(0);
    this.player.setDepth(10);
    this.physics.add.overlap(this.player, this.coins, this.collectCoin, null, this);

    this.ring = this.add.graphics();
    this.ring.setDepth(50);
    this.ring.setBlendMode(Phaser.BlendModes.ADD);
    this.ring.lineStyle(2, 0xffffff, 1);
    this.ring.strokeCircle(0, 0, 24);
    this.ring.setPosition(this.player.x, this.player.y);

    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Score: 0", {
      fontSize: "16px",
      fill: "#ffffff"
    });
    this.scoreText.setScrollFactor(0);
    this.scoreText.setDepth(100);

    this.uiCamera = this.cameras.add(0, 0, 480, 360);
    this.cameras.main.ignore(this.scoreText);
    this.uiCamera.ignore([this.bg, this.player, this.coins, this.ring]);
  }

  update() {
    this.player.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(false);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-right", true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    } else {
      this.player.anims.play("idle", true);
    }
    this.ring.setPosition(this.player.x, this.player.y);
  }

  collectCoin(player, coin) {
    coin.setTint(0xff3366);
    this.tweens.add({
      targets: coin,
      scale: 0,
      duration: 200,
      onComplete: () => coin.destroy()
    });
    const floater = this.add.text(coin.x, coin.y, "+10", {
      fontSize: "16px",
      fill: "#ffeeaa"
    });
    floater.setOrigin(0.5);
    floater.setDepth(50);
    this.tweens.add({
      targets: floater,
      y: floater.y - 30,
      alpha: 0,
      duration: 500,
      onComplete: () => floater.destroy()
    });
    this.tweens.add({
      targets: this.ring,
      scale: 1.5,
      duration: 200,
      yoyo: true
    });
    this.score += 10;
    this.scoreText.setText("Score: " + this.score);
    this.tweens.add({
      targets: this.scoreText,
      scale: 1.4,
      duration: 100,
      yoyo: true
    });
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```

# --solutions--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.bg = this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12,
      repeat: -1
    });
    this.player.anims.play("idle");

    this.coins = this.physics.add.group();
    [
      [80, 80],
      [400, 80],
      [80, 280],
      [400, 280],
      [240, 80]
    ].forEach(([x, y]) => {
      const coin = this.coins.create(x, y, "coin");
      coin.setScale(2);
      coin.setDepth(10);
      coin.anims.play("coin-spin");
    });

    this.bg.setDepth(0);
    this.player.setDepth(10);
    this.physics.add.overlap(this.player, this.coins, this.collectCoin, null, this);

    this.ring = this.add.graphics();
    this.ring.setDepth(50);
    this.ring.setBlendMode(Phaser.BlendModes.ADD);
    this.ring.lineStyle(2, 0xffffff, 1);
    this.ring.strokeCircle(0, 0, 24);
    this.ring.setPosition(this.player.x, this.player.y);

    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Score: 0", {
      fontSize: "16px",
      fill: "#ffffff"
    });
    this.scoreText.setScrollFactor(0);
    this.scoreText.setDepth(100);

    this.uiCamera = this.cameras.add(0, 0, 480, 360);
    this.cameras.main.ignore(this.scoreText);
    this.uiCamera.ignore([this.bg, this.player, this.coins, this.ring]);
  }

  update() {
    this.player.setVelocity(0);
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-120);
      this.player.setFlipX(false);
      this.player.anims.play("walk-left", true);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(120);
      this.player.setFlipX(true);
      this.player.anims.play("walk-right", true);
    } else if (this.cursors.up.isDown) {
      this.player.setVelocityY(-120);
      this.player.anims.play("walk-up", true);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(120);
      this.player.anims.play("walk-down", true);
    } else {
      this.player.anims.play("idle", true);
    }
    this.ring.setPosition(this.player.x, this.player.y);
  }

  collectCoin(player, coin) {
    coin.setTint(0xff3366);
    this.tweens.add({
      targets: coin,
      scale: 0,
      duration: 200,
      onComplete: () => coin.destroy()
    });
    const floater = this.add.text(coin.x, coin.y, "+10", {
      fontSize: "16px",
      fill: "#ffeeaa"
    });
    floater.setOrigin(0.5);
    floater.setDepth(50);
    this.tweens.add({
      targets: floater,
      y: floater.y - 30,
      alpha: 0,
      duration: 500,
      onComplete: () => floater.destroy()
    });
    this.tweens.add({
      targets: this.ring,
      scale: 1.5,
      duration: 200,
      yoyo: true
    });
    this.score += 10;
    this.scoreText.setText("Score: " + this.score);
    this.tweens.add({
      targets: this.scoreText,
      scale: 1.4,
      duration: 100,
      yoyo: true
    });
    if (this.score >= 50) {
      const overlay = this.add.graphics();
      overlay.fillStyle(0x000000, 1);
      overlay.fillRect(0, 0, 480, 360);
      overlay.setDepth(200);
      overlay.setAlpha(0);
      overlay.setScrollFactor(0);
      this.cameras.main.ignore(overlay);
      this.tweens.add({ targets: overlay, alpha: 0.6, duration: 500 });
      const restartText = this.add.text(240, 180, "Press SPACE to restart", {
        fontSize: "18px",
        fill: "#ffffff"
      });
      restartText.setOrigin(0.5);
      restartText.setDepth(201);
      restartText.setScrollFactor(0);
      this.cameras.main.ignore(restartText);
      this.input.keyboard.once("keydown-SPACE", () => this.scene.restart());
    }
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
