---
id: 66faa20000000000000000e2
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `preload` method to `MainScene`. Load the hero walk sprite sheet under the key `"hero"` from `/curriculum-assets/phaser/spritesheets/hero-walk.png`, and the coin sprite sheet under `"coin"` from `/curriculum-assets/phaser/spritesheets/coin.png`. Both sheets use `frameWidth: 16` and `frameHeight: 16`.

# --hints--

You should call `this.load.spritesheet` with the key `"hero"`.

```js
assert.match(code, /this\.load\.spritesheet\s*\(\s*["']hero["']/);
```

You should call `this.load.spritesheet` with the key `"coin"`.

```js
assert.match(code, /this\.load\.spritesheet\s*\(\s*["']coin["']/);
```

Both spritesheets should declare `frameWidth: 16` and `frameHeight: 16`.

```js
assert.match(code, /frameWidth:\s*16,\s*frameHeight:\s*16/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

--fcc-editable-region--

--fcc-editable-region--
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
