---
id: 66faa20000000000000000e3
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Add a `create` method to `MainScene`. Inside it, paint a midnight-blue background rectangle at the centre of the scene with `this.bg = this.add.rectangle(240, 180, 480, 360, 0x223344);`. Then add the hero as a physics sprite at `(240, 180)` with frame `0`, scale it 2× and bind it to the world bounds. Finish by capturing cursor keys into `this.cursors`.

# --hints--

You should assign a rectangle to `this.bg`.

```js
assert.match(code, /this\.bg\s*=\s*this\.add\.rectangle\s*\(/);
```

You should assign a physics sprite to `this.player` using the `"hero"` texture.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\s*\(\s*240\s*,\s*180\s*,\s*["']hero["']/);
```

You should make the player collide with the world bounds and scale it by 2.

```js
assert.match(code, /this\.player\.setCollideWorldBounds\(\s*true\s*\)/);
assert.match(code, /this\.player\.setScale\(\s*2\s*\)/);
```

You should capture cursor keys into `this.cursors`.

```js
assert.match(code, /this\.cursors\s*=\s*this\.input\.keyboard\.createCursorKeys\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

--fcc-editable-region--

--fcc-editable-region--
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
