---
id: 66faa20000000000000000e5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Add an `update` method that drives the player. Reset the velocity to zero each frame, then read the cursors: left/right move horizontally at `±120` and play `"walk-left"` / `"walk-right"`; up/down move vertically at `±120` and play `"walk-up"` / `"walk-down"`. Use `setFlipX(false)` for left and `setFlipX(true)` for right so the same row reads as both directions. When no key is held, fall back to `"idle"`.

# --hints--

You should reset the player's velocity to zero each frame.

```js
assert.match(code, /this\.player\.setVelocity\(\s*0\s*\)/);
```

You should drive horizontal movement with `this.cursors.left` / `this.cursors.right` at ±120.

```js
assert.match(code, /this\.cursors\.left\.isDown[\s\S]*?setVelocityX\(\s*-120\s*\)/);
assert.match(code, /this\.cursors\.right\.isDown[\s\S]*?setVelocityX\(\s*120\s*\)/);
```

You should drive vertical movement with `this.cursors.up` / `this.cursors.down` at ±120.

```js
assert.match(code, /this\.cursors\.up\.isDown[\s\S]*?setVelocityY\(\s*-120\s*\)/);
assert.match(code, /this\.cursors\.down\.isDown[\s\S]*?setVelocityY\(\s*120\s*\)/);
```

You should fall back to the `"idle"` animation when no cursor key is held.

```js
assert.match(code, /else\s*\{\s*this\.player\.anims\.play\(\s*["']idle["']/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Particle Pickup Effect with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.spritesheet(
      "hero",
      "/curriculum-assets/phaser/spritesheets/hero-walk.png",
      { frameWidth: 16, frameHeight: 16 }
    );
    this.load.spritesheet(
      "coin",
      "/curriculum-assets/phaser/spritesheets/coin.png",
      { frameWidth: 16, frameHeight: 16 }
    );
  }

  create() {
    this.bg = this.add.rectangle(240, 180, 480, 360, 0x223344);
    this.player = this.physics.add.sprite(240, 180, "hero", 0);
    this.player.setCollideWorldBounds(true);
    this.player.setScale(2);
    this.cursors = this.input.keyboard.createCursorKeys();

    this.anims.create({
      key: "idle",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 0 }),
      frameRate: 6,
      repeat: -1
    });
    this.anims.create({
      key: "walk-down",
      frames: this.anims.generateFrameNumbers("hero", { start: 0, end: 3 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-up",
      frames: this.anims.generateFrameNumbers("hero", { start: 4, end: 7 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-left",
      frames: this.anims.generateFrameNumbers("hero", { start: 8, end: 11 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "walk-right",
      frames: this.anims.generateFrameNumbers("hero", { start: 12, end: 15 }),
      frameRate: 8,
      repeat: -1
    });
    this.anims.create({
      key: "coin-spin",
      frames: this.anims.generateFrameNumbers("coin", { start: 0, end: 7 }),
      frameRate: 12,
      repeat: -1
    });
    this.player.anims.play("idle");
  }

  update() {
--fcc-editable-region--

--fcc-editable-region--
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 360,
  parent: "game-container",
  pixelArt: true,
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene]
};

const game = new Phaser.Game(config);
```
