---
id: 65c646d4148ae3b2d1cbcac4
title: Step 30
challengeType: 20
dashedName: step-30
---

# --description--

Note that, your search returns something like `80: <__main__.TreeNode object at 0x108b3e0>`. This is the default string representation when printing an instance of a class.

To change that to print a useful value, define another method named `__str__` in the `TreeNode` class. It takes a single argument `self`.

After defining `__str__` you'll get an exception in the console because the `__str__` method doesn't return anything yet. You'll work on the method body in the next step.

# --hints--

You should define a method `__str__` that takes a single argument `self`. Remember to use `pass`.

```js
const lines = code.split('\n').map(line =>
  line.endsWith('\r') ? line.slice(0, -1) : line
);
const indentation = line => line.length - line.trimStart().length;
const methodName = line => {
  const trimmedLine = line.trim();
  if (
    !trimmedLine.startsWith('def') ||
    ![' ', '\t'].includes(trimmedLine[3])
  )
    return '';

  const openingParenthesis = trimmedLine.indexOf('(');
  return openingParenthesis === -1
    ? ''
    : trimmedLine.slice(3, openingParenthesis).trim();
};

const classIndex = lines.findIndex(line => line.trim() === 'class TreeNode:');
assert(classIndex >= 0);

const classIndentation = indentation(lines[classIndex]);
const classEnd = lines.findIndex(
  (line, index) =>
    index > classIndex &&
    line.trim().startsWith('class ') &&
    indentation(line) <= classIndentation
);
const classLines = classEnd === -1 ? lines.length : classEnd;

const initIndex = lines.findIndex(
  (line, index) =>
    index > classIndex &&
    index < classLines &&
    indentation(line) > classIndentation &&
    methodName(line) === '__init__'
);
const strIndex = lines.findIndex(
  (line, index) =>
    index > classIndex &&
    index < classLines &&
    indentation(line) > classIndentation &&
    methodName(line) === '__str__'
);

assert(initIndex >= 0);
assert(strIndex > initIndex);

const strHeader = lines[strIndex].trim();
const openingParenthesis = strHeader.indexOf('(');
const closingParenthesis = strHeader.lastIndexOf(')');
assert(methodName(strHeader) === '__str__');
assert(
  strHeader.slice(openingParenthesis + 1, closingParenthesis).trim() === 'self'
);
assert(strHeader.slice(closingParenthesis + 1).trim().startsWith(':'));

const strIndentation = indentation(lines[strIndex]);
const methodEnd = lines.findIndex(
  (line, index) =>
    index > strIndex &&
    index < classLines &&
    line.trim() &&
    indentation(line) <= strIndentation
);
const strBodyEnd = methodEnd === -1 ? classLines : methodEnd;
assert(
  lines
    .slice(strIndex + 1, strBodyEnd)
    .some(line => indentation(line) > strIndentation && line.trim() === 'pass')
);
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
