---
id: bad87fee1348bd9aeda08826
title: jQuery を使用してターゲットの要素を id で指定する
challengeType: 6
forumTopicId: 18317
required:
  - 
    link: 'https://cdnjs.cloudflare.com/ajax/libs/animate.css/3.2.0/animate.css'
dashedName: target-elements-by-id-using-jquery
---

# --description--

ターゲットの要素を id 属性で指定することもできます。

まず、`$("#target3")` セレクターを使用して、id `target3` を持つ `button` 要素をターゲットとして選択してください。

CSS 宣言と同様に、id 名の前に `#` を入力することに注意してください。

次に、jQuery の `.addClass()` 関数を使用して、クラス `animated` と `fadeOut` を追加してください。

id `target6` を持つ `button` 要素をフェードアウトさせる方法を次に示します。

```js
$("#target6").addClass("animated fadeOut");
```

# --hints--

`id` が `target3` の `button` 要素を選択し、jQuery の `addClass()` 関数を使用してクラス `animated` を付けます。

```js
assert($('#target3').hasClass('animated'));
```

id が `target3` の要素をターゲットとして選択し、jQuery の `addClass()` 関数を使用してクラス `fadeOut` を適用します。

```js
assert(
  ($('#target3').hasClass('fadeOut') || $('#target3').hasClass('fadeout')) &&
    code.match(/\$\(\s*.#target3.\s*\)/g)
);
```

jQuery のみを使用して、これらのクラスを要素に追加してください。

```js
assert(!code.match(/class.*animated/g));
```

# --seed--

## --seed-contents--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");

  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<script>
  $(document).ready(function() {
    $("button").addClass("animated bounce");
    $(".well").addClass("animated shake");
    $("#target3").addClass("animated fadeOut");
  });
</script>

<!-- Only change code above this line -->

<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
