---
id: 699f412114c87221f16a99e3
title: Step 6
challengeType: 1
dashedName: step-6
---

# --description--

Another benefit of using TypeScript is that it will catch errors if you pass to many or to few arguments to a function call.

To illustrate this, update your `square` function call to take two arguments of `5` and `10`.

# --hints--

Your `square` function call should have two arguments of `5` and `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
