---
id: 699f4ec939f10da62fb56c3e
title: Step 15
challengeType: 1
dashedName: step-15
---

# --description--

Now it is time to test out your function.

Create a variable called `raisedResult` and assign it the result of calling `raiseTo` with the arguments `2` and `3`.

Then, log `raisedResult` to the console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Your `raisedResult` variable should be a number.

```js
assert.isNumber(raisedResult);
```

You should log `raisedResult` to the console.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
