---
id: 699f5063d2393c674b1f45dc
title: Step 16
challengeType: 1
dashedName: step-16
---

# --description--

As you recall from a prior step, if you call a function with the wrong number of arguments, TypeScript will throw and error.

There is a way to mark parameters as optional by adding a `?` after the parameter name. This allows you to call the function with fewer arguments than parameters.

Here is an example:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

In this example, if the `greeting` parameter is not provided when calling the `greet` function, it will be `undefined`, and the function will return a default greeting.

Update your `raiseTo` function to make the `exponent` parameter optional. 

# --hints--

Your `raiseTo` function should have the `exponent` parameter marked as optional.

```js
assert.match(
  code,
  /function\s+raiseTo\(base:\s*number,\s*exponent\?:\s*number\):\s*number/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```
