---
id: 699f4639ae0016caaf81b7c6
title: Step 10
challengeType: 1
dashedName: step-10
---

# --description--

Now it is time to test out your function.

Start by creating a variable called `avgResult` and assigning it the function call of `getAverage`. The argument for that function call should be `[2, 14, 26, 8]`.

Then log the value of `avgResult` to the console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

You should have a variable called `avgResult`.

```js
assert.isDefined(avgResult);
```

You should assign `getAverage([2, 14, 26, 8])` to your `avgResult` variable.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Your `avgResult` variable should be a number.

```js
assert.isNumber(avgResult);
```

You should log `avgResult` to the console.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
