---
id: 699f41b8da20f816b5c7dece
title: Step 7
challengeType: 1
dashedName: step-7
---

# --description--

If you open up the console, you should see the following error message:

```md
Expected 1 arguments, but got 2.
```

To resolve the error, update your `square` function call to take only one argument of `5`.

# --hints--

Your `square` function call should have only one argument of `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```

