---
id: 699f3cb8d057f29954205d7d
title: Step 4
challengeType: 1
dashedName: step-4
---

# --description--

As you recall from earlier lessons, you can add type annotations to function parameters like this:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

In this example, TypeScript will display an error if you try to call the `add` function with arguments that are not numbers.

Update your `square` function to include a type annotation for the `num` parameter. The type should be `number`.

# --hints--

You should add a type annotation to the `num` parameter in the `square` function. The type should be a `number`.

```js
assert.match(code, /function\s+square\(num:\s*number\)/);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```
