---
id: 699f4a3e7b4c6bfa538cc4f7
title: Step 13
challengeType: 1
dashedName: step-13
---

# --description--

If you open up the console, you will see the following error message:

```md
Type 'string' is not assignable to type 'number'.
```

To resolve this error, you need to change the return back to `sum / numbers.length`. 

# --hints--

Your `getAverage` function should return a number.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Your `getAverage` function should return the average of the numbers in the array.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
