---
id: 699f4861dd5b5f164e1bda4a
title: Step 11
challengeType: 1
dashedName: step-11
---

# --description--

Another benefit of using TypeScript is ensuring that the return types are enforced and what you expect. Sometimes TypeScript can infer the return type of a function, but you can also explicitly specify it.

To illustrate this, change the return type of the `getAverage` function to the string `"something"`. 

# --hints--

Your `getAverage` function should return a string.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Your `getAverage` function should return the string `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
