---
id: 6436810f38ca75013a371cb4
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

In this step, you will do 2 things: 

* Give a title to the project by using the `title` tag. Inside the tags, write `FCC Forum Leaderboard`.
* Use the `link` tag to link the `styles.css` file to the HTML file.

# --hints--

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your title should be `FCC Forum Leaderboard`.

```js
const title = document.querySelector('title');
const titleText = title.textContent.trim();
assert.match(titleText, /^FCC\s+Forum\s+Leaderboard\s*$/, `unexpected title text: "${titleText}"`);
```

Your code should have a `link` element. with `href` attribute set to `styles.css`

```js
const html = document.documentElement.outerHTML;
const linkRegex = /<link.*href="(\/\.)?styles\.css".*>/;
const linkMatch = html.match(linkRegex);
assert(linkMatch !== null, 'HTML document should contain a link tag with href of "styles.css"');
```

Your `link` element should have a `rel` attribute with a value of `stylesheet`.

```js
const link = document.querySelector('head > link');
assert.equal(link.getAttribute('rel'), 'stylesheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8" />
</head>
--fcc-editable-region--
<body>
</body>
</html>
```
