---
id: 64368ab6404fe201e3abe8fe
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

You need to add 2 more tags inside the body tag:

* The `header` tag will help to display the logo of freeCodeCamp and the title of the page.
* The `main` tag will contain the leaderboard information.

# --hints--

Your code should have a `header` element

```js
const header = document.querySelector('header');
assert.exists(header);
```

Your code should have a `main` element

```js
const main = document.querySelector('main');
assert.exists(main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
--fcc-editable-region--
<body>

</body>
--fcc-editable-region--
<script src="./script.js"></script>
</html>
```

