---
id: 643690119241c60274a48a5b
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Inside the `nav` tag, do the following :

* Create an `image` tag with class `fcc-logo`.
* Add the `src` attribute and set it to `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.
* Add the `alt` attribute and set it to `freeCodeCamp logo`.

# --hints--

Your code should have an `img` element inside the `nav` tag

```js
const img = document.querySelector('nav > img');
assert.exists(img);
```

Your `img` element should have a class named `fcc-logo`

```js
const img = document.querySelector('nav > img');
assert.equal(img.classList.contains('fcc-logo'), true);
```

Your `img` element should have `src` attribute set to `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`

```js
const img = document.querySelector('nav img');
assert.exists(img);
assert.equal(img.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Your `img` element should have the `alt` attribute set to `freeCodeCamp logo`

```js
const img = document.querySelector('nav img');
assert.exists(img);
assert.equal(img.getAttribute('alt').trim(), 'freeCodeCamp logo');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
--fcc-editable-region--
    <nav>

    </nav>
--fcc-editable-region--
    <h1 class="title">Latest topics</h1>
  </header>
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```
