---
id: 6436810f38ca75013a371cb4
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

In this step, you will do 2 things: 

* give a title to the project by using the `title` tag. Inside the tags, write `FCC Forum Leaderboard`.
* use the `link` tag to link the `styles.css` file to the html file.

# --hints--

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your code should have a `link` element. Make sure to include `href` and `rel` attributes.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8" />
</head>
--fcc-editable-region--
<body>
</body>
</html>
```
