---
id: 643698c37a381602b2876f7d
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

You can't see the logo yet because it's white, this will be fixed using the `styles.css` file.
As you can see that a color palette is already made. You just need to use it.

Target the `body` tag in the CSS and set its `background-color` to `var(--main-bg-color)`

# --hints--

Your `body` selector should have a `background-color` property set to `var(--main-bg-color)`. Make sure to add a semicolon `;` at the end of the code.

```js
const background = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert.equal(background, 'var(--main-bg-color)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```

```css
:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

--fcc-editable-region--

--fcc-editable-region--
```
