---
id: 6436ba8fc66991055c33f897
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

You will target the `nav` tag and give it a `background-color` of `var(--dark-navy)`. Then you will add `padding` of `10px 0` to the `nav` tag. This will add padding of 10px to the top and bottom of the `nav` tag.

# --hints--

Your `nav` tag selector should have a `background-color` property set to `var(--dark-navy)`.

```js
const background = new __helpers.CSSHelp(document).getStyle('nav')?.getPropertyValue('background-color');
assert.equal(background, 'var(--dark-navy)');
```

Your `nav` tag selector should have a `padding` property set to `10px 0`.

```js
const cssHelper = new __helpers.CSSHelp(document);
const paddingTop = cssHelper.getStyle('nav')?.getPropertyValue('padding-top');
const paddingLeft = cssHelper.getStyle('nav')?.getPropertyValue('padding-left');
const paddingRight = cssHelper.getStyle('nav')?.getPropertyValue('padding-right');
const paddingBottom = cssHelper.getStyle('nav')?.getPropertyValue('padding-bottom');
assert.equal(paddingTop, '10px');
assert.equal(paddingBottom, '10px');
assert.equal(paddingLeft, '0px');
assert.equal(paddingRight, '0px');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```

```css
:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

--fcc-editable-region--

--fcc-editable-region--
```
