---
id: 64368f7d44ecc0023cd0026a
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

In this step, you need to add 2 tags inside the `header` tag :

* the `nav` tag will contain the logo
* the `h1` tag will contain the title of the page. Inside this tag write `Latest topics`
* on the `h1` tag, add the class `title`

# --hints--

Your code should have a `nav` element

```js
const nav = document.querySelector('nav');
assert.exists(nav);
```

Your code should have a `h1` element. Make sure to add the h1 tag after the nav tag.

```js
const h1 = document.querySelector('header > h1');
assert.exists(h1);
assert.equal(h1.textContent, 'Latest topics');
```

Your `h1` element should have a class named `title`

```js
const h1 = document.querySelector('header > h1');
assert.exists(h1);
assert(h1.classList.contains('title'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
--fcc-editable-region--
  <header>

  </header>
--fcc-editable-region--
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```
