---
id: 64368a8e73bb0001b1e856e8
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Now, we will include the `script` tag below the closing `body` tag to link the `script.js` file to the html file.

# --hints--

Your code should have a `script` element. Make sure to include `src` attribute.

```js
const script = document.querySelector('script[src="script.js"], script[src="./script.js"]');
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
--fcc-editable-region--
<body>
</body>

--fcc-editable-region--
</html>
```
