---
id: 64368a8e73bb0001b1e856e8
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

To link the `script.js` file to the HTML file, you should place the script tag after the body tag's closing tag.

# --hints--

Your code should have a `script` tag.

```js
assert(code.match(/<script.*>.*<\/script>/g));
```

Your `script` tag should have a `src` attribute with value `script.js`

```js
assert(code.match(/<script.*src\s*=\s*['"]\s*(\.\/)?script\.js\s*['"]\s*>.*<\/script>/g));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
--fcc-editable-region--
<body>
</body>

--fcc-editable-region--
</html>
```
