---
id: 6436b5c2b3d8bf05032daa87
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

As you can see that we have changed the background color of the body, but the logo is too big for the project. You need to resize it and keep it at the center. For that follow the steps below:

* Target the class named `fcc-logo` and set its `width` to `210px`.
* To keep the image at the center, use `margin: auto;` and `display: block;`.

# --hints--

Your `.fcc-logo` class selector should have a `width` property set to `210px`.

```js
const width = new __helpers.CSSHelp(document).getStyle('.fcc-logo')?.getPropertyValue('width');
assert.equal(width, '210px');
```

Your `.fcc-logo` class selector should have a `margin` property set to `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('.fcc-logo')?.getPropertyValue('margin');
assert.equal(margin, 'auto');
```

Your `.fcc-logo` class selector should have a `display` property set to `block`.

```js
const display = new __helpers.CSSHelp(document).getStyle('.fcc-logo')?.getPropertyValue('display');
assert.equal(display, 'block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```

```css
:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

--fcc-editable-region--

--fcc-editable-region--
```
