---
id: 64387a4863757e04541b84a0
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

In this step, you will add some text and assign ids on the `th` elements. For that, follow these steps:

* The text inside the first `th` element should be `Topics`
* keep the second `th` element empty
* the text inside the third `th` element should be `Replies`, give it a `id` of `replies`.
* the text inside the fourth `th` element should be `Views`, give it a `id` of `views`.
* the text inside the fifth `th` element should be `Activity`, give it a `id` of `activity`.

# --hints--

Your first `th` element should have the text `Topics`.

```js
const tr = document.querySelector('thead > tr');
const thList = tr.querySelectorAll('th');
assert.equal(thList[0].textContent.trim(), 'Topics', 'Expected "Topics" text in first th element');
```

Your second `th` element should not have any text. It should be blank

```js
const tr = document.querySelector('thead > tr');
const thList = tr.querySelectorAll('th');
assert.equal(thList[1].textContent.trim(), '', 'Expected empty text in second th element');
```

Your third `th` element should have the text `Replies` and an `id` of `replies`.

```js
const tr = document.querySelector('thead > tr');
const thList = tr.querySelectorAll('th');
assert.equal(thList[2].textContent.trim(), 'Replies');
assert.equal(thList[2].id, 'replies');
```

Your fourth `th` element should have text `Views` and an `id` of `views`.

```js
const tr = document.querySelector('thead > tr');
const thList = tr.querySelectorAll('th');
assert.equal(thList[3].textContent.trim(), 'Views', 'Expected "Views" text in fourth th element');
assert.equal(thList[3].id, 'views');
```

Your fifth `th` element should have text that says `Activity` and an `id` of `activity`.

```js
const tr = document.querySelector('thead > tr');
const thList = tr.querySelectorAll('th');
assert.equal(thList[4].textContent.trim(), 'Activity');
assert.equal(thList[4].id, 'activity');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
--fcc-editable-region--
        <th></th>
        <th></th>
        <th></th>
        <th></th>
        <th></th>
--fcc-editable-region--
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

```
