---
id: 64385b3099343c019f813ca9
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

Now target the `title` class and set the `text-align` property to `center` and `margin` property to `20px 0`.Also, set the `color` property to `var(--light-grey)`.

# --hints--

Your `.title` class selector should have a `text-align` property set to `center`.

```js
const align = new __helpers.CSSHelp(document).getStyle('.title')?.getPropertyValue('text-align');
assert.equal(align, 'center');
```

Your `.title` class selector should have a `margin` property set to `20px 0`. Don't forget to use the semicolon `;` at the end of each property.

```js
const cssHelper = new __helpers.CSSHelp(document);
const marginTop = cssHelper.getStyle('.title')?.getPropertyValue('margin-top');
const marginBottom = cssHelper.getStyle('.title')?.getPropertyValue('margin-bottom');
const marginleft = cssHelper.getStyle('.title')?.getPropertyValue('margin-left');
const marginright = cssHelper.getStyle('.title')?.getPropertyValue('margin-right');
assert.equal(marginTop, '20px');
assert.equal(marginBottom, '20px');
assert.equal(marginleft, '0px');
assert.equal(marginright, '0px');
```

Your `.title` class selector should have a `color` property set to `var(--light-grey)`.

```js
const color = new __helpers.CSSHelp(document).getStyle('.title')?.getPropertyValue('color');
assert.equal(color, 'var(--light-grey)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
