---
id: 64385e49183c7102fa79d931
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Add the `table` element to the `main` element. To do that, follow these steps:

Create a `table` element inside the `main` element.
Inside the `table` element, create a `thead` element that will contain the table headings.
Inside the `thead` element, create a `tbody` element that will contain the table data. Also, add an `id` attribute to the `tbody` element and set its value to `posts-container`. 

# --hints--

Your code should have a `table` element.

```js
const table = document.querySelector('main > table');
assert.exists(table);
```

Your code should have a `thead` element. Make sure to add it inside the `table` element.

```js
const thead = document.querySelector('main > table > thead');
assert.exists(thead);
```

Your code should have a `tbody` element. Make sure to add it inside the `table` element, next to `thead` element.

```js
const tbody = document.querySelector('main > table > thead + tbody');
assert.exists(tbody);
```

Your `tbody` element should have a id named `posts-container`

```js
const tbody = document.querySelector('#posts-container');
assert.exists(tbody);
assert.equal(tbody.id, 'posts-container');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  --fcc-editable-region--
  <main>

  </main>
  --fcc-editable-region--
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

```
