---
id: 64387d286d6a8d048649068e
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

As you can see, the table is not looking good. You need to add some styles to it. For that, follow these steps:

* target the `table` element and give it a `width` of `80%`.
* add `color` property to the `table` element and set it to `var(--dark-grey)`.
* add the `margin` property to `auto`.
* set the `table-layout` property to `fixed`.
* set the `border-collapse` property to `collapse`.

# --hints--

Your `table` element should have a `width` of `80%`.

```js
const table = new __helpers.CSSHelp(document).getStyle('table')?.getPropertyValue('width');
assert.equal(table, '80%');
```

Your `table` element should have the `color` property set to `var(--dark-grey)`.

```js
const color = new __helpers.CSSHelp(document).getStyle('table')?.getPropertyValue('color');
assert.equal(color, 'var(--dark-grey)');
```

Your `table` element should have the `margin` property set to `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('table')?.getPropertyValue('margin');
assert.equal(margin, 'auto');
```

Your `table` element should have the `table-layout` property set to `fixed`.

```js
const tableLayout = new __helpers.CSSHelp(document).getStyle('table')?.getPropertyValue('table-layout');
assert.equal(tableLayout, 'fixed');
```

Your `table` element should have property `border-collapse` set to `collapse`.

```js
const tableLayout = new __helpers.CSSHelp(document).getStyle('table')?.getPropertyValue('border-collapse');
assert.equal(tableLayout, 'collapse');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

--fcc-editable-region--
--fcc-editable-region--

```
