---
id: 6438875d81fc1f05460ff51f
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Next, you will target all the `td` elements and style them according to these steps:

* Add a `border-bottom` of `1px solid var(--dark-grey)`.
* Add a `padding` of `20px 0`.

# --hints--

Your `td` element should have a `border-bottom` of `1px solid var(--dark-grey)`.

```js
const table = new __helpers.CSSHelp(document).getStyle('td')?.getPropertyValue('border-bottom');
assert.equal(table, '1px solid var(--dark-grey)');
```

Your `td` element should have a `padding` of `20px 0`.

```js
const td = new __helpers.CSSHelp(document).getStyle('td')?.getPropertyValue('padding-top');
assert.equal(td, '20px', 'Expected "20px" padding-top on td element');

const td2 = new __helpers.CSSHelp(document).getStyle('td')?.getPropertyValue('padding-bottom');
assert.equal(td2, '20px', 'Expected "20px" padding-bottom on td element');

const td3 = new __helpers.CSSHelp(document).getStyle('td')?.getPropertyValue('padding-left');
assert.equal(td3, '0px', 'Expected "0px" padding-left on td element');

const td4 = new __helpers.CSSHelp(document).getStyle('td')?.getPropertyValue('padding-right');
assert.equal(td4, '0px', 'Expected "0px" padding-right on td element');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

--fcc-editable-region--
--fcc-editable-region--

```
