---
id: 64388585eea91605083429ec
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Now you will target all of the `th` elements and style them by following these steps:

* Add a `border-bottom` of `2px solid var(--dark-grey)`.
* Add a `padding-bottom` of `10px`.
* Add a `font-size` of `1.2rem`.

# --hints--

Your `th` element should have a `border-bottom` of `2px solid var(--dark-grey)`.

```js
const table = new __helpers.CSSHelp(document).getStyle('th')?.getPropertyValue('border-bottom');
assert.equal(table, '2px solid var(--dark-grey)');
```

Your `th` element should have a `padding-bottom` of `10px`.

```js
const table = new __helpers.CSSHelp(document).getStyle('th')?.getPropertyValue('padding-bottom');
assert.equal(table, '10px');
```

Your `th` element should have a `font-size` of `1.2rem`.

```js
const table = new __helpers.CSSHelp(document).getStyle('th')?.getPropertyValue('font-size');
assert.equal(table, '1.2rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

--fcc-editable-region--
--fcc-editable-region--

```
