---
id: 64388945b59341057e955fbb
title: Step 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, you will be using the freeCodeCamp forum API. Requesting data from an API is an asynchronous operation, so you will need to use an asynchronous function with the `async` keyword.

Here is an example of an asynchronous function.

```js
async function example() {
  console.log("this is an example");
}
```

Create an `async` function named `fetchData`.

# --hints--

You should create a function named `fetchData` with the `async` keyword at the start of the function.

```js
assert.match(code, /async\s*function\s*fetchData()/);
```

`fetchData` should be defined.

```js
assert.isDefined(fetchData);
```

`fetchData` should be a function.

```js
assert.isFunction(fetchData);
```

`fetchData` should be an empty function.

```js
assert.match(code, /async\s*function\s*fetchData\s*\(\)\s*{\s*}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

td{
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0px;
}

```

```js

--fcc-editable-region--
--fcc-editable-region--
```
