---
id: 643851f146326e0152a69d2c
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Now, you can see some gaps all around due to the default margin and padding of the project. you can fix that by targeting the universal selector `*` and setting the `margin` and `padding` to `0`.

# --hints--

Your universal selector `*` should have a `margin` property set to `0`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('*')?.getPropertyValue('margin');
assert.ok(margin === '0' || margin === '0px');
```

Your universal selector `*` should have a `padding` property set to `0`.

```js
const padding = new __helpers.CSSHelp(document).getStyle('*')?.getPropertyValue('padding');
assert.ok(padding === '0' || padding === '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main></main>
</body>
<script src="./script.js"></script>
</html>
```

```css

--fcc-editable-region--

--fcc-editable-region--

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

```
