---
id: 64387773a8e9d103b03f1394
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Now you will add the table headings. For that, follow these steps:

* Create a `tr` element inside the `thead` element. `tr` represents a table row. 
* Inside the `tr` element, create 5 `th` elements. `th` represents a table heading.

# --hints--

Your code should have a `tr` element inside the `thead` element.

```js
const tr = document.querySelector('thead > tr');
assert.exists(tr);
```

Your code should have 5 `th` element inside the `tr` element.

```js
const tr = document.querySelector('thead > tr');
assert.exists(tr);

const thList = tr.querySelectorAll('th');
assert.equal(thList.length, 5, 'Expected 5 th elements inside tr element');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
--fcc-editable-region--
    <thead>

    </thead>
--fcc-editable-region--
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

```
