---
id: 643c99404743b30b77bd736d
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

In this step, you will use the `postsContainer` variable to print sample text to the page. You will use the `innerHTML` property to print the text to the page. The `innerHTML` property is used to set or return the HTML content of an element. You can use it to set the content of the `postsContainer` variable to the text `Hello World`. Here is an example:

```js
const element = document.getElementById('id');
element.innerHTML += 'Hello World';
```

You will target the `postsContainer` variable and set its `innerHTML` property to the text `Hello World`. Like this 
  
```js
postsContainer.innerHTML += 'Hello World';
```

Don't forget to add the `+` operator to the `innerHTML` property. This will append the text to the existing content of the `postsContainer` variable. Otherwise, the text will replace the existing content.

# --hints--

You should use the `postsContainer` variable to print `Hello World` to the page. Don't forget to add the `+` operator to the `innerHTML` property.

```js
assert.match(code, /postsContainer\.innerHTML\s*\+=\s*['"]Hello World['"]/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

td{
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0px;
}

```

```js
const FORUM_LATEST = "https://forum-proxy.freecodecamp.rocks/latest";
const FORUM_TOPIC_URL = "https://forum.freecodecamp.org/t/";
const FORUM_CATEGORY_URL = "https://forum.freecodecamp.org/c/";
const AVATAR_URL = "https://sea1.discourse-cdn.com/freecodecamp";

const postsContainer = document.getElementById('posts-container');

async function fetchData(){ 
  let res = await fetch(FORUM_LATEST);
  let data = await res.json();
  console.log(data);
  showLatestPosts(data);
}

fetchData();

function showLatestPosts(data){
  const {topic_list, users} = data;
  const {topics} = topic_list;

--fcc-editable-region--
  topics.forEach((item)=>{
    const { id, title, views, posts_count, slug, posters, excerpt, category_id, bumped_at} = item;

  });
--fcc-editable-region--

}

```
