---
id: 643c9e408d1a2a0bc1eb4a87
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

In total, `Hello World` will be printed 30 times. This is because the `forEach` method is called on the `topics` array, which has 30 items. The `forEach` method will call the callback function once for each item in the array. The callback function will then print `Hello World` to the project. You don't need `Hello World` to be printed 30 times, you need the content from the API to be displayed on the page as a table content.

All the properties you are destructuring from `item` object are variables that you will use to display the content on the page. But you cannot display variables inside a string. you need to use template literals to display the variables inside a string. Here is an example:

```js
const name = "John";
const age = 30;
const job = "developer";
const city = "New York";

console.log(`My name is ${name}. I am ${age} years old. I am a ${job} and I live in ${city}.`);
```

You will remove `Hello World` and replace it with a template literal. It is a pair of backticks (``). Also please remove the single or double quotes from the template literal. Here's an example of a template literal:

```js
postsContainer.innerHTML += ``
```

# --hints--

Your code should remove `Hello World` and replace it with a template literal.

```js
assert(/postsContainer\.innerHTML \+= `/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

td{
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0px;
}

```

```js
const FORUM_LATEST = "https://forum-proxy.freecodecamp.rocks/latest";
const FORUM_TOPIC_URL = "https://forum.freecodecamp.org/t/";
const FORUM_CATEGORY_URL = "https://forum.freecodecamp.org/c/";
const AVATAR_URL = "https://sea1.discourse-cdn.com/freecodecamp";

const postsContainer = document.getElementById('posts-container');

async function fetchData(){ 
  let res = await fetch(FORUM_LATEST);
  let data = await res.json();
  console.log(data);
  showLatestPosts(data);
}

fetchData();

function showLatestPosts(data){
  const {topic_list, users} = data;
  const {topics} = topic_list;

--fcc-editable-region--
  topics.forEach((item)=>{
    const { id, title, views, posts_count, slug, posters, excerpt, category_id, bumped_at} = item;

    postsContainer.innerHTML += "Hello World";

  });
--fcc-editable-region--

}

```
