---
id: 643c0ab6779047014ca69b4f
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

The `fetch` method is used to fetch data from the API. You will use the `await` keyword to handle the asynchronous nature of the `fetch()` method. The `await` keyword will wait for the `fetch` method to return the data from the API. An example of the `await` keyword is shown below.

```js
async function example() {
  const data = await fetch("https://api-to-fetch-data.com");
  console.log(data);
}
```

To have a deeper understanding, follow these steps

* Using `let`, create a variable inside `fetchData()` named `res` and assign it to the `fetch()` method.
* Make sure to add the `await` keyword before the `fetch` method.

# --hints--

You should declare a variable named `res` inside the `async function`

```js
assert.match(code, /async\s*function\s*fetchData\(\)\s*{\s*let\s*res/);
```

You should use the `await` keyword before the `fetch()` method.

```js
assert.match(code, /async\s*function\s*fetchData\(\)\s*{\s*let\s*res\s*=\s*await/);
```

You should use the `fetch()` method after the `await` keyword.

```js
assert.match(code, /async\s*function\s*fetchData\(\)\s*{\s*let\s*res\s*=\s*await\s*fetch\(/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

td{
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0px;
}

```

```js
--fcc-editable-region--
async function fetchData(){ 
  
}
--fcc-editable-region--
```
