---
id: 6443c4d659395d0175cd121a
title: Step 45
challengeType: 0
dashedName: step-45
---

# --description--

The `time` argument is a string. You should use the `Date` object to convert it into a date object. Here is an example:

```js
const date = new Date(time);
console.log(date); // Sat Apr 03 2021 15:24:51 GMT+0800 (Singapore Standard Time)
```

After the `const currentTime` line, create a new constant named `lastPost` and assign `new Date(time)` to it.

# --hints--

You should create a new constant named `lastPost` and assign `new Date(time)` to it.

```js
assert.match(code, /const\s+lastPost\s*=\s*new\s+Date\s*\(\s*time\s*\)\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

td{
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0px;
}

```

```js
const FORUM_LATEST = "https://forum-proxy.freecodecamp.rocks/latest";
const FORUM_TOPIC_URL = "https://forum.freecodecamp.org/t/";
const FORUM_CATEGORY_URL = "https://forum.freecodecamp.org/c/";
const AVATAR_URL = "https://sea1.discourse-cdn.com/freecodecamp";

const postsContainer = document.getElementById('posts-container');

--fcc-editable-region--
function showTime(time){
  const currentTime = new Date();

  
}
--fcc-editable-region--

async function fetchData(){ 
  let res = await fetch(FORUM_LATEST);
  let data = await res.json();
  console.log(data);
  showLatestPosts(data);
}

fetchData();

function showLatestPosts(data){
  const {topic_list, users} = data;
  const {topics} = topic_list;

  topics.forEach((item)=>{
    const { id, title, views, posts_count, slug, posters, excerpt, category_id, bumped_at} = item;

    postsContainer.innerHTML += `
    <tr>
      <td>
        <p>${title}</p>
      </td>
      <td></td>
      <td>${posts_count - 1}</td>
      <td>${views}</td>
      <td></td>
    </tr>`;
  });

}

```
