---
id: 645f06144bc0dc0143a5a3a2
title: Step 100
challengeType: 0
dashedName: step-100
---

# --description--

Every item inside the `topics` array is an object which carries one property called `user_id`. You will use this id to display the user profile picture in our project. First, create a function named `displayUsers`

# --hints--

You need to create a function named `displayUsers`

```js
assert.match(code, /\s*function\s*displayUsers()/);
```

`displayUsers` should be defined.

```js
assert.isDefined(displayUsers);
```

`displayUsers` should be a function.

```js
assert.isFunction(displayUsers);
```

`displayUsers` should be an empty function.

```js
assert.match(code, /\s*function\s*displayUsers\s*\(\s*\)\s*{\s*}/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <title> FCC Forum Leaderboard</title>
  <link rel="stylesheet" href="./styles.css" />
</head>
<body>
  <header>
    <nav>
      <img class="fcc-logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg"
        alt="freeCodeCamp logo" />
    </nav>
    <h1 class="title">Latest topics</h1>
  </header>
  <main>
  <table>
    <thead>
      <tr>
        <th>Topics</th>
        <th></th>
        <th id="replies">Replies</th>
        <th id="views">Views</th>
        <th id="activity">Activity</th>
      </tr>
    </thead>
    <tbody id="posts-container"></tbody>
  </table>
</main>
</body>
<script src="./script.js"></script>
</html>
```

```css

*{
  margin: 0px;
  padding: 0px
}

:root {
  --main-bg-color: #2a2a40;
  --black: #000;
  --dark-navy: #0a0a23;
  --dark-grey: #d0d0d5;
  --medium-grey: #dfdfe2;
  --light-grey: #f5f6f7;
  --peach: #f28373;
  --salmon-color: #f0aea9;
  --light-blue: #8bd9f6;
  --light-orange: #f8b172;
  --light-green: #93cb5b;
  --golden-yellow: #f1ba33;
  --gold: #f9aa23;
  --green: #6bca6b;
}

body {
  background-color: var(--main-bg-color);
}

.fcc-logo{
  width: 210px;
  margin: auto;
  display: block
}

nav {
  background-color: var(--dark-navy);
  padding: 10px 0;
}

.title{
  text-align: center;
  margin: 10px 0;
  color: var(--light-grey);
}

table{
  width: 80%;
  color: var(--dark-grey);
  margin: auto;
  table-layout: fixed;
  border-collapse: collapse;
}

th:first-child {
  width: 70%;
  text-align: start;
}

th{
  border-bottom: 2px solid var(--dark-grey);
  padding-bottom: 10px;
  font-size: 1.2rem;
}

td{
  border-bottom: 1px solid var(--dark-grey);
  padding: 20px 0px;
}

td:not(:first-child) {
  text-align: center;
}

.post-title{
  font-size: 1.2rem;
  color: var(--medium-grey);
  text-decoration: none;
}

.category{
  padding: 3px;
  color: var(--black);
  text-decoration: none;
  display: block;
  width: fit-content;
  margin-top: 2px; 
  margin-left: 0;
  margin-right: 0;
  margin-bottom: 5px;
}

.career{
  background-color: var(--salmon-color);
}

.feedback{
  background-color: var(--light-blue);
}

.html-css{
  background-color: var(--light-blue);
}

.support{
  background-color: var(--light-orange);
}

.General{
  background-color: var(--light-green);
}

.javascript{
  background-color: var(--golden-yellow);
}

.backend{
  background-color: var(--gold);
}

.python{
  background-color: var(--green);
}

.motivation{
  background-color: var(--peach);
}
```

```js
const FORUM_LATEST = "https://forum-proxy.freecodecamp.rocks/latest";
const FORUM_TOPIC_URL = "https://forum.freecodecamp.org/t/";
const FORUM_CATEGORY_URL = "https://forum.freecodecamp.org/c/";
const AVATAR_URL = "https://sea1.discourse-cdn.com/freecodecamp";

const postsContainer = document.getElementById('posts-container');

const categoryList = {
  299: { category: "Career Advice", className: "career" },
  409: { category: "Project Feedback", className: "feedback" },
  417: { category: "freeCodeCamp Support", className: "support" },
  421: { category: "JavaScript", className: "javascript" },
  423: { category: "HTML - CSS", className: "html-css" },
  424: { category: "Python", className: "python" },
  432: { category: "You Can Do This!", className: "motivation" },
  560: { category: "Backend Development", className: "backend" },
};

function forumCategories(id){
  let selectedCategory = {}

  if(categoryList.hasOwnProperty(id)){
    const {className, category} = categoryList[id];

    selectedCategory.className = className;
    selectedCategory.category = category;
  }
  else{
    selectedCategory.className = 'General';
    selectedCategory.category = 'General';
    selectedCategory.id = 1;
  }
  const url = `${FORUM_CATEGORY_URL}${selectedCategory.className}/${id}`;
  const linkText = selectedCategory.category;
  const linkClass = `category ${selectedCategory.className}`
  const linkTarget = '_blank';

  return `<a href="${url}" class="${linkClass}" target="${linkTarget}">
    ${linkText}
  </a>`
}

function showTime(time){
  const currentTime = new Date();
  const lastPost = new Date (time);

  const timeDifference = currentTime - lastPost;
  const MS_PER_MINUTE = 1000 * 60;

  const minutesAgo = Math.floor(timeDifference / MS_PER_MINUTE);
  const hoursAgo = Math.floor(minutesAgo / 60);
  const daysAgo = Math.floor(hoursAgo / 24);

  if (minutesAgo < 60) { 
    return `${minutesAgo}m ago`;
  };

  if(hoursAgo < 24){
    return `${hoursAgo}h ago`; 
  }

  return `${daysAgo}d ago`;
  
}

function countingViews(views){
  const thousands = Math.floor(views / 1000);

  if(views >= 1000){
    return `${thousands}k`;
  };

  return views;

}

--fcc-editable-region--

--fcc-editable-region--

async function fetchData(){ 
  let res = await fetch(FORUM_LATEST);
  let data = await res.json();
  console.log(data);
  showLatestPosts(data);
}

fetchData();

function showLatestPosts(data){
  const {topic_list, users} = data;
  const {topics} = topic_list;

  topics.forEach((item)=>{
    const { id, title, views, posts_count, slug, posters, excerpt, category_id, bumped_at} = item;

    postsContainer.innerHTML += `
    <tr>
      <td>
        <p class="post-title">${title}</p>
        ${forumCategories(category_id)}
      </td>
      <td></td>
      <td>${posts_count - 1}</td>
      <td>${countingViews(views)}</td>
      <td>${showTime(bumped_at)}</td>
    </tr>`
  });
}

```
