---
id: 613297a923965e0703b64796
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Você já deve estar familiarizado com o elemento `meta`. Ele é usado para especificar informações sobre a página como título, descrição, palavras-chave e autor.

Dê à página um elemento `meta` com um valor apropriado para `charset`.

O atributo `charset` especifica a codificação de caracteres da página e, hoje em dia, `UTF-8` é a única codificação suportada pela maioria dos navegadores.

# --hints--

Você deve criar um elemento `meta` ao redor do elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Você deve dar à tag `meta` um `charset` de `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
