---
id: 613297a923965e0703b64796
title: Etapa 2
challengeType: 0
dashedName: step-2
---

# --description--

Você já deve estar familiarizado com a tag `meta`. Ela é usada para especificar informações sobre a página como título, descrição, palavras-chave e autor.

Dê à página uma tag `meta` com um valor apropriado para `charset`.

O atributo `charset` especifica a codificação de caracteres da página e, hoje em dia, `UTF-8` é a única codificação suportada pela maioria dos navegadores.

# --hints--

Você deve criar um elemento `meta` ao redor do elemento `head`.

```js
// TODO: Once builder is fixed so head info is not in body
assert.exists(document.querySelector('head > meta'));
```

Você deve dar à tag `meta` um `charset` de `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset'), 'UTF-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
    background: #f5f6f7;
    color: #1b1b32;
    font-family: Helvetica;
    margin: 0;
}
```
