---
id: 6133d11ef548f51f876149e3
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

A navegação é uma parte fundamental da acessibilidade. Os leitores de tela dependem de você para fornecer a estrutura de sua página. Isso é feito com elementos HTML semânticos.

Adicione um elemento `header` e um elemento `main` à sua página.

O elemento `header` será usado para introduzir a página, bem como fornecer um menu de navegação.

O elemento `main` contém o conteúdo principal de sua página.

# --hints--

Você deve adicionar um elemento `header` ao `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Você deve adicionar um novo elemento `main` ao `body`.

```js
assert.exists(document.querySelector('body > main'));
```

O elemento `header` deve vir antes do elemento `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
