---
id: 613e275749ebd008e74bb62e
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Uma propriedade útil de um _SVG_ (gráficos vetoriais escaláveis) é o atributo `path` (caminho), que permite que a imagem seja dimensionada sem afetar a resolução da imagem resultante.

Atualmente, o `img` está assumindo que é o tamanho padrão, o que é muito grande. O CSS tem uma função `max`, que retorna o maior valor de um conjunto de valores separados por vírgulas. Por exemplo:

```css
img {
  width: max(250px, 25vw);
}
```

No exemplo acima, a largura da imagem será de 250px se a largura da viewport for menor que 1000 pixels. Se a largura da viewport for maior que 1000 pixels, a largura da imagem será de 25vw. Isto porque 25vw é igual a 25% da largura da viewport.

Dimensione a imagem usando `id` como seletor e definindo a `width` como sendo, no máximo, `100px` ou `18vw`.

# --hints--

Você deve usar o seletor `#logo` para marcar o elemento `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Você deve dar à `img` uma `width` de `max(100px, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(100px, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
