---
id: 613e2546d0594208229ada50
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

No `header`, forneça um contexto sobre a página, aninhando um elemento `img`, um `h1` e um `nav`.

A `img` deve apontar para `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg` e ter no `id` o valor `logo`.

O `h1` deve conter o texto `HTML/CSS Quiz`.

# --hints--

Você deve adicionar um elemento `img` ao elemento `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Você deve adicionar um elemento `h1` ao elemento `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Você deve adicionar um elemento `nav` ao elemento `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Você deve colocar os elementos `img`, `h1` e `nav` nessa ordem.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Você deve dar ao elemento `img` um atributo `src` equivalente a `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Você deve dar ao elemento `img` um atributo `id` equivalente a `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Você deve atribuir ao elemento `h1` o texto `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
