---
id: 61408f155e798909b6908712
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Para permitir a navegação na página, adicione uma lista não ordenada com os três itens da lista a seguir:

- `INFO`
- `HTML`
- `CSS`

O texto da lista de itens deve estar envolvido pelas tags de âncora.

# --hints--

Você deve aninhar um elemento `ul` dentro de `nav`.

```js
assert.equal(document.querySelectorAll('nav > ul')?.length, 1);
```

Você deve aninhar três elementos `li` dentro do elemento `ul`.

```js
assert.equal(document.querySelectorAll('nav > ul > li')?.length, 3);
```

Você deve aninhar um elemento `a` dentro de cada elemento `li`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.length, 3);
```

Você deve dar ao primeiro elemento `a` o texto `INFO`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[0]?.textContent, 'INFO');
```

Você deve dar ao segundo elemento `a` o texto `HTML`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[1]?.textContent, 'HTML');
```

Você deve dar ao terceiro elemento `a` o texto `CSS`.

```js
assert.equal(document.querySelectorAll('nav > ul > li > a')?.[2]?.textContent, 'CSS');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
--fcc-editable-region--
      <nav>

      </nav>
--fcc-editable-region--
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}
```
