---
id: 6140883f74224508174794c0
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Faça com que o `header` ocupe toda a largura de seu contêiner pai, defina sua `height` como `50px` e defina a `background-color` para `#1b1b32`. Em seguida, defina o `display` para usar o _Flexbox_.

# --hints--

Você deve usar o seletor de elemento `header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

Você deve dar ao `header` uma `width` de `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

Você deve dar ao `header` uma `height` de `50px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

Você deve dar ao `header` uma `background-color` de `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

Você deve dar ao `header` um `display` de `flex`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(100px, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
