---
id: 614ccc21ea91ef1736b9b578
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Boas-vindas à primeira parte do questionário com acessibilidade. Como você está se tornando um desenvolvedor HTML e CSS experiente, já iniciaremos com o boilerplate básico.

Inicie esta jornada de acessibilidade fornecendo um atributo `lang` ao seu elemento `html`. Isto ajudará os leitores de tela a identificar o idioma da página.

# --hints--

Você deve dar ao elemento `html` um atributo `lang`. _Dica: você pode usar o valor `en` para o inglês._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
