---
id: 5d822fd413a79914d39e98d0
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Crea un elemento `div` all'interno del `body` con la classe `background-buildings`. Sarà il contenitore per un gruppo di edifici.

# --hints--

Dovresti creare un elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

L'elemento `div` dovrebbe essere all'interno del `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

L'elemento `div` dovrebbe avere la classe `background-buildings`

```js
assert([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

