---
id: 5d822fd413a79914d39e98cd
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

In CSS, è possibile selezionare tutto con un asterisco. Aggiungi un bordo a tutto il contenuto della pagina usando il selettore `*` e impostando `border` su `1px solid black`. Questo è un trucco che aiuta a visualizzare la posizione degli elementi e la loro dimensione. Lo rimuoverai più tardi.

# --hints--

Dovresti usare il selettore `*`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Dovresti usare la proprietà `border` per lo stile di tutti gli elementi.

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('border'));
```

Tutti gli elementi dovrebbero avere un bordo con il valore `1px solid black`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles?.border, '1px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
