---
id: 5d822fd413a79914d39e98e7
title: Step 31
challengeType: 0
dashedName: step-31
---

# --description--

Assegna ai sei nuovi elementi i seguenti valori di `width` e `height`: `10%` e `60%` a `.fb1`, `10%` e `40%` a `.fb2`, `10%` e `35%` a `.fb3`, `8%` e `45%` a `.fb4`, `10%` e `33%` a `.fb5`, `9%` e `38%` a `.fb6`.

# --hints--

Dovresti creare un selettore `.fb1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb1'));
```

Dovresti assegnare al selettore `.fb1` una proprietà `width` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.width, '10%');
```

Dovresti dare al selettore `.fb1` una proprietà `height` del `60%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb1')?.height, '60%');
```

Dovresti creare un selettore `.fb2`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb2'));
```

Dovresti assegnare al selettore `.fb2` una proprietà `width` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.width, '10%');
```

Dovresti dare al selettore `.fb2` una proprietà `height` del `40%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb2')?.height, '40%');
```

Dovresti creare un selettore `.fb3`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb3'));
```

Dovresti assegnare al selettore `.fb3` una proprietà `width` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.width, '10%');
```

Dovresti dare al selettore `.fb3` una proprietà `height` del `35%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb3')?.height, '35%');
```

Dovresti creare un selettore `.fb4`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb4'));
```

Dovresti assegnare al selettore `.fb4` una proprietà `width` dell'`8%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.width, '8%');
```

Dovresti dare al selettore `.fb4` una proprietà `height` del `45%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb4')?.height, '45%');
```

Dovresti creare un selettore `.fb5`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb5'));
```

Dovresti assegnare al selettore `.fb5` una proprietà `width` del `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.width, '10%');
```

Dovresti dare al selettore `.fb5` una proprietà `height` del `33%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb5')?.height, '33%');
```

Dovresti creare un selettore `.fb6`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.fb6'));
```

Dovresti assegnare al selettore `.fb6` una proprietà `width` del `9%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.width, '9%');
```

Dovresti dare al selettore `.fb6` una proprietà `height` del `38%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.fb6')?.height, '38%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div></div>
      <div></div>
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
      <div class="bb2"></div>
      <div class="bb3"></div>
      <div></div>
      <div class="bb4"></div>
      <div></div>
      <div></div>
    </div>

    <div class="foreground-buildings">
      <div class="fb1"></div>
      <div class="fb2"></div>
      <div class="fb3"></div>
      <div class="fb4"></div>
      <div class="fb5"></div>
      <div class="fb6"></div>
    </div>
  </body>
</html>
```

```css
:root {
  --building-color1: #aa80ff;
  --building-color2: #66cc99;
  --building-color3: #cc6699;
  --building-color4: #538cc6;
}

* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
  display: flex;
  align-items: flex-end;
  justify-content: space-evenly;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}

.bb1a {
  width: 70%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1b {
  width: 80%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1c {
  width: 90%;
  height: 10%;
  background-color: var(--building-color1);
}

.bb1d {
  width: 100%;
  height: 70%;
  background-color: var(--building-color1);
}

.bb2 {
  width: 10%;
  height: 50%;
  background-color: var(--building-color2);
}

.bb3 {
  width: 10%;
  height: 55%;
  background-color: var(--building-color3);
}

.bb4 {
  width: 11%;
  height: 58%;
  background-color: var(--building-color4);
}

.foreground-buildings {
  width: 100%;
  height: 100%;
  position: absolute;
  top: 0;
}
--fcc-editable-region--

--fcc-editable-region--

```

