---
id: 5d822fd413a79914d39e98cc
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Aggiungi un elemento `title` a `head` e dai al progetto il titolo `City Skyline`. Inoltre, annida un elemento `link` autoconcludente nell'elemento `head`. Dagli un attributo `rel` con valore di `stylesheet` e un attributo `href` con valore di `styles.css`.

# --hints--

Il codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

L'elemento `title` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Il progetto dovrebbe avere il titolo `City Skyline`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'city skyline')
```

Fai attenzione all'uso delle maiuscole/minuscole e all'ortografia nel titolo.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'City Skyline');
```

Il codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/)
```

Non dovresti cambiare i tag `head` esistenti. Assicurati di non aver eliminato il tag di chiusura.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

L'elemento `link` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<link[\w\W\s]+\/>/i));
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
