---
id: 5d822fd413a79914d39e98ca
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Aggiungi i tag di apertura e chiusura `html` sotto `DOCTYPE` in modo da avere un posto dove iniziare a inserire del codice. Assicurati di impostare la lingua sull'inglese.

# --hints--

La dichiarazione `DOCTYPE` dovrebbe essere all'inizio dell'HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

L'elemento `html` dovrebbe avere una tag di apertura con un attributo `lang` di `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

L'elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

I tag `html` dovrebbero essere nell'ordine corretto.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Dovresti avere un solo elemento `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
