---
id: 5d822fd413a79914d39e98cf
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Puoi vedere il `body` (è il rettangolo più interno della pagina); il rettangolo intorno a esso è l'elemento `html`. Fai in modo che il `body` riempia l'intero viewport impostando `height` su `100vh`. Rimuovi la proprietà `margin` predefinita dal `body` impostando `margin` a `0`. Infine, imposta la proprietà `overflow` su `hidden` per nascondere qualsiasi cosa si estenda oltre il viewport.

# --hints--

Dovresti usare il selettore `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

L'elemento `body` dovrebbe avere una proprietà `height` di `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

L'elemento `body` dovrebbe avere una proprietà `margin` di `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

L'elemento `body` dovrebbe avere la proprietà `overflow` impostata su `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--

```

