---
id: 65789506b30453080f77470c
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

So far, you have already met different data types:

- Immutable data types, such as integers, strings, tuples, and Booleans.
- Mutable data types, such as lists, and dictionaries.

A dictionary is identified by a pair of curly braces, `{}`.

Start by creating a variable called `copper` and assign it an empty dictionary using a pair of curly braces, in the same way you would create an empty list with a pair of square brackets.

# --hints--

You should have a variable called `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

Your `copper` variable should have the value of an empty dictionary. Use a pair of curly braces for that.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
