---
id: 61537c9eecea6a335db6da79
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Під елементом `.header` створіть новий елемент `div` з `id` зі значенням `gallery`.

В цьому елементі `#gallery` створіть десять елементів `img`.

# --hints--

Ви повинні створити другий елемент `div` в своєму елементі `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length === 2);
```

Ваш новий елемент `div` повинен йти після елемента `.header`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[0]?.classList?.contains('header'));
```

Ваш новий елемент `div` повинен мати `id` зі значенням `gallery`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.[1]?.id === 'gallery');
```

Ваш елемент `#gallery` повинен мати десять елементів `img`.

```js
assert(document.querySelector('#gallery')?.querySelectorAll('img')?.length === 10);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
  </body>
</html>
--fcc-editable-region--
```

```css

```
