---
id: 615392916d83fa4f02f7e2cf
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Мабуть, ви помітили, що всі зображення змістились в один ряд.

Властивість `flex-wrap` визначає як ваші предмети повинні поводитись, якщо flex-контейнер замалий. Встановлення цієї властивості на `wrap` дозволить вашим предметам загортатись в наступному рядку/стовпці (залежить від головної осі), де `nowrap` запобіжить обгортці елементів. Коли встановлено `nowrap`, предмети зморщуються або переповнюються.

Надайте селектору `#gallery` властивість `flex-wrap` зі значенням `wrap`. Ви повинні побачити, як ваші зображення утворюють чотириколонковий макет. Це тому, що в попередньому кроці ви встановили їхній `width` на `25%`.

# --hints--

Ваш селектор `#gallery` повинен мати властивість `flex-wrap` зі значенням `wrap`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery')?.flexWrap === 'wrap');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: Arial;
  background: #EBE7E7;
}

.header {
  text-align: center;
  padding: 32px;
  background: #E0DDDD;
}

--fcc-editable-region--
#gallery {
  display: flex;
  flex-direction: row;
}
--fcc-editable-region--

#gallery img {
  width: 25%;
  height: 300px;
}
```
