---
id: 61537c9eecea6a335db6da79
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Під елементом `.header` створіть новий елемент `div` та призначте йому `class` зі значенням `gallery`. Цей `div` діятиме як контейнер для зображень галереї.

В цьому елементі `.gallery` створіть дев'ять елементів `img`.

# --hints--

Ви повинні створити елемент `div` в своєму елементі `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Ваш новий елемент `div` повинен мати `class` зі значенням `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Ваш новий елемент `div` повинен йти після елемента `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Ваш елемент `.gallery` повинен мати дев'ять елементів `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
