---
id: 61537c5f81f0cf325b4a854c
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

В межах елемента `body` створіть `div` з `class` зі значенням `header`.

Всередині елемента `.header` вкладіть елемент `h1` з текстом `CSS FLEXBOX PHOTO GALLERY`.

# --hints--

Ви повинні мати елемент `div` в межах елемента `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('div'))
```

Ваш елемент `div` повинен мати `class` зі значенням `header`.

```js
assert(document?.querySelector('body')?.querySelector('div')?.classList?.contains('header'))
```

Ваш елемент `.header` повинен мати елемент `h1`.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Ваш `h1` повинен мати текст `CSS FLEXBOX PHOTO GALLERY`. Пам'ятайте, що вживання великої літери важливе.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'CSS FLEXBOX PHOTO GALLERY')
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
