---
id: 6153897c27f6334716ee5abe
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Вирівняйте текст `.header` за центром. Зробіть текст великими літерами, використовуючи властивість `text-transform` зі значенням `uppercase`.

Надайте йому відступ зі значенням `32px` з всіх боків. Встановіть значення кольору фону на `#0a0a23` та тексту на `#fff`.

Додайте `border-bottom` зі значенням `4px solid #fdb347`.


# --hints--

Ви повинні мати селектор `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Ваш селектор `.header` повинен мати властивість `text-align` зі значенням `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign === 'center');
```

Ваш селектор `.header` повинен мати властивість `text-transform` зі значенням `uppercase`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform === 'uppercase');
```

Ваш селектор `.header` повинен мати властивість `padding` зі значенням `32px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.padding === '32px');
```

Ваш селектор `.header` повинен мати властивість `background-color` зі значенням `#0a0a23`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor === 'rgb(10, 10, 35)');
```

Ваш селектор `.header` повинен мати властивість `color` зі значенням `#fff`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.header')?.color === 'rgb(255, 255, 255)');
```

Ваш селектор `.header` повинен мати властивість `border-bottom` зі значенням `4px solid #fdb347`.

```js
console.log(new __helpers.CSSHelp(document).getStyle('.header'))
assert(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom === '4px solid rgb(253, 179, 71)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
