---
id: 61537bb9b1a29430ac15ad38
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

В межах елемента `head` додайте елемент `title` з текстом `Photo Gallery` та елемент `link`, щоб додати свій файл `styles.css` до сторінки.

# --hints--

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ваш проєкт повинен мати титул `Photo Gallery`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Пам’ятайте, що вживання великої літери та написання мають значення.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
