---
id: 6153893900438b4643590367
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Вирівняйте текст `.header` за центром. Make the text uppercase using the `text-transform` property with `uppercase` as the value.

Give it a padding of `32px` on all sides. Set the background to `#0a0a23` and the text to `#fff` as the color values.

Додайте `border-bottom` зі значенням `4px solid #fdb347`.

# --hints--

Ви повинні мати селектор `.header`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

Селектор `.header` повинен мати властивість `text-align` зі значенням `center`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

Селектор `.header` повинен мати властивість `text-transform` зі значенням `uppercase`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

Селектор `.header` повинен мати властивість `padding` зі значенням `32px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

Селектор `.header` повинен мати властивість `background-color` зі значенням `#0a0a23`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

Селектор `.header` повинен мати властивість `color` зі значенням `#fff`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

Селектор `.header` повинен мати властивість `border-bottom` зі значенням `4px solid #fdb347`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
