---
id: 61537a8054753e2f1f2a1574
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

В межах елемента `head` додайте теґ `meta`, що містить `name` зі значенням `viewport` та `content` зі значенням `width=device-width, initial-scale=1`.

Також додайте теґ `meta`, що містить `charset` зі значенням `UTF-8`.

# --hints--

Ви повинні мати два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Один елемент `meta` повинен мати `name` зі значенням `viewport` та `content` зі значенням `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

Інший елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
