---
id: 61537d86bdc3dd343688fceb
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Потім надайте кожному `img` атрибут `src` відповідно до його порядку в документі. Перший `img` повинен мати `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. Решта повинні бути такими самими, окрім заміни `1` на номер `img` в документі.

# --hints--

Всі дев'ять елементів `img` повинні мати атрибут `src`.

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

Ваш перший елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`.

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Ваш другий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg`.

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Ваш третій елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg`.

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Ваш четвертий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg`.

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

Ваш п'ятий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg`.

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Ваш шостий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg`.

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Ваш сьомий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg`.

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Ваш восьмий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg`.

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Ваш дев'ятий елемент `img` повинен мати атрибут `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg`.

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
