---
id: 615389bd81347947ea7ba896
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Flexbox – це одновимірний макет CSS, який може керувати тим, як елементи розміщуються та вирівнюються всередині контейнера.

Щоб використати його, надайте елементу властивість `display` зі значенням `flex`. Це зробить елемент <em>гнучким контейнером</em>. Будь-які прямі дочірні елементи гнучкого контейнера називаються <em>гнучкими елементами</em>.

Створіть селектор `.gallery` та зробіть його гнучким контейнером.

# --hints--

Ви повинні мати селектор `.gallery`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Ваш селектор `.gallery` повинен мати властивість `display` зі значенням `flex`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.display === 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
