---
id: 6153908a366afb4d57185c8d
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Flexbox має головну та поперечну осі. Головна вісь визначається властивістю `flex-direction`, яка має чотири можливі значення:

- `row` (за замовчуванням): горизонтальна вісь з гнучкими елементами зліва направо
- `row-reverse`: горизонтальна вісь з гнучкими елементами справа наліво
- `column`: вертикальна вісь з гнучкими елементами зверху вниз
- `column-reverse`: вертикальна вісь з гнучкими елементами знизу вгору

**Примітка**. Осі та напрямки відрізнятимуться залежно від напрямку тексту. Показані значення – для напрямку тексту зліва направо.

Спробуйте різні значення, щоб побачити, як вони впливають на макет.

Коли ви закінчите, встановіть явний `flex-direction` зі значенням `row` для елемента `.gallery`.

# --hints--

Ваш селектор `.gallery` повинен мати властивість `flex-direction` зі значенням `row`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection === 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

--fcc-editable-region--
.gallery {
  display: flex;
}
--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
