---
id: 61537c5f81f0cf325b4a854c
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Додайте елемент `header` в елемент `body` та призначте йому клас `header`.

Всередині `header` створіть `h1` з текстом `css flexbox photo gallery`.

# --hints--

Ви повинні мати елемент `header` в межах елемента `body`.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Ваш елемент `header` повинен мати `class` зі значенням `header`.

```js
assert(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Ваш елемент `header` повинен мати елемент `h1` всередині.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

Ваш елемент `h1` повинен мати текст `css flexbox photo gallery` всередині.

```js
assert(document?.querySelector('.header')?.querySelector('h1')?.innerText === 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
