---
id: 615f171d05def3218035dc85
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Зараз ваші зображення занадто великі, тому ви не помітите змін CSS.

Створіть селектор `#gallery img`, щоб націлити свої зображення. Надайте йому властивість `width` зі значенням `25%`.

Також встановіть властивість `height` на `300px`, щоб зберегти однаковий розмір зображень.

# --hints--

Ви повинні мати селектор `#gallery img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#gallery img'))
```

Ваш селектор `#gallery img` повинен мати властивість `width` зі значенням `25%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery img')?.width === '25%')
```

Ваш селектор `#gallery img` повинен мати властивість `height` зі значенням `300px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#gallery img')?.height === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="header">
      <h1>CSS FLEXBOX PHOTO GALLERY</h1>
    </div>
    <div id="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg"/>
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/10.jpg"/>
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
