---
id: 616971b3cd990186b66c99a1
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Agora que seu arquivo CSS externo está configurado, você pode começar a estilizar a página.

Como lembrete, aqui está como apontar para um elemento de parágrafo e alinhá-lo à direita:

```css
p {
  text-align: right;
}
```

Cria uma nova regra de CSS vinculada ao elemento `h1` e defina sua propriedade como `text-align` com o valor `center`.

# --hints--

Você deve usar o seletor `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

O seletor `h1` deve definir a propriedade `text-align` como `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
