---
id: 616965351e74d4689eb6de30
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Você pode ter vários elementos `meta` de fechamento automático em uma página da web. O que distingue um elemento `meta` do outro é o atributo. Você deve adicionar um novo elemento meta para cada atributo que você deseja especificar.

Adicione outro elemento `meta` de fechamento automático dentro do elemento `head`. Dê a ele o atributo `name` definido como `viewport` e o atributo `content` definido como `width=device-width, initial-scale=1.0` para que a página tenha a mesma aparência em todos os dispositivos.

# --hints--

Você deve ter dois elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

O novo elemento `meta` deve ser de fechamento automático.

```js
assert(code.match(/<\/meta>/i) === null);
```

O novo elemento `meta` deve ter o atributo `name` definido como `viewport` e o atributo `content` definido como `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
