---
id: 61696ef7ac756c829f9e4048
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Neste projeto, você trabalhará com um arquivo CSS externo para estilizar a página. Já criamos um arquivo de `styles.css` para você. Mas antes que você possa usá-lo, você precisará vinculá-lo à página.

Adicione um elemento `link` dentro do elemento `head`. Dê a ele o atributo `rel` definido como `"stylesheet"` e o atributo `href` definido como `"styles.css"`.

# --hints--

Você deve ter um elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

O elemento `link` deve ser um elemento nulo, sem uma tag final `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

O elemento `link` deve ter o atributo `rel` com o valor de `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor de `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

O elemento `link` deve estar dentro do elemento `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
