---
id: 616968c2c94c8071b349c146
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Agora você está pronto para começar a adicionar conteúdo à página.

Dentro do elemento `body`, aninhe um elemento `h1` com o texto `CSS Color Markers`.

# --hints--

O código deve ter um elemento `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Você deve acrescentar uma tag de abertura para `h1`.

```js
assert(code.match(/<h1\s*>/i));
```

Você deve acrescentar uma tag de fechamento para `h1`.

```js
assert(code.match(/<\/h1\s*>/i));
```

O elemento `h1` deve estar dentro do elemento `body`.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

O elemento `h1` deve conter o texto `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Lembre-se, maiúsculas, minúsculas e a ortografia são importantes para o `h1`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
