---
id: 61695d1fbc003856628bf561
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Para dizer aos navegadores como codificar caracteres na sua página, defina o `charset` como `utf-8`. `utf-8` é um conjunto de caracteres universal que inclui quase todos os caracteres de todos os idiomas humanos.

Dentro do elemento `head`, aninhe um elemento `meta` com o atributo `charset` definido para `utf-8`. Assim como os elementos img, os elementos `meta` são de fechamento automático e não precisam de tags de fechamento.

# --hints--

Você deve ter apenas um elemento `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

O elemento `meta` deve ser de fechamento automático.

```js
assert(code.match(/<\/meta>/i) === null);
```

O elemento `meta` deve ter o atributo `charset` com o valor `utf-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
