---
id: 61695d1fbc003856628bf561
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

O atributo `charset` especifica a codificação de caracteres usada pelo documento. `utf-8` (Formato de Transformação Unicode, de 8 bits) é um padrão de codificação de caracteres usado para comunicação eletrônica.

Dentro do elemento `head`, aninhe um elemento `meta` com o atributo `charset` definido como `"utf-8"`.

# --hints--

Você deve ter apenas um elemento `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

O elemento `meta` deve ser um elemento nulo, sem uma tag final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

O elemento `meta` deve ter o atributo `charset` com o valor `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

O elemento `meta` deve estar dentro do elemento `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
