---
id: 61696ef7ac756c829f9e4048
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Neste projeto, você trabalhará com um arquivo CSS externo para estilizar a página. Já criamos um arquivo de `styles.css` para você. Mas antes que você possa usá-lo, você precisará vinculá-lo à página.

Adicione um elemento `link` dentro do elemento `head`. Dê a ele um atributo `rel` com o valor `stylesheet`, um atributo `type` com o valor `text/css` e um atributo `href` com o valor `styles.css`.

# --hints--

O código deve ter um elemento `link`.

```js
assert(code.match(/<link/i)?.length === 1);
```

O elemento `link` deve ser de fechamento automático.

```js
assert(code.match(/<\/link>/i) === null);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
const head = code.match(/<head>(.|\r|\n)*<\/head>/i)?.[0];
assert(head.match(/<link/i)?.length === 1)
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
assert(code.match(/<link[\s\S]*?rel=('|"|`)stylesheet\1/gi)?.length === 1);
```

O elemento `link` deve ter o atributo `type` com o valor `text/css`.

```js
assert(code.match(/<link[\s\S]*?type=('|"|`)text\/css\1/gi)?.length === 1);
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`. Lembre-se, letras maiúsculas e minúsculas são levadas em conta quando arquivo externos são vinculados.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```
`css`
