---
id: 61696ef7ac756c829f9e4048
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Neste projeto, você trabalhará com um arquivo CSS externo para estilizar a página. Já criamos um arquivo de `styles.css` para você. Mas antes que você possa usá-lo, você precisará vinculá-lo à página.

Adicione um elemento `link` dentro do elemento `head`. Dê a ele o atributo `rel` definido como `stylesheet` e o atributo `href` definido como `styles.css`.

# --hints--

Você deve ter um elemento `link` de autofechamento.

```js
assert(document.querySelectorAll('link').length === 1);
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert.exists(document.querySelector('head > link'));
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
