---
id: 616d50b93ba424d6282c99cf
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

A cor de fundo foi aplicada, mas como o elemento `div` do marcador não tem conteúdo dentro dele, ele não tem qualquer altura por padrão.

Na regra do CSS `.marker`, defina a propriedade `height` como `25px` e a propriedade `width` como `200px`

# --hints--

A regra do CSS `.marker` deve ter a propriedade `width` definida para `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

A regra do CSS `.marker` deve ter uma propriedade `height` definida para `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
