---
id: 616d3a67ccf800ad94ec89ae
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Agora, você vai adicionar alguns elementos que mais pra frente serão estilizados em marcadores de cor.

Primeiro, dentro do elemento `body`, adicione um elemento `div` e defina seu atributo `class` para `container`. Certifique-se que o elemento `div` está abaixo do elemento `h1`.

# --hints--

Você deve acrescentar uma tag de abertura para `div`.

```js
assert(code.match(/<div.*>/i));
```

Você deve acrescentar uma tag de fechamento para `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

O elemento `div` deve estar dentro do elemento `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

O elemento `div` deve estar abaixo do elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

O elemento `div` deve ter o atributo `class` definido como `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
