---
id: 616d525007b8c5d8b3308b1c
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

O marcador ficaria melhor se estivesse centralizado na página. Uma maneira fácil de fazer isso é com a <dfn>propriedade abreviada</dfn> `margin` .

No último projeto, você definiu a área de margem dos elementos separadamente com propriedades como `margin-top` e `margin-left`. A propriedade abreviada `margin` facilita a definição de várias áreas de margem ao mesmo tempo.

Para centralizar seu marcador na página, defina a sua propriedade `margin` como `auto`. Isso define `margin-top`, `margin-right`, `margin-bottom` e `margin-left` como `auto`.

# --hints--

A regra do CSS `.marker` deve ter uma propriedade `margin` definida como `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
