---
id: 616d595270d902f0e7086e18
title: Passo 15
challengeType: 0
dashedName: step-15
---

# --description--

Embora você tenha três elementos separados que o marcador `div`, eles se parecem com um único retângulo grande. Você deve adicionar um pouco de espaço entre eles para tornar mais fácil ver cada elemento.

Quando a propriedade abreviada `margin` tiver dois valores, ela define `margin-top` e `margin-bottom` para o primeiro valor, e `margin-left` e `margin-right` para o segundo valor.

Na regra de CSS `.marker`, defina a propriedade `margin` para `10px auto`.

# --hints--

A regra do CSS `.marker` deve ter uma propriedade `margin` definida para `10px auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
