---
id: 616d4a84b756d9c4b8255093
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

É hora de adicionar um pouco de cor ao marcador. Lembre-se de que uma maneira de adicionar cor a um elemento é usar uma <dfn>palavra-chave de cor</dfn> como `black`, `cyan` ou `yellow`.

Como lembrete, aqui está como vincular a classe `freecodecamp`:

```css
.freecodecamp {

}
```

Crie uma nova regra de CSS vinculada à classe `marker` e defina sua propriedade `background-color` para `red`.

**Observação:** você não verá nenhuma alteração após adicionar o CSS.

# --hints--

Você deve usar um seletor de classe para a classe `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

A regra de CSS `.marker` deve ter a propriedade `background-color` definida para `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
