---
id: 61764c602bee6974e7790f35
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Na escola, você pode ter aprendido que vermelho, amarelo e azul são as cores primárias e aprendeu a criar novas cores misturando essas cores. No entanto, este é um modelo obsoleto.

Hoje em dia, existem dois modelos de cores: o modelo <dfn>aditivo</dfn> RGB (vermelho, verde, azul), usado em dispositivos eletrônicos, e o modelo <dfn>subtrativo</dfn> CMYK (ciano, magenta, amarelo, preto), usado em impressão. Neste projeto, você trabalhará com o modelo RGB.

Primeiro, adicione a `class` `one` ao primeiro elemento do marcador `div`.

# --hints--

Adicione a class `one` ao primeiro elemento do marcador `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

O primeiro marcador `div` deve ter os marcadores de classe `marker` e `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" type="text/css" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
