---
id: 6176519636a76b810ab1219a
title: Passo 20
challengeType: 0
dashedName: step-20
---

# --description--

Crie uma regra de CSS para a classe `two` e defina sua propriedade `background-color` para `green`.

Além disso, crie uma regra separada de CSS para a classe `three` e defina sua propriedade `background-color` para `blue`.

# --hints--

Você deve usar um seletor de classe para a classe `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

A regra de CSS `.two` deve ter sua propriedade `background-color` definida para `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

Você deve usar um seletor de classe para a classe `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

A regra de CSS `.three` deve ter a sua propriedade `background-color` definida para `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
