---
id: 617b97abd9f3f61d1590b815
title: Passo 26
challengeType: 0
dashedName: step-26
---

# --description--

No modelo aditivo de cor RGB, <dfn>cores primárias</dfn> são cores que, quando combinadas, criam branco puro. Mas, para que isso aconteça, cada cor tem que estar em sua maior intensidade.

Antes de combinar cores, defina seu marcador verde de volta ao verde puro. Defina os valores vermelho, verde e azul no máximo de `255` para a função `rgb` na regra de CSS `.two`.

# --hints--

A regra do CSS `.two` deve ter uma propriedade `background-color` definida para `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 127, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 255);
}

```
