---
id: 617b65579ce424bf5f02ca73
title: Passo 21
challengeType: 0
dashedName: step-21
---

# --description--

Existem dois modelos de cores: o modelo <dfn>aditivo</dfn> RGB (vermelho, verde, azul), usado em dispositivos eletrônicos, e o modelo <dfn>subtrativo</dfn> CMYK (ciano, magenta, amarelo, preto), usado em impressão.

Neste projeto, você trabalhará com o modelo RGB. Isso significa que as cores começam como pretas e mudam quando são introduzidos diferentes níveis de vermelho, verde e azul. Uma maneira fácil de ver isso é com a função do CSS `rgb`.

Crie uma nova regra do CSS vinculada à classe `container` e defina sua propriedade `background-color` para `rgb(0, 0, 0)`.

# --hints--

Você deve usar um seletor de classe para a classe alvo `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

A regra do CSS `.container` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
