---
id: 617b954d9f4f6217a749380e
title: Passo 25
challengeType: 0
dashedName: step-25
---

# --description--

Agora, adicione um pouco mais de espaço vertical entre os seus marcadores e a borda do elemento `container` em que estão inseridos.

Na regra do CSS `.container`, use a propriedade `padding` para adicionar `10px` de espaçamento no topo e na parte inferior e defina o espaçamento esquerdo e direito para `0`. Funciona parecido com a propriedade abreviada `margin` que você usou anteriormente.

# --hints--

Você não deve remover a propriedade `background-color` e seu valor da regra do CSS de `.container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

A regra do CSS `.container` deve ter uma propriedade `padding` definida para `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
