---
id: 617b8b38f32bf2080a140675
title: Passo 22
challengeType: 0
dashedName: step-22
---

# --description--

Uma função é um trecho de código que pode receber uma entrada e executar uma ação específica. A função do CSS `rgb` aceita valores, ou <dfn>argumentos</dfn>, para vermelho, verde e azul, e produz uma cor:

```css
rgb(red, green, blue);
```

Cada valor vermelho, verde e azul é um número de `0` a `255`. `0` significa que há 0% dessa cor e é preta. `255` significa que há 100% dessa cor.

Na regra do CSS `.one`, substitua a palavra-chave de cor `red` pela função `rgb`. Atualize a função `rgb` para que o valor para vermelho seja `255`, o valor do verde `0` e o valor de azul `0`.

# --hints--

A regra do CSS `.one` não deve usar a palavra-chave de cor `red` para definir a propriedade `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

A regra do CSS `.one` deve ter uma propriedade `background-color` definida para `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
