---
id: 617bbb6b97a83f6d1f7d6e38
title: Passo 31
challengeType: 0
dashedName: step-31
---

# --description--

Agora que você está familiarizado com as cores secundárias, você aprenderá como criar <dfn>cores terciárias</dfn>. As cores terciárias são criadas combinando uma cor primária com uma cor secundária.

Para criar a cor terciária laranja, atualize a função `rgb` na regra do CSS `.one` de modo que o vermelho esteja no valor máximo, e defina verde para `127`.

# --hints--

A regra do CSS `.one` deve ter uma propriedade `background-color` definida para `rgb(255, 127, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 255, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 255);
}

.three {
  background-color: rgb(255, 0, 255);
}

```
