---
id: 617bd6ec666b1da2587e4e37
title: Passo 34
challengeType: 0
dashedName: step-34
---

# --description--

Existem mais três cores terciárias: chartreuse verde (verde + amarelo), azure (azul + ciano), e rosa (vermelho + magenta).

Para criar a cor chartreuse verde, atualize a função `rgb` na regra `.one` do CSS de modo que o vermelho esteja em `127` e defina verde para seu nível máximo.

Para criar a cor azure, atualize a função `rgb` na regra `.two` do CSS de modo que o verde esteja em `127` e azul esteja em seu nível máximo.

E para rosa, que às vezes é chamado de rosa brilhante, atualize a função de `rgb` na regra `.three` do CSS para que o azul esteja em `127` e vermelho esteja no valor máximo.

# --hints--

A regra do CSS `.one` deve ter uma propriedade `background-color` definida para `rgb(127, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(127, 255, 0)');
```

A regra do CSS `.two` deve ter uma propriedade `background-color` definida para `rgb(0, 127, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 127, 255)');
```

A regra do CSS `.three` deve ter uma propriedade `background-color` definida como `rgb(255, 0, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(255, 0, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 127, 0);
}

.two {
  background-color: rgb(0, 255, 127);
}

.three {
  background-color: rgb(127, 0, 255);
}
--fcc-editable-region--

```
