---
id: 617b8e0d93a8d10d9a90e720
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

Observe que a `background-color` do seu marcador ainda está vermelha. Isso porque você definiu o valor vermelho da função `rgb` com o valor máximo de `255` (ou seja, 100% vermelho) e definiu os valores verde e azul para `0`.

Agora use a função `rgb` para definir as outras cores.

Na regra do CSS `.two` use a função `rgb` para definir o `background-color` para o valor máximo de verde, e `0` para os outros valores. Na regra do CSS `.three` use a função `rgb` para definir o `background-color` para o valor máximo de azul, e `0` para os outros valores.

# --hints--

A regra do CSS `.two` não deve usar a palavra-chave de cor `green` para definir a propriedade `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

A regra do CSS `.two` deve ter uma propriedade `background-color` definida para `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

A regra do CSS `.three` não deve usar a palavra-chave de cor `blue` para definir a propriedade `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

A regra do CSS `.three` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
