---
id: 618a00ed1ca871a2b3aca0cb
title: Passo 35
challengeType: 0
dashedName: step-35
---

# --description--

Agora que você passou por todas as cores primárias, secundárias e terciárias do circulo cromático, vai ser mais fácil entender outros conceitos de teoria de cores e como eles afetam o design.

Primeiro, nas regras `.one`, `.two` e `.three` do CSS, ajuste os valores na função `rgb` para que a `background-color` de cada elemento esteja definida como preto puro. Lembre-se de que a função `rgb` usa um modelo de cor aditivo, onde as cores começam a ser pretas e mudam conforme os valores de vermelho, verde e azul.

# --hints--

A regra do CSS `.one` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

A regra do CSS `.two` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

A regra do CSS `.three` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
