---
id: 618a0927005553b74bfa05ae
title: Passo 36
challengeType: 0
dashedName: step-36
---

# --description--

O circulo cromático é um círculo onde cores, ou <dfn>matizes</dfn>, semelhantes ficam próximas e cores diferentes ficam mais distantes. Por exemplo, vermelho puro está entre as cores da rosa e da laranja.

Duas cores que são opostas uma da outra na roda de cores são chamadas de <dfn>cores complementares</dfn>. Se duas cores complementares forem combinadas, produzirão cinza. Mas quando elas são colocados lado a lado, essas cores produzem um forte contraste visual e parecem mais brilhantes.

Defina o valor vermelho no máximo de `255` para produzir vermelho puro na função `rgb` da regra do CSS `.one`. Na função `rgb` da regra de CSS `.two`, defina os valores de verde e azul para o valor máximo de `255` para produzir ciano.

# --hints--

A regra do CSS `.one` deve ter uma propriedade `background-color` definida para `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

A regra do CSS `.two` deve ter uma propriedade `background-color` definida para `rgb(0, 255, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
