---
id: 619b7fd56aa2253778bcf5f7
title: Passo 48
challengeType: 0
dashedName: step-48
---

# --description--

O modelo de cores de <dfn>HSL</dfn>, ou matiz, saturação e luz, é outra forma de representar cores.

A função do CSS hsl aceita 3 valores: um número de 0 a 360 para o matiz, uma porcentagem de 0 a 100 para saturação e uma porcentagem de 0 a 100 para a luz.

Se você imaginar um circulo cromático, a cor vermelha fica a 0 graus, verde é a 120 graus e azul está a 240 graus.

Saturação é a intensidade de uma cor de 0%, ou cinza, até 100% para cor pura. Você deve adicionar o sinal `%` de percentual aos valores de saturação e luz.

A luz é o nível de brilho no qual uma cor aparece, de 0%, ou preta completa, até 100%, branco total, com 50% de neutralidade.

Na regra do CSS`.blue`, use a função `hsl` para alterar a propriedade `background-color` para azul puro. Defina a matiz como `240`, a saturação como `100%`, e a luz como `50%`.

# --hints--

A regra do CSS `.blue` deve ter uma propriedade `background-color` definida para `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
