---
id: 619b761916dac02643940022
title: Passo 46
challengeType: 0
dashedName: step-46
---

# --description--

Uma maneira muito comum de aplicar cor a um elemento com CSS é com <dfn>hexadecimal</dfn> ou valores hexadecimais. Embora os valores hexadecimais pareçam complicados, eles são apenas uma outra forma de valores RGB.

Os valores de cores hexadecimais começam com um caractere `#` e recebem seis caracteres de 0-9 e A-F. O primeiro par de caracteres representa vermelho, o segundo par representa verde, e o terceiro par representa azul. Por exemplo, `#4B5320`.

No seletor de classe `.green` do CSS, defina a propriedade `background-color` para um código de cor hexadecimal com os valores `00` para vermelho, `FF` para verde e `00` azul.

# --hints--

A regra de CSS `.green` deve ter a sua propriedade `background-color` definida para `#00FF00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
