---
id: 619b7424f43ec9215e538afe
title: Passo 44
challengeType: 0
dashedName: step-44
---

# --description--

Em seguida, mude a classe `two` para `green` no segundo marcador `div` e a classe `three` para `blue` no terceiro marcador `div`.

# --hints--

O segundo marcador `div` não deve ter a classe `two`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(!containerSecondChild?.classList?.contains('two'));
```

O segundo marcador `div` deve ter os marcadores de classe `marker` e `green`.

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert(containerSecondChild?.classList?.contains('marker') && containerSecondChild?.classList?.contains('green'));
```

O terceiro marcador `div` não deve ter a classe `three`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(!containerThirdChild?.classList?.contains('three'));
```

O terceiro marcador `div` deve ter os marcadores de classe `marker` e `blue`.

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert(containerThirdChild?.classList?.contains('marker') && containerThirdChild?.classList?.contains('blue'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      --fcc-editable-region--
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}

```
