---
id: 619b7c3c83de403126b69c1e
title: Passo 47
challengeType: 0
dashedName: step-47
---

# --description--

Você já pode estar familiarizado com valores decimais, ou base 10, que vão de 0 a 9. Valores hexadecimais ou de base 16 vão de 0 a 9 e, depois, de A a F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Com cores hexadecimais, `00` é 0% dessa cor e `FF` é 100%. Então, `#00FF00` traduz 0% vermelho, 100% verde e 0% azul, sendo igual a `rgb(0, 255, 0)`.

Diminua a intensidade de verde definindo o valor verde da cor hexadecimal para `7F`.

# --hints--

A regra de CSS `.green` deve ter a sua propriedade `background-color` definida para `#007F00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
