---
id: 619b74fa777a2b2473c94f82
title: Passo 45
challengeType: 0
dashedName: step-45
---

# --description--

Atualize o seletor de classe `.two` do CSS para que ele vincule à nova classe `green`. E atualize o seletor de classe `.three` para que ele vincule à nova classe `blue`.

# --hints--

Você deve usar um seletor de classe `.two`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

Você deve usar um seletor de classe para a classe `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

A regra do CSS `.green` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Você deve usar um seletor de classe `.three`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

Você deve usar um seletor de classe para a classe alvo `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

A regra do CSS `.blue` deve ter uma propriedade `background-color` definida para `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
