---
id: 61a489b8579e87364b2d2cdb
title: Passo 49
challengeType: 0
dashedName: step-49
---

# --description--

Você aprendeu algumas maneiras de definir cores lisas no CSS, mas você também pode usar uma transição de cor, ou <dfn>gradient</dfn>, em um elemento.

Um gradiente é quando uma cor muda para outra. A função CSS `linear-gradient` permite que você controle a direção da transição ao longo de uma linha e quais cores são usadas.

Uma coisa a lembrar é que a função `linear-gradient` realmente cria um elemento `image` e que ela geralmente é pareada com a propriedade `background` que pode aceitar uma imagem como um valor.

Na regra do CSS `.red`, altere a propriedade `background-color` para `background`.

# --hints--

A regra do CSS `.red` deve ter uma propriedade `background` com o valor `rgb(255, 0, 0)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
