---
id: 61a493ead935c148d2b76312
title: Passo 50
challengeType: 0
dashedName: step-50
---

# --description--

A função `linear-gradient` é muito flexível -- aqui está a sintaxe básica que você vai usar neste tutorial:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` é a direção da linha usada para a transição. `color1` e `color2` são argumentos de cor, que são as cores que serão usadas na transição em si. Elas podem ser de qualquer tipo de cor, incluindo palavras-chave de cores, hex, `rgb`, ou `hsl`.

Agora você aplicará um gradiente vermelho a verde ao longo de uma linha de 90 graus ao primeiro marcador.

Primeiro, na regra do CSS `.red`, defina a propriedade `background` como `linear-gradient()`, e aplique nele o valor `90deg` como a `gradientDirection`.

# --hints--

A regra do CSS `.red` deve ter uma propriedade `background` com o valor `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
