---
id: 61a49d15bdbb5e57cc6fd280
title: Passo 54
challengeType: 0
dashedName: step-54
---

# --description--

Color-stops permitem que você ajuste onde as cores são colocadas ao longo da linha gradiente. São uma unidade de comprimento como `px` ou porcentagens que seguem uma cor na função `linear-gradient`.

Por exemplo, neste gradiente vermelho preto, a transição do vermelho para o preto ocorre no ponto 90% ao longo da linha gradiente. Então, o vermelho ocupa a maior parte do espaço disponível:

```css
linear-gradient(90deg, red 90%, black);
```

Na função`linear-gradient`, adicione um color-stop de `75%` depois do primeiro argumento de cor vermelha. Não adicione pontos de cores para os outros argumentos de cores.

# --hints--

A regra do CSS `.red` deve ter uma propriedade `background` definida para `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
