---
id: 61a5c906ab73313316e342f0
title: Passo 65
challengeType: 0
dashedName: step-65
---

# --description--

Se nenhum argumento `gradientDirection` é fornecido para a função `linear-gradient`, ele organiza as cores de cima para baixo, ou ao longo de uma linha de 180 graus, por padrão.

Limpe seu código um pouco mais removendo o argumento de `gradientDirection` de ambas as funções `linear-gradient`.

# --hints--

Você deve remover os argumentos `gradientDirection` das funções `linear-gradient` em suas regras de CSS `.red` e `.green`.

```js
assert(!code.match(/linear-gradient\s*\(\s*180deg/gi));
```

A regra do CSS `.red` deve ter uma propriedade `background` definida para `linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

A regra do CSS `.green` deve ter uma propriedade `background` com o valor `linear-gradient(#55680D, #71F53E, #116C31)`.

```js
assert.include(['linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))', 'rgba(0,0,0,0)linear-gradient(rgb(85,104,13),rgb(113,245,62),rgb(17,108,49))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.green')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}
--fcc-editable-region--

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
