---
id: 61adc60b69cd4b87739d2174
title: Passo 71
challengeType: 0
dashedName: step-71
---

# --description--

Crie uma nova regra do CSS vinculada à classe `sleeve`. Defina a propriedade `width` como `110px` e a propriedade `height` para `25px`.

# --hints--

Você deve usar um seletor de classe para a classe alvo `sleeve`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve'));
```

O seletor `.sleeve` deve ter a propriedade `width` definida para `110px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.width === '110px');
```

O seletor `.sleeve` deve ter a propriedade `height` definida para `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
