---
id: 61b093219e7fc020b43b1bb4
title: Passo 74
challengeType: 0
dashedName: step-74
---

# --description--

Outra maneira de definir a opacidade para um elemento é com o <dfn>canal alfa</dfn>. Semelhante à propriedade `opacity`, o canal alfa controla o nível de transparência ou de opacidade de uma cor.

Você já definiu a opacidade da capa com uma cor nomeada e a propriedade `opacity`, mas você pode adicionar um canal alfa para as outras propriedades de cor do CSS.

Dentro da regra `.sleeve`, remova a propriedade `opacity` e seu valor.

# --hints--

Na regra do CSS `.sleeve`, remova a propriedade `opacity` e seu valor.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
  opacity: 0.5;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
