---
id: 61b0936d9e7fc020b43b1bb8
title: Passo 78
challengeType: 0
dashedName: step-78
---

# --description--

Parece que sua capa desapareceu, mas não se preocupe -- ela ainda está lá. Ocorre que a nova tampa `div` está ocupando toda a largura do marcador e empurrando a capa para a próxima linha.

Isso ocorre porque a propriedade padrão `display` para elementos `div` é o `block`. Então, quando dois elementos `block` estão próximos um do outro, eles empilham como blocos de verdade. Por exemplo, os elementos do marcador estão todos empilhados em cima um do outro.

Para posicionar dois elementos `div` na mesma linha, defina suas propriedades `display` para `inline-block`.

Crie uma nova regra para vincular as duas classes `cap` e `sleeve` e defina `display` para `inline-block`.

# --hints--

Você deve usar um seletor de classe para vincular ambas as classes `cap` e `sleeve`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

A regra do CSS deve ter uma propriedade `display` definida como `inline-block`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
