---
id: 5f3c866d0fc037f7311b4ac8
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

That's closer, but the price didn't stay over on the right. This is because `inline-block` elements only take up the width of their content. To spread them out, add a `width` property to the `flavor` and `price` class selectors that have a value of `50%` each.

# --hints--

You should set the `width` property to `50%` in your `.flavor` selector.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert(flavorWidth === '50%');
```

You should set the `width` property to `50%` in your `.price` selector.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert(priceWidth === '50%');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Hello World</title>
    <script src="https://unpkg.com/react@18/umd/react.development.js"></script>
    <script src="https://unpkg.com/react-dom@18/umd/react-dom.development.js"></script>

    <script src="https://unpkg.com/@babel/standalone/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
  </head>
  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { SomeComponent } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <SomeComponent />
      );
    </script>
  </body>
</html>
```

```jsx
export const SomeComponent = () => {
  const [count, setCount] = React.useState(0);
  return (
    <div>
--fcc-editable-region--
      <h1>Some Component</h1>
      <p>Count: {count}</p>
      <button onClick={() => setCount(count + 1)}>Increment</button>
--fcc-editable-region--
    </div>
  );
};

```

