---
id: 5f33071498eb2472b87ddee4
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. 

Start by adding your `<!DOCTYPE html>` and `html` element with a `lang` attribute set to `en`.

# --hints--

You should have the `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

You should have an opening `<html lang="en">` with the language set to english.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

You should have a closing `</html>` tag.

```js
assert.match(code, /<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
