---
id: 5dc174fcf86c76b9248c6eb2
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Os elementos em HTML têm tags de abertura, como `<h1>`, e tags de fechamento, como `</h1>`.

O texto de um elemento vai entre as tags de abertura e de fechamento.

Encontre o elemento `h1` e altere seu texto para:

`CatPhotoApp`

# --hints--

O texto `CatPhotoApp` deve estar presente no código. Não se esqueça de verificar a ortografia.

```js
assert(code.match(/catphotoapp/i));
```

O elemento `h1` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

O elemento `h1` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/h1\>/));
```

Você tem mais de um elemento `h1`. Remova o elemento `h1` a mais.

```js
assert(document.querySelectorAll('h1').length === 1);
```

O elemento `h1` deve conter o texto `CatPhotoApp`. Você omitiu o texto, digitou o texto incorretamente ou ele não está entre as tags de abertura e fechamento do elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Parece que você está usando uma extensão de navegador que está modificando a página. Não se esqueça de desativar todas as extensões do navegador.

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

