---
id: 5dc2385ff86c76b9248c6eb7
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

O HTML5 tem alguns elementos que identificam diferentes áreas de conteúdo. Essas elementos tornam seu HTML mais fácil de ler e ajudam com a otimização dos mecanismos de busca (SEO) e com a acessibilidade.

Identifique a seção principal desta página adicionando uma tag `<main>` antes do elemento `h1` e uma tag de fechamento `</main>` depois do elemento `p`.

# --hints--

O elemento `main` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('main'));
```

O elemento `main` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/main\>/));
```

A tag de abertura do elemento `main` deve estar abaixo da tag de abertura do elemento `body`. Eles estão na ordem errada.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

A tag de abertura do elemento `main` deve estar acima do elemento `h1`. Eles estão na ordem errada.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

A tag de fechamento do elemento `main` deve estar abaixo do elemento `p`. Eles estão na ordem errada.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert(
  mainNode.contains(pNode) &&
    pNode.textContent.toLowerCase().match(/see more cat photos in our gallery/)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
  </body>
</html>
```

