---
id: 5dc23991f86c76b9248c6eb8
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

No passo anterior, você aninhou o elemento `h2`, o comentário e o elemento `p` dentro do elemento `main`. Indentar os elementos aninhados com mais dois espaços do que seus elementos pais melhora a legibilidade:

```html
<ul> 
  <li> This `li` element indented </li>
  <li> This `li` element is also indented </li>
</ul>
```

Adicione mais dois espaços na frente de `h1`, comentário e elementos `p` para que seu HTML seja mais legível.

# --hints--

O código deve ter um elemento `h2` com o texto `Cat Photos`. Você pode ter excluído acidentalmente esse elemento, está faltando as tags de abertura e fechamento dele ou o texto foi alterado.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Você não deve adicionar os elementos `ul` ou `li` do exemplo.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

O elemento `h2` deve estar abaixo da tag de abertura do elemento `main` e a tag de abertura deve estar 6 espaços após o início da linha.

```js
assert(code.toLowerCase().match(/<main\>\s*\n\s{6}<h2>/));
```

O código deve ter um comentário. Você excluiu o comentário de um passo anterior.

```js
assert(code.match(/<!--.*-->/));
```

O texto do comentário deve ser `TODO: Add link to cat photos`. Não altere o texto nem o espaçamento do comentário.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

O comentário deve estar abaixo do elemento `h2` e começar 6 espaços após o início da linha.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

O código deve ter um elemento `p`. Você removeu o elemento `p` de um passo anterior.

```js
assert(document.querySelector('p'));
```

O texto do elemento `p` deve ser `Click here to view more cat photos.` Não altere o texto, o espaçamento ou a pontuação do elemento `p`.

```js
assert(
  document
    .querySelector('p')
    .innerText.toLowerCase()
    .match(/click\s+here\s+to\s+view\s+more\s+cat\s+photos\.?$/)
);
```

O elemento `p` deve estar abaixo do comentário e sua tag de abertura deve estar 6 espaços após o início da linha.

```js
assert(code.toLowerCase().match(/-->\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Click here to view more cat photos.</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```

