---
id: 5dc23f9bf86c76b9248c6eba
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Você pode adicionar imagens a um site da web usando o elemento `img`. Elementos `img` têm uma tag de abertura, mas não têm a tag de fechamento. A tag de um elemento que não precisa de uma tag de fechamento é conhecida como <dfn>tag de fechamento automático</dfn>.

Adicione um elemento `img` abaixo do elemento `p`. Neste momento, nenhuma imagem será exibida no navegador.

# --hints--

O elemento `img` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('img'));
```

O elemento `img` não deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(!code.match(/<\/img\>/));
```

Você deve ter apenas um elemento `img`. Remova os elementos adicionais.

```js
assert(document.querySelectorAll('img').length === 1);
```

O elemento `img` deve estar abaixo do elemento `p`. Eles estão na ordem errada.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```

