---
id: 5dc24165f86c76b9248c6ebc
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Todos os elementos `img` devem ter um atributo `alt`. O texto de um atributo `alt` é usado por leitores de tela para melhorar a acessibilidade. Ele é exibido se a imagem não puder ser carregada. Por exemplo, `<img src="cat.jpg" alt="A cat">` tem um atributo `alt` com o texto `A cat`.

Dentro do elemento `img`, adicione um atributo `alt` com este texto:

`A cute orange cat lying on its back`

# --hints--

O código deve ter um elemento `img`. Você removeu o elemento `img` de um passo anterior.

```js
assert(document.querySelector('img'));
```

O elemento `img` não tem um atributo `alt`. Verifique se há um espaço depois do nome da tag de abertura e/ou se há espaços antes de todos os nomes dos atributos.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

O valor do atributo `alt` do elemento `img` está definido com outro valor que não é 'A cute orange cat lying on its back'. Certifique-se de que o valor do atributo `alt` esteja cercado com aspas.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

