---
id: 5dc24073f86c76b9248c6ebb
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

<dfn>Atributos</dfn> em HTML são palavras especiais usadas dentro da tag de abertura de um elemento para controlar o comportamento dele. O atributo `src` em um elemento `img` especifica o URL da imagem (onde a imagem está localizada). Um exemplo de um elemento `img` usando o atributo `src` é: `<img src="https://www.example.com/the-image.jpg">`.

Dentro do elemento existente `img`, adicione um atributo `src` com este URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

O código deve ter um elemento `img`. Você pode ter removido o elemento `img` ou não cercou o valor do atributo `src` com aspas.

```js
assert(document.querySelector('img'));
```

O elemento `img` deve ter um atributo `src`. Você omitiu o atributo ou tem um erro de digitação. Certifique-se de que haja um espaço entre o nome do elemento e o nome do atributo.

```js
assert(document.querySelector('img').src);
```

O atributo `src` do elemento `img` deve estar definido como '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Você omitiu o URL ou tem um erro de digitação. As verificações distinguem maiúsculas e minúsculas no URL.

```js
assert(document.querySelector('img').src === 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Embora você tenha definido o `src` do elemento `img` com o URL correto, é recomendável sempre cercar o valor de um atributo com aspas.

```js
assert(!/\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

