---
id: 5ddbd81294d8ddc1510a8e56
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

O texto de um link deve ser colocado entre as tags de abertura e fechamento de um elemento de âncora (`a`). Por exemplo, `<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>` é um link com o texto `click here to go to freeCodeCamp.org`.

Adicione o texto `link to cat pictures` ao elemento de âncora. Ele se tornará o texto do link.

# --hints--

O elemento de âncora (`a`) deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('a'));
```

O elemento de âncora (`a`) deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/a\>/));
```

O elemento de âncora (`a`) deve conter o texto `link to cat pictures`. Certifique-se de colocar o texto do link entre as tags de abertura e de fechamentoa do elemento de âncora (`a`).

```js
assert(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' ') ===
    'link to cat pictures'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

