---
id: 5dfa22d1b521be39a3de7be0
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Você pode transformar qualquer texto em um link, como o texto dentro de um elemento `p`.

``` html
<p>I think <a href="https://www.freecodecamp.org">freeCodeCamp</a> is great.</p>
```

No texto do elemento `p`, transforme as palavras `cat photos` em um link adicionando as tags de abertura e fechamento de um elemento de âncora (`a`) ao redor dessas palavras. Em seguida, defina o valor do atributo `href` como `https://freecatphotoapp.com`


# --hints--

Você deve colocar um novo elemento de âncora (`a`) dentro de cada elemento `p`.

```js
const nestedAnchor = document.querySelector(`p > a`);
assert.isNotNull(nestedAnchor)
```

O valor de `href` do link deve ser `https://freecatphotoapp.com`. Você omitiu o valor de `href` ou tem um erro de digitação.

```js
assert(
  document.querySelector('p > a').href === 'https://freecatphotoapp.com/'
);
```

O texto do link deve ser `cat photos`. Você omitiu o texto ou tem um erro de digitação.

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isTrue(innerContent.trim() === 'cat photos');
```

The text inside your anchor element has extra leading or trailing whitespace. The only space in the anchor text should be between the word `cat` and the word `photos`.

```js
const nestedAnchor = document.querySelector('p > a');
const innerContent = nestedAnchor.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

After nesting the anchor (`a`) element, the only `p` element content visible in the browser should be `See more cat photos in our gallery.` Double check the text, spacing, or punctuation of both the `p` and nested anchor element.

```js
assert.match(code, /<p>see more <a[^>]*>cat photos<\/a> in our gallery\.?<\/p>/i)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

