---
id: 5dfa22d1b521be39a3de7be0
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Você pode transformar qualquer texto em um link, como o texto dentro de um elemento `p`.

``` html
<p>Para mim, o <a href="https://www.freecodecamp.org/">freeCodeCamp</a> é o máximo.</p>
```

No texto do elemento `p`, transforme as palavras `cat photos` em um link adicionando as tags de abertura e fechamento de um elemento de âncora (`a`) ao redor dessas palavras. Em seguida, defina o valor do atributo `href` como `https://freecatphotoapp.com/`


# --hints--

Você deve colocar um novo elemento de âncora (`a`) dentro de cada elemento `p`.

```js
assert($('p > a').length);
```

O valor de `href` do link deve ser `https://freecatphotoapp.com`. Você omitiu o valor de `href` ou tem um erro de digitação.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

O texto do link deve ser `cat photos`. Você omitiu o texto ou tem um erro de digitação.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Após aninhar o elemento de âncora (`a`), o único conteúdo do elemento `p` visível no navegador deve ser `See more cat photos in our gallery.` Verifique o texto novamente em termos de espaçamento e pontuação, tanto do elemento `p` quanto do elemento de âncora aninhado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/see more cat photos in our gallery\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

