---
id: 5dfa22d1b521be39a3de7be0
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Você pode ver as palavras `cat photos` sublinhadas ao lado da imagem na visualização agora. Esse é o seu link. Clique nele à vontade. No texto do elemento `p`, transforme as palavras `cat photos` em um link para `https://freecatphotoapp.com`. Quando terminar, remova a tag de âncora e o texto antigos abaixo do parágrafo.

# --hints--

O código deve conter apenas um elemento de âncora (`a`). Remova os elemento elementos de âncora adicionais.

```js
assert(document.querySelectorAll('a').length === 1);
```

O elemento de âncora (`a`) deve estar aninhado dentro do elemento `p`.

```js
assert($('p > a').length);
```

O texto do link deve ser `cat photos`. Você omitiu o texto ou tem um erro de digitação.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com' &&
    nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

Após aninhar o elemento de âncora (`a`), o único conteúdo do elemento `p` visível no navegador deve ser `Click here to view more cat photos.` Verifique o texto novamente em termos de espaçamento e pontuação, tanto do elemento `p` quanto do elemento de âncora aninhado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more cat photos.</p>
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

