---
id: 5dfb6a35eacea3f48c6300b4
title: Passo 24
challengeType: 0
dashedName: step-24
---

# --description--

Um elemento de legenda da figura (`figcaption`) é usado para adicionar uma legenda para descrever a imagem contida dentro do elemento `figure`. Por exemplo, `<figcaption>A cute cat</figcaption>` adiciona a legenda `A cute cat`.

Após a imagem dentro do elemento `figure`, adicione um elemento `figcaption` com o texto definido como:

`Cats love lasagna.`

# --hints--

O elemento `figcaption` deve ter uma tag de abertura. As tags de abertura têm a seguinte sintaxe: `<elementName>`.

```js
assert(document.querySelector('figcaption'));
```

O elemento `figcaption` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/figcaption\>/));
```

O elemento `figcaption` da lasanha deve estar dentro do elemento `figure`.

```js
assert(
  document.querySelector('figure > figcaption') &&
    document.querySelector('figure > figcaption')
);
```

O elemento `img` da lasanha deve estar dentro do elemento `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

O elemento `figcaption` dentro do elemento `figure` deve estar abaixo do elemento `img`. Eles estão na ordem errada.

```js
assert(
  document.querySelector('figcaption').previousElementSibling.nodeName === 'IMG'
);
```

O texto do elemento `figcaption` deve ser `Cats love lasagna.` Você omitiu o texto ou tem um erro de digitação.

```js
assert(
  document.querySelector('figcaption').innerText.match(/Cats love lasagna.?$/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

        </figure>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

