---
id: 5dfb655eeacea3f48c6300b3
title: Passo 23
challengeType: 0
dashedName: step-23
---

# --description--

O elemento `figure` representa um conteúdo auto-contido e permitirá que você associe uma imagem com uma legenda.

Aninhe a imagem que você acabou de adicionar dentro de um elemento `figure`.

# --hints--

O elemento `figure` deve ter uma tag de abertura. As tags de abertura têm a seguinte sintaxe: `<elementName>`.

```js
assert(document.querySelector('figure'));
```

O elemento `figure` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/figure\>/));
```

O elemento `figure` deve estar logo acima da tag de fechamento de fechamento do segundo elemento `section`.

```js
assert($('section')[1].lastElementChild.nodeName === 'FIGURE');
```

O elemento `img` da lasanha deve estar dentro do elemento `figure`.

```js
assert(
  document.querySelector('figure > img') &&
    document.querySelector('figure > img').getAttribute('src').toLowerCase() ===
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

