---
id: 5ef9b03c81a63668521804dc
title: Passo 44
challengeType: 0
dashedName: step-44
---

# --description--

Você pode usar botões de opção para perguntas onde você queira que o usuário dê apenas uma resposta entre várias opções.

Aqui está um exemplo de um botão de opção com a opção `cat`: `<input type="radio"> cat`. Lembre-se de que os elementos `input` são de <dfn>fechamento automático</dfn>.

Antes da entrada de texto, adicione um botão de opção com a opção definida como:

`Indoor`

# --hints--

Você deve criar um elemento `input` para o botão de opção. Confira a sintaxe.

```js
assert($('form > input').length >= 2);
```

O elemento `input` deve ter uma tag de abertura, mas não uma de fechamento.

```js
assert($('form > input') && !code.match(/<\/input\>/g));
```

Você deve adicionar apenas um elemento `input` para o botão de opção. Remova os elementos adicionais.

```js
assert($('form > input').length === 2);
```

O novo elemento `input` deve estar acima do elemento `input` existente com o atributo `type` definido como `text`. Eles estão na ordem errada.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

O novo elemento `input` não tem um atributo `type`. Verifique se há um espaço depois do nome da tag de abertura.

```js
assert($('input')[0].hasAttribute('type'));
```

O novo elemento `input` deve ter apenas um atributo `type`. Remova os elementos adicionais.

```js
assert($('input')[0]
  .getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

O novo elemento `input` deve ter o atributo `type` com o valor `radio`. Você omitiu o valor ou tem um erro de digitação. Lembre-se de que os valores dos atributos devem estar cercados com aspas.

```js
assert(
  $('input')[0]
    .getAttribute('type')
    .match(/^radio$/i)
);
```

Embora você tenha definido o atributo `type` do elemento `input` com o valor `radio`, é recomendável sempre cercar o valor de um atributo com aspas.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

O texto do botão `radio` deve ser `Indoor` e estar localizado depois do botão em vez de antes dele.

```js
const radioInputElem = $('input')[0];
assert(!radioInputElem.previousSibling.nodeValue.match(/Indoor/i));
```

O texto `Indoor` deve estar localizado diretamente à direita do seu botão `radio`. Você omitiu o texto ou tem um erro de digitação.

```js
const radioInputElem = $('input')[0];
assert(
  radioInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/\s*Indoor/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

