---
id: 5ef9b03c81a63668521804d6
title: Passo 35
challengeType: 0
dashedName: step-35
---

# --description--

Agora, adicione um formulário para a web para coletar informações de usuários.

Depois do cabeçalho `Cat Form`, adicione um elemento `form`.

# --hints--

O elemento `form` deve ter uma tag de abertura e uma de fechamento. Pode ser que você não tenha inserido uma delas, as duas, ou que elas estejam na ordem errada.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

As tags do elemento `form` não estão na ordem correta.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

O elemento `form` dentro do último elemento `section` deve estar abaixo do elemento `h2`. O elemento `h2` e o elemento `form` estão na ordem errada.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

O elemento `form` não deve ter conteúdo. Remova os elementos HTML ou o texto entre as tags do elemento `form`.

```js
assert($('form')[0].innerHTML.trim().length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

