---
id: 5ef9b03c81a63668521804e5
title: Passo 61
challengeType: 0
dashedName: step-61
---

# --description--

Para que uma caixa de seleção ou botão de opção esteja selecionado por padrão, você precisa adicionar o atributo `checked` ao elemento. Não há a necessidade de definir um valor para o atributo `checked`. Em vez disso, apenas adicione a palavra `checked` (selecionado) ao elemento `input`, certificando-se de que haja um espaço entre esse atributo e os outros.

Faça com que o primeiro botão de opção e a primeira caixa de seleção estejam selecionados por padrão.

# --hints--

Certifique-se de que ainda existam dois botões de opção e três caixas de seleção dentro de seus respectivos elementos `fieldset`.

```js
assert(
  $('input[type="radio"]').length === 2 &&
    $('fieldset > input[type="checkbox"]').length === 3
);
```

O primeiro botão de opção, com a `id` definida como `indoor`, deve ter o atributo `checked`.

```js
assert($('input[type="radio"]')[0].hasAttribute('checked'));
```

O segundo botão de opção não deve ter o atributo `checked`.

```js
assert(!$('input[type="radio"]')[1].hasAttribute('checked'));
```

A primeira caixa de seleção não tem o atributo `checked`.

```js
assert($('input[type="checkbox"]')[0].hasAttribute('checked'));
```

A segunda caixa de seleção não deve ter o atributo `checked`.

```js
assert(!$('input[type="checkbox"]')[1].hasAttribute('checked'));
```

A terceira caixa de seleção não deve ter o atributo `checked`.

```js
assert(!$('input[type="checkbox"]')[2].hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
            <input id="loving" type="checkbox" name="personality" value="loving"> <label for="loving">Loving</label>
            <input id="lazy" type="checkbox" name="personality" value="lazy"> <label for="lazy">Lazy</label>
            <input id="energetic" type="checkbox" name="personality" value="energetic"> <label for="energetic"> Energetic</label>
          </fieldset>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

