---
id: 5ef9b03c81a63668521804e3
title: Passo 58
challengeType: 0
dashedName: step-58
---

# --description--

Adicione outra caixa de seleção depois daquela que você acaba de adicionar. O valor do atributo `id` deve ser `lazy` e o valor do atributo `name` deve ser o mesmo que o da última caixa de seleção.

Adicione também um elemento `label` à direita da nova caixa de seleção com o texto `Lazy`. Certifique-se de associar o elemento `label` com a nova caixa de seleção usando o atributo `for`.

# --hints--

Você precisa adicionar uma nova caixa de seleção.

```js
assert($('input[type="checkbox"]').length === 2);
```

A nova caixa de seleção deve ter um atributo `id` com o valor `lazy` e um atributo `name` com o valor `personality`. Verifique se há um espaço depois do nome da tag de abertura e/ou se há espaços antes de todos os nomes dos atributos.

```js
const checkboxes = [...$('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

A nova caixa de seleção deve vir depois da primeira. Elas estão na ordem errada.

```js
const checkboxes = [...$('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

No lado direito da nova caixa de seleção, deve haver um elemento `label` com o texto `Lazy`.

```js
const nextElementSibling = $('input[type="checkbox"]')[1].nextElementSibling;
assert(
  nextElementSibling.nodeName === 'LABEL' &&
    nextElementSibling.innerText.replace(/\s+/g, '').match(/^Lazy$/i)
);
```

O novo `label` deve ter um atributo `for` com o mesmo valor que o atributo `id` da nova caixa de seleção. Você omitiu o valor ou tem um erro de digitação.

```js
assert(
  $('input[type="checkbox"]')[1].nextElementSibling.getAttribute('for') ===
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

