---
id: 5ef9b03c81a63668521804d2
title: Passo 27
challengeType: 0
dashedName: step-27
---

# --description--

O código de uma lista ordenada (`ol`) é semelhante ao de uma lista não ordenada, mas a lista de itens em uma lista ordenada é numerada quando eles são exibidos.

Após o último elemento `h3` do segundo elemento `section`, adicione uma lista ordenada com esses três itens de lista:

`flea treatment` `thunder` `other cats`

# --hints--

O elemento `ol` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

O elemento `ol` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/ol>/));
```

O elemento `ol` deve estar acima da tag de fechamento do segundo elemento `section`. Eles estão na ordem errada.

```js
assert($('main > section')[1].lastElementChild.nodeName === 'OL');
```

Os três elementos `li` devem estar dentro do elemento `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Você deve ter os três elementos `li` com os textos `flea treatment`, `thunder` e `other cats`, não importando a ordem.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Você deve ter apenas um elemento `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

