---
id: 5ef9b03c81a63668521804e2
title: Passo 54
challengeType: 0
dashedName: step-54
---

# --description--

Os formulários normalmente usam caixas de seleção para perguntas que tenham mais de uma resposta. Por exemplo, aqui temos uma caixa de seleção com a opção `tacos`: `<input type="checkbox"> tacos`.

Abaixo do elemento `legend` que você acaba de adicionar, adicione um elemento `input` com o atributo `type` definido como `checkbox` e dê a ele a opção:

`Loving`

# --hints--

O elemento `input` de sua caixa de seleção deve ter uma tag de abertura, mas não uma de fechamento.

```js
assert($('fieldset > input') && !code.match(/<\/input\>/g));
```

Você deve adicionar apenas um elemento input para a caixa de seleção. Remova os elementos adicionais.

```js
assert($('fieldset > input').length === 1);
```

O novo elemento `input` deve estar abaixo do elemento `legend` com o texto `What's your cat's personality?`. Eles estão na ordem errada.

```js
const existingLegendElem = $('fieldset > legend')[1];
assert(
  existingLegendElem &&
    existingLegendElem.nextElementSibling.nodeName === 'INPUT'
);
```

O novo elemento `input` não tem um atributo `type`. Verifique se há um espaço depois do nome da tag de abertura.

```js
assert($('fieldset > input')[0].hasAttribute('type'));
```

O novo elemento `input` deve ter o atributo `type` com o valor `checkbox`. Você omitiu o valor ou tem um erro de digitação. Lembre-se de que os valores dos atributos devem estar cercados com aspas.

```js
assert(
  $('fieldset > input')[0]
    .getAttribute('type')
    .match(/^checkbox$/i)
);
```

Embora você tenha definido o atributo `type` do elemento `input` com o valor `checkbox`, é recomendável sempre cercar o valor de um atributo com aspas.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

O texto `Loving` deve estar localizado diretamente à direita da caixa de seleção. Certifique-se de que haja um espaço entre o elemento e o texto. Você omitiu o texto ou tem um erro de digitação.

```js
const checkboxInputElem = $('input[type="checkbox"]')[0];
assert(
  checkboxInputElem.nextSibling.nodeValue.replace(/\s+/g, ' ').match(/ Loving/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

