---
id: 5ef9b03c81a63668521804da
title: Passo 42
challengeType: 0
dashedName: step-42
---

# --description--

Use o elemento `button` para criar um botão clicável. Por exemplo, `<button>Click Here</button>` cria um botão com o texto `Click Here`.

Adicione um elemento `button` com o texto `Submit` abaixo do elemento `input`. O comportamento padrão de clicar em um botão de formulário sem qualquer atributo envia o formulário para o local especificado no atributo `action`.

# --hints--

O elemento `button` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('button'));
```

O elemento `button` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/button\>/));
```

O texto do elemento `button` deve ser `Submit`. Você omitiu o texto ou tem um erro de digitação.

```js
assert(document.querySelector('button').innerText.toLowerCase() === 'submit');
```

O elemento `button` deve estar abaixo do elemento `input`. Eles estão na ordem errada.

```js
const collection = [...document.querySelectorAll('input, button')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('INPUT') < collection.indexOf('BUTTON'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>

--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

