---
id: 5efb2c990dc218d6c85f89b2
title: Passo 43
challengeType: 0
dashedName: step-43
---

# --description--

Mesmo que você tenha adicionado o botão abaixo do campo (input) de texto, eles aparecem um ao lado de outro na página. Isso ocorre porque os elementos `input` e `button` são <dfn>elementos em linha</dfn> (inline), que não aparecem em novas linhas.

O botão que você adicionou enviará o formulário por padrão. No entanto, confiar no comportamento padrão pode causar confusão. Adicione o atributo `type` com o valor `submit` ao `button` para deixar claro que ele é um botão de envio.

# --hints--

O elemento `button` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelector('button'));
```

O elemento `button` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/button\>/));
```

O elemento `button` não tem um atributo `type`. Verifique se há um espaço depois do nome da tag de abertura.

```js
assert($('button')[0].hasAttribute('type'));
```

O novo elemento `button` deve ter o atributo `type` com o valor `submit`. Você omitiu o valor ou tem um erro de digitação. Lembre-se de que os valores dos atributos devem estar cercados com aspas.

```js
assert(
  $('button')[0]
    .getAttribute('type')
    .match(/^submit$/i)
);
```

Embora você tenha definido o atributo `type` do elemento `button` com o valor `submit`, é recomendável sempre cercar o valor de um atributo com aspas.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```

