---
id: 5f05a1d8e233dff4a68508d8
title: Passo 47
challengeType: 0
dashedName: step-47
---

# --description--

Crie outro botão de opção abaixo do primeiro. Coloque-o dentro de um elemento `label` com `Outdoor` sendo o texto do `label`. Dê ao botão de opção um atributo `id` com o valor `outdoor`.

# --hints--

Apenas o `input` do botão de opção `Indoor` original deve ter um `id` definido como `indoor`. Você pode reiniciar o passo para recuperar o HTML original, se necessário.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Você não deve fazer nenhuma alteração no botão de opção `Indoor`. Você pode reiniciar o passo para recuperar o HTML original, se necessário.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Você deve adicionar apenas um elemento `input` dentro do elemento `label`. Certifique-se de que o novo `label` tenha uma tag de abertura e uma de fechamento.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Você não deve adicionar atributos às tags de abertura de `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

O novo `input` deve ter o atributo `id` com o valor `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

O novo `input` deve ter o atributo `type` com o valor `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Você não deve adicionar novos elementos além de um `input` dentro de um `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Não deve haver texto à esquerda do novo `input`.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

O texto do label para o novo botão de opção deve ser `Outdoor`. Isso inclui a capitalização.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

O novo botão de opção e o novo label devem estar logo abaixo/após o botão de opção e o label de `Indoor`. Não deve haver outras tags entre eles.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

