---
id: 5f07c98cdb9413cbd4b16750
title: Passo 17
challengeType: 0
dashedName: step-17
---

# --description--

É hora de adicionar uma nova seção. Adicione um segundo elemento `section` abaixo do elemento `section` existente.

# --hints--

O elemento `section` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 2);
```

Você deve acrescentar apenas uma tag de abertura para `section`. Remova as tags adicionais.

```js
assert(document.querySelectorAll('section').length === 2);
```

O elemento `section` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/section>/g).length >= 2);
```

Você deve acrescentar apenas uma tag de fechamento para `section`. Remova as tags adicionais.

```js
assert(code.match(/<\/section>/g).length === 2);
```

O segundo elemento `section` não deve estar dentro do primeiro elemento `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 0);
```

Os dois elementos `section` devem estar dentro das tags de abertura e fechamento do elemento `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(foundElems.length === 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

