---
id: 5f07be6ef7412fbad0c5626b
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Antes de adicionar conteúdo novo, você deve usar um elemento `section` para separar o conteúdo de fotos de gatos do resto do conteúdo.

Faça com que os elementos `h2`, de comentário, `p` e âncora (`a`) fiquem aninhados em um elemento `section`.

# --hints--

O elemento `section` deve ter uma tag de abertura. As tags de abertura têm a seguinte sintaxe: `<elementName>`.

```js
assert(document.querySelector('section'));
```

O elemento `section` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/section\s*>/i));
```

Todo o elemento `section` deve estar dentro das tags de abertura e fechamento do elemento `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

O elemento `h2`, o comentário, o elemento `p` e o elemento de âncora (`a`) devem estar entre as tags de abertura e fechamento do elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElements.length === 3);
```

O elemento `h1` não deve estar dentro do elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```

