---
id: 5f07be6ef7412fbad0c5626b
title: Passo 16
challengeType: 0
dashedName: step-16
---

# --description--

Antes de adicionar conteúdo novo, você deve usar um elemento `section` para separar o conteúdo de fotos de gatos do resto do conteúdo.

Pegue todos os elementos atualmente localizados no elemento `main` e aninhe-os em um elemento `section`.

# --hints--

O elemento `section` deve ter uma tag de abertura. As tags de abertura têm a seguinte sintaxe: `<elementName>`.

```js
assert(document.querySelector('section'));
```

O elemento `section` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/section\>/));
```

Todo o elemento `section` deve estar dentro das tags de abertura e fechamento do elemento `main`.

```js
assert(document.querySelector('section').parentNode.nodeName === 'MAIN');
```

O elemento `h2`, o comentário, o elemento `p` e o elemento de âncora (`a`) devem estar entre as tags de abertura e fechamento do elemento `section`.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElems = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert(foundElems.length === 3);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Click here to view more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
    </main>
--fcc-editable-region--
  </body>
</html>
```

