---
id: 5f07fb1579dc934717801375
title: Passo 33
challengeType: 0
dashedName: step-33
---

# --description--

É hora de adicionar uma nova seção. Adicione um terceiro elemento `section` abaixo do segundo elemento `section` existente.

# --hints--

O elemento `section` deve ter uma tag de abertura. As tags de abertura têm essa sintaxe: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

Você deve acrescentar apenas uma tag de abertura para `section`. Remova as tags adicionais.

```js
assert(document.querySelectorAll('section').length === 3);
```

O elemento `section` deve ter uma tag de fechamento. As tags de fechamento têm um caractere `/` logo após o caractere `<`.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

Você deve acrescentar apenas uma tag de fechamento para `section`. Remova as tags adicionais.

```js
assert(code.match(/<\/section>/g).length === 3);
```

Todos os elementos `section` devem estar dentro das tags de abertura e fechamento do elemento `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

O último elemento `section` não deve ter conteúdo. Remova os elementos HTML ou o texto dentro do elemento `section`.

```js
assert($('main > section')[2].children.length === 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```

