---
id: 62dabe2ef403a12d5d295273
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora que você transformou o texto de `cat photos` dentro do elemento `p` em um link, você não precisa do segundo link abaixo do elemento `p`. Exclua todo o elemento de âncora abaixo do elemento `p`.

# --hints--

O código deve ter apenas um elemento de parágrafo (`p`).

```js
assert.isTrue(document.querySelectorAll('p').length === 1);
```

O código deve ter apenas um elemento de âncora (`a`).

```js
assert.isTrue(document.querySelectorAll('a').length === 1);
```

Não faça alterações aos elementos que você criou no passo anterior.

```js
assert.strictEqual(document.querySelector('p').innerHTML, 'See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.')
```

O texto `cat photos` não deve estar abaixo do elemento `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
