---
id: 62dabe2ef403a12d5d295273
title: Etapa 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora que você transformou o texto de `cat photos` dentro do elemento `p` em um link, você não precisa do segundo link abaixo do elemento `p`. Exclua todo o elemento de âncora abaixo do elemento `p`.

# --hints--

O código deve ter apenas um elemento de âncora (`a`).

```js
assert(document.querySelectorAll('a').length === 1);
```

O elemento `p` ainda deve ter o elemento de âncora (`a`) dentro dele.

```js
assert($('p > a').length);
```

O valor de `href` do link deve ser `https://freecatphotoapp.com`. Você omitiu o valor de `href` ou tem um erro de digitação.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.getAttribute('href') === 'https://freecatphotoapp.com'
);
```

O texto do link deve ser `cat photos`. Você omitiu o texto ou tem um erro de digitação.

```js
const nestedAnchor = $('p > a')[0];
assert(
  nestedAnchor.innerText.toLowerCase().replace(/\s+/g, ' ') === 'cat photos'
);
```

O único conteúdo do elemento `p` visível no navegador deve ser `Click here to view more cat photos.` Verifique o texto novamente em termos de espaçamento e pontuação, tanto do elemento `p` quanto do elemento de âncora aninhado.

```js
const pText = document
  .querySelector('p')
  .innerText.toLowerCase()
  .replace(/\s+/g, ' ');
assert(pText.match(/click here to view more cat photos\.?$/));
```

O texto `cat photos` não deve estar abaixo do elemento `p`.

```js
const pNextSibling = document.querySelector('p').nextSibling;
assert(
  pNextSibling.nodeName === '#text' && !/cat\s*photos/.test(pNextSibling.nodeValue) ||
  pNextSibling.nodeName === 'IMG'
);
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <main>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Click here to view more <a href="https://freecatphotoapp.com">cat photos</a>.</p>
      <a href="https://freecatphotoapp.com">link to cat pictures</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
