---
id: 587d7fb4367417b2b2512bff
title: Step 5
challengeType: 2
forumTopicId: 301525
dashedName: step-5
---

# --description--

A `version` is one of the required fields of your package.json file. This field describes the current version of your project. Here's an example:

```json
"version": "1.2.0",
```

# --instructions--

Add a `version` to the package.json file of your project.

# --hints--

package.json should have a valid "version" key

```js
(getUserInput) =>
  $.get(getUserInput('url') + '/_api/package.json').then(
    (data) => {
      var packJson = JSON.parse(data);
      assert(packJson.version, '"version" is missing');
    },
    (xhr) => {
      throw new Error(xhr.responseText);
    }
  );
```

