---
id: 64dc8ea01436383a88256d7f
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

The puzzle starts with the disks piled up on the first rod, in decreasing size. You need to populate your first list with numbers representing the various disk sizes.

Instead of adding items manually to the first list, generate the sequence of numbers from `3` to `1` by using the `range()` function and assign it to `rods['A']`.

The syntax is `range(start, stop, step)`, where `start` is your starting integer, `stop` is the last integer (not inclusive), and `step` is the difference between one number and the following one.

# --hints--

You should use the `range()` function to assign `rods['A']` a sequence of numbers. The syntax for the `range` function is `range(start, stop, step)`.

```js
({ test: () => assert(__pyodide.runPython(`
  a = __locals.get('rods')
  type(a['A']) is range
  `))
})
```

You should use the `range()` function to assign `rods['A']` the sequence of numbers from `3` to `1`.

```js
({ test: () => assert(__pyodide.runPython(`
 a = __locals.get('rods')
 a['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
