---
id: 64dc985e6720a23edac63b51
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

The allowed disk movements between the rods display a repetition pattern. Allowed movements between the rods repeat every three moves. 
For example, the first, the fourth and the seventh moves are allowed between `A` and `C`, that is, between the source and the target rods. In this case, the remainder of the division between the move number and 3 is equal to 1.

Inside the `for` loop you wrote in the previous step, remove the print statement and add an `if` statement to execute when `(i + 1) % 3 == 1`. Then, print the f-string `f'Move {i + 1} allowed between {source} and {target}'`. Please, note the use of `i + 1` as move number, since `i` equals zero for the first iteration.

# --hints--

You should have an `if` statement to execute when `(i + 1) % 3 == 1`. 

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

You should print the string `f'Move {i + 1} allowed between {source} and {target}'`. 

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, target, auxiliary):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'C', 'B')
```
