---
id: 64dce5f6ff907c2b2c47bb05
title: Step 23
challengeType: 20
dashedName: step-23
---

# --description--

When `target` is empty, the disk should be moved necessarily from `source` to `target`.

After the declaration of `forward`, add an `if` statement to check if `rods[target]` is empty. If it is, change `forward` to `True`.

# --hints--

You should have an `if` statement to check if `rods[target]` is empty.

```js
({ test: () => assert.match(code, /if\s+not\s+rods\s*\[\s*target\s*\]\s*:/) })
```

You should set `forward` to `True` inside your new `if` statement. You should not change the first declaration of `forward`.

```js
({ test: () => assert.match(code, /forward\s*=\s*False\s+if\s+not\s+rods\s*\[\s*target\s*\]\s*:\s+forward\s*=\s*True/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
--fcc-editable-region--
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
--fcc-editable-region--
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
