---
id: 64ddd02838f36117cd9738eb
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

Outside the `else` statement, add a comment that says `display our progress` and print the content of the lists to check that everything is working.

# --hints--

You should print the `rods` object outside the `else` statement.

```js
({ test: () => assert.match(code, /(?<=rods\[\s*source\s*\]\.append\s*\(\s*rods\[\s*target\s*\]\.pop\s*\(\s*\)\s*\)\s+)print\s*\(\s*rods\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, target, auxiliary):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if len(rods[target]) == 0:
                forward = True
            elif len(rods[source]) != 0 and len(rods[target]) != 0 and rods[source][-1] < rods[target][-1]:
                forward = True
--fcc-editable-region--                
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
--fcc-editable-region--
        elif (i + 1) % 3 == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif (i + 1) % 3 == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'C', 'B')
```
