---
id: 64ddd280d69fa818ca782d50
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

As you can see, disk 1 is going back and forth every three moves. This happens because you still need to take care of movements between the other rods.
Instead of repeating the same code you wrote during the previous few steps and changing the rods, it would be better to move that code inside a function to call in each conditional statement. Declare an empty function named `make_allowed_move` and don't forget the `pass` keyword.

# --hints--

You should declare an empty function named `make_allowed_move`. Remember to use the `pass` keyword inside the function body and the correct indentation.

```js
({ test: () => {
		assert(__pyodide.runPython(`
        import inspect
        inspect.isfunction(__locals.get('make_allowed_move'))
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--

def move(n, source, target, auxiliary):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if len(rods[target]) == 0:
                forward = True
            elif len(rods[source]) != 0 and len(rods[target]) != 0 and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
        elif (i + 1) % 3 == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif (i + 1) % 3 == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'C', 'B')
```
