---
id: 64ddd65848a12919d7e1c7d0
title: Step 29
challengeType: 20
dashedName: step-29
---

# --description--

It's time to move some code from the `move` function to the `make_allowed_move` function. Move the code nested inside the outermost `if`, except the first `print` statement, inside your new function. Pay close attention to the indentation.

# --hints--

Test 1

```js
({test: () => assert(true)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def make_allowed_move(rod1, rod2):
    pass

def move(n, source, target, auxiliary):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        if (i + 1) % 3 == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if len(rods[target]) == 0:
                forward = True
            elif len(rods[source]) != 0 and len(rods[target]) != 0 and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
--fcc-editable-region--            
        elif (i + 1) % 3 == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif (i + 1) % 3 == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'C', 'B')
```
