---
id: 64df202aa1342114cd077920
title: Step 37
challengeType: 20
dashedName: step-37
---

# --description--

If the number of disks is even and the remainder equals `1`, the move is allowed between the source rod and the auxiliary rod. Add an `else` statement to print the allowed movement and call `make_allowed_move` with the correct arguments.

If you look at the output, you can see that the execution stops at the third move because of an `IndexError`. This happens because the code is still incomplete and needs an `else` statement you will be writing soon. To make it work, turn your `make_allowed_move` call into a comment.

# --hints--

You should add an `else` statement. Within that, you should print the allowed movement and write the `make_allowed_move` call with the correct arguments as a comment.

```js
({ test: () => {
  assert.match(code, /else\s*:\s+print\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*auxiliary\s*\}\1\s*\)\s+#\s*make_allowed_move\(\s*source\s*,\s*auxiliary\s*\)/);
	} })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
number_of_moves = 2 ** NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def make_allowed_move(rod1, rod2):    
    forward = False
    if not rods[rod2]:
        forward = True
    elif rods[rod1] and rods[rod2] and rods[rod1][-1] < rods[rod2][-1]:
        forward = True              
    if forward:
        print(f'Moving disk {rods[rod1][-1]} from {rod1} to {rod2}')
        rods[rod2].append(rods[rod1].pop())
    else:
        print(f'Moving disk {rods[rod2][-1]} from {rod2} to {rod1}')
        rods[rod1].append(rods[rod2].pop())
    
    # display our progress
    print(rods, '\n')

def move(n, source, target, auxiliary):
    # display starting configuration
    print(rods, '\n')   
    for i in range(number_of_moves):
--fcc-editable-region--     
        if (i + 1) % 3 == 1:
            if n % 2 != 0:
                print(f'Move {i + 1} allowed between {source} and {target}')
                make_allowed_move(source, target)
            
--fcc-editable-region--            
        elif (i + 1) % 3 == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
            make_allowed_move(source, auxiliary)
        elif (i + 1) % 3 == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')
            make_allowed_move(auxiliary, target)
           
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'C', 'B')
```
