---
id: 65560ffdb7d05d248e012280
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside the TreeNode class, create the `__init__` method so that we can initialize the attributes of the object. Leave it empty for now and don't add any parameters.

# --hints--

Use the `def` keyword to define the `__init__` method.

```js
({ test: () => assert.match(code, /def/) });
```

Your method name should be `__init__`.

```js
({ test: () => assert.match(code, /__init__/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
--fcc-editable-region--
        
        pass
--fcc-editable-region--
```
