---
id: 65560f9380be92226084ef46
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--
ddd
In this project, we would create a Binary Search Tree. A binary search tree is a data structure in which each node has at most two children, with the left child containing values less than the node and the right child containing values greater than the node, allowing for efficient searching and sorting operations.

Begin by defining an empty class `TreeNode`. The `TreeNode` class represents a `node` in a binary search tree(BST).

# --hints--

You should use the `class` keyword to declare an empty class named `TreeNode`. Don't forget to add the colon at the end.

```js
({
  test: () => {
    assert(__pyodide.runPython(`
      import inspect
      a = __locals.get("TreeNode")
      inspect.isclass(a)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
    pass
```

