---
id: 65561022956c1024e7184add
title: Step 3
challengeType: 20
dashedName: step-3
---

# --description--

The `__init__` method takes two parameters: `self` (which represents the instance of the class being created) and `key` (the value to be stored in the node). Add those two parameters to the `__init__` method.

# --hints--

You should add the `self` parameter in the method.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,/) });
```

You should add the `key` parameter in the method.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
class TreeNode:
--fcc-editable-region--
    def __init__():
        pass
--fcc-editable-region--
```
