---
id: 65560ffdb7d05d248e012280
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Inside the TreeNode class, create a constructor method `__init__` The init method is a special method in Python that is called when a new object of the class is created. Leave it empty for now and don't add any parameters.

# --hints--

Use the `def` keyword to define the `__init__` method.

```js
({ test: () => assert.match(code, /def/) });
```

Your method name should be `__init__`.

```js
({ test: () => assert.match(code, /def\s+__init__/) });
```

Don't forget the `()` after the method name.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\)/) });
```

Don't forget the `:` at the end.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\):\s*$/) });
```

The definition should be properly indented with 4 spaces before the def keyword.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\(\):/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    pass
--fcc-editable-region--

--fcc-editable-region--
```

