---
id: 655a194276dfa11460f7b5e3
title: Step 24
challengeType: 20
dashedName: step-24
---

# --description--

Create an `if` statement that checks two conditions:

- If `node` is `None`: This indicates that the search has reached the end of a branch without finding the key.
- If `node.key == key`: This means that the key has been found in the current node.

This serves as the base case for the recursive search.

# --hints--

The first part of the condition should check if `node` is `None`.

```js
({ test: () => assert.match(code, /if\s+node\s+is\s+None/) });
```

The second part of the condition should check if `node.key` is equal to `key`.

```js
({ test: () => assert.match(code, /node\.key\s*==\s*key/) });
```

Use `or` to combine the two conditions.

```js
({ test: () => assert.match(code, /if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
--fcc-editable-region--
        
--fcc-editable-region--
            # Do not change code below this line
            pass
```
