---
id: 655a51a705c97d7a9294ab2a
title: Step 47
challengeType: 20
dashedName: step-47
---

# --description--

Moving back to the helper function `_min_value`.

When there are two children of the node we want to delete, we need to choose the new node from the children. We are using the in-order successor method for it.

The in-order successor method chooses the smallest element from the right subtree and places that element in place of the deleted node.

For now, define the function `_min_value` with the arguments `self`, `node`.

# --hints--

Define the function `_min_value` with the arguments `self`, `node`.

```js
({ test: () => assert.match(code, /def _min_value\(self, node\):/) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right

            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)
        return node
--fcc-editable-region--
            
--fcc-editable-region--
        # do not change code below this line
        pass
```
