---
id: 655a4614304cd36031cb4e75
title: Step 31
challengeType: 20
dashedName: step-31
---

# --description--

Assign the result of the `_delete` operation back to `self.root`.

This is done because the deletion operation might result in a new root(for example, if the node to be deleted is the current root).

# --hints--

Assign the result of the `_delete` operation to `self.root`.

```js
({ test: () => assert.match(code, /self\.root\s*=\s*self\._delete\(self\.root,\s*key\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
--fcc-editable-region--
            self._delete(self.root, key)
--fcc-editable-region--
```
