---
id: 655a44442b60ee5a28df8ee5
title: Step 28
challengeType: 20
dashedName: step-28
---

# --description--

If the second `if` statement is not `true`, it means that the target key is greater than or equal to the key of the current node.

In this case, outside the `if` block, recursively call the method `_search` with the right child of the current node (`self._search(node.right, key)`). Also, return this value.

This step is taken because, in a binary search tree, if the target key is greater than the current node's key, the search continues in the right subtree.

# --hints--

You should recursively call the `_search` method using `self.search`.

```js
({ test: () => assert.match(code, /self\._search/) });
```

You should have two parameters in the call to the `_search` method: `node.right` and `key`.

```js
({ test: () => assert.match(code, /self\._search\(node\.right,\s*key\)/) });
```

Return the result of the call to the `_search` method.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(node\.right,\s*key\)/)
});
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
--fcc-editable-region--
            pass
--fcc-editable-region--
```
