---
id: 655a132b56e0160cfeca08d4
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

Inside the `__init__` method, initialize an instance attribute called `root` to `None`. 

The `root` attribute represents the root node of the binary search tree. Since this is the constructor when a new `BinarySearchTree` object is created, it starts with an empty tree, so the `root` attribute is set to `None`.

# --hints--

You should initialize the `root` attribute to `None` using `self.root`.

```js
assert.match(code, /self\.root\s*=\s*None/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
--fcc-editable-region--
        
--fcc-editable-region--
        
```
