---
id: 655a493ab909a96f7c316cd5
title: Step 42
challengeType: 20
dashedName: step-42
---

# --description--

Under the `if` statement, create an `elif` statement that checks if the right child of the current node (`node.right`) is `None`.

This condition checks if there is no right child, the left child is returned as the replacement.

# --hints--

You should create an `elif` statement that checks if the right child of the current node is `None`. 

```js
assert.match(code, /elif\s+node\.right\s+is\s+None/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right
--fcc-editable-region--
                
--fcc-editable-region--
                    # do not edit below this line
                    pass
```
