---
id: 655a1a6d8c44db154b00c909
title: Step 26
challengeType: 20
dashedName: step-26
---

# --description--

Create another `if` statement that checks if the target key is less than the key of the current node.

# --hints--

You should create another `if` statement to check if `key` < `node.key`.

```js
const after_split = code.split('def search(self, key):')[1];
assert.match(after_split, /if\s+key\s+<\s+node\.key/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
                
--fcc-editable-region--
            # Only change code above this line
            pass
```
