---
id: 655a14a6fe4cd50f38d01dd3
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

Now we need to check if the `node` parameter is `None`. If it is, this means that the method has reached a leaf node or an empty spot in the tree where the new node should be inserted.

Inside the method body, put a condition that checks `if node is None`.

# --hints--

Use an `if` condition for logic.

```js
({ test: () => assert.match(code, /if/) });
```

Check if the `node` parameter is `None`.

```js
({ test: () => assert.match(code, /node is None/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--
    def _insert(self, node, key):
            
--fcc-editable-region--
                pass
```
