---
id: 655a515f4b85ce79464fe5e8
title: Step 46
challengeType: 20
dashedName: step-46
---

# --description--

Finally, return the current node (`return node`). This is the return value of the `_delete` method.

The deletion operation may have modified the structure of the tree, and the updated node is returned.

# --hints--

You should return the current node.

```js
console.log(
  'splitted',
  code.split('node.key = self._min_value(node.right)')[1]
);
const after_split = code.split('node.key = self._min_value(node.right)')[1];
assert.match(after_split, /return\s+node/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right

                elif node.right is None:
                    return node.left

                node.key = self._min_value(node.right)
                node.right = self._delete(node.right, node.key)
--fcc-editable-region--
                pass
--fcc-editable-region--
```
