---
id: 655a16db1ccc5e132b5bc44d
title: Step 22
challengeType: 20
dashedName: step-22
---

# --description--

The `search` method returns the result of the `_search` method.

Append `return` to the call to the `_search` method.

# --hints--

Append return to the call to the `_search` method.

```js
({ test: () => assert.match(code, /return\s+self\._search\(self\.root,\s*key\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
--fcc-editable-region--
        self._search(self.root, key)
        pass
--fcc-editable-region--
```
