---
id: 655a19e924364f14dffc44e6
title: Step 25
challengeType: 20
dashedName: step-25
---

# --description--

If either condition is true, the method returns the current node.

# --hints--

You should return `node` from your `if` conditional block.

```js
const after_split = code.split('def search(self, key):')[1];
assert.match(after_split, /return\s+node/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
--fcc-editable-region--

```
