---
id: 655a552337ee62882fdeee79
title: Step 55
challengeType: 20
dashedName: step-55
---

# --description--

Now, let's work on the `_inorder_traversal` method.

Create a method `_inorder_traversal` that takes three parameters:

- `self`
- `node`: The current node being considered during the traversal.
- `result`: The list to which the keys are appended in sorted order.

# --hints--

you should create a method `_inorder_traversal` that takes three parameters- `self`, `node`, and `result`.

```js
assert.match(code, /def\s+_inorder_traversal\(self,\s*node,\s*result\)/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right

            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)
        return node

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

    def inorder_traversal(self):
        result = []
        self._inorder_traversal(self.root, result)
        return result
--fcc-editable-region--

--fcc-editable-region--
        pass
```
