---
id: 655a158c7e80c810f6eff1e1
title: Step 18
challengeType: 20
dashedName: step-18
---

# --description--

Recursively call the `_insert` method on the right child of the current node. Assign the result to the right attribute of the current node.

# --hints--

You should call the `self._insert` method using the parameters `node.right` and `key`.

```js
({ test: () => assert.match(code, /self\._insert\(node\.right,\s*key\)/) });
```

Assign the result to the right attribute (`node.right`)of the current node.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(node\.right,\s*key\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
--fcc-editable-region--
                
--fcc-editable-region--
                pass
```
