---
id: 655a56dc4749dc906fac6802
title: Step 62
challengeType: 20
dashedName: step-62
---

# --description--

Insert the values using a `for` loop in the `bst` object.

In the body of `for`, use the `insert` method to insert and sort.

# --hints--

You should have a `for node in nodes` loop.

```js
assert.match(code, /for\s+node\s+in\s+nodes/);
```

You should use the `insert` method to insert and sort.

```js
assert.match(code, /bst\.insert/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right

                elif node.right is None:
                    return node.left

                node.key = self._min_value(node.right)
                node.right = self._delete(node.right, node.key)
            return node

        def _min_value(self, node):
            while node.left is not None:
                node = node.left
            return node.key

        def inorder_traversal(self):
            result = []
            self._inorder_traversal(self.root, result)
            return result

        def _inorder_traversal(self, node, result):
            if node:
                self._inorder_traversal(node.left, result)
                result.append(node.key)

bst = BinarySearchTree()
nodes = [50, 30, 20, 40, 70, 60, 80]
--fcc-editable-region--

--fcc-editable-region--
```
