---
id: 655a4ffc762d117470b94e3b
title: Step 44
challengeType: 20
dashedName: step-44
---

# --description--

If neither of the above conditions is met, it means the node has both left and right children.

To choose the successor, we need to find the minimum value in the right subtree.
The smallest value will be the in-order successor of the current node.

The `_min_value` method is a helper method that finds the smallest value in a given subtree. We will define it later on.

For now, add the following line `node.key = self._min_value(node.right)` outside of the `elif` block.

# --hints--

You should add the following line outside of the `elif` block:

```js
assert.match(code, /node\.key\s*=\s*self\._min_value\(node\.right\)/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right

                elif node.right is None:
                    return node.left
--fcc-editable-region--
                
--fcc-editable-region--
                    # do not change below this line
                    pass
```
