---
id: 655a4985e6cc1a712eec132e
title: Step 43
challengeType: 20
dashedName: step-43
---

# --description--

If the previous condition is met, it means there is no right child, and the left child is returned as the replacement.

Return the left child of the current node.

# --hints--

You should return the left child of the current node.

```js
assert.match(code, /return\s+node\.left/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right

                elif node.right is None:
--fcc-editable-region--
                
--fcc-editable-region--
```
