---
id: 655a173e5b8adc13b761ed74
title: Step 23
challengeType: 20
dashedName: step-23
---

# --description--

The `_search` method is a private method meant for internal use within the BinarySearchTree class.

Define the `_search` method with these two parameters:

- `node`: The current node being examined during the search.
- `key`: The key (value) that is being searched for in the binary search tree.

# --hints--

Check method spelling, it should be`_search`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

`_search` should take three parameters: `self`, `node`, and `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)
--fcc-editable-region--
        
--fcc-editable-region--
        # Do not change below this line
        pass
```
