---
id: 655a13aa8af2510de51f2e1c
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--

The `insert` method takes two parameters:

- `self`: This parameter represents the instance of the class, which allows the method to access and modify the attributes and methods of that instance.
- `key`: This parameter is the key value that you want to insert into the binary search tree.

Add these two parameters to the function definition.

Note that this method would be called by the user.

# --hints--

The `insert` method should contain the `self` parameter.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

The insert method should contain the `key` parameter.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
    def insert():
--fcc-editable-region--
           # Only change code above this line
            pass
```
