---
id: 655a486b3c0fef6a5bc7158f
title: Step 38
challengeType: 20
dashedName: step-38
---

# --description--

If `true`, the recursively call the method with itself with the right child of the current `node` and `key` as parameters.

Also, assign the value to the right node

# --hints--

Call the method `_delete` with the right child of the current node and the key as arguments.

```js
assert.match(code, /self\._delete\(node\.right,\s*key\)/);
```

You should assign the result of the `_delete` method to the right child(`node.right`) of the current node.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(node\.right,\s*key\)/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
--fcc-editable-region--
            
--fcc-editable-region--
```
