---
id: 655a173e5b8adc13b761ed74
title: Step 23
challengeType: 20
dashedName: step-23
---

# --description--

Coming back to the `_search`, define the method with these three parameters `self`,`node` and `key`.

# --hints--

You should declare a `_search` method. Please, note the underscore at the beginning.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

`_search` should take three parameters: `self`, `node`, and `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)
--fcc-editable-region--
    
--fcc-editable-region--
        # Do not change below this line
        pass
```
