---
id: 655a46fce0ce5a638c180e36
title: Step 33
challengeType: 20
dashedName: step-33
---

# --description--

Let's work on the base case.

Create an `if` statement that checks if the current `node` is `None`.

# --hints--

You should create an `if` statement that checks if the current `node` is `None`.

```js
assert.match(code, /if\s+node\s+is\s+None/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
--fcc-editable-region--
            pass
--fcc-editable-region--
```
