---
id: 655a1479ba7e7c0ee6c1acdd
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

Now we will define the `_insert` method which is a helper function and does the actual insertion. This method is recursive, meaning it calls itself to traverse the tree until the appropriate location for the new node is found

Define the `_insert` function with the following three parameters:

- `self`: This parameter represents the instance of the class, allowing the method to access class attributes and other methods.
- `node`: This parameter represents the current node in the recursive traversal of the tree.
- `key`: This parameter is the key value that you want to insert into the binary search tree.

# --hints--

The `_insert` method should have three parameters.

```js
({ test: () => assert.match(code, /def _insert\(self, node, key\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

--fcc-editable-region--

--fcc-editable-region--
        pass
```
