---
id: 655a13744e82580d7ee3073d
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

Next, we would define a mechanism to insert nodes in the tree. For that, define and empty `insert` method within the `BinarySearchTree` class.

# --hints--

You should define the `insert` method using the `def` keyword.

```js
({ test: () => assert.match(code, /def/) });
```

The method name should be `insert`.

```js
({ test: () => assert.match(code, /insert\s*\(\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None
--fcc-editable-region--

--fcc-editable-region--
            pass
```
