---
id: 655a15b683445611528cccf1
title: Step 19
challengeType: 20
dashedName: step-19
---

# --description--

After the insertion process is complete, whether it's a new node at a leaf node or modifications in the left or right subtrees, return the current node to update the tree structure at the higher levels of the recursive call stack.


# --hints--

You should return the current node.

```js
({ test: () => assert.match(code, /return\s*node/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
--fcc-editable-region--
                
--fcc-editable-region--
```
