---
id: 655a227e57aabb25d1f9c987
title: Step 20
challengeType: 20
dashedName: step-20
---

# --description--

Now, it's time to work on the search functionality.

Create a function named `search` inside the `BinarySearchTree` class. The search method would take two parameters: `self` and `key`.

`self` refers to the instance of the `BinarySearchTree` class on which the method is called.
`key` is the value that the user wants to search for in the BST.

# --hints--

Your method should be named `search` and defined using `def`.

```js
({ test: () => assert.match(code, /def\s+search/) });
```

The `search` method should take two parameters: `self` and `key`.

```js
({ test: () => assert.match(code, /def\s+search\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

--fcc-editable-region--
            
--fcc-editable-region--
            pass  
```
