---
id: 655a167ea3e96512bf1343ce
title: Step 21
challengeType: 20
dashedName: step-21
---

# --description--

Inside the search method, call the private helper method `_search` with the following parameters:

- `self.root`: This is the root of the binary search tree. The search starts from the root.
- `key`: This is the value that the user wants to find in the BST.

We will define the private helper method in the next steps.

# --hints--

You should call the private helper method `_search` using `self`.
    
```js
({ test: () => assert.match(code, /self\._search/) });
```

You should call the private helper method `_search` using `self.root` and `key`.

```js
({ test: () => assert.match(code, /self\._search\(self\.root,\s*key\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
--fcc-editable-region--
                
--fcc-editable-region--
                pass
```
