---
id: 655a50f992ba7177aff2b718
title: Step 45
challengeType: 20
dashedName: step-45
---

# --description--

After finding the minimum value, recursively delete the node with the minimum value from the right subtree.

This step ensures that the node with the minimum value is removed from the tree while maintaining the BST property.

# --hints--

You should call the _delete method recursively with the right child of the current node as the first argument.

```js
assert.match(code, /self\._delete\(node\.right/);
```

You should call the _delete method recursively with `node.key` as the second argument.

```js
assert.match(code, /self\._delete\(node\.right\,\s*node\.key/);
```

You should assign the return value of the recursive call to the right child of the current node.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(node\.right\,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right

                elif node.right is None:
                    return node.left

                node.key = self._min_value(node.right)
--fcc-editable-region--
                pass
--fcc-editable-region--
```
