---
id: 655a10e4a620fc091ba43b3d
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Inside the `init` method, assign the value of the `key` parameter to the key attribute of the node using `self.key`.
This means that the `key` attribute of the `TreeNode` instance will be set to the value passed during the object's creation.

# --hints--

You should assign the value of the `key` parameter to the key attribute of the node using `self.key`.

```js
assert.match(code, /self\.key\s*=\s*key/);
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self,key):
--fcc-editable-region--

        pass
--fcc-editable-region--
```
