---
id: 655a51ff908edc7c12c3a92c
title: Step 48
challengeType: 20
dashedName: step-48
---

# --description--

To find the smallest value in the right subtree, create a `while` loop in the function body to iterate through the left children of the given node until it reaches the leftmost (smallest) node in the subtree.

# --hints--

You should use the condition `node.left is not None` in the `while` loop.

```js
({ test: () => assert.match(code, /while node\.left is not None:/) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right

            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)
        return node

    def _min_value(self, node):
--fcc-editable-region--
            
--fcc-editable-region--
            # do not change code below this line
            pass
```
