---
id: 655a151cc6041f0ff7d24ded
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

Now we would recursively traverse the tree and insert the values using the principle for binary trees:

- Values smaller than the key are placed in the left sub-tree
- Values greater than the key are placed in the right sub-tree

Create a condition that checks if the `key` < `node.key`.

# --hints--

You should create another if statement to check if `key < node.key`.

```js
({ test: () => assert.match(e.code.original["main.py"], /^(\s+)if\s+node\s+is\s+None\s*:.*?\1if\s+key\s*<\s*node\.key\s*:/ms) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
--fcc-editable-region--
            
--fcc-editable-region--
            pass
```
