---
id: 655a42ece011fa55cfa6b183
title: Step 27
challengeType: 20
dashedName: step-27
---

# --description--

If `key < node.key` is `True`, it means that the target key is smaller than the current node's key, and the search needs to continue in the left subtree.

Now, return the result of calling the `_search` method with the left child of the current node and key as the arguments.

# --hints--

You should recursively call the `_search` method using `self._search`.
    
```js
({ test: () => assert.match(code, /self\._search/) });
```

You should pass `node.left` and `key` as the arguments to your `_search` call.

```js
({ test: () => assert.match(code, /self\._search\(node\.left,\s*key\)/) });
```

You should return the result of the `_search` method call.

```js    
({ test: () => assert.match(code, /return\s+self\._search\(node\.left,\s*key\)/) });    
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

    def search(self, key):
        return self._search(self.root, key)

    def _search(self, node, key):
        if node is None or node.key == key:
            return node
--fcc-editable-region--
        if key < node.key:                
--fcc-editable-region--
```
