---
id: 655a12c2ab0bcc0c2ba30e16
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

For the `BinarySearchTree` class, create a constructor method `__init__`

This method takes one parameter, `self`, which represents the instance of the class being created.

# --hints--

You should define the `__init__` method using the `def` keyword.

```js
({ test: () => assert.match(code, /def/) });
```

The method should have a single parameter, `self`.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
class BinarySearchTree:
--fcc-editable-region--
        
--fcc-editable-region--
        pass
```
