---
id: 655a142ae611b30e5df0ac16
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

Now, inside the `insert` method, we would call another method `_insert()` that performs the actual insertion.

Begin by assigning `self._insert(self.root, key)` to `self.root`.

Note that,

- `self.root` passes the root node of the tree as the first argument. This is the starting point for the insertion process.
- `key`: passes the `key` value you want to insert as the second argument.

# --hints--

You should recursively call the `_insert()` method using `self._insert`

```js
({ test: () => assert.match(code, /self\._insert/)})
```

You should use the `self.root` parameter in the `_insert()` method.

```js
({ test: () => assert.match(code, /self\.root,/)})
```

You should use the `key` parameter in the `_insert()` method.

```js
({ test: () => assert.match(code, /key/)})
```

Assign the return value of `_insert()` to `self.root`.

```js
({ test: () => assert.match(code, /self\.root\s*=\s*self\._insert/)})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
    def insert(self,key):
--fcc-editable-region--
```
