---
id: 655a536e99be288210f01451
title: Step 51
challengeType: 20
dashedName: step-51
---

# --description--

The inorder_traversal method is responsible for performing an in-order traversal of the binary search tree (BST) and returning the keys of the nodes in sorted order.

In-order traversal is a depth-first binary tree traversal algorithm that visits the left subtree, the current node, and then the right subtree.

Create a method `inorder_traversal` that takes one argument `self`.

# --hints--

You should define a helper method `_inorder_traversal` that takes the argument `self`.

```js
({ test: () => assert.match(code, /def\s+inorder_traversal\(self\)/) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
        def __init__(self):
            self.root = None

        def insert(self,key):
            self.root = self._insert(self.root, key)

        def _insert(self, node, key):
            if node is None:
                return TreeNode(key)
            if key < node.key:
                node.left = self._insert(node.left, key)
            elif key > node.key:
                node.right = self._insert(node.right, key)
            return node

        def search(self, key):
            return self._search(self.root, key)

        def _search(self, node, key):
            if node is None or node.key == key:
                return node
            if key < node.key:
                return self._search(node.left, key)
            return self._search(node.right, key)

        def delete(self, key):
            self.root = self._delete(self.root, key)

        def _delete(self, node, key):
            if node is None:
                return node
            if key < node.key:
                node.left = self._delete(node.left, key)
            elif key > node.key:
                node.right = self._delete(node.right, key)
            else:
                if node.left is None:
                    return node.right

                elif node.right is None:
                    return node.left

                node.key = self._min_value(node.right)
                node.right = self._delete(node.right, node.key)
            return node

        def _min_value(self, node):
            while node.left is not None:
                node = node.left
            return node.key
--fcc-editable-region--

--fcc-editable-region--
            # do not change code below this line
            pass
```
