---
id: 655a14e50e2b450f94602f26
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

If the `node` is `None`, a new `TreeNode` is created with the provided key, and it becomes the new leaf node, effectively inserting the key into the tree.

Return the `TreeNode(key)` if the previous condition is `true`.

# --hints--

You should return the TreeNode with the provided key.

```js
({ test: () => assert.match(code, /return TreeNode\(key\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def insert(self,key):
        self.root = self._insert(self.root, key)

    def _insert(self, node, key):
        if node is None:
--fcc-editable-region--
                
--fcc-editable-region--
            pass
```
