---
id: cf1111c1c11feddfaeb3bdef
title: Add Two Numbers with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` is a data type in JavaScript which represents numeric data.

Now let's try to add two numbers using JavaScript.

JavaScript uses the `+` symbol as an addition operator when placed between two numbers.

**Example:**

```js
const MY_VAR = 5 + 10;
```

`MY_VAR` now has the value `15`.

# --instructions--

Change the `0` so that SUM will equal `20`.

# --hints--

`SUM` should equal `20`.

```js
assert(SUM === 20);
```

You should use the `+` operator.

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'SUM = '+z;})(SUM);
```

## --seed-contents--

```js
const SUM = 10 + 0;
```

# --solutions--

```js
const SUM = 10 + 10;
```
