---
id: 66352b2418c108072d0a062e
title: From Flowcharts to JavaScript
challengeType: 1
dashedName: from-flowcharts-to-js
---

# --description--

Are flowcharts mastered, and are you ready to bring them to life?  This is where the magic happens!  In "From Flowcharts to JS (JavaScript)," we'll take your flowchart knowledge and translate it into the powerful language of code.  We'll revisit those flowchart symbols, but this time, we'll see how they translate into lines of JavaScript. Transform your flowcharts into working JavaScript creations to bridge the gap between visual logic and real-world code!  
Let's code!
**For Example:-**

```js
console.log("Hello User.");
```

<h2>Hinglish</h2>
Kya aapne flowcharts ko samajh liya hai, aur kya aap unhe zindagi mein laane ke liye tayaar hain? Yahaan jaadu hota hai! "From Flowcharts to JS (JavaScript)," mein hum aapke flowchart gyaan ko lekar aapke saamne code ki taqatwar bhasha mein translate karenge. Hum un flowchart symbols ko dobara dekhenge, lekin is baar, hum dekhenge ki ve JavaScript ke lines mein kaise translate hote hain. Apne flowcharts ko kaam karne waale JavaScript creations mein badal kar visual logic aur real-world code ke beech ka fasla mitayein!  
Chaliye code karte hain!

Udaharan ke liye:

```js
console.log("Hello user");
```

# --instructions--

Print "your name" by using `console.log()`.

# --hints--

You should use `console.log()` to print your name.

```js
assert.doesNotThrow(() => { console.log("Hello JavaScript"); }, Error, "Expected 'Hello JavaScript' to be printed.");
```

# --seed--
## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello JavaScript");

```
