---
id: cf1111c1c12feddfaeb2bdef
title: Generate Random Whole Numbers within a Range
challengeType: 1
forumTopicId: 18187
dashedName: generate-random-whole-numbers-within-a-range
---

# --description--

You can generate a random whole number in the range from zero to a given number. You can also pick a different lower number for this range.

You'll call your minimum number `min` and your maximum number `max`.

This formula gives a random whole number in the range from `min` to `max`. Take a moment to read it and try to understand what this code is doing:

```js
Math.floor(Math.random() * (max - min + 1)) + min
```

# --instructions--

Create a function called `randomRange` that takes a range `myMin` and `myMax` and returns a random whole number that's greater than or equal to `myMin` and less than or equal to `myMax`.

# --hints--

The lowest random number that can be generated by `randomRange` should be equal to your minimum number, `myMin`.

```js
assert.strictEqual(calcMin, 5);
```

The highest random number that can be generated by `randomRange` should be equal to your maximum number, `myMax`.

```js
assert.strictEqual(calcMax, 15);
```

The random number generated by `randomRange` should be an integer, not a decimal.

```js
assert.strictEqual(randomRange(0, 1) % 1, 0);
```

`randomRange` should use both `myMax` and `myMin`, and return a random number in your range.

```js
assert.match(__helpers.removeJSComments(code), /myMax/g);
assert.match(__helpers.removeJSComments(code), /myMin/g);
assert.match(__helpers.removeJSComments(code), /Math.floor/g);
assert.match(__helpers.removeJSComments(code), /Math.random/g);
```

# --seed--

## --seed-contents--

```js
function randomRange(myMin, myMax) {
  return 0;
}
```

# --solutions--

```js
function randomRange(myMin, myMax) {
  return Math.floor(Math.random() * (myMax - myMin + 1)) + myMin;
}
```
