---
id: cf1231c1c11feddfaeb5bdef
title: Multiply Two Numbers with JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

We can also multiply one number by another.

JavaScript uses the `*` symbol for multiplication of two numbers.

**Example**

```js
const MY_VAR = 13 * 13;
```

`MY_VAR` would have the value `169`.

# --instructions--

Change the `0` so that `PRODUCT` will equal `80`.

# --hints--

The variable `PRODUCT` should be equal to 80.

```js
assert(PRODUCT === 80);
```

You should use the `*` operator.

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'PRODUCT = '+z;})(PRODUCT);
```

## --seed-contents--

```js
const PRODUCT = 8 * 0;
```

# --solutions--

```js
const PRODUCT = 8 * 10;
```
