---
id: 661a6798129b0c40e6cc6fed
title: Swapping Variable-values
challengeType: 1
dashedName: swapping-var-val
---

# --description--

If you have a pen in your right hand and a pencil in your left hand, swapping them would mean putting the pen in your left hand and the pencil in your right hand. Similarly, variables in JavaScript can hold a value temporarily while you swap it with another.


# --instructions--

Write a program that asks the user for two numbers and then swaps their values. Use a temporary variable to hold one value while swapping the other. Finally, print the swapped values.

# --hints--

Use a third variable named `temp`.

```js
assert(code.match(temp));
```

After swapping, `a` and `b` should be  `99` and `45` respectively.

```js
assert(A===99 && B===45);
```

# --seed--

## --seed-contents--

```js

var a = 45;
var b = 99;
// Only change code below this line



 // Only change code above this line
console.log(a);
console.log(b);
```

# --solutions--

```js
var a = 12;
var b = 4;

var temp = a;
a = b;
b = temp;

console.log(value1);
console.log(value2);

``` 
