---
id: 6612dad1c726f003e1d542b1
title: Swapping Variables using Operators
challengeType: 1
dashedName: swapping-using-opereators
---

# --description--

Write a program to take two inputs from the user and swap them without using a temporary or third variable.

# --instructions--

Swap Them Without Using A Third Variable instructions.

# --hints--

`2,3` should return `3,2`

```js
assert(a,b===3,2)
```

# --seed--
## --seed-contents--

```js
let a=2;
let b=3

```

# --solutions--

```js
let a= 2
let b= 3
a = a + b;
b = a - b;
a = a - b;
```
