---
id: 657ef2a86d4e545cec9a85fb
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Uppercase characters in camel case or pascal case indicate the start of new words.

Inside the loop body, use an `if` statement in conjunction with the `.isupper()` string method to check for uppercase characters and move `pass` inside the new `if` statement.

# --hints--

You should write a new `if` statement with `char.isupper()` as the condition. Don't forget the colon at the end.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +if\s+char\.isupper\s*\(\s*\)\s*:/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
