---
id: 61fd6b7c83dbf54a08cf0498
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro do elemento `div`, adicione três elementos `span`. Dê a cada um deles um atributo `class` definido como `year` e adicione os textos a seguir (em ordem): `2019`, `2020` e `2021`.

# --hints--

O elemento `div` deve ter três elementos `span`.

```js
assert(document.querySelector('div')?.children?.length === 3);
```

Cada elemento `span` deve ter o atributo `class` definido como `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert(span?.classList?.contains('year')));
```

O primeiro `span` deve ter o texto `2019`.

```js
assert(document.querySelector('div')?.children?.[0]?.textContent === '2019');
```

O segundo `span` deve ter o texto `2020`.

```js
assert(document.querySelector('div')?.children?.[1]?.textContent === '2020');
```

O terceiro `span` deve ter o texto `2021`.

```js
assert(document.querySelector('div')?.children?.[2]?.textContent === '2021');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
