---
id: 61fd71d596e8f253b9408b39
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

O elemento `tr` é usado para indicar uma linha de tabela. Adicione um elemento `tr` dentro do elemento `thead`. No novo elemento `tr`, adicione um elemento `td` seguido por três elementos `th`.

O elemento `td` indica uma célula de dados, enquanto o elemento `th` indica uma célula do cabeçalho.

# --hints--

O elemento `thead` deve ter um elemento `tr`.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

O elemento `tr` deve ter um elemento `td` como primeiro filho.

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

O elemento `tr` deve ter três elementos `th` depois do elemento `td`.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
