---
id: 61fd75ea7f663457612dba02
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Dentro de cada um dos novos elementos `th`, aninhe um elemento `span` com o atributo `class` definido como `sr-only year`. Dê a eles os seguintes textos (em ordem): `2019`, `2020` e `2021`.

Dê ao terceiro elemento `th` a `class` com o valor de `current`.

Deixe o elemento `td` vazio. Esse elemento existe apenas para garantir que a tabela tenha um layout de quatro colunas e associar os cabeçalhos às colunas corretas.

# --hints--

Cada um dos elementos `th` deve ter um elemento `span`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.length === 1);
  assert(th?.children?.[0]?.localName === 'span');
});
```

Cada um dos novos elementos `span` deve ter o atributo `class` com o valor `sr-only year`.

```js
const ths = [...document.querySelectorAll('th')];
ths?.forEach(th => {
  assert(th?.children?.[0]?.classList?.contains('sr-only'));
  assert(th?.children?.[0]?.classList?.contains('year'));
});
```

O primeiro elemento `span` deve ter o texto `2019`.

```js
assert(document.querySelectorAll('th')?.[0]?.children?.[0]?.textContent === '2019');
```

O segundo elemento `span` deve ter o texto `2020`.

```js
assert(document.querySelectorAll('th')?.[1]?.children?.[0]?.textContent === '2020');
```

O terceiro elemento `span` deve ter o texto `2021`.

```js
assert(document.querySelectorAll('th')?.[2]?.children?.[0]?.textContent === '2021');
```

O terceiro elemento `th` deve ter o atributo `class` com o valor `current`.

```js
assert(document.querySelector('table')?.querySelectorAll('th')?.[2]?.classList?.contains('current'));
```

O elemento `td` deve estar vazio.

```js
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.textContent === '');
assert(document.querySelector('table')?.querySelectorAll('td')?.[0]?.children?.length === 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
              <tr>
                <td></td>
                <th></th>
                <th></th>
                <th></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
