---
id: 61fd6343fb42013d99bcd7f3
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do elemento `body`, aninhe um elemento `section` e um elemento `main`.

# --hints--

O primeiro elemento dentro do elemento `body` deve ser seu elemento `main`.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

Você deve ter um elemento `section` dentro do elemento `main`.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
