---
id: 61fd66c687e610436494c6f1
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro do elemento `section`, adicione um elemento `h1` e um elemento `span`.

# --hints--

O elemento `section` deve ter um elemento `h1`.

```js
assert(document.querySelector('section')?.children?.[0]?.localName === 'h1');
```

O elemento `h1` deve ter um elemento `span`.

```js
assert(document.querySelector('h1')?.children?.[0]?.localName === 'span');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
