---
id: 61fd5a93fd62bb35968adeab
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Defina o HTML com uma declaração `DOCTYPE` e com os elementos `html` (indicando aqui que o documento está em inglês), `head` e `body`.

Dê ao elemento `head` os elementos `meta` apropriados para o `charset` e a `viewport`, um elemento `title` com o título apropriado e um elemento `link` para a sua folha de estilo.

# --hints--

O código deve ter uma declaração `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Seu código deve ter um elemento `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

O elemento `html` deve ter uma tag de abertura com um atributo `lang` de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

O código deve ter um elemento `head` dentro do elemento `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

O código deve ter um elemento `body` dentro do elemento `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

O elemento `head` deve vir antes do elemento `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Você deve ter dois elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Um elemento `meta` deve ter o atributo `name` definido como `viewport` e o atributo `content` definido como `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

O outro elemento `meta` deve ter o atributo `charset` com o valor `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

O código deve ter um elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

O elemento `title` deve ter algum texto.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

O código deve ter um elemento `link`.

```js
assert(/<link/.test(code))
```

O elemento `link` deve estar dentro do elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i))
```

O elemento `link` deve ter o atributo `rel` com o valor `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

O elemento `link` deve ter o atributo `href` com o valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
