---
id: 61fd67a656743144844941cb
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Leitores de tela anunciam os elementos do HTML com base no fluxo do documento. No fim das contas, queremos que o balanço patrimonial tenha o título de "Balance Sheet" e um subtítulo de "AcmeWidgetCorp". No entanto, essa ordem não faz sentido se anunciada por um leitor de tela.

Dê ao `span` existente o atributo `class` definido como `flex` e adicione dois elementos `span` dentro dele. Dê ao primeiro o texto `AcmeWidgetCorp`. Dê ao segundo o texto `Balance Sheet`. Você usará o CSS para inverter a ordem do texto na página, mas a ordem do HTML fará mais sentido para um leitor de tela.

# --hints--

O elemento `span` existente deve ter o atributo `class` definido como `flex`.

```js
assert(document.querySelector('h1')?.children?.[0]?.classList?.contains('flex'));
```

O elemento `span` existente deve ter dois novos elementos `span` dentro dele.

```js
assert(document.querySelector('.flex')?.children?.[0]?.localName === 'span');
assert(document.querySelector('.flex')?.children?.[1]?.localName === 'span');
```

Os novos elementos `span` não devem ter um atributo `class`.

```js
assert(!document.querySelector('.flex')?.children?.[0]?.classList?.length);
assert(!document.querySelector('.flex')?.children?.[1]?.classList?.length);
```

O primeiro novo elemento `span` deve ter o texto `AcmeWidgetCorp`.

```js
assert(document.querySelector('.flex')?.children?.[0]?.textContent === 'AcmeWidgetCorp');
```

O segundo novo elemento `span` deve ter o texto `Balance Sheet`.

```js
assert(document.querySelector('.flex')?.children?.[1]?.textContent === 'Balance Sheet');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
