---
id: 61fd719788899952e67692b9
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Os elementos `thead` and `tbody` são usados para indicar qual porção da tabela é o cabeçalho e qual porção contém os dados primários ou o conteúdo.

Adicione um `thead` e um `tbody` à primeira `table`, abaixo do elemento `caption`.

# --hints--

O primeiro elemento `table` deve ter um elemento `thead`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

O primeiro elemento `table` deve ter um elemento `tbody`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

O elemento `thead` deve estar imediatamente abaixo do elemento `caption`.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

O elemento `tbody` deve estar imediatamente abaixo do elemento `thead`.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
