---
id: 67e3a6b7f60b4085588189e6
title: Build a Tic Tac Toe Game
challengeType: 25
dashedName: build-a-tic-tac-toe-game
demoType: onClick
---

# --description--

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a `Board` component that renders `9` `button.square` elements in a 3x3 grid.
2. Clicking `button.square` elements should alternate between displaying an `X` or `O` within the element.
3. Once a player has won the game, clicking on any `button.square` should not cause any further changes.
4. You should create a `button#reset` element that resets the game when clicked.
5. A message should be displayed indicating either `X` or `O` as the winner, or neither if the result is a draw.

# --hints--

You should have two `input` elements of type `text`.

```js
const inputs = document.querySelectorAll('input[type="text"]');
assert.equal(inputs.length, 2);
```


# --before-all--

```js
function spyOn(
  obj,
  method,
) {
 const original = obj[method];
 const calls = [];
 const results = [];
 const fn = (...args) => {
   calls.push(args);
   const result = original(...args);
   results.push(result);
   return result;
 };
 obj[method] = fn;
 fn.calls = calls;
 fn.returns = results;
 fn.restore = () => {
   obj[method] = original;
 };
 return fn;
}

async function adjustAllRequired() {
  const inps = document.querySelectorAll("input[required], select[required], textarea[required]");
  inps.forEach((inp) => {
    let setter = Object.getOwnPropertyDescriptor(window.HTMLInputElement.prototype, "value").set;
    switch (inp.type) {
      case "text":
        setter.call(inp, "required text");
        inp.dispatchEvent(new Event("change", { bubbles: true}));
        break;
      case "email":
        setter.call(inp, "required-email@freecodecamp.org");
        inp.dispatchEvent(new Event("change", { bubbles: true}));
        break
      case "number":
        setter.call(inp, 1);
        inp.dispatchEvent(new Event("change", { bubbles: true}));
        break;
      case "checkbox":
        setter = Object.getOwnPropertyDescriptor(window.HTMLInputElement.prototype, "checked").set;
        setter.call(inp, true);
       inp.dispatchEvent(new Event("click", { bubbles: true}));
        break;
    }
    inp.dispatchEvent(new Event("change", { bubbles: true}));
  });
}
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Tic-Tac-Toe</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Board } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Board />);
    </script>
</body>

</html>
```

```css

```

```jsx
export function Board() {

}
```

# --solutions--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>Tic-Tac-Toe</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Board } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Board />);
    </script>
</body>

</html>
```

```css
* {
    font-family: "Secular One", sans-serif;
    font-weight: 400;
    font-style: normal;
}

.board {
    display: flex;
    flex-direction: column;
    align-items: center;
}

.status {
    margin: 10px;
    font-size: 1.2em;
}

.board-row {
    display: flex;
}

.square {
    width: 60px;
    height: 60px;
    font-size: 1.5em;
    margin: 5px;
    background: #fff;
    border: 1px solid #999;
    cursor: pointer;
    border-radius: 5px;
}

.square:focus {
    outline: none;
}

.reset-button {
    margin-top: 20px;
    padding: 10px 20px;
    font-size: 1em;
    cursor: pointer;
}
```

```jsx
export function Board() {
  const [squares, setSquares] = React.useState(Array(9).fill(null));
  const [isXNext, setIsXNext] = React.useState(true);
  const winner = calculateWinner(squares);
  const isDraw = squares.every(square => square !== null) && !winner;

  function handleClick(index) {
    if (squares[index] || winner || isDraw) return;
    const nextSquares = squares.slice();
    nextSquares[index] = isXNext ? 'X' : 'O';
    setSquares(nextSquares);
    setIsXNext(!isXNext);
  };

  function resetGame() {
    setSquares(Array(9).fill(null));
    setIsXNext(true);
  };

  function renderSquare(index) {
    return <Square value={squares[index]} onClick={() => handleClick(index)} />;
  };

  return (
    <div className="board">
      <h1>Tic-Tac-Toe</h1>
      <div className="status">
        {winner ? `Winner: ${winner}` : isDraw ? "It's a Draw!" : `Next Player: ${isXNext ? 'X' : 'O'}`}
      </div>
      <div className="board-row">
        {renderSquare(0)}
        {renderSquare(1)}
        {renderSquare(2)}
      </div>
      <div className="board-row">
        {renderSquare(3)}
        {renderSquare(4)}
        {renderSquare(5)}
      </div>
      <div className="board-row">
        {renderSquare(6)}
        {renderSquare(7)}
        {renderSquare(8)}
      </div>
      <button className="reset-button" onClick={resetGame}>Reset Game</button>
    </div>
  );
};

function Square({ value, onClick }) {
  return (
    <button className="square" onClick={onClick}>
      {value}
    </button>
  );
};

function calculateWinner(squares) {
  const lines = [
    [0, 1, 2], [3, 4, 5], [6, 7, 8],
    [0, 3, 6], [1, 4, 7], [2, 5, 8],
    [0, 4, 8], [2, 4, 6]
  ];
  for (let line of lines) {
    const [a, b, c] = line;
    if (squares[a] && squares[a] === squares[b] && squares[a] === squares[c]) {
      return squares[a];
    }
  }
  return null;
};
```
