---
id: 61fd6cc9475a784b7776233e
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Debajo de tu elemento `div` existente, agrega un nuevo elemento `div` con un `class` establecido en `table-wrap`. Este será el contenedor para tus tablas.

Dentro de eso, agrega tres elementos `table`. Tú usarás CSS para estilizarlas en una sola tabla, pero usando tablas separadas ayudarás a los lectores de pantalla a entender el flujo del documento.

# --hints--

Debes crear un nuevo elemento `div`.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

Tu nuevo elemento `div` debe tener el `class` establecido en `table-wrap`.

```js
assert(document.querySelector('.table-wrap')?.localName === 'div');
```

Tu elemento `.table-wrap` debe ir después de tu `div` existente.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap'));
```

Tu elemento `.table-wrap` debe tener tres elementos `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert(children?.length === 3);
children.forEach(child => assert(child?.localName === 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
