---
id: 61fd66c687e610436494c6f1
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro de tu elemento `section`, agrega un elemento `h1` con un elemento `span` anidado.

# --hints--

Tu elemento `section` debe tener un elemento `h1`.

```js
assert(document.querySelector('section')?.children?.[0]?.localName === 'h1');
```

Tu elemento `h1` debe tener un elemento `span`.

```js
assert(document.querySelector('h1')?.children?.[0]?.localName === 'span');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
