---
id: 61fd719788899952e67692b9
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

Los elementos `thead` y `tbody` son usados para indicar que parte de tu tabla es la cabecera, y que parte contiene los datos principales o el contenido.

Agrega un `thead` y un `tbody` a tu primer `table`, debajo del elemento `caption`.

# --hints--

Tu primer elemento `table` debe tener un elemento `thead`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('thead'));
```

Tu primer elemento `table` debe tener un elemento `tbody`.

```js
assert(document.querySelectorAll('table')?.[0]?.querySelector('tbody'));
```

Tu elemento `thead` debe estar inmediatamente debajo de tu elemento `caption`.

```js
assert(document.querySelector('caption')?.nextElementSibling?.localName === 'thead');
```

Tu elemento `tbody` debe estar inmediatamente debajo de tu elemento `thead`.

```js
assert(document.querySelector('thead')?.nextElementSibling?.localName === 'tbody');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
