---
id: 61fd71d596e8f253b9408b39
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

El elemento `tr` se utiliza para indicar una fila de la tabla. Agrega un elemento `tr` dentro de tu elemento `thead`. En tu nuevo elemento `tr`, agrega un elemento `td`, seguido de tres elementos `th`.

El elemento `td` indica una celda de datos, mientras que el elemento `th` indica una celda de cabecera.

# --hints--

Tu elemento `thead` debe tener un elemento `tr`.

```js
assert(document.querySelector('thead')?.children?.[0]?.localName === 'tr');
```

Tu elemento `tr` debe tener un elemento `td` como su primer hijo.

```js
assert(document.querySelector('tr')?.children?.[0]?.localName === 'td');
```

Tu elemento `tr` debe tener tres elementos `th`, después del elemento `td`.

```js
assert(document.querySelector('tr')?.children?.[1]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[2]?.localName === 'th');
assert(document.querySelector('tr')?.children?.[3]?.localName === 'th');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
          </table>
          <table>
          </table>
          <table>
          </table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
