---
id: 61fd6343fb42013d99bcd7f3
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu elemento `body`, anida un elemento `section` dentro de un elemento `main`.

# --hints--

El primer elemento dentro del elemento `body` debe ser tu elemento `main`.

```js
assert(document.querySelector('body')?.children?.[0]?.localName === 'main');
```

Debes tener un elemento `section` dentro de tu elemento `main`.

```js
assert(document.querySelector('main')?.children?.[0]?.localName === 'section');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
