---
id: 689f7996426c1534fa5bea66
title: Build a Golf Score Translator
challengeType: 26
dashedName: build-a-golf-score-translator
demoType: onClick
---

# --description--

In the game of Golf, each hole has a `par`, meaning, the average number of `strokes` a golfer is expected to make in order to sink the ball in the hole to complete the play. Depending on how far above or below `par` your `strokes` are, there is a different nickname.

Your function will be passed `par` and `strokes` arguments. Return the correct string according to this table which lists the strokes in order of priority; top (highest) to bottom (lowest):

<table><thead><tr><th>Strokes</th><th>Return</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

In this lab, you will write a function that converts the `par` and `strokes` to their nickname. 

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**User Stories:**

1. You should create a function named `golfScore`.
1. `golfScore` should take two numeric arguments, which is the par for the course and the amount of strokes made
1. `golfScore` should return a string.

# --hints--

You should create a function named `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` should take two parameters.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` should return a string

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(4, 1)` should return the string `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(4, 2)` should return the string `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` should return the string `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(4, 3)` should return the string `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(4, 4)` should return the string `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(1, 1)` should return the string `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(5, 5)` should return the string `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(4, 5)` should return the string `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(4, 6)` should return the string `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(4, 7)` should return the string `Go Home!`

```js
assert.strictEqual(golfScore(4,7), 'Go Home!');
```

`golfScore(5, 9)` should return the string `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```


# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```

