---
id: 6735a7370e0ae93a4577c689
title: Design a Magazine Layout Using CSS Grid
challengeType: 14
dashedName: design-a-magazine-layout-using-css-grid
# demoType can either be 'onClick' or 'onLoad'. If the project or lab doesn't have a preview, delete the property
demoType: onClick
---

# --description--
In this lab, you will design a magazine layout using CSS Grid, including concepts like grid rows and grid columns.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab. 

**User Stories:**

1. You should have a `main` element with the class `magazine-cover`.

1. Your `.magazine-cover` should contain a `header` element for the title and `section` elements for each article and the cover image.

1. You should set the display property of the `.magazine-cover` to `grid.`

1. The `.magazine-cover` should define a grid layout with the following structure:

  - The title spanning the top row.
  - Two feature articles and a cover image in the second row.
  - Two small articles and the cover image in the third row.

1. The `.magazine-cover` should have a `grid-template-columns` property that divides the space into three equal parts.

1. The `.magazine-cover` should have a `grid-template-rows` property. It should be set to `auto` for the first row, and divide equally the space occupied by the remaining two rows.

1. You should add a gap between grid items and set the width of the `.magazine-cover` to occupy most of the viewport width.

**Note:** Be sure to link your stylesheet in your HTML and apply your CSS.

# --hints--

You should have a `main` element.

```js
assert.exists(document.querySelector('main'));
```

You should have a `.magazine-cover` class.

```js
assert.exists(document.querySelector('.magazine-cover'));
```

You should have a `header` element for the title.

```js
assert.exists(document.querySelector('header'));
```

You should have a `section` elements for each article and the cover image, with classes `feature-article`, `small-article1`, `small-article2`, and `cover-image` respectively.

```js
const featureArticle = document.querySelector('section.feature-article');
const smallArticle1 = document.querySelector('section.small-article1');
const smallArticle2 = document.querySelector('section.small-article2');
const coverImage = document.querySelector('section.cover-image');
```

The display property of the `.magazine-cover` should be set to `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.magazine-cover')).display, 'grid');
```

Ensure that you have a `grid-template-areas` property with `title` spanning all three columns.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.equal(magazineCoverStyle.gridTemplateAreas.includes('"title title title"'), true);
```

The second row of the grid template should contain two `feature article` and a `cover image`.

```js
assert.equal(magazineCoverStyle.gridTemplateAreas.includes('"feature-article feature-article cover-image"'), true);
```

Your third row of the grid template should contain `small-article1`, `small-article2`, and `cover-image`.

```js 
assert.equal(magazineCoverStyle.gridTemplateAreas.includes('"small-article1 small-article2 cover-image"'), true);
```

The `.magazine-cover` should be divided by three equal columns.

```js
assert.equal(magazineCoverStyle.gridTemplateColumns, '1fr 1fr 1fr');
```

You shoould have the `grid-template-rows` property set to `auto` for the first row and divide equally the remaining two rows.

```js
assert.equal(magazineCoverStyle.gridTemplateRows, 'auto 1fr 1fr');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Magazine Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Adventure Magazine</title>
  <link rel="stylesheet" href="styles.css">
</head>
<body>
  <main class="magazine-cover">
      <header class="title">
          <h1>Adventure Awaits</h1>
      </header>
      <section class="feature-article">
          <h2>Top 10 Exotic Destinations for 2024</h2>
          <p>Discover the most breathtaking places to visit this year, from hidden beaches to mountain retreats. Our guide takes you through the best spots for your next adventure.</p>
      </section>
      <section class="small-article1">
          <h3>Gear Up: Must-Have Gadgets</h3>
          <p>Check out the latest tech and gear to make your travels more exciting and comfortable.</p>
      </section>
      <section class="small-article2">
          <h3>Meet the Explorers</h3>
          <p>Get to know the modern adventurers who are pushing the boundaries of exploration.</p>
      </section>
      <section class="cover-image">
          <img src="https://cdn.freecodecamp.org/curriculum/labs/magazine-cover.png" alt="Adventure Cover Image">
      </section>
  </main>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 0;
  padding: 0;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.magazine-cover {
  display: grid;
  grid-template-areas: 
      "title title title"
      "feature-article feature-article cover-image"
      "small-article1 small-article2 cover-image";
  grid-template-columns: 1fr 1fr 1fr;
  grid-template-rows: auto 1fr 1fr;
  gap: 10px;
  width: 80%;
  background-color: #fff;
}

.title {
  grid-area: title;
  background-color: #1B1B32;
  color: white;
  padding: 10px;
  text-align: center;
}

.feature-article {
  grid-area: feature-article;
  background-color: #f9f9f9;
  padding: 10px;
}

.small-article1, .small-article2 {
  background-color: #e9e9e9;
  padding: 10px;
}

.cover-image {
  grid-area: cover-image;
  display: flex;
  justify-content: center;
  align-items: center;
  background-color: #ddd;
  padding: 10px;
}

.cover-image img {
  max-width: 100%;
  height: auto;
} 
```
