---
id: 60b69a66b6ddb80858c51580
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

La proprietà `z-index` viene utilizzata per creare dei "livelli" per gli elementi HTML. Se hai familiarità con gli strumenti di modifica delle immagini, potresti aver già lavorato con dei livelli. Questo è un concetto simile.

Il valore corrisponde alla stratificazione dei livelli, quindi gli elementi con un valore `z-index` superiore appariranno sopra gli elementi con un valore `z-index` inferiore. Ciò può essere combinato con il posizionamento della lezione precedente per creare effetti unici.

Poiché l'elemento `back-wall` dovrà apparire "dietro" gli altri elementi che creerai, assegna all'elemento `back-wall` un `z-index` di `-1`.

# --hints--

Il selettore `#back-wall` dovrebbe avere una proprietà `z-index` impostata su `-1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.zIndex === '-1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
