---
id: 60b69a66b6ddb80858c51588
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Usando un selettore di id, definisci lo stile dell'elemento con l'id `white-hat`. Assegnagli delle proprietà `width` e `height` con il valore `0` e una proprietà `border-style` del valore di `solid`.

# --hints--

Dovresti usare un selettore `#white-hat`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat'));
```

Il selettore `#white-hat` dovrebbe avere una proprietà `width` impostata su `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.width === '0px');
```

Il selettore `#white-hat` dovrebbe avere una proprietà `height` impostata su `0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.height === '0px');
```

Il selettore `#white-hat` dovrebbe avere una proprietà `border-style` impostata su `solid`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#white-hat')?.borderStyle === 'solid');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        <div id="white-hat"></div>
        <div id="black-mask">
          <div class="eyes left"></div>
          <div class="eyes right"></div>
        </div>
        <div id="gray-instrument">
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
          <div class="black-dot"></div>
        </div>
        <div id="tan-table"></div>
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}

#offwhite-character {
  width: 300px;
  height: 550px;
  background-color: GhostWhite;
  position: absolute;
  top: 20%;
  left: 17.5%;
}

--fcc-editable-region--

--fcc-editable-region--

```
