---
id: 60b69a66b6ddb80858c51583
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Crea quattro elementi `div` all'interno dell'elemento `offwhite-character`. Assegna a questi elementi `div` i seguenti valori `id`, in ordine: `white-hat`, `black-mask`, `gray-instrument`, `tan-table`.

# --hints--

Dovresti aggiungere quattro elementi `div` all'interno dell'elemento `.offwhite-character`.

```js
assert(document.querySelectorAll('#offwhite-character div').length === 4);
```

Il primo nuovo elemento `div` dovrebbe avere un `id` con il valore `white-hat`.

```js
assert(document.querySelectorAll('#offwhite-character div')[0]?.getAttribute('id') === 'white-hat');
```

Il secondo nuovo elemento `div` dovrebbe avere un `id` con il valore `black-mask`.

```js
assert(document.querySelectorAll('#offwhite-character div')[1]?.getAttribute('id') === 'black-mask');
```

Il terzo nuovo elemento `div` dovrebbe avere un `id` con il valore `gray-instrument`.

```js
assert(document.querySelectorAll('#offwhite-character div')[2]?.getAttribute('id') === 'gray-instrument');
```

Il quarto nuovo elemento `div` dovrebbe avere un `id` con il valore `tan-table`.

```js
assert(document.querySelectorAll('#offwhite-character div')[3]?.getAttribute('id') === 'tan-table');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
    <div class="characters">
      <div id="offwhite-character">
        --fcc-editable-region--

        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  width: 100%;
  height: 60%;
  position: absolute;
  top: 0;
  left: 0;
  z-index: -1;
}
```
