---
id: 60b69a66b6ddb80858c5157e
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Assegna all'elemento `#back-wall` una proprietà `width` del `100%` e una proprietà `height` del `60%`.

# --hints--

Dovresti impostare la proprietà `width` del selettore `#back-wall` al `100%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.width === '100%');
```

Dovresti impostare la proprietà `height` del selettore `#back-wall` al `60%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('#back-wall')?.height === '60%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    <link rel="stylesheet" href="./styles.css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.8.2/css/all.css">
  </head>
  <body>
    <div id="back-wall"></div>
  </body>
</html>
```

```css
body {
  background-color: rgb(184, 132, 46);
}

#back-wall {
  background-color: #8B4513;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
