---
id: 60b80da8676fb3227967a731
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Vai avanti e collega il tuo file CSS, anche se non hai ancora scritto alcun CSS.

Aggiungi un elemento `link` con un attributo `rel` di `stylesheet` e un `href` con il valore `styles.css`.

# --hints--


Il codice dovrebbe avere un elemento `link`.

```js
assert.match(code, /<link/)
```

L'elemento `link` dovrebbe essere un elemento autoconcludente.

```js
assert(code.match(/<link[\w\W\s]+\/>/i));
```

L'elemento `link` dovrebbe essere all'interno dell'elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/>[\w\W\s]*<\/head>/i))
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/)
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <title>Picasso Painting</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
