---
id: 5f3326b143638ee1a09ff1e3
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

爲了準備創建一些實際的內容，在 `head` 元素下面添加一個 `body` 元素。

# --hints--

你應該有一個 `<body>` 開始標籤。

```js
assert(code.match(/<body>/i));
```

你應該有一個 `</body>` 結束標籤。

```js
assert(code.match(/<\/body>/i));
```

你不應該改變你的 `head` 元素。 確認你沒有刪除結束標籤。

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

你的 `body` 元素應該在你的 `head` 元素之後。

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

