---
id: 5f3313e74582ad9d063e3a38
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `html` 元素中添加一個 `head` 元素，接下來你可以添加 `title` 元素。 `title` 元素的文本應該是 `Cafe Menu`。

# --hints--

你需要一個開始標籤 `<head>`

```js
assert(code.match(/<head>/i));
```

需要一個結束標籤 `</head>`

```js
assert(code.match(/<head>/i));
```

需要一個開始標籤 `<title>`

```js
assert(code.match(/<title>/i));
```

需要一個結束標籤 `</title>`

```js
assert(code.match(/<\/title>/i));
```

所有的 `title` 元素都應嵌套在 `head` 元素中。

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

你的 `title` 元素內應該有文本 `Cafe Menu`。 您可能需要檢查您的拼寫。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
