---
id: 5f33310c1851c6c4da013250
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

是時候添加一些菜單內容。 在現有的 `header` 元素下面添加一個 `main`元素。 它最終將包含有關咖啡館提供的咖啡和甜點的價格信息。

# --hints--

你應該有一個 `<main>` 開始標籤。

```js
assert(code.match(/<main>/i));
```

你應該有一個 `</main>` 結束標籤。

```js
assert(code.match(/<\/main>/i));
```

你不應該改變你的 `header` 元素。 確認你沒有刪除結束標籤。

```js
assert($('header').length === 1);
```

你的 `main` 元素應該在你的 `header` 元素之後。

```js
const main = document.querySelectorAll('main')[0];
assert(main.previousElementSibling.tagName === 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <header>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </header>
--fcc-editable-region--
  </body>
</html>
```
