---
id: 5f331e55dfab7a896e53c3a1
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` 是網頁上幾個不可見附加信息之一，但它對搜索引擎和網頁如何顯示很有用。

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8` to tell the browser how to encode characters for the page.

# --hints--

你應該有一個 `meta` 標籤。

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 標籤應該有一個 `charset` 屬性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 屬性的值應該爲 `utf-8`。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
