---
id: 5f33294a6af5e9188dbdb8f3
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

咖啡廳的名稱是 `CAMPER CAFE`。 你的 `body` 元素中應該有一個 `h1` 元素。 用大寫字母顯示咖啡店名，使其突出顯示。

# --hints--

你應該有一個 `<h1>` 開始標籤。

```js
assert(code.match(/<h1>/i));
```

你應該有一個 `</h1>` 結束標籤。

```js
assert(code.match(/<\/h1>/i));
```

你不應該改變你現有的 `body`元素。

```js
assert($('body').length === 1);
```

你的 `h1` 元素應該嵌套在你的 `body` 元素中。

```js
assert($('h1')[0].parentElement.tagName === "BODY");
```

你的 `h1` 元素應該有文字 `CAMPER CAFE`。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
