---
id: 5f33071498eb2472b87ddee4
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

In this project, you will learn the basics of CSS (Cascading Style Sheets) by building a cafe menu. CSS 是一種用於設置 HTML 文檔樣式的語言。 它描述了 HTML 元素在屏幕上的顯示方式。

正如你在 Cat Photo App 的最後幾個步驟中所學習的那樣，開始構建你的網頁需要一個基本結構。 每個 HTML 文檔都應包含 `DOCTYPE` 聲明和 `html` 元素。 `DOCTYPE` 告訴瀏覽器文檔的 HTML 版本。 而 `html` 元素代表根元素，它包含所有其他元素。

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

添加 `<!DOCTYPE html>` 標記和具有值爲 `en` 的 `lang` 屬性的 `html` 元素。

# --hints--

你應該有 `<!DOCTYPE html>` 聲明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個 `<html lang="en">` 開始標籤，語言設置爲英語。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你應該有一個 `</html>` 結束標籤。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
