---
id: 5f33294a6af5e9188dbdb8f3
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

是時候添加一些菜單內容。 在現有的 `body` 元素裏面添加一個 `main` 元素。 它最終將包含有關咖啡館提供的咖啡和甜點的價格信息。

# --hints--

你的代碼應該有一個 `<main>` 開始標籤。

```js
assert(code.match(/<main>/i));
```

你的代碼應該有一個 `</main>` 結束標籤。

```js
assert(code.match(/<\/main>/i));
```

你不應該改變你的 `body` 元素。 確認你沒有不小心刪除結束標籤。

```js
assert($('body').length === 1);
```

你的 `main` 元素應該在你的 `body` 元素之後。

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
