---
id: 5f332a88dc25a0fd25c7687a
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

咖啡廳的名稱是 `CAMPER CAFE`。 在 `main` 元素中添加一個 `h1` 元素。 給它命名爲 “cafe”，用大寫字母，使其突出顯示。

# --hints--

你應該有一個 `<h1>` 開始標籤。

```js
assert(code.match(/<h1>/i));
```

你應該有一個 `</h1>` 結束標籤。

```js
assert(code.match(/<\/h1>/i));
```

你不應該改變你現有的 `main` 元素。

```js
assert($('main').length === 1);
```

你的 `h1` 元素應該嵌套在你的 `main` 元素中。

```js
assert($('h1')[0].parentElement.tagName === "MAIN");
```

你的 `h1` 元素應該有文字 `CAMPER CAFE`。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

