---
id: 5f332a88dc25a0fd25c7687a
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

爲了讓遊客知道這家咖啡館成立於 2020 年，在 `h1` 元素下面添加一個 `p`元素，文字爲 `Est. 2020`。

# --hints--

你應該有一個 `<p>` 開始標籤。

```js
assert(code.match(/<p>/i));
```

你應該有一個 `</p>` 結束標籤。

```js
assert(code.match(/<\/p>/i));
```

你不應該改變你現有的 `h1` 元素。 確認你沒有刪除結束標籤。

```js
assert($('h1').length === 1);
```

你的 `p` 元素應該在 `h1` 元素的下面。

```js
assert($('p')[0].previousElementSibling.tagName === 'H1');
```

你的 `p` 元素應該有文本 `Est. 2020`。

```js
assert(code.match(/<p>Est. 2020<\/p>/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
--fcc-editable-region--
  </body>
</html>
```

