---
id: 5f33071498eb2472b87ddee4
title: 步驟 1
challengeType: 0
dashedName: step-1
---

# --description--

正如你在 Cat Photo App 的最後幾個步驟中所學習的那樣，開始構建你的網頁需要一個基本結構。

添加 `<!DOCTYPE html>` 標記和具有值爲 `en` 的 `lang` 屬性的 `html` 元素。

# --hints--

需要先聲明 `DOCTYPE`。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個帶有值爲 `en` 的 `lang` 屬性的 `<html>` 開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你需要一個 `<html>` 結束標籤。 別忘了結束標籤是在開始的 `<` 後有一個 `/`。

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
