---
id: 5f33071498eb2472b87ddee4
title: 第1步：
challengeType: 0
dashedName: step-1
---

# --description--

之前在Cat Photo App的最後幾個步驟中提過，構建網頁前需要一個基本的結構。

添加 `<!DOCTYPE html>` 標記和具有值爲 `en` 的 `lang` 屬性的 `html` 元素。

# --hints--

需要先聲明`DOCTYPE`

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你應該有一個帶有值爲 `en` 的 `lang` 屬性的 `<html>` 開始標籤。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

需要一個閉合標籤 `<html>` 別忘了閉合標籤的尖括號 `<` 後需要有一個 `/`

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
