---
id: 5f331e55dfab7a896e53c3a1
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` 是網頁上幾個不可見附加信息之一，但它對搜索引擎和網頁如何顯示很有用。

在`head`元素中，嵌套一個 `meta` 元素，其屬性名爲 `charset`，值爲`utf-8`，以告訴瀏覽器如何對頁面進行編碼。

# --hints--

你應該有一個 `meta` 標籤。

```js
assert.match(code, /<meta.*?\/?>/is);
```

`meta` 元素是一個空元素，它不應該有一個結束標籤 `</meta>`。

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 標籤應該有一個 `charset` 屬性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 屬性的值應該爲 `utf-8`。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
