---
id: 5f332b23c2045fb843337579
title: 步驟 7
challengeType: 0
dashedName: step-7
---

# --description--

由於上一步添加的 `p` 元素提供了關於咖啡館的補充信息，因此將 `h1` 和 `p` 元素都嵌套在一個 `header`元素中。

# --hints--

你應該有一個 `<header>` 開始標籤。

```js
assert(code.match(/<header>/i));
```

你應該有一個 `</header>` 結束標籤。

```js
assert(code.match(/<\/header>/i));
```

你的 `h1` 元素應該嵌套在你的 `header` 元素中。

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[0].tagName === 'H1');
```

你的 `p` 元素應該嵌套在你的 `header`元素中。

```js
const header = document.querySelectorAll('header')[0];
assert(header.children[1].tagName === "P");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <h1>CAMPER CAFE</h1>
    <p>Est. 2020</p>
--fcc-editable-region--
  </body>
</html>
```

