---
id: 5f34a1fd611d003edeafd681
title: 步驟 19
challengeType: 0
dashedName: step-19
---

# --description--

棕色的背景使閱讀文本變得困難。 將 `body` 元素的背景顏色更改爲 `burlywood`，使其具有一些顏色，但你仍然可以閱讀文本。

# --hints--

你應該將 `background-color` 屬性設置爲 `burlywood`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'burlywood');
assert(hasBackground);
```

你的 `body` 元素應該有一個 `burlywood` 背景。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'burlywood');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: brown;
}
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
```
