---
id: 5f344fc1520b6719f2e35605
title: 步驟 8
challengeType: 0
dashedName: step-8
---

# --description--

菜單上將有兩個部分，一個是咖啡，一個是甜點。 在 `main` 元素中添加一個 `section` 元素，這樣你就有一個地方來放置所有可用的咖啡。

# --hints--

你應該有一個 `<section>` 開始標籤。

```js
assert(code.match(/<section\s*>/i));
```

你應該有一個 `</section>` 結束標籤。

```js
assert(code.match(/<\/section\s*>/i));
```

你不應該改變你現有的 `main` 元素。 確認你沒有刪除結束標籤。

```js
assert($('main').length === 1);
```

你的 `section`元素應該在你的 `main` 元素中。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert(section.parentElement === main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
