---
id: 5f3477cb2e27333b1ab2b955
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

現在你需要鏈接 `styles.css` 文件，以便再次應用這些樣式。 在 `head` 元素中，添加一個 `link` 元素。 設置其 `rel` 屬性值爲 `"stylesheet"`、`href` 屬性值爲 `"styles.css"`。

# --hints--

你的代碼應該有一個 `link` 元素。

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

`link` 元素是一個空元素，它不應該有一個結束標籤 `</link>`。

```js
assert.notMatch(code, /<\/link>/i);
```

你不應該改變你現有的 `head` 元素。 請確保你沒有刪除結束標記。

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

你應該有一個 `link` 元素。

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

你的 `link` 元素應該在你的 `head` 元素中。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 `link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 `link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
