---
id: 5f3477aefa51bfc29327200b
title: 步驟 14
challengeType: 0
dashedName: step-14
---

# --description--

你已經通過在 `style` 標籤內編寫 CSS 來設置了三個元素的樣式。 這樣做可以，但由於會有更多的樣式，最好把所有的樣式放在一個單獨的文件中並鏈接到它。

我們已經爲你創建了一個單獨的 `styles.css` 文件，並將編輯器視圖切換到該文件。 你可以通過編輯器頂部的標籤在不同的文件之間轉換。

首先，將你創建的樣式改寫到 `styles.css` 文件中。 請確保排除開頭和結尾的 `style` 標籤。

# --hints--

你的 `styles.css` 文件應該有 `h1, h2, p` 類型選擇器。

```js
(getUserInput) => {
  assert(getUserInput('editableContents').replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
}
```

你的選擇器應將 `text-align` 屬性設置爲 `center`。

```js
(getUserInput) => {
  assert(getUserInput('editableContents').match(/text-align:\s*center;?/));
}
```

代碼不應該包含 `<style>` 標籤。

```js
  assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
