---
id: 5f344f9c805cd193c33d829c
title: 步驟 11
challengeType: 0
dashedName: step-11
---

# --description--

你可以在 `style` 元素中指定一個樣式，併爲此設置一個屬性，像這樣：

```css
element {
 property: value;
}
```

通過設置 `text-align` 屬性值爲 `center`，讓你的 `h1` 元素居中，。

# --hints--

你的 `style` 元素中應該有一個 `h1` 元素。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

你的 `text-align` 屬性的值應該是 `center`。

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

你的 `h1` 選擇器應該將 `text-align` 屬性設置爲 `center`。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

