---
id: 5f344fbc22624a2976425065
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

在 `section` 元素中創建一個 `h2` 元素，並在它裏面填寫文字 `Coffee`。

# --hints--

你應該有一個 `<h2>` 開始標籤。

```js
assert(code.match(/<h2\s*>/i));
```

你應該有一個 `</h2>` 結束標籤。

```js
assert(code.match(/<\/h2\s*>/i));
```

你不應該改變你現有的 `section` 元素。 確認你沒有刪除結束標籤。

```js
assert($('section').length === 1);
```

你的 `h2` 元素應該在 `section` 元素中。

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

你的 `h2` 元素內應該有文本 `Coffee`。

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
