---
id: 5f3477cbcb6ba47918c1da92
title: 步驟 17
challengeType: 0
dashedName: step-17
---

# --description--

爲了使頁面樣式在移動端看起來與在桌面或筆記本電腦上相似，你需要添加一個帶有特殊 `content` 屬性的 `meta` 元素。

在 `head` 元素中添加以下內容：

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

你應該有兩個 `meta` 元素。

```js
assert(code.match(/<meta.*?\/?>/g).length === 2);
```

你的 `meta` 元素應該有一個 `name` 屬性，其值爲 `viewport`。

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/name=('|")viewport\1/) || meta[1].outerHTML.match(/name=('|")viewport\1/));
```

你的 `meta` 元素應該有一個 `content` 屬性，其值爲 `width=device-width, initial-scale=1.0`。

```js
const meta = $('meta');
assert(meta[0].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1].outerHTML.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
