---
id: 5f3477ae9675db8bb7655b30
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

在上一步中，你使用 <dfn>類型選擇器</dfn> 來設置 `h1` 元素的樣式。 將 `h2` 和 `p` 元素居中，通過在現有的 `style` 元素中給它們添加一個新的選擇器來實現。

# --hints--

你不應更改現有的 `h1` 選擇器。

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

你不應該添加一個新的 `style` 標籤。 將新的 CSS 規則添加到現有的 `style` 標籤中。

```js
const hasManyStyleTags = document.querySelectorAll('style').length > 1;
assert(!hasManyStyleTags);
```

你應該添加一個新的 `h2` 選擇器。

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

你應該添加一個新的 `p` 選擇器。

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

你的 `h1` 元素應該有一個 `text-align` 值爲 `center`。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

你的 `h2` 元素應該有一個 `text-align` 值爲 `center`。

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

你的 `p` 元素應該有一個 `text-align` 值爲 `center`。

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

