---
id: 5f3477ae34c1239cafe128be
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

你現在有三個樣式完全相同的類型選擇器。 你可以通過創建一個選擇器列表將同一組樣式添加到許多元素上。 每個選擇器都用逗號分隔，如下所示：

```css
selector1, selector2 {
  property: value;
}
```

刪除三個現有的類型選擇器，並將它們替換爲一個選擇器列表，該列表使 `h1`、`h2` 和 `p` 元素的文本居中。

# --hints--

你應該對所有三個元素使用一個類型選擇器，`h1, h2, p`。 請一定按照這個順序。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

`style` 元素中應該只有一個選擇器。

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

