---
id: 5f344fad8bf01691e71a30eb
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

到目前爲止，在你創建的內容的演示和外觀方面，你一直受到限制。 爲了開始控制，在 `head` 元素中添加一個 `style` 元素。

# --hints--

你應該有一個 `<style>` 開始標籤。

```js
assert(code.match(/<style\s*>/i));
```

你應該有一個 `</style>` 結束標籤。

```js
assert(code.match(/<\/style\s*>/));
```

你的 `style` 元素都應嵌套在 `head` 元素中。

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
