---
id: 5f356ed69db0a491745e2bb6
title: 步驟 27
challengeType: 0
dashedName: step-27
---

# --description--

要給 `div` 元素應用類屬性，需要移除 `id` 屬性，並在 `div` 元素的開始標籤中添加 `class` 屬性。 確保將 class 的值設置爲 `menu`。

# --hints--

你的 `div` 仍然應該渲染。 確保你沒有改變 `<div>` 標籤。

```js
assert($('div').length === 1);
```

你的 `div` 元素應該有 `menu` 類。

```js
assert($('div').attr('class').includes('menu'));
```

你的 `div` 元素不應該再有 `menu` id。

```js
assert(!$('div#menu').length);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

