---
id: 5f35e5c44359872a137bd98f
title: 步驟 28
challengeType: 0
dashedName: step-28
---

# --description--

由於咖啡館主要銷售的產品是咖啡，因此你可以使用咖啡豆的圖像作爲頁面背景。

刪除 `body` 類型選擇器中的註釋及其內容。 現在添加一個 `background-image` 屬性並將其值設置爲 `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`。

# --hints--

你應該刪除註釋掉的 `background-color` 屬性。

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

你的 `body` 選擇器不應有任何註釋。

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

你應該將 `background-image` 屬性設置爲 `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

你的 `body` 元素應該有咖啡豆的背景圖片。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');
console.log(bodyBackground);
assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

