---
id: 5f35e5c4321f818cdc4bed30
title: 步驟 29
challengeType: 0
dashedName: step-29
---

# --description--

看起來很好。 是時候開始添加一些菜單項了。 在 `Coffee` 標題下添加一個空的 `article` 元素。 它將包含你當前提供的每種咖啡的風味和價格。

# --hints--

你應該有一個 `<article>` 開始標籤。

```js
assert(code.match(/<article>/i));
```

你應該有一個 `</article>` 結束標籤。

```js
assert(code.match(/<\/article>/i));
```

你不應更改現有的 `h2` 元素。

```js
assert($('h2').length === 1);
```

你的 `article` 元素應該在你的 `h2` 元素之後。

```js
const article = $('article')[0];
assert(article.previousElementSibling.tagName === 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

