---
id: 5f3c866de7a5b784048f94b1
title: 步驟 36
challengeType: 0
dashedName: step-36
---

# --description--

這正是你想要的，但如果口味和價格在同一條線上，那就太好了。 `p` 元素是 <dfn>塊級</dfn> 元素，因此它們佔據了其父元素的整個寬度。

要將它們放在同一行，你需要對 `p` 元素應用一些樣式，因此它們更像<dfn>內聯</dfn>元素。 將值爲 `item` 的 `class` 屬性添加到 `Coffee` 標題下的第一個 `article` 元素。

# --hints--

你應該將 `item` class 應用於你的 `article` 元素。

```js
assert(code.match(/<article\s*class=('|")item\1\s*>/i))
```

你應該只有一個 `item` class 元素。

```js
assert($('.item').length === 1);
```

你的第一個 `article` 元素應該具有 `item` class。

```js
assert($('article')[0].className === 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
