---
id: 5f3c866d5414453fc2d7b480
title: 步驟 31
challengeType: 0
dashedName: step-31
---

# --description--

從現有的咖啡/價格對開始，使用 `article` 元素添加以下咖啡和價格，每個元素內部有兩個嵌套的 `p` 元素。 和之前一樣，第一個 `p` 元素的文本應該包含咖啡風味，第二個 `p` 元素的文本應該包含價格。

```bash
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

# --hints--

你應該有五個 `article` 元素。

```js
assert($('article').length === 5);
```

每個 `article` 元素應該有兩個 `p` 元素。

```js
const articles = $('article');
assert(articles[0].children.length === 2);
assert(articles[1].children.length === 2);
assert(articles[2].children.length === 2);
assert(articles[3].children.length === 2);
assert(articles[4].children.length === 2);
```

你的第一個 `article` 元素應該包含 `p` 元素，其中包含文本 `French Vanilla` 和 `3.00`。

```js
const children = $('article')[0].children;
assert(children[0].innerText.match(/French Vanilla/i));
assert(children[1].innerText.match(/3\.00/i));
```

你的第二個 `article` 元素應具有 `p` 元素，其中包含文本 `Caramel Macchiato` 和 `3.75`。

```js
const children = $('article')[1].children;
assert(children[0].innerText.match(/Caramel Macchiato/i));
assert(children[1].innerText.match(/3\.75/i));
```

你的第三個 `article` 元素應該包含 `p` 元素，其中包含文本 `Pumpkin Spice` 和 `3.50`。

```js
const children = $('article')[2].children;
assert(children[0].innerText.match(/Pumpkin Spice/i));
assert(children[1].innerText.match(/3\.50/i));
```

你的第四個 `article` 元素應該包含 `p` 元素，其中包含文本 `Hazelnut` 和 `4.00`。

```js
const children = $('article')[3].children;
assert(children[0].innerText.match(/Hazelnut/i));
assert(children[1].innerText.match(/4\.00/i));
```

你的第五個 `article` 元素應該包含帶有文本 `Mocha` 和 `4.50` 的 `p` 元素。

```js
const children = $('article')[4].children;
assert(children[0].innerText.match(/Mocha/i));
assert(children[1].innerText.match(/4\.50/i));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
