---
id: 5f769702e6e33127d14aa120
title: 步驟 35
challengeType: 0
dashedName: step-35
---

# --description--

現在將文本與具有 `price` class 的元素的對齊方式設置爲 `right`。

# --hints--

你應該有一個 `price` class 選擇器。

```js
assert(code.match(/\.price\s*{/i));
```

你的 `price` class 選擇器應將 `text-align` 屬性設置爲 `right`。

```js
assert(code.match(/\.price\s*{\s*text-align:\s*right;?/i));
```

你的 `.price` 元素應該靠右對齊。

```js
assert($('.price').css('text-align') === 'right');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

--fcc-editable-region--

--fcc-editable-region--

```

