---
id: 69de4308e652aaeb5ccdb03e
title: Step 1
challengeType: 20
dashedName: step-1
# demoType can either be 'onClick' or 'onLoad'. If the project or lab doesn't have a preview, delete the property
# demoType: onClick
---

# --description--

In this workshop, you will learn how to work with strings in Python using different string methods.

Strings are used to store text data like `names`, `emails`, or `addresses`. For example:

```python
name = "" #empty string
```

An empty string is a string that has no characters inside the quotes.

In this step, you will create a variable and assign it an empty string.

Create a variable named `email` and assign it an `empty string`.

# --hints--

You should define a variable called `email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("email")`)) })
```

You should assign a empty `string` value to `email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("email")
not node.is_empty() and (
  node.is_equivalent("email = ''") or 
  node.is_equivalent('email = ""')
)
`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
