---
id: 69de54575ee98e22b24527e8
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

In the previous step, you created an empty string.

A string can store text when it is written inside quotes. For example:

```py
name = 'Kiruthika'
```

Now, update the `email` variable so that it contains the email address `'alice.johnson@company.com'`.

# --hints--

You should define a variable called `email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("email")`)) })
```

You should assign the exact value `'alice.johnson@company.com'` to `email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("email")
not node.is_empty() and (
  node.is_equivalent("email = 'alice.johnson@company.com'")
)
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
email = ''
--fcc-editable-region--
```
