---
id: 69e20a0451c7a3ef189514d3
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now that you know the position of `"@"` in the email, you can extract the username - the part before `"@"`.

Python allows you to extract a portion of a string using slicing. The syntax is:

```py
string[start:end]
```

If you leave out `start`, it defaults to `0`. For example:

```py
text = "hello world"
first_word = text[:5]
print(first_word)  # Output: hello
```

Use slicing on the `email` string with the `at_position` index to extract the username (everything before `"@"`) and store it in the `username` variable, then print it.

# --hints--

You should create a variable named `username`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("username")`)) })
```

You should use slicing with `at_position` to assign `username`.

```js
({ test: () => assert(runPython(`"email[:at_position]" in _code`)) })
```

You should print `username`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(username)")`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"
at_position = email.find("@")
--fcc-editable-region--

--fcc-editable-region--
```
