---
id: 69e20a0451c7a3ef189514d3
title: Step 4
challengeType: 20
dashedName: step-4
---

# --description--

Now that you know the position of `@` in the email, you can extract the username, which is placed before the at sign.

In the previous workshop, you extracted a portion of a string using slicing. As a reminder, the syntax is:


# --hints--

You should create a variable named `username`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable('username')`)) })
```

You should use slicing with `at_position` to assign `username`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable('username')
not node.is_empty() and node.is_equivalent("username = email[:at_position]")
`)) })
```

You should print `username` with the label `'Username: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Username: ' + username)")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

--fcc-editable-region--

--fcc-editable-region--
```
