---
id: 69e6299710319e7cb8ebdc37
title: Step 6
challengeType: 20
dashedName: step-6
---

# --description--

Real-world data is often messy with extra spaces and inconsistent casing.

The `.strip()` method removes whitespace from both ends of a string. The `.title()` method capitalises the first letter of each word. You can chain them together:

```python
messy = "  hello world  "
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Create a variable `messy_name` and assign it `"  sARaH dAVis  "`. Then create `clean_name` by calling `.strip().title()` on `messy_name`. 

Then print both like this:

```python
print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
```

# --hints--

You should create a variable named `messy_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("messy_name")`)) })
```

You should assign `"  sARaH dAVis  "` to `messy_name`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("messy_name")
not node.is_empty() and (
  node.is_equivalent('messy_name = "  sARaH dAVis  "') or
  node.is_equivalent("messy_name = '  sARaH dAVis  '")
)
`)) })
```

You should create a variable named `clean_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("clean_name")`)) })
```

You should use `.strip()` and `.title()` on `messy_name`.

```js
({ test: () => assert(runPython(`"messy_name.strip().title()" in _code`)) })
```

You should print the original name.

```js
({ test: () => assert(runPython(`"Original name:" in _code and "messy_name" in _code`)) })
```

You should print the cleaned name.

```js
({ test: () => assert(runPython(`"Cleaned name:" in _code and "clean_name" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

--fcc-editable-region--

--fcc-editable-region--
```
