---
id: 69e6284f3b7659c5a05c2a79
title: Step 5
challengeType: 20
dashedName: step-5
---

# --description--

You can check whether a substring exists inside a string using the `in` keyword. It returns `True` or `False`. For example:

```py
sentence = 'I love Python'
result = 'Python' in sentence
print(result)  # Output: True
```

Create two variables:

- `has_at` - checks if `'@'` is in `email`
- `has_com` - checks if `'.com'` is in `email`

Print `has_at` with the label `'Has @ symbol: '` and `has_com` with the label `'Has .com: '`.

# --hints--

You should use `'@' in email` to assign `has_at`.

```js
({ test: () => assert(runPython(`has_at == ('@' in email)`)) })
```

You should use `'.com' in email` to assign `has_com`.

```js
({ test: () => assert(runPython(`has_com == ('.com' in email)`)) })
```

You should print `has_at` with the label `'Has @ symbol: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Has @ symbol: ' + str(has_at))")`)) })
```

You should print `has_com` with the label `'Has .com: '`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print('Has .com: ' + str(has_com))")`)) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'

at_position = email.find('@')
print('Email: ' + email)
print('Position of @: ' + str(at_position))

username = email[:at_position]
print('Username: ' + username)

--fcc-editable-region--

--fcc-editable-region--
```
