---
id: 69e62ce96df237884da25282
title: Step 7
challengeType: 20
dashedName: step-7
---

# --description--

The `.lower()` method converts all characters in a string to lowercase. This is useful for normalising email addresses:

```python
email = "  HELLO@EXAMPLE.COM  "
clean = email.strip().lower()
print(clean)  # Output: hello@example.com
```

Create a variable `messy_email` and assign it `"  SARAH.DAVIS@COMPANY.COM  "`.

Then create `clean_email` by calling `.strip().lower()` on `messy_email`.

Then print both like this:

```python
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")
```

# --hints--

You should create a variable named `messy_email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("messy_email")`)) })
```

You should assign `"  SARAH.DAVIS@COMPANY.COM  "` to `messy_email`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("messy_email")
not node.is_empty() and (
  node.is_equivalent('messy_email = "  SARAH.DAVIS@COMPANY.COM  "') or
  node.is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
)
`)) })
```

You should create a variable named `clean_email`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("clean_email")`)) })
```

You should use `.strip().lower()` on `messy_email` to assign `clean_email`.

```js
({ test: () => assert(runPython(`"messy_email.strip().lower()" in _code`)) })
```

You should print the original email.

```js
({ test: () => assert(runPython(`"Original email:" in _code and "messy_email" in _code`)) })
```

You should print the cleaned email.

```js
({ test: () => assert(runPython(`"Cleaned email:" in _code and "clean_email" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "
clean_name  = messy_name.strip().title()
print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")

--fcc-editable-region--

--fcc-editable-region--
```
