---
id: 69e75dc335b6ec9903a87d3e
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

You can also use `.replace()` to remove unwanted characters. For example, removing dashes from a phone number:

```py
phone = "456-147"
clean = phone.replace("-", "")
print(clean)  # Output: 456147
```

Create a variable `phone` and assign it `"123-456-7890"`. Then create `phone_clean` by replacing `"-"` with `""`.

Then, print the `phone` variable with the label `"Phone with dashes: "`. 
Also, print the `phone_clean` variable with the label `"Phone without dashes: "`.

# --hints--

You should create a variable named `phone`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("phone")`)) })
```

You should assign `"123-456-7890"` to `phone`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("phone")
not node.is_empty() and (
  node.is_equivalent('phone = "123-456-7890"') or
  node.is_equivalent("phone = '123-456-7890'")
)
`)) })
```

You should create a variable named `phone_clean`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("phone_clean")`)) })
```

You should use `.replace("-","")` on `phone` to assign `phone_clean`.

```js
({ test: () => assert(runPython(`'phone.replace("-","")' in _code or "phone.replace('-','')" in _code`)) })
```

You should print `phone` with the label "Phone with dashes".

```js
({ test: () => assert(runPython(`"Phone with dashes: " in _code and "phone" in _code`)) })
```

You should print `phone_clean` with the label "Phone without dashes". 

```js
({ test: () => assert(runPython(`"Phone without dashes: " in _code and "phone_clean" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)

--fcc-editable-region--

--fcc-editable-region--
```
