---
id: 69e76e91402d7850510a5f18
title: Step 13
challengeType: 20
dashedName: step-13
---

# --description--

The `.join(list)` method is the opposite of `.split()`. It joins a list of strings into one string using a separator. For example:

```py
parts = ["apple", "banana", "cherry"]
joined = " - ".join(parts)
print(joined)  # Output: apple - banana - cherry
```

Using `address_parts`, create a variable `rejoined` by joining with `" | "`.

Then print `rejoined` with the label `"Rejoined: "`.

# --hints--

You should create a variable named `rejoined`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("rejoined")`)) })
```

You should use `" | ".join(address_parts)` to assign `rejoined`.

```js
({ test: () => assert(runPython(`'" | ".join(address_parts)' in _code or "' | '.join(address_parts)" in _code`)) })
```

You should print `rejoined` with the label `"Rejoined: "`.

```js
({ test: () => assert(runPython(`"Rejoined:" in _code and "rejoined" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)


phone = "123-456-7890"
phone_clean = phone.replace('-','')
print("Phone with dashes: " + phone)
print("Phone without dashes: " + phone_clean)

full_address = "123 Main Street, Springfield, IL"
address_parts = full_address.split(", ")
print("Address parts: " + str(address_parts))

--fcc-editable-region--

--fcc-editable-region--
```
