---
id: 69e73a141addc6f442924b0d
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

You already used `.find()` and slicing on the `email` variable. Now you will use the same technique on `clean_email`.

Find the position of `"@"` in `clean_email` and store it in `at_pos`. Then use slicing to extract everything before `"@"` and store it in `email_user`.

Then print the `clean_email` with this lable `"Email: "`
Also, print `email_user` with the lable `"Username: "`

# --hints--

You should create a variable named `at_pos`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("at_pos")`)) })
```

You should use `.find("@")` on `clean_email` to assign `at_pos`.

```js
({ test: () => assert(runPython(`"clean_email.find(" in _code and '"@"' in _code or "'@'" in _code`)) })
```

You should create a variable named `email_user`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("email_user")`)) })
```

You should use slicing with `at_pos` to assign `email_user`.

```js
({ test: () => assert(runPython(`"clean_email[:at_pos]" in _code`)) })
```

You should print `clean_email` with the label.

```js
({ test: () => assert(runPython(`_Node(_code).has_call('print("Email: " + clean_email)')`)) })
```

You should print `email_user` with the label.

```js
({ test: () => assert(runPython(`_Node(_code).has_call('print("Username: " + email_user)')`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

--fcc-editable-region--

--fcc-editable-region--
```
