---
id: 69e7766b93d27261775d1cf1
title: Step 15
challengeType: 20
dashedName: step-15
---

# --description--

In previous steps you used `+` to build print statements. Python also has f-strings which make it easier to include variables inside strings. You write an `f` before the opening quote and wrap variables in `{}`:

```python
name = "Sarah"
print(f"Hello, {name}")  # Output: Hello, Sarah
```

Now use f-strings to print the following variables with their labels:

- `employee_code` with the label `"Code: "`
- `starts_with_emp` with the label `"Starts with EMP: "`
- `clean_email` with the label `"Email: "`
- `ends_with_com` with the label `"Ends with .com: "`

# --hints--

You should print `employee_code` using an f-string with the label `"Code: "`.

```js
({ test: () => assert(runPython(`
any("Code:" in line and "employee_code" in line and line.strip().startswith("print(f") for line in _code.splitlines())
`)) })
```

You should print `starts_with_emp` using an f-string with the label `"Starts with EMP: "`.

```js
({ test: () => assert(runPython(`
any("Starts with EMP:" in line and "starts_with_emp" in line and line.strip().startswith("print(f") for line in _code.splitlines())
`)) })
```

You should print `clean_email` using an f-string with the label `"Email: "`.

```js
({ test: () => assert(runPython(`
any("Email:" in line and "clean_email" in line and line.strip().startswith("print(f") for line in _code.splitlines())
`)) })
```

You should print `ends_with_com` using an f-string with the label `"Ends with .com: "`.

```js
({ test: () => assert(runPython(`
any("Ends with .com:" in line and "ends_with_com" in line and line.strip().startswith("print(f") for line in _code.splitlines())
`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)

phone = "123-456-7890"
phone_clean = phone.replace("-", "")
print("Phone with dashes: " + phone)
print("Phone without dashes: " + phone_clean)

full_address = "123 Main Street, Springfield, IL"
address_parts = full_address.split(", ")
print("Address parts: " + str(address_parts))

rejoined = " | ".join(address_parts)
print("Rejoined: " + rejoined)

employee_code = "EMP-2024-SD"
starts_with_emp = employee_code.startswith("EMP")
ends_with_com = clean_email.endswith(".com")

--fcc-editable-region--

--fcc-editable-region--
```
