---
id: 69e736c39dfa403496e8071e
title: Step 8
challengeType: 20
dashedName: step-8
---

# --description--

The `.upper()` method converts all characters to uppercase. This is useful for name badges or headings:

```py
name = "sarah davis"
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Using `clean_name`, create a variable named `name_badge` by calling `.upper()` on it. 

Then print the `name_badge` with the label `"Name badge: "`

# --hints--

You should create a variable named `name_badge`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("name_badge")`)) })
```

You should use `.upper()` on `clean_name` to assign `name_badge`.

```js
({ test: () => assert(runPython(`"clean_name.upper()" in _code`)) })
```

You should print `name_badge` with the label.

```js
({ test: () => assert(runPython(`"Name badge: " in _code and "name_badge" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

--fcc-editable-region--

--fcc-editable-region--
```
