---
id: 69e73f4f5f986657647bbd63
title: Step 10
challengeType: 20
dashedName: step-10
---

# --description--


The `.replace(old, new)` method replaces every occurrence of a substring with another. You can also chain it with other methods. For example:

```py
text = "john.doe"
name = text.replace(".", " ").title()
print(name)  # Output: John Doe
```

Using `email_user`, create a variable named `display_name` by replacing `"."` with `" "` and then calling `.title()` on the result.

Then, print the `display_name` variable with the label `"Display name: "`.

# --hints--

You should create a variable named `display_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("display_name")`)) })
```

You should use `.replace(".", " ")` on `email_user`.

```js
({ test: () => assert(runPython(`'email_user.replace(".", " ")' in _code or "email_user.replace('.', ' ')" in _code`)) })
```

You should chain the `.title()` method after `.replace()` when assigning to display_name.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("display_name")
"title()" in str(node)
`)) })
```

You should print `display_name` with the label.

```js
({ test: () => assert(runPython(`"Display name: " in _code and "display_name" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]

print("Email: " + clean_email)
print("Username: " + email_user)

--fcc-editable-region--

--fcc-editable-region--
```
