---
id: 69e76d1997661be7f0040e2d
title: Step 12
challengeType: 20
dashedName: step-12
---

# --description--

The `.split(separator)` method splits a string into a list of substrings at each occurrence of the separator. For example:

```py
fruits = "apple, banana, cherry"
parts = fruits.split(", ")
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

Create a variable `full_address` and assign it `"123 Main Street, Springfield, IL"`. Then create `address_parts` by splitting on `", "`.

Then print `address_parts` with the label `"Address parts: "`. Since `address_parts` is a list, you will need to convert it to a string using `str()` before concatenating it with the label.

# --hints--

You should create a variable named `full_address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_address")`)) })
```

You should assign `"123 Main Street, Springfield, IL"` to `full_address`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("full_address")
not node.is_empty() and (
  node.is_equivalent('full_address = "123 Main Street, Springfield, IL"') or
  node.is_equivalent("full_address = '123 Main Street, Springfield, IL'")
)
`)) })
```

You should create a variable named `address_parts`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address_parts")`)) })
```

You should use `.split(", ")` on `full_address` to assign `address_parts`.

```js
({ test: () => assert(runPython(`"full_address.split" in _code and "address_parts" in _code`)) })
```

You should print `address_parts` with the label `"Address parts: "`.

```js
({ test: () => assert(runPython(`"Address parts:" in _code and "address_parts" in _code`)) })
```

You should use `str()` to convert `address_parts` when printing.

```js
({ test: () => assert(runPython(`"str(address_parts)" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)

phone = "123-456-7890"
phone_clean = phone.replace('-','')
print("Phone with dashes: " + phone)
print("Phone without dashes: " + phone_clean)

--fcc-editable-region--

--fcc-editable-region--
```
