---
id: 69e774b82fe0bfdb684f95b4
title: Step 14
challengeType: 20
dashedName: step-14
---

# --description--

Sometimes you need to check how a string starts or ends. Python provides two methods for this:

- `.startswith(prefix)` returns `True` if the string begins with the given prefix
- `.endswith(suffix)` returns `True` if the string ends with the given suffix

For example:

```python
code = "EMP-001"
print(code.startswith("EMP"))  # True
print(code.endswith("001"))    # True
```

Create a variable `employee_code` and assign it `"EMP-2024-SD"`. Then:

- Create `starts_with_emp` by checking if `employee_code` starts with `"EMP"`
- Create `ends_with_com` by checking if `clean_email` ends with `".com"`

# --hints--

You should create a variable named `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

You should assign `"EMP-2024-SD"` to `employee_code`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_variable("employee_code")
not node.is_empty() and (
  node.is_equivalent('employee_code = "EMP-2024-SD"') or
  node.is_equivalent("employee_code = 'EMP-2024-SD'")
)
`)) })
```

You should create a variable named `starts_with_emp`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("starts_with_emp")`)) })
```

You should use `.startswith("EMP")` on `employee_code`.

```js
({ test: () => assert(runPython(`"employee_code.startswith" in _code and "EMP" in _code`)) })
```

You should create a variable named `ends_with_com`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("ends_with_com")`)) })
```

You should use `.endswith(".com")` on `clean_email`.

```js
({ test: () => assert(runPython(`"clean_email.endswith" in _code and ".com" in _code`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)

phone = "123-456-7890"
phone_clean = phone.replace("-", "")
print("Phone with dashes: " + phone)
print("Phone without dashes: " + phone_clean)

full_address = "123 Main Street, Springfield, IL"
address_parts = full_address.split(", ")
print("Address parts: " + str(address_parts))

rejoined = " | ".join(address_parts)
print("Rejoined: " + rejoined)

--fcc-editable-region--

--fcc-editable-region--
```
