---
id: 69eb3158c33d9a229a79951c
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

The `.count(substring)` method returns the number of times a substring appears in a string. For example:

```py
text = "banana"
print(text.count("a"))  # Output: 3
```

Create two variables:

- `dot_count` — count how many `"."` appear in `clean_email`
- `dash_count` — count how many `"-"` appear in `employee_code`

Then print both variables using f-strings with the labels `"Dots in email: "` and `"Dashes in code: "`.

# --hints--

You should create a variable named `dot_count`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("dot_count")`)) })
```

You should use `.count(".")` on `clean_email` to assign `dot_count`.

```js
({ test: () => assert(runPython(`
any("clean_email.count" in line and "dot_count" in line for line in _code.splitlines())
`)) })
```

You should create a variable named `dash_count`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("dash_count")`)) })
```

You should use `.count("-")` on `employee_code` to assign `dash_count`.

```js
({ test: () => assert(runPython(`
any("employee_code.count" in line and "dash_count" in line for line in _code.splitlines())
`)) })
```

You should print `dot_count` with the label `"Dots in email: "`.

```js
({ test: () => assert(runPython(`
any("Dots in email:" in line and "dot_count" in line and line.strip().startswith("print(f") for line in _code.splitlines())
`)) })
```

You should print `dash_count` with the label `"Dashes in code: "`.

```js
({ test: () => assert(runPython(`
any("Dashes in code:" in line and "dash_count" in line and line.strip().startswith("print(f") for line in _code.splitlines())
`)) })
```

# --seed--

## --seed-contents--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)

phone = "123-456-7890"
phone_clean = phone.replace("-", "")
print("Phone with dashes: " + phone)
print("Phone without dashes: " + phone_clean)

full_address = "123 Main Street, Springfield, IL"
address_parts = full_address.split(", ")
print("Address parts: " + str(address_parts))

rejoined = " | ".join(address_parts)
print("Rejoined: " + rejoined)

employee_code = "EMP-2024-SD"
starts_with_emp = employee_code.startswith("EMP")
ends_with_com = clean_email.endswith(".com")

print(f"Code: {employee_code}")
print(f"Starts with EMP: {starts_with_emp}")
print(f"Email: {clean_email}")
print(f"Ends with .com: {ends_with_com}")

--fcc-editable-region--

--fcc-editable-region--
```


# --solutions--

```py
email = "alice.johnson@company.com"

at_position = email.find("@")
print("Email: " + email)
print("Position of @: " + str(at_position))

username = email[:at_position]
print("Username: " + username)

has_at = "@" in email
has_com = ".com" in email
print("Has @ symbol: " + str(has_at))
print("Has .com: " + str(has_com))

messy_name = "  sARaH dAVis  "

messy_email = "  SARAH.DAVIS@COMPANY.COM  "

clean_name = messy_name.strip().title()
clean_email = messy_email.strip().lower()

print("Original name: '" + messy_name + "'")
print("Cleaned name: '" + clean_name + "'")
print("Original email: '" + messy_email + "'")
print("Cleaned email: '" + clean_email + "'")

name_badge = clean_name.upper()
print("Name badge: " + name_badge)

at_pos = clean_email.find("@")
email_user = clean_email[:at_pos]
display_name = email_user.replace(".", " ").title()

print("Email: " + clean_email)
print("Username: " + email_user)
print("Display name: " + display_name)

phone = "123-456-7890"
phone_clean = phone.replace("-", "")
print("Phone with dashes: " + phone)
print("Phone without dashes: " + phone_clean)

full_address = "123 Main Street, Springfield, IL"
address_parts = full_address.split(", ")
print("Address parts: " + str(address_parts))

rejoined = " | ".join(address_parts)
print("Rejoined: " + rejoined)

employee_code = "EMP-2024-SD"
starts_with_emp = employee_code.startswith("EMP")
ends_with_com = clean_email.endswith(".com")

print(f"Code: {employee_code}")
print(f"Starts with EMP: {starts_with_emp}")
print(f"Email: {clean_email}")
print(f"Ends with .com: {ends_with_com}")

dot_count = clean_email.count(".")
dash_count = employee_code.count("-")
print(f"Dots in email: {dot_count}")
print(f"Dashes in code: {dash_count}")
```
