---
id: 65687a005aba3ea815b84e68
title: Step 16
challengeType: 20
dashedName: step-16
---

# --description--

Remove the `print` call from the `main` function.

# --hints--

You should not have a `print` call within the `main` function.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;
        const no_comments = __helpers.python.removeComments(function_body);
        assert.notMatch(no_comments, /print\(/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)
--fcc-editable-region--

main()
```
