---
id: 66faa1000000000000000014
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Time to put a player on screen. `this.physics.add.sprite(x, y, key)` creates a sprite with arcade physics enabled — meaning it can have velocity, collide with other bodies, and respect world bounds.

Inside `create`, add a sprite at position `(240, 580)` using the `"player"` texture and assign it to `this.player`.

# --hints--

`this.player` should be a physics sprite at `(240, 580)` with the `"player"` texture.

```js
assert.match(code, /this\.player\s*=\s*this\.physics\.add\.sprite\(\s*240\s*,\s*580\s*,\s*([\'"`])player\1\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image("player", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><polygon points='16,2 30,30 2,30' fill='%2300ff88'/></svg>");
    this.load.image("enemy", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><rect x='2' y='2' width='28' height='28' fill='%23ff3366'/></svg>");
    this.load.image("bullet", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='6' height='16'><rect width='6' height='16' fill='%23ffff00'/></svg>");
    this.load.image("star", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='40' height='40'><circle cx='10' cy='10' r='1' fill='white'/><circle cx='30' cy='25' r='1' fill='white'/><circle cx='20' cy='35' r='1' fill='white'/></svg>");
  }

  create() {
--fcc-editable-region--

--fcc-editable-region--
  }

  update(time) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: MainScene
};

const game = new Phaser.Game(config);

```
