---
id: 66faa1000000000000000006
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

With the scene class in place, it's time to configure Phaser. The configuration object tells Phaser what kind of renderer to use, how big the canvas should be, and which scenes to run.

Below the `MainScene` class, declare a `const` named `config` and assign it an empty object `{}`.

# --hints--

You should declare a `const` named `config`.

```js
assert.match(code, /const\s+config\s*=\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
  }

  update(time) {
  }
}

--fcc-editable-region--

--fcc-editable-region--

```
