---
id: 66faa100000000000000000d
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

The game starts — but no scene runs because the config doesn't reference one yet. Phaser needs an explicit `scene` entry in the config to know what to launch.

Inside `config` (after `backgroundColor`), add a `scene` property with the value `MainScene`.

# --hints--

Your `config` object should have a `scene` property set to `MainScene`.

```js
assert.match(code, /scene\s*:\s*MainScene\b/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
  }

  update(time) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
--fcc-editable-region--

--fcc-editable-region--
};

const game = new Phaser.Game(config);

```
