---
id: 66faa1000000000000000002
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Every Phaser scene needs a constructor that calls `super()` with the scene's key. The key is how Phaser identifies the scene when switching between them.

Inside `MainScene`, add a `constructor` method. Inside it, call `super("MainScene")`.

# --hints--

Your `MainScene` class should have a `constructor` method.

```js
assert.match(code, /constructor\s*\(\s*\)\s*\{/);
```

Your constructor should call `super("MainScene")`.

```js
assert.match(code, /super\(\s*(['"`])MainScene\1\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
--fcc-editable-region--

--fcc-editable-region--
}


```
