---
id: 66faa100000000000000000f
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

By default, arcade physics applies gravity to every body. For a top-down space game you don't want that — set gravity to zero so ships float in space.

Inside the `physics` block, add an `arcade` property whose value is an object with `gravity` set to `{ y: 0 }`.

# --hints--

Your `physics` property should have an `arcade` object with `gravity: { y: 0 }`.

```js
assert.match(code, /arcade\s*:\s*\{\s*gravity\s*:\s*\{\s*y\s*:\s*0\s*\}\s*\}/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
  }

  update(time) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
--fcc-editable-region--

--fcc-editable-region--
  },
  scene: MainScene
};

const game = new Phaser.Game(config);

```
