---
id: 66faa1000000000000000011
title: Step 17
challengeType: 0
dashedName: step-17
---

# --description--

Load the enemy sprite the same way, with the key `"enemy"` and this data URI:

```js
"data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><rect x='2' y='2' width='28' height='28' fill='%23ff3366'/></svg>"
```

# --hints--

You should call `this.load.image("enemy", ...)` with the exact enemy data URI inside `preload`.

```js
const preloadMatch = code.match(/preload\s*\(\s*\)\s*\{([\s\S]*?)\n\s*\}\s*\n\s*create\s*\(/);
assert.isNotNull(preloadMatch);
assert.match(preloadMatch[1], /this\.load\.image\(\s*(['"`])enemy\1\s*,\s*(["`])data:image\/svg\+xml;utf8,<svg xmlns='http:\/\/www\.w3\.org\/2000\/svg' width='32' height='32'><rect x='2' y='2' width='28' height='28' fill='%23ff3366'\/><\/svg>\2\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image("player", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><polygon points='16,2 30,30 2,30' fill='%2300ff88'/></svg>");
--fcc-editable-region--

--fcc-editable-region--
  }

  create() {
  }

  update(time) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: MainScene
};

const game = new Phaser.Game(config);

```
