---
id: 66faa100000000000000001e
title: Step 30
challengeType: 0
dashedName: step-30
---

# --description--

In `update`, fire a bullet when space is pressed and enough time has passed.

Add an `if` block that checks both `this.fireKey.isDown` and `time > this.lastFired`. Inside the block, call `this.fireBullet()`.

# --hints--

You should check `this.fireKey.isDown && time > this.lastFired`.

```js
assert.match(code, /this\.fireKey\.isDown\s*&&\s*time\s*>\s*this\.lastFired/);
```

You should call `this.fireBullet()` inside the `if` block.

```js
assert.match(code, /this\.fireBullet\(\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image("player", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><polygon points='16,2 30,30 2,30' fill='%2300ff88'/></svg>");
    this.load.image("enemy", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><rect x='2' y='2' width='28' height='28' fill='%23ff3366'/></svg>");
    this.load.image("bullet", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='6' height='16'><rect width='6' height='16' fill='%23ffff00'/></svg>");
    this.load.image("star", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='40' height='40'><circle cx='10' cy='10' r='1' fill='white'/><circle cx='30' cy='25' r='1' fill='white'/><circle cx='20' cy='35' r='1' fill='white'/></svg>");
  }

  create() {
    this.player = this.physics.add.sprite(240, 580, "player");
    this.player.setCollideWorldBounds(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.fireKey = this.input.keyboard.addKey("SPACE");
    this.lastFired = 0;
    this.bullets = this.physics.add.group();
  }

  update(time) {
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-300);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(300);
    } else {
      this.player.setVelocityX(0);
    }
    if (this.cursors.up.isDown) {
      this.player.setVelocityY(-300);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(300);
    } else {
      this.player.setVelocityY(0);
    }
--fcc-editable-region--

--fcc-editable-region--
  }

  fireBullet() {
    const bullet = this.bullets.create(this.player.x, this.player.y - 30, "bullet");
    bullet.setVelocityY(-500);
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: MainScene
};

const game = new Phaser.Game(config);

```
