---
id: 66faa1000000000000000040
title: Step 64
challengeType: 0
dashedName: step-64
---

# --description--

Trigger the win when the score crosses the threshold.

Inside `hitEnemy`, after the `scoreText` update, add an `if` block that checks `this.score >= SCORE_TO_WIN`. Inside the block, call `this.scene.start("GameOverScene", { score: this.score, win: true })` so the game-over scene knows the player won.

# --hints--

You should check `this.score >= SCORE_TO_WIN` inside `hitEnemy`.

```js
assert.match(code, /this\.score\s*>=\s*SCORE_TO_WIN/);
```

`hitEnemy` should call `this.scene.start("GameOverScene", { score: this.score, win: true })`.

```js
assert.match(code, /this\.scene\.start\(\s*([\'"`])GameOverScene\1\s*,\s*\{[\s\S]*?score\s*:\s*this\.score[\s\S]*?win\s*:\s*true[\s\S]*?\}\s*\)/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js

const SCORE_TO_WIN = 100;
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
    this.load.image("player", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><polygon points='16,2 30,30 2,30' fill='%2300ff88'/></svg>");
    this.load.image("enemy", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='32' height='32'><rect x='2' y='2' width='28' height='28' fill='%23ff3366'/></svg>");
    this.load.image("bullet", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='6' height='16'><rect width='6' height='16' fill='%23ffff00'/></svg>");
    this.load.image("star", "data:image/svg+xml;utf8,<svg xmlns='http://www.w3.org/2000/svg' width='40' height='40'><circle cx='10' cy='10' r='1' fill='white'/><circle cx='30' cy='25' r='1' fill='white'/><circle cx='20' cy='35' r='1' fill='white'/></svg>");
  }

  create() {
    this.background = this.add.tileSprite(240, 320, 480, 640, "star");
    this.player = this.physics.add.sprite(240, 580, "player");
    this.player.setCollideWorldBounds(true);
    this.cursors = this.input.keyboard.createCursorKeys();
    this.fireKey = this.input.keyboard.addKey("SPACE");
    this.lastFired = 0;
    this.bullets = this.physics.add.group();
    this.enemies = this.physics.add.group();
    this.time.addEvent({
      delay: 1000,
      callback: this.spawnEnemy,
      callbackScope: this,
      loop: true
    });
    this.physics.add.overlap(this.bullets, this.enemies, this.hitEnemy, null, this);
    this.score = 0;
    this.scoreText = this.add.text(10, 10, "Score: 0", { fontSize: "20px", fill: "#fff" });
    this.lives = 3;
    this.livesText = this.add.text(380, 10, "Lives: 3", { fontSize: "20px", fill: "#fff" });
    this.physics.add.overlap(this.player, this.enemies, this.hitPlayer, null, this);
  }

  update(time) {
    this.background.tilePositionY -= 2;
    if (this.cursors.left.isDown) {
      this.player.setVelocityX(-300);
    } else if (this.cursors.right.isDown) {
      this.player.setVelocityX(300);
    } else {
      this.player.setVelocityX(0);
    }
    if (this.cursors.up.isDown) {
      this.player.setVelocityY(-300);
    } else if (this.cursors.down.isDown) {
      this.player.setVelocityY(300);
    } else {
      this.player.setVelocityY(0);
    }
    if (this.fireKey.isDown && time > this.lastFired) {
      this.fireBullet();
      this.lastFired = time + 250;
    }
    this.bullets.children.each((bullet) => {
      if (bullet.y < 0) bullet.destroy();
    });
    this.enemies.children.each((enemy) => {
      if (enemy.y > 640) enemy.destroy();
    });
  }

  fireBullet() {
    const bullet = this.bullets.create(this.player.x, this.player.y - 30, "bullet");
    bullet.setVelocityY(-500);
  }

  spawnEnemy() {
    const x = Phaser.Math.Between(30, 450);
    const enemy = this.enemies.create(x, -30, "enemy");
    enemy.setVelocityY(150);
  }

  hitEnemy(bullet, enemy) {
    bullet.destroy();
    enemy.destroy();
    this.score += 10;
    this.scoreText.setText("Score: " + this.score);
--fcc-editable-region--

--fcc-editable-region--
  }

  hitPlayer(player, enemy) {
    enemy.destroy();
    this.lives -= 1;
    this.livesText.setText("Lives: " + this.lives);
    this.cameras.main.flash(150, 255, 0, 0);
    if (this.lives <= 0) {
      this.scene.start("GameOverScene", { score: this.score });
    }
  }
}

class GameOverScene extends Phaser.Scene {
  constructor() {
    super("GameOverScene");
  }

  create(data) {
    this.add.text(240, 280, "GAME OVER", { fontSize: "48px", fill: "#ff0000" }).setOrigin(0.5);
    this.add.text(240, 340, "Score: " + data.score, { fontSize: "24px", fill: "#fff" }).setOrigin(0.5);
    this.add.text(240, 400, "Press R to restart", { fontSize: "18px", fill: "#fff" }).setOrigin(0.5);
    this.input.keyboard.on("keydown-R", () => {
      this.scene.start("MainScene");
    });
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
  physics: {
    default: "arcade",
    arcade: { gravity: { y: 0 } }
  },
  scene: [MainScene, GameOverScene]
};

const game = new Phaser.Game(config);

```
