---
id: 66faa100000000000000000e
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Phaser ships with two physics systems. You'll use **arcade physics**, the lightweight option, perfect for retro-style shooters.

Inside `config`, between `backgroundColor` and `scene`, add a `physics` property. Its value should be an object with `default` set to `"arcade"`.

# --hints--

Your `config` object should have a `physics` property.

```js
assert.match(code, /physics\s*:\s*\{/);
```

Your `physics` property should have `default` set to `"arcade"`.

```js
assert.match(code, /default\s*:\s*(['"`])arcade\1/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
  }

  update(time) {
  }
}

const config = {
  type: Phaser.AUTO,
  width: 480,
  height: 640,
  parent: "game-container",
  backgroundColor: "#000033",
--fcc-editable-region--

--fcc-editable-region--
  scene: MainScene
};

const game = new Phaser.Game(config);

```
