---
id: 66faa1000000000000000008
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

The `width` property controls how wide the game canvas is, in pixels.

Add a `width` property to `config` with the value `480`.

# --hints--

Your `config` object should have `width` set to `480`.

```js
assert.match(code, /width\s*:\s*480\b/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
  }

  update(time) {
  }
}

const config = {
  type: Phaser.AUTO,
--fcc-editable-region--

--fcc-editable-region--
};


```
