---
id: 66faa1000000000000000005
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Finally, add the `update` method. Phaser calls this on every frame (about 60 times per second). It receives the current time as its first argument, which you'll use later for shooting cooldowns.

Below `create`, add an `update` method that takes a `time` parameter.

# --hints--

Your `MainScene` class should have an `update` method that takes a `time` parameter.

```js
assert.match(code, /update\s*\(\s*time\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```html

<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Learn Game Development by Building a Space Shooter with Phaser</title>
    <link rel="stylesheet" href="./styles.css" />
    <script src="https://cdn.jsdelivr.net/npm/phaser@3.80.1/dist/phaser.min.js"></script>
  </head>
  <body>
    <div id="game-container"></div>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

body {
  background-color: #000;
  display: flex;
  justify-content: center;
  align-items: center;
  min-height: 100vh;
}

#game-container {
  box-shadow: 0 0 20px rgba(255, 255, 255, 0.3);
}

```

```js
class MainScene extends Phaser.Scene {
  constructor() {
    super("MainScene");
  }

  preload() {
  }

  create() {
  }

--fcc-editable-region--

--fcc-editable-region--
}


```
