---
id: 60f0286404aefb0562a4fdf9
title: Hatua ya 4
challengeType: 0
dashedName: step-4
---

# --description--

Unapaswa kuongeza kipengele cha `title` na `meta` kwenye `head`. Upe mradi wako jina la `Registration Form`, na ipe sifa ya `charset` yenye thamani ya `UTF-8` kwa kipengele cha `meta`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Kipengele cha `title` kinapaswa kuwa ndani ya kipengele cha `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Mradi wako unapaswa kuwa na kichwa cha `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Kumbuka, ukubwa wa herufi na tahajia ni jambo la muhimu katika kichwa.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
