---
id: 60f027c87bc98f050395c139
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Kisha, ongeza tagi za kufungua na kufunga za `head` na `body` ndani ya kipengele cha `html`.

# --hints--

Unapaswa kuwa na tagi ya kufungua ya `head`.

```js
assert(code.match(/<head\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Unapaswa kuwa na tagi ya kufungua ya `body`.

```js
assert(code.match(/<body\s*>/i));
```

Unapaswa kuwa na tagi ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Vipengele vya `head` na `body` vinapaswa kuwa ndugu.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

Kipengele cha `head` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Kipengele cha `body` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
