---
id: 60f027099a15b00485563dd2
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Chini ya `DOCTYPE`, ongeza kipengele cha `html` chenye sifa ya `lang` iliyowekwa kuwa `en`, ili upate nafasi kuanza kuweka msimbo.

# --hints--

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Kipengele chako cha `html` kinafaa kuwa na tagi ya ufunguzi yenye sifa ya `lang` iliyo na thamani ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na tagi ya kufunga.

```js
assert(code.match(/<\/html\s*>/));
```

Tagi zako `html` zinapaswa kuwa katika mpangilio sahihi.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Unapaswa kuwa na kipengele kimoja cha `html` tu.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
