---
id: 60f030d388cb74067cf291c3
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Chini ya kichwa, tumia maandishi yafuatayo ndani ya kipengele cha aya ili kuwahimiza watumiaji kujisajili:

```md
Please fill out this form with the required information
```

# --hints--

Unapaswa kuongeza kipengele cha `p` kwenye `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Unapaswa kuweka kipengele cha `p` chini ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Unapaswa kukipa kipengele cha `p` maandishi ya `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
