---
id: 60f1a9cbd23023082e149fee
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Hiyo inaonekana bora zaidi. Sasa, fanya mandharinyuma kuwa rahisi machoni, kwa kubadilisha `body` `background-color` hadi `#1b1b32`. Kisha, ili kuona maandishi, badilisha `color` hadi `#f5f6f7`.

# --hints--

Unapaswa kuongeza `background-color` ndani ya kichaguzi cha kipengele `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Unapaswa kuipa `background-color` thamani ya `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Unapaswa kuongeza `color` ndani ya kichaguzi cha kipengele `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Unapaswa kuipa `color` thamani ya `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
