---
id: 60f1922fcbd2410527b3bd89
title: Hatua ya 8
challengeType: 0
dashedName: step-8
---

# --description--
The `vh` unit stands for viewport height, and is equal to 1% of the `height` of the viewport. This makes it relative to the viewport height.

Ni wakati wa kuboresha mradi na CSS. Anza kwa kuipa `body` `width` ya `100%`, na `height` ya `100vh`.

# --hints--

Unapaswa kutumia kipengele kichaguzi cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Unapaswa kuipa `body` `width` ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Unapaswa kuipa `body` `height` ya `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
