---
id: 60f1a5e2d2c23707a4f9a660
title: Hatua ya 9
challengeType: 0
dashedName: step-9
---

# --description--

Sasa, ondoa upau wa kusogeza ulio mlalo, kwa kuweka `body` chaguo-msingi `margin` ulioongezwa na baadhi ya vivinjari hadi `0`.

# --hints--

Unapaswa kuongeza `margin` ndani ya kichaguzi cha kipengele `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

Unapaswa kuipa thamani `margin` ya `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
