---
id: 60f5d2776c854e069560fbe6
title: Hatua ya 14
challengeType: 0
dashedName: step-14
---

# --description--

`fieldset` ya kwanza itashikilia sehemu za jina, barua pepe na nenosiri. Anza kwa kuongeza vipengele vinne vya `label` kwenye `fieldset` ya kwanza.

# --hints--

Lazima kuwe na vipengele vitatu vya `fieldset` kwa jumla.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

Vipengele vya `fieldset` vyote vinapaswa kuwa watoto wa moja kwa moja wa kipengele cha `form`.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

Vipengele vinne vya `label` vyote vinapaswa kuwa watoto wa moja kwa moja wa kipengele cha kwanza cha `fieldset`. Hakikisha kuwa umefunga vipengele vya `label`.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Vipengele viwili vya mwisho vya `fieldset` vinapaswa kuwa tupu.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Hakikisha kuwa umefunga vipengele vya `label`.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```



# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
--fcc-editable-region--
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>

      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
