---
id: 60f84ec41116b209c280ba91
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

Ukiwa na `type="password"` unaweza kutumia sifa ya `pattern` kufafanua usemi wa kawaida ambao nenosiri lazima lilingane ili kuchukuliwa kuwa halali.

Ongeza sifa ya `pattern` kwa kipengele cha `input` cha nenosiri ili kuhitaji ingizo linalolingana: `[a-z0-5]{8,}`

Ya hapo juu ni usemi wa kawaida unaolingana na herufi ndogo nane au zaidi au tarakimu `0` hadi `5`. Kisha, ondoa sifa ya `minlength`, na uijaribu.

# --hints--

Unapaswa kukipa kipengele cha `input` cha nenosiri sifa ya`pattern`.

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.pattern);
```

Unapaswa kuipa sifa ya `pattern` thamani ya `[a-z0-5]{8,}`.

```js
assert.equal(document.querySelector('input[type="password"]')?.pattern, '[a-z0-5]{8,}');
```

Unapaswa kuondoa sifa ya `minlength` kutoka kwa kipengele cha `input` cha nenosiri.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, -1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" minlength="8" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
