---
id: 60f80e0081e0f2052ae5b505
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Kubainisha sifa ya `type` ya kipengele cha fomu ni muhimu kwa kivinjari kujua ni aina gani ya data inapaswa kutarajia. Ikiwa `type` haijabainishwa, kivinjari kitakuwa chaguomsingi ya `text`.

Vipe vipengele viwili vya kwanza `input` ya sifa `type` ya `text`, cha tatu sifa ya `type` ya `email`, na ya nne sifa ya `type` ya `password`.

Aina ya `email` inaruhusu barua pepe zilizo na `@` pekee na `.` kwenye kikoa. Aina ya `password` huficha ingizo, na kuonya ikiwa tovuti haitumii HTTPS.

# --hints--

Unapaswa kukipa kipengele cha kwanza cha `input` sifa ya `type` ya `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Unapaswa kukipa kipengele cha pili cha `input` sifa ya `type` ya `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Unapaswa kukipa kipengele cha tatu cha `input` sifa ya `type` ya `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Unapaswa kukipa kipengele cha nne cha `input` sifa ya `type` ya `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
