---
id: 60fabf0dd4959805dbae09e6
title: Hatua ya 33
challengeType: 0
dashedName: step-33
---

# --description--

Ongeza `label` nyingine baada ya ya kwanza, yenye maandishi `Input your age (years):`. Kisha, weka `input` na `type` ya `number`.

Kisha, ongeza sifa ya `min` kwenye `input` yenye thamani ya `13` kwa sababu watumiaji walio na umri wa chini ya miaka 13 hawafai kujisajili. Pia, watumiaji pengine hawatakuwa na zaidi ya umri wa miaka 120; ongeza sifa ya `max` yenye thamani ya `120`.

Sasa, mtu akijaribu kuwasilisha fomu iliyo na thamani nje ya masafa, onyo litaonekana, na fomu haitawasilishwa. Jaribu.

# --hints--

Unapaswa kuongeza `label` kwa `fieldset` ya tatu, baada ya `label` iliyopo.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

Unapaswa kuipa `label` maandishi ya `Input your age (years):`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

Unapaswa kuipa `label` `input` yenye `type` ya `number`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

Unapaswa kuipa `input` sifa ya `min` yenye thamani ya `13`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

Unapaswa kuipa `input` sifa ya `max` yenye thamani ya `120`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>

      </fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
