---
id: 62cc5b1779e4d313466f73c5
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Weka kipengee cha kujifunga cha `link` ndani ya kipengele cha `head`. Ipe sifa ya `rel` yenye thamani ya `stylesheet` na sifa ya `href` yenye thamani ya `styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.exists(document.querySelector('link'));
```

Hupaswi kubadilisha tagi zako za `head` zilizopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Unapaswa kuwa na kipengele kimoja cha kujifunga cha `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `rel` yenye thamani `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
