---
id: 62ff8b9dab5ac88e4d3d43a3
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Kwa kufuata mbinu bora za ufikivu, unganisha vipengele vya `input` na vipengele vya `label` kwa kutumia sifa ya `for`.

Tumia `first-name`, `last-name`, `email`, na `new-password` kama thamani husika za sifa ya `id`.

# --hints--

Kipengee chako cha kwanza cha `input` kinapaswa kuwa na `id` ya `first-name`.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Kipengee chako cha pili cha `input` kinapaswa kuwa na `id` ya `last-name`.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Kipengee chako cha tatu cha `input` kinapaswa kuwa na `id` ya `email`.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Kipengee chako cha nne cha `input` kinapaswa kuwa na `id` ya `new-password`.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Kipengele cha kwanza cha `label` kinafaa kuwa na sifa `for` yenye thamani ya `first-name`.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Kipengele cha pili cha `label` kinafaa kuwa na sifa `for` yenye thamani ya `last-name`.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Kipengele cha tatu cha `label` kinafaa kuwa na sifa `for` yenye thamani ya `email`.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Kipengele cha nne cha `label` kinafaa kuwa na sifa `for` yenye thamani ya `new-password`.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
