---
id: 63541ef4f96cd82e8e6c788a
title: Hatua ya 12
challengeType: 0
dashedName: step-12
---

# --description--

Sifa ya `method` inabainisha jinsi ya kutuma data ya fomu kwa URL iliyobainishwa katika sifa ya `action`. Data ya fomu inaweza kutumwa kupitia ombi la `GET` kama vigezo vya URL (na `method="get"`) au kupitia ombi la `POST` kama data. katika ombi (na `method="post"`).

Weka sifa ya `method` itume data ya fomu yako kupitia ombi la `POST`.

# --hints--

Hupaswi kuongeza kipengele kipya cha `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Kipengele chako cha `form` kinafaa kuwa na sifa ya `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Sifa ya `method` yako inapaswa kuwekwa kuwa `post`.

```js
assert.match(document.querySelector('form')?.getAttribute('method'), /^post$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
