---
id: 63b607af6fcdb119aae9b16a
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Crie um elemento `fieldset` com o `id` definido como `breakfast`.

Dentro desse elemento, crie uma `legend` com o texto `Breakfast` e uma `div` vazia com a `class` definida como `input-container`.

# --hints--

Você deve criar um elemento `fieldset` no `form`.

```js
assert.exists(document.querySelector('form fieldset'));
```

O elemento `fieldset` deve vir depois do elemento `input`.

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

O elemento `fieldset` deve ter o atributo `id` definido como `breakfast`.

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

Você deve criar um elemento `legend` no `fieldset`.

```js
assert.exists(document.querySelector('form fieldset legend'));
```

O elemento `legend` deve ter o texto `Breakfast`.

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

Você deve criar um elemento `div` no `fieldset`.

```js
assert.exists(document.querySelector('form fieldset div'));
```

O elemento `div` deve vir depois do elemento `legend`.

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

O elemento `div` deve ter uma `class` definida como `input-container`.

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
