---
id: 63b60af1a0b9f7238a9dd294
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

O menu de seleção precisa de opções para cada um dos elementos `fieldset` de alimentos e exercícios que você criou nos passos anteriores. Use o elemento `option` para criar uma opção para cada `fieldset`. O atributo `value` de cada opção deve ser o `id` do `fieldset`. O texto de cada opção deve ser o texto de `legend`.

Defina a opção `Breakfast` como a opção `selected`.

# --hints--

Você deve criar cinco elementos `option` dentro do elemento `select`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.length, 5);
```

O primeiro `option` deve ter o texto `Breakfast`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.textContent, 'Breakfast');
```

O primeiro `option` deve ter o atributo `value` definido como `breakfast`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.value, 'breakfast');
```

O segundo `option` deve ter o texto `Lunch`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.textContent, 'Lunch');
```

O segundo `option` deve ter o atributo `value` definido como `lunch`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.value, 'lunch');
```

O terceiro `option` deve ter o texto `Dinner`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.textContent, 'Dinner');
```

O terceiro `option` deve ter o atributo `value` definido como `dinner`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.value, 'dinner');
```

O quarto `option` deve ter o texto `Snacks`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.textContent, 'Snacks');
```

O quarto `option` deve ter o atributo `value` definido como `snacks`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.value, 'snacks');
```

O quinto `option` deve ter o texto `Exercise`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.textContent, 'Exercise');
```

O quinto `option` deve ter o atributo `value` definido como `exercise`.

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.value, 'exercise');
```

O primeiro `option` deve ser selecionado por padrão.

```js
assert.isTrue(document.querySelectorAll('.controls select option')?.[0]?.getAttributeNames()?.includes('selected'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">

              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
body {
  font-family: "Lato", Arial;
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button:hover {
  cursor: pointer;
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid black;
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: #006627;
}

.deficit {
  color: #B30000;
}
```

```js

```
