---
id: 682b8dcac698311f81552da4
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

In the next few steps, you will add the form elements responsible for collecting the user's preferred contact method. 

Start by adding another `fieldset` element with a `legend` element nested inside. Your `legend` element should have the text `Preferred Contact Method`.

# --hints--

You should have a total of three `fieldset` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 3);
```

You should have a `legend` element inside of your third `fieldset` element.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

Your `legend` element should have the text `Preferred Contact Method`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Preferred Contact Method");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

```css
```
