---
id: 682b7958b95e2c15e9509d7f
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

The next step is to add the associated `input` element for the student's information.

Start by adding an `input` element with a `type` attribute set to `"text"`. Then add a `name` and `id` attribute both set to `student-name`. 

Next, add a `placeholder` attribute set to `Ex. Jane Doe`. And finally, add a `required` attribute. 

# --hints--

You should have an `input` element.

```js
assert.exists(document.querySelector("fieldset input"));
```

Your `input` element should have a `type` attribute set to `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Your `input` element should have a `name` attribute set to `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Your `input` element should have an `id` attribute set to `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Your `input` element should have a `placeholder` attribute set to `Ex. Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "Ex. Jane Doe");
```

Your `input` element should have a `required` attribute.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          --fcc-editable-region--
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```

```css
```
