---
id: 682b70b753f55e11e55b3b3b
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

The first section of the form will focus on the student's information like their name and grade.

Inside the `form` element, add a `fieldset` element. Inside that `fieldset` element, add a `legend` element with the text `Student Information`. 

Below the `legend` element, add a `label` element with the text `Full Name: ` and a `for` attribute with the value of `student-name`.

# --hints--

You should have a `fieldset` element.

```js
assert.exists(document.querySelector("fieldset"));
```

Your `fieldset` element should be inside of the `form` element.

```js
assert.exists(document.querySelector("form fieldset"));
```

You should have a `legend` element.

```js
assert.exists(document.querySelector("legend"));
```

Your `legend` element should be inside of the `fieldset` element.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Your `legend` element should have the text `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

You should have a `label` element.

```js
assert.exists(document.querySelector("label"));
```

Your `label` element should have the text `Full Name: `. Don't forget the space after the colon.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Your `label` element should have a `for` attribute.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

Your `label` element's `for` attribute should have the value `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Your `label` element should be below the `legend` element.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      --fcc-editable-region--
      <form>
        
      </form>
      --fcc-editable-region--
    </main>
  </body>
</html>
```
