---
id: 682b94c8db82ed229cac60df
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

The next step is to add the `label` and `input` elements for the email contact method.

Start by adding a label with a class of `"contact-method"` and a `for` attribute set to `"email"`. The label should text should be `Email: `.

Below your `label` element, add an radio button with `id` and `value` attributes set to `"email"`. The `name` attribute should be set to `"contact-method"` and the class should be set to `"contact-method-radio-btn"`. 

Lastly, make sure this radio button is checked by default.

# --hints--

You should have a total of four `label` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Your fourth `label` element should have the text `Email: `. Don't forget the space after the colon.

```js
assert.match(code, /Email:\s+<\/label>/);
```

Your fourth `label` element should have a class of `"contact-method"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label:last-of-type")?.getAttribute("class"), "contact-method");
```

Your fourth `label` element should have a `for` attribute.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Your fourth `label` element's `for` attribute should have the value `"email"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "email");
```

You should have four `input` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("input"), 4);
```

Your fourth `input` element should be below your fourth `label` element.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Your fourth `input` element should have a `type` attribute set to `"radio"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "radio");
```

Your fourth `input` element should have a `name` attribute set to `"contact-method"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "contact-method");
```

Your fourth `input` element should have an `id` attribute set to `"email"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "email");
```

Your fourth `input` element should have an `value` attribute set to `"email"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("value"), "email");
```

Your fourth `input` element should have a class set to `"contact-method-radio-btn"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("class"), "contact-method-radio-btn");
```

Your fourth `input` element should have a `checked` attribute.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>Preferred Contact Method</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

```css
```
