---
id: 682bf6feceec294193a637f8
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

The last section of the form will be for users to add additional notes they would like to submit.

Start by adding a `fieldset` element with a `legend` element nested inside. Your `legend` element should have the text `Additional Notes`.

# --hints--

You should have a total of four `fieldset` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 4);
```

You should have a `legend` element inside of your fourth `fieldset` element.

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

Your `legend` element should have the text `Additional Notes`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Additional Notes");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        <fieldset>
          <legend>Preferred Contact Method</legend>
          <label class="contact-method" for="email">Email: </label>
          <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

          <label class="contact-method" for="phone">Phone: </label>
          <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

```css
```
