---
id: 682bf4e05394cb3e8f2973cd
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

The next step is to add the `label` and `input` elements for the phone contact method.

Start by adding a label with a class of `"contact-method"` and a `for` attribute set to `"phone"`. The label should text should be `Phone: `.

Below your `label` element, add an radio button with `id` and `value` attributes set to `"phone"`. The `name` attribute should be set to `"contact-method"` and the class should be set to `"contact-method-radio-btn"`. 

# --hints--

You should have a total of five `label` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("label"), 5);
```

Your fifth `label` element should have the text `Phone: `. Don't forget the space after the colon.

```js
assert.match(code, /Phone:\s+<\/label>/);
```

Your fifth `label` element should have a class of `"contact-method"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label:last-of-type")?.getAttribute("class"), "contact-method");
```

Your fifth `label` element should have a `for` attribute.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Your fifth `label` element's `for` attribute should have the value `"phone"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label:last-of-type")?.getAttribute("for"), "phone");
```

You should have five `input` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Your fifth `input` element should be below your fifth `label` element.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Your fifth `input` element should have a `type` attribute set to `"radio"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "radio");
```

Your fifth `input` element should have a `name` attribute set to `"contact-method"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "contact-method");
```

Your fifth `input` element should have an `id` attribute set to `"phone"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "phone");
```

Your fifth `input` element should have an `value` attribute set to `"phone"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("value"), "phone");
```

Your fifth `input` element should have a class set to `"contact-method-radio-btn"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("class"), "contact-method-radio-btn");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        <fieldset>
          <legend>Preferred Contact Method</legend>
          <label class="contact-method" for="email">Email: </label>
          <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
