---
id: 682bfa6c8811114340434be0
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

The next step is to add the form elements responsible for collecting the additional notes from users. 

Inside your last `fieldset` element, add a `label` element with a `for` attribute set to `"notes"`.  Your `label` text should be `Any specific concerns or topics you'd like to discuss?`.

Then below your `label` element, add a `textarea` element with `id` and `name` attributes set to `"notes"`. Your `textarea` element should also have a `rows` attribute set to `"4"` and a `cols` attribute set to `"50"`.

# --hints--

You should have a total of six `label` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("label"), 6);
```

Your sixth `label` element should have a `for` attribute set to `"notes"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label:last-of-type")?.getAttribute("for"), "notes");
```

Your sixth `label` element should have the text `Any specific concerns or topics you'd like to discuss?`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label:last-of-type")?.innerText.trim(), "Any specific concerns or topics you'd like to discuss?");
```

You should have a `textarea` element.

```js
assert.exists(document.querySelector("textarea"));
```

Your `textarea` element should be below your sixth `label` element.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const textareaEl = document.querySelector("fieldset:last-of-type textarea");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(textareaEl));
```

Your `textarea` element should have a `name` attribute set to `"notes"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(inputEl?.getAttribute("name"), "notes");
```

Your `textarea` element should have a `id` attribute set to `"notes"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(inputEl?.getAttribute("id"), "notes");
```

Your `textarea` element should have a `rows` attribute set to `"4"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(inputEl?.getAttribute("rows"), "4");
```

Your `textarea` element should have a `cols` attribute set to `"50"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(inputEl?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        <fieldset>
          <legend>Preferred Contact Method</legend>
          <label class="contact-method" for="email">Email: </label>
          <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

          <label class="contact-method" for="phone">Phone: </label>
          <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>Additional Notes</legend>

        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```

```css
```
