---
id: 682b5e4426902e82d5395ae6
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to add custom styles to radio buttons by building a parent teacher conference form. The HTML boilerplate has been provided for you.

Start by adding a `main` element with a class called `container`.

# --hints--

You should have a `main` element.

```js
assert.exists(document.querySelector("main"));
```

Your `main` element should have a class called `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
