---
id: 682b842d83a5b81ad0607435
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

Now it is time to add the form elements to collect the students grade information.

Start by adding another `label` element with the text `Student Grade: ` and `for` attribute set to `"grade"`.

Then below your `label` element, add an `input` element with the `type` attribute set to `"number"`. The `name` and `id` attributes should be set to `"grade"`. The `placeholder` attribute should be set to `"Ex. 4"`. Lastly, your number `input` should be required. 

# --hints--

You should have a total of two `label` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Your second `label` element should have the text `Student Grade: `. Don't forget the space after the colon.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Your second `label` element should have a `for` attribute.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

Your second `label` element's `for` attribute should have the value `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

You should have two `input` elements on the page.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Your second `input` element should be below your second `label` element.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Your second `input` element should have a `type` attribute set to `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Your second `input` element should have a `name` attribute set to `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Your second `input` element should have an `id` attribute set to `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Your second `input` element should have a `placeholder` attribute set to `"Ex. 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "Ex. 4");
```

Your second `input` element should have a `required` attribute.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />

          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
