---
id: 682b6ed4fc6724107bed6cce
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Now it is time to add the `form` and `input` elements which will represent the parent and student information.

Start by adding a `form` element below the `p` element. 

# --hints--

You should have a `form` element.

```js
assert.exists(document.querySelector(".container form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      --fcc-editable-region--

      --fcc-editable-region--
    </main>
  </body>
</html>
```
