---
id: 682cee64a819857c94c08950
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

It would be nice if the main container had a different background color than the body. Up until this point, you have been used to working with named colors like `red` or `blue` but there are different color models you can work with like hex codes.

Hexadecimal color values, or hex codes, are alphanumeric(a-z0-9) strings used to represent hundreds of different colors. They provide more flexible color options in your designs and you will learn more about how these work in a future module.

Here is an example of a couple of hex values:

```css
/* Red hex code */
/* #FF0000 */

/* Blue hex code */
/* #0000FF */
```

Inside your `.container` selector, set the background color to `#ffffff1a`.

# --hints--

Your `.container` selector should have a `background-color` set to `#ffffff1a`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".container")?.backgroundColor, "rgba(255, 255, 255, 0.1)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        <fieldset>
          <legend>Preferred Contact Method</legend>
          <label class="contact-method" for="email">Email: </label>
          <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

          <label class="contact-method" for="phone">Phone: </label>
          <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
        </fieldset>

        <fieldset>
          <legend>Additional Notes</legend>
          <label for="notes">
            Any specific concerns or topics you'd like to discuss?
          </label>
          <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
        </fieldset>

        <button class="submit-btn" type="submit">Submit Form</button>
      </form>
    </main>
  </body>
</html>
```

```css
body {
  background-color: midnightblue;
  color: whitesmoke;
}

--fcc-editable-region--
.container {
  width: 80%;
  max-width: 600px;
  margin: 20px auto;

}
--fcc-editable-region--
```
