---
id: 682c0629545c0e47c65b136c
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

The last part of the form is to add the submit button.

Below your last `fieldset` element, add a `button` element with a `type` attribute set to `"submit"` and a class `"submit-btn"`. The button text should be `Submit Form`.

# --hints--

You should have a `button` element.

```js
assert.exists(document.querySelector("button"));
```

Your `button` element should have a `type` attribute set to `"submit"`.

```js
assert.strictEqual(document.querySelector("button")?.getAttribute("type"), "submit");
```

Your `button` element should have a class of `"submit-btn"`.

```js
assert.exists(document.querySelector("button.submit-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        <fieldset>
          <legend>Preferred Contact Method</legend>
          <label class="contact-method" for="email">Email: </label>
          <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

          <label class="contact-method" for="phone">Phone: </label>
          <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
        </fieldset>

        <fieldset>
          <legend>Additional Notes</legend>
          <label for="notes">
            Any specific concerns or topics you'd like to discuss?
          </label>
          <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
        </fieldset>

        --fcc-editable-region--

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
