---
id: 682c0c9a871ffc518bbdfd4a
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

Now it is time to work on adding styles to your page.

Start by setting the `background-color` for the `body` element to `MidnightBlue`. Then, set the `color` property in the `body` element to `whitesmoke`. 

# --hints--

You should have a `body` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

The `body` element should have a `background-color` of `MidnightBlue`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("background-color").toUpperCase(), ["MidnightBlue", "#191970", "RGB(25, 25, 112)", "HSL(240, 100%, 27%)"]);
```

The `body` element should have a `color` of `whitesmoke`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("color").toUpperCase(), ["whitesmoke", "#F5F5F5", "RGB(245, 245, 245)", "HSL(0, 0%, 96%)"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Parent Teacher Conference Form</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main class="container">
      <h1 class="title center">Parent Teacher Conference Form</h1>
      <p class="description center">
        Please fill out the form below to help schedule your parent-teacher
        conference.
      </p>

      <form>
        <fieldset>
          <legend>Student Information</legend>
          <label for="student-name">Full Name: </label>
          <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="Ex. Jane Doe"
          />


          <label for="grade">Student Grade: </label>
          <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="Ex. 4"
          />
        </fieldset>

        <fieldset>
          <legend>Parent/Guardian Information</legend>
          <label for="parent-name">Parent/Guardian Information: </label>
          <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="Ex. Nancy Doe"
          />
        </fieldset>

        <fieldset>
          <legend>Preferred Contact Method</legend>
          <label class="contact-method" for="email">Email: </label>
          <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

          <label class="contact-method" for="phone">Phone: </label>
          <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
        </fieldset>

        <fieldset>
          <legend>Additional Notes</legend>
          <label for="notes">
            Any specific concerns or topics you'd like to discuss?
          </label>
          <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
        </fieldset>

        <button class="submit-btn" type="submit">Submit Form</button>
      </form>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
