---
id: 615f41c979787462e76dab90
title: Schritt 20
challengeType: 0
dashedName: step-20
---

# --description--

Deine neue Klasse besitzt noch kein Styling. Erstelle einen `.bold`-Selektor und gib ihm eine `font-weight`-Eigenschaft mit dem Wert `800`, um den Text fettgedruckt darzustellen.

Fahre fort und entferne auch die `font-weight`-Eigenschaft aus deinem `h1`-Selektor.

# --hints--

Du solltest einen `.bold`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold'));
```

Dein `.bold`-Selektor sollte eine `font-weight`-Eigeschaft mit dem Wert `800` enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('.bold')?.fontWeight === '800');
```

Dein `h1`-Selektor sollte keine `font-weight`-Eigenschaft enthalten.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === "");
```

Du solltest deinen `h1`-Selektor nicht entfernen.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--
h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px;
}

p {
  margin: 0;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

--fcc-editable-region--
```
