---
id: 615f38279e5c3d53692ea441
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Umschließe deine `h1` und `p`-Elemente mit einem `div`-Element. Gib diesem `div` ein `class`-Attribut, das auf `label` gesetzt ist.

# --hints--

Du solltest ein neues `div`-Element erstellen.

```js
assert(document.querySelector('div'));
```

Dein neues `div`-Element sollte das `class`-Attribut auf `label` gesetzt haben.

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

Deine `h1` und `p`-Elemente sollten sich innerhalb deines `.label`-Elements befinden.

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

--fcc-editable-region--
<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
--fcc-editable-region--
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
