---
id: 615f4bfb9de4a16703b56eb6
title: Schritt 23
challengeType: 0
dashedName: step-23
---

# --description--

Jetzt können wir die horizontalen Abstände mit `flex` hinzufügen. Füge in deinem `p`-Selektor eine `display`-Eigenschaft, die auf `flex` eingestellt ist, und eine `justify-content`-Eigenschaft hinzu, die auf `space-between` eingestellt ist.

# --hints--

Dein `p`-Selektor sollte eine `display`-Eigenschaft auf `flex` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('p')?.display === 'flex');
```

Dein `p`-Selektor sollte eine `justify-content`-Eigenschaft auf `space-between` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('p')?.justifyContent === 'space-between');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1 class="bold">Nutrition Facts</h1>
    <div class="divider"></div>
    <p>8 servings per container</p>
    <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

--fcc-editable-region--
p {
  margin: 0;
}
--fcc-editable-region--

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}
```
