---
id: 615f357957e370510f21ea16
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Erstelle einen `body`-Selektor und weise ihm eine `font-family` von `Open Sans` mit einem Fallback von `sans-serif` zu.

Remember that fonts with spaces in the name must be wrapped in quotes for CSS.

# --hints--

Du solltest einen `body`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Dein `body`-Selektor sollte eine `font-family`-Eigenschaft auf `"Open Sans", sans-serif` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
