---
id: 615f378014c2da526a109c81
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

The font is a bit small. Erstelle einen `html`-Selektor und setze die Schriftart auf eine Größe von `16px`.

# --hints--

Du solltest einen `html`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('html'));
```

Dein `html`-Selektor sollte eine `font-size`-Eigenschaft auf `16px` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('html')?.fontSize === '16px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

body {
  font-family: 'Open Sans', sans-serif;
}
```
