---
id: 615f34948891834dd77655a6
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Füge ein zweites `p`-Element mit dem Text `Serving size 2/3 cup (55g)` hinzu.

# --hints--

Du solltest ein zweites `p`-Element haben.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Dein zweites `p`-Element sollte sich innerhalb deines `body`-Elements befinden.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Dein zweites `p`-Element sollte nach deinem vorhandenen `p`-Element stehen.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Dein zweites `p`-Element sollte den Text `Serving size 2/3 cup (55g)` enthalten.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>

</body>
</html>
--fcc-editable-region--
```

```css

```
