---
id: 615f2abbe7d18d49a1e0e1c8
title: Schritt 1
challengeType: 0
dashedName: step-1
---

# --description--

Wir haben dir eine grundlegende HTML-Boilerplate zur Verfügung gestellt.

Erstelle innerhalb deines `body`-Elements ein `h1`-Element und weise ihm den Text `Nutrition Facts` zu.

# --hints--

Du solltest ein neues `h1`-Element hinzufügen.

```js
assert.exists(document.querySelector('h1'));
```

Dein `h1`-Element sollte sich innerhalb deines `body`-Elements befinden.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Dein `h1`-Element sollte den Text `Nutrition Facts` enthalten.

```js
assert(document.querySelector('h1')?.innerText === 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>

</body>
</html>
--fcc-editable-region--
```

```css

```
