---
id: 615f34ecc1091b4fd5a8a484
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Füge innerhalb des `head`-Elements ein `link`-Element hinzu, bei dem das `rel`-Attribut auf `stylesheet` und das `href`-Attribut auf `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` gesetzt sind.

Auf diese Weise wird die `Open Sans`-Schriftart-Familie mit den Schriftstärke-Einstellungen `400`, `700` und `800` importiert.

Füge außerdem ein `link`-Element hinzu, um deine `styles.css`-Datei zu verknüpfen.

# --hints--

Dein Code sollte zwei selbstabschließende `link`-Elemente enthalten.

```js
assert(document.querySelectorAll('link').length === 2);
```

Beide `link`-Elemente benötigen ein `rel`-Attribut mit dem Wert `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines der `link`-Elemente sollte ein `href`-Attribut mit `./styles.css` als Wert enthalten.

```js
assert(code.match(/<link[\s\S]*?href\s*=\s*('|"|`)(\.\/)?styles\.css\1/g)?.length === 1);
```

Eines der `link`-Elemente sollte ein `href`-Attribut mit `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` als Wert enthalten.

```js
const links = [...document.querySelectorAll('link')]
assert(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
