---
id: 615f7de4487b64919bb4aa5e
title: Schritt 53
challengeType: 0
dashedName: step-53
---

# --description--

Erstelle nach deinem letzten `.divider` ein neues `p`-Element mit dem Text `Cholesterol 0mg 0%`. Umschließe den Text `Cholesterol` mit einem `span`-Element und weise diesem `span`-Element die `class` `bold` zu. Umschließe den Text `0%` mit einem weiteren `span`-Element, wobei die `class` auf `bold` gesetzt ist. Bette schließlich das `Cholesterol`-`span`-Element und den Text `0mg` in einem weiteren `span`-Element zur Ausrichtung ein.

# --hints--

Du solltest ein neues `p`-Element am Ende deines `.daily-value`-Elements erstellen.

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.localName === 'p');
```

Dein neues `p`-Element sollte den Text `Cholesterol 0mg 0%` haben.

```js
const text = document.querySelector('.daily-value.small-text')?.lastElementChild?.textContent?.trim()?.replace(/\n/g, ' ')?.replace(/\s\s+/g, ' ');
assert(text === 'Cholesterol 0mg 0%');
```

Dein neues `p`-Element sollte genau drei `span`-Elemente besitzen.

```js
assert(document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelectorAll('span')?.length === 3);
```

Der Text `Cholesterol` sollte in einem `span`-Element eingebettet sein.

```js
const spans = [...document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelectorAll('span')];
const cholesterolSpan = spans.filter(span => span?.innerHTML?.match(/^[\s\n]*Cholesterol[\s\n]*$/));
assert(cholesterolSpan.length === 1);
```

Dein `span`-Element mit `Cholesterol` sollte das `class`-Attribut auf `bold` eingestellt haben.

```js
const spans = [...document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelectorAll('span')];
const cholesterolSpan = spans.filter(span => span?.innerHTML?.match(/^[\s\n]*Cholesterol[\s\n]*$/));
assert(cholesterolSpan[0]?.classList?.contains('bold'));
```

Der Text `0%` sollte in einem `span`-Element eingebettet sein.

```js
const spans = [...document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelectorAll('span')];
const zeroPercentSpan = spans.filter(span => span?.innerHTML?.match(/^[\s\n]*0%[\s\n]*$/));
assert(zeroPercentSpan.length === 1);
```

Bei deinem `span`-Element mit `0%` sollte das `class`-Attribut auf `bold` eingestellt sein.

```js
const spans = [...document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelectorAll('span')];
const zeroPercentSpan = spans.filter(span => span?.innerHTML?.match(/^[\s\n]*0%[\s\n]*$/));
assert(zeroPercentSpan[0]?.classList?.contains('bold'));
```

Dein `span`-Element mit `Cholesterol` und dein `0mg`-Text sollten in einem `span`-Element eingebettet sein.

```js
const spans = [...document.querySelector('.daily-value.small-text')?.lastElementChild?.querySelectorAll('span')];
const cholesterolZeroSpan = spans.filter(span => span?.innerText === 'Cholesterol 0mg');
assert(cholesterolZeroSpan.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <header>
      <h1 class="bold">Nutrition Facts</h1>
      <div class="divider"></div>
      <p>8 servings per container</p>
      <p class="bold">Serving size <span>2/3 cup (55g)</span></p>
    </header>
    <div class="divider large"></div>
    <div class="calories-info">
      <div class="left-container">
        <h2 class="bold small-text">Amount per serving</h2>
        <p>Calories</p>
      </div>
      <span>230</span>
    </div>
    <div class="divider medium"></div>
    --fcc-editable-region--
    <div class="daily-value small-text">
      <p class="bold right no-divider">% Daily Value *</p>
      <div class="divider"></div>
      <p><span><span class="bold">Total Fat</span> 8g</span> <span class="bold">10%</span></p>
      <p class="indent no-divider">Saturated Fat 1g <span class="bold">5%</span></p>
      <div class="divider"></div>
      <p class="indent no-divider"><span><i>Trans</i> Fat 0g</span></p>
      <div class="divider"></div>
    </div>    
    --fcc-editable-region--
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

header h1 {
  text-align: center;
  margin: -4px 0;
  letter-spacing: 0.15px
}

p {
  margin: 0;
  display: flex;
  justify-content: space-between;
}

.divider {
  border-bottom: 1px solid #888989;
  margin: 2px 0;
}

.bold {
  font-weight: 800;
}

.large {
  height: 10px;
}

.large, .medium {
  background-color: black;
  border: 0;
}

.medium {
  height: 5px;
}

.small-text {
  font-size: 0.85rem;
}


.calories-info {
  display: flex;
  justify-content: space-between;
  align-items: flex-end;
}

.calories-info h2 {
  margin: 0;
}

.left-container p {
  margin: -5px -2px;
  font-size: 2em;
  font-weight: 700;
}

.calories-info span {
  margin: -7px -2px;
  font-size: 2.4em;
  font-weight: 700;
}

.right {
  justify-content: flex-end;
}

.indent {
  margin-left: 1em;
}

.daily-value p:not(.no-divider) {
  border-bottom: 1px solid #888989;
}
```
