---
id: 683ec95f2f8781355556ff81
title: Step 9
challengeType: 20
dashedName: step-9
---

# --description--

You need to test out things so you will understand the flow of the algorithm at this initial stage.

To do that, you need to first break out of the loop and return a custom message in the function, at least for now.

That's because the current implementation will only allow one iteration, so if the condition is not met, there will be an infinite loop.

Just after the `if`statement, use the `break` keyword in the block of the `while` loop, and in the block of the function, return the message `This shows because the value is farther than the middle value to the right or left`.

You will understand the purpose of this message in a few steps.

# --hints--

You should use the `break` keyword in the body of the `while` loop. Don't forget to look out for the right indentation level.

```js
({
    test: () => assert(runPython(
      `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().has_stmt("break")`)
    )
})
```

You should return the message `This shows because the value is farther than the middle value to the right or left` in the `binary_search` function block. Don't forget to look out for the right indentation level.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").has_return("'This shows because the value is farther than the middle value to the right or left'")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1

    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--

--fcc-editable-region--
```
