---
id: 683ec95f2f8781355556ff83
title: Step 11
challengeType: 20
dashedName: step-11
---

# --description--

You can see that the first call returns `[3]`. That's because `3` is the middle value in `[1, 2, 3, 4, 5]`.

On the the hand, the second call returns `'This shows because the value is farther than the middle value to the right or left'`. This happens because, after checking the middle value, the loop currently lacks the logic to narrow the search range by updating the `low` or `high` variables. It only acts as a single midpoint check and then finishes without further searching.

To allow the binary search to continue narrowing its search, add an `elif` block that checks if `value` is greater than `value_at_middle`. Add the `pass` keyword inside the `elif` for now.

# --hints--

You should add an `elif` with the condition `value > value_at_middle`. Don't forget to add the `pass` keyword into the `elif` block.

```js
({
    test: () => assert(runPython(`
        _Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].is_equivalent("if value == value_at_middle:\\n    return path_to_target\\nelif value > value_at_middle:\\n    pass")
    `))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--

--fcc-editable-region--          
        break
    return "This shows because the value is farther than the middle value to the right or left"

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4))
```
