---
id: 6564642ce403fe29cc23fee7
title: Step 1
challengeType: 20
dashedName: step-1
---

# --description--

A Python module is a file containing code designed to perform specific tasks. The Python standard library contains many modules that you can import and use in your code. You can achieve this by using the `import` statement followed by the name of the module.

Start this project by importing the `string` module.

# --hints--

You should import the `string` module.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
