---
id: 616968c2c94c8071b349c146
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Now you're ready to start adding content to the page.

Within the `body`, nest an `h1` element with the text `CSS Color Markers`.

# --hints--

Your code should have an `h1` element.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

You should have an opening `h1` tag.

```js
assert(code.match(/<h1\s*>/i));
```

You should have a closing `h1` tag.

```js
assert(code.match(/<\/h1\s*>/i));
```

Your `h1` element should be within the `body` element.

```js
assert(document.querySelector('h1')?.parentElement?.localName === 'body');
```

Your `h1` element should have the text `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Remember, the casing and spelling matter for the `h1` text.

```js
const h1 = document.querySelector('h1');
assert.equal(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
