---
id: 619b74fa777a2b2473c94f82
title: Step 45
challengeType: 0
dashedName: step-45
---

# --description--

Update the CSS class selector `.two` so it targets the new `green` class. And update the `.three` class selector so it targets the new `blue` class.

# --hints--

Your code should no longer have a `.two` class selector.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

You should use a class selector to target the class `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

Your `.green` CSS rule should have a `background-color` property set to `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Your code should no longer have a `.three` class selector.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

You should use a class selector to target the class `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Your `.blue` CSS rule should have a `background-color` property set to `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
