---
id: 61695c4aad56f95497c19583
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Пам’ятайте, що елемент `title` надає пошуковим системам додаткову інформацію про сторінку. It also displays the content of that `title` element in two more ways:

* in the title bar when the page is open
* in the browser tab for the page when you hover on it. Even if that tab is not active, once you hover on the tab, the `title` text is displayed.

В елемент `head` вкладіть елемент `title` з текстом `Colored Markers`.

# --hints--

Ви повинні мати початковий теґ `title`.

```js
assert(code.match(/<title\s*>/i));
```

Ви повинні мати кінцевий теґ `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Ваш проєкт повинен мати титул `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Пам’ятайте, що вживання великої літери та написання мають значення для тексту титулу.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
