---
id: 61695ab9f6ffe951c16d03dd
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Вкладіть елемент `head` в елемент `html`. Відразу після елемента `head` додайте елемент `body`.

# --hints--

Ви повинні мати початковий теґ `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий теґ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий теґ `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий теґ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Має бути тільки один початковий теґ `head`.

```js
assert(code.match(/<head\s*>/ig).length === 1);
```

Має бути тільки один кінцевий теґ `head`.

```js
assert(code.match(/<\/head\s*>/ig).length === 1);
```

Має бути тільки один початковий теґ `body`.

```js
assert(code.match(/<body\s*>/ig).length === 1);
```

Має бути тільки один кінцевий теґ `body`.

```js
assert(code.match(/<\/body\s*>/ig).length === 1);
```

Ваш елемент `head` повинен бути порожнім.

```js
assert(code.match(/<head\s*>\s*<\/head\s*>/i));
```

Ваш елемент `body` повинен бути порожнім.

```js
assert(code.match(/<body\s*>\s*<\/body\s*>/i));
```

Ваш елемент `body` повинен бути розміщеним після елемента `head`.

```js
assert(code.match(/<\/head\s*>\s*<body\s*>/i));
```

Ваш елемент `head` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Ваш елемент `body` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
