---
id: 616965351e74d4689eb6de30
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Наприкінці використайте теґ `<meta>` `viewport`, щоб переконатися, що ваша сторінка виглядає однаково на всіх пристроях.

Вкладіть самозакривний елемент `meta` в `head`. Надайте йому атрибут `name` зі значенням `viewport` та атрибут `content` зі значенням `width=device-width, initial-scale=1.0`.

# --hints--

Ви повинні мати два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ваш новий елемент `meta` повинен бути самозакривним елементом.

```js
assert(code.match(/<\/meta>/i) === null);
```

Ваш новий елемент `meta` повинен мати атрибут `name` зі значенням `viewport` та атрибут `content` зі значенням `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
