---
id: 61696ef7ac756c829f9e4048
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

В цьому проєкті ви працюватимете із зовнішнім CSS-файлом, щоб стилізувати сторінку. Ми вже створили для вас файл `styles.css`. Але перш ніж ви зможете користуватись ним, потрібно прив’язати його до сторінки.

Вкладіть елемент `link` в елемент `head`. Надайте йому атрибут `rel` зі значенням `stylesheet` та атрибут `href` зі значенням `styles.css`.

# --hints--

Ви повинні мати один самозакривний елемент `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  --fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
