---
id: 616965351e74d4689eb6de30
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Ви можете мати декілька самозакривних елементів `meta` на вебсторінці. Відрізнити елементи `meta` можна за допомогою атрибутів. Ви повинні додати новий елемент meta до кожного атрибута, який хочете вказати.

Додайте ще один самозакривний елемент `meta` в межах `head`. Надайте йому атрибут `name` зі значенням `viewport` та атрибут `content` зі значенням `width=device-width, initial-scale=1.0`, щоб сторінка виглядала на всіх пристроях одинаково.

# --hints--

Ви повинні мати два елементи `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ваш новий елемент `meta` повинен бути самозакривним елементом.

```js
assert(code.match(/<\/meta>/i) === null);
```

Ваш новий елемент `meta` повинен мати атрибут `name` зі значенням `viewport` та атрибут `content` зі значенням `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const reValidContent = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/;
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && reValidContent.test(m?.getAttribute('content')));
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
