---
id: 61695d1fbc003856628bf561
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Щоб повідомити браузерам як кодувати символи на вашій сторінці, встановіть `charset` на `utf-8`. `utf-8` — це універсальний набір символів, що містить майже всі символи всіх людських мов.

Всередині елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `utf-8`. Пам'ятайте, що елементи `meta` є самозакривними та не потребують кінцевих теґів.

# --hints--

Ви повинні мати один елемент `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 1);
```

Ваш елемент `meta` повинен бути самозакривним елементом.

```js
assert(code.match(/<\/meta>/i) === null);
```

Ваш інший елемент `meta` повинен мати атрибут `charset` зі значенням `utf-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
