---
id: 616971b3cd990186b66c99a1
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Тепер, коли ваш зовнішній CSS-файл налаштовано, ви можете почати стилізувати сторінку.

Нагадуємо, як націлити на елемент абзацу та вирівняти його за правим краєм:

```css
p {
  text-align: right;
}
```

Створіть нове CSS-правило, націлене на елемент `h1`, та встановіть значення його властивості `text-align` на `center`.

# --hints--

Ви повинні використати селектор `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ваше CSS-правило `h1` повинне мати властивість `text-align` зі значенням `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
