---
id: 61695c4aad56f95497c19583
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Пам’ятайте, що елемент `title` надає пошуковим системам додаткову інформацію про сторінку. Він також показує вміст цього елементу `title` ще двома способами:

* на смужці заголовка, якщо сторінка відкрита
* на вкладці браузера, якщо навести мишкою на сторінку. Навіть якщо вкладка неактивна, як тільки ви наведете курсором на неї, буде показано текст `title`.

В елемент `head` вкладіть елемент `title` з текстом `Colored Markers`.

# --hints--

Ви повинні мати початковий теґ `title`.

```js
assert(code.match(/<title\s*>/i));
```

Ви повинні мати кінцевий теґ `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Ваш проєкт повинен мати титул `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Пам’ятайте, що вживання великої літери та написання мають значення для тексту титулу.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
