---
id: 616d50b93ba424d6282c99cf
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Зверніть увагу, що ваш маркер не має жодного кольору. Насправді фоновий колір було застосовано, але оскільки маркер-елемент `div` порожній, то немає ніякої висоти за замовчуванням.

В CSS-правилі `.marker` встановіть значення властивості `width` на `200px` та значення властивості `height` на `25px`.

# --hints--

Ваше CSS-правило `.marker` повинне мати властивість `width` зі значенням `200px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.width === '200px');
```

Ваше CSS-правило `.marker` повинне мати властивість `height` зі значенням `25px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.height === '25px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  background-color: red;
}
--fcc-editable-region--

```
