---
id: 616d3a67ccf800ad94ec89ae
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Зараз ви додасте деякі елементи, які згодом оформите в кольорові маркери.

Спершу додайте елемент `div` в елемент `body` та встановіть його атрибут `class` на `container`. Переконайтеся, що елемент `div` знаходиться під елементом `h1`.

# --hints--

Ви повинні мати початковий теґ `div`.

```js
assert(code.match(/<div.*>/i));
```

Ви повинні мати кінцевий теґ `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

Ваш елемент `div` повинен бути в межах елемента `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Ваш елемент `div` повинен бути після елемента `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Ваш елемент `div` повинен мати атрибут `class` зі значенням `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
