---
id: 616d525007b8c5d8b3308b1c
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Ваш маркер мав би кращий вигляд, якби був по центру сторінки. Це легко зробити завдяки <dfn>скороченій властивості</dfn> `margin`.

В останньому проєкті ви окремо встановлювали поля елементів з допомогою таких властивостей, як `margin-top` та `margin-left`. Скорочена властивість `margin` дозволяє легко встановлювати декілька полів одночасно.

Щоб зцентрувати маркер на сторінці, встановіть його властивість `margin` на `auto`. Це встановлює `margin-top`, `margin-right`, `margin-bottom` та `margin-left` на `auto`.

# --hints--

Ваше CSS-правило `.marker` повинне мати властивість `margin` зі значенням `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
