---
id: 616d595270d902f0e7086e18
title: Крок 15
challengeType: 0
dashedName: step-15
---

# --description--

Хоча у вас є три окремі маркер-елементи `div`, вони виглядають як один великий прямокутник. Ви повинні додати між ними деякий проміжок, щоб легше було бачити кожен елемент.

Якщо скорочена властивість `margin` має два значення, вона встановлює для `margin-top` та `margin-bottom` перше значення, а для `margin-left` та `margin-right` друге.

В CSS-правилі `.marker` встановіть властивість `margin` зі значенням `10px auto`.

# --hints--

Ваше CSS-правило `.marker` повинне мати властивість `margin` зі значенням `10px auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}
--fcc-editable-region--

```
