---
id: 6176519636a76b810ab1219a
title: Крок 20
challengeType: 0
dashedName: step-20
---

# --description--

Створіть CSS-правило, націлене на клас `two`, та встановіть значення його властивості `background-color` на `green`.

Також створіть окреме CSS-правило, націлене на клас `three`, та встановіть значення його властивості `background-color` на `blue`.

# --hints--

Ви повинні використати селектор класу, щоб націлити клас `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

Ваше CSS-правило `.two` повинне мати властивість `background-color` зі значенням `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

Ви повинні використати селектор класу, щоб націлити клас `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

Ваше CSS-правило `.three` повинне мати властивість `background-color` зі значенням `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
