---
id: 61764f23ea21477b76f3b80f
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Потім видаліть властивість `background-color` та її значення з CSS-правила `.marker`.

# --hints--

Ваше CSS-правило `.marker` не повинне мати властивості `background-color`.

```js
assert(!code.includes('background-color') && !new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}
--fcc-editable-region--

```
