---
id: 61764c602bee6974e7790f35
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

Щоб надати маркерам різні кольори, вам потрібно буде додати до кожного унікальний клас. До елемента можна додати декілька класів, перерахувавши їх в атрибуті `class` та розділивши їх пробілом. Наприклад, нижче додано класи `animal` та `dog` до елемента `div`.

```html
<div class="animal dog">
```

Якщо ви додаєте декілька класів до елемента HTML, стилі перших класів можуть бути перевизначені наступними класами.

Щоб почати, додайте клас `one` до першого елемента-маркера `div`.


# --hints--

Ви повинні додати клас `one` до першого маркер-елемента `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Ваш перший маркер `div` повинен мати класи `marker` та `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
