---
id: 617bc3386dc7d07d6469bf20
title: Крок 32
challengeType: 0
dashedName: step-32
---

# --description--

Зверніть увагу: щоб створити оранжевий, вам потрібно змінити значення `rgb`, а саме збільшити інтенсивність червоного та зменшити інтенсивність зеленого. Це тому що оранжевий колір є комбінацією червоного та жовтого, та знаходиться між двома кольорами на колірному колі.

Щоб створити третинний яскраво-зелений, поєднайте блакитний з зеленим. Оновіть функцію `rgb` в CSS-правилі `.two`, щоб зелений мав максимальне значення, та встановіть для синього значення `127`.

# --hints--

Ваше CSS-правило `.two` повинне мати властивість `background-color` зі значенням `rgb(0, 255, 127)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 255, 255);
}
--fcc-editable-region--

.three {
  background-color: rgb(255, 0, 255);
}

```
