---
id: 617b954d9f4f6217a749380e
title: Крок 25
challengeType: 0
dashedName: step-25
---

# --description--

Тепер додайте трохи більше вертикального простору між вашими маркерами та краєм елемента `container`, в якому вони знаходяться.

В CSS-правилі `.container` використайте скорочену властивість `padding`, щоб додати `10px` верхнього і нижнього відступу та `0` лівого і правого відступу. Це працює подібно до скороченої властивості `margin`, яку ви використовували раніше.

# --hints--

Ви не повинні видаляти властивість `background-color` та її значення з CSS-правила `.container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

Ваше CSS-правило `.container` повинне мати властивість `padding` зі значенням `10px 0`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.padding === '10px 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.container {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 127, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
