---
id: 617b65579ce424bf5f02ca73
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Існує дві основні колірні моделі: <dfn>адитивна</dfn> модель RGB (червоний, зелений, синій), яку використовують в електронних пристроях, та <dfn>субтрактивна</dfn> модель CMYK (блакитний, пурпуровий, жовтий, чорний), яку використовують в друці.

В цьому проєкті ви працюватимете з моделлю RGB. Це означає, що кольори починаються з чорного та змінюються, коли вводяться різні рівні червоного, зеленого та синього. Це легко побачити з допомогою CSS-функції `rgb`.

Створіть нове CSS-правило, націлене на клас `container` та встановіть його `background-color` на чорний з `rgb(0, 0, 0)`.

# --hints--

Ви повинні використати селектор класу, щоб націлити клас `container`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

Ваше CSS-правило `.container` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
