---
id: 617b8b38f32bf2080a140675
title: Крок 22
challengeType: 0
dashedName: step-22
---

# --description--

Функція – це фрагмент коду, який може приймати вхідні дані та виконувати певну дію. CSS-функція `rgb` приймає значення, або <dfn>аргументи</dfn>, для червоного, зеленого та синього та створює колір:

```css
rgb(red, green, blue);
```

Кожне червоне, зелене та синє значення – це число від `0` до `255`. `0` означає, що є 0% цього кольору, та є чорним. `255` означає, що є 100% цього кольору.

В CSS-правилі `.one` замініть ключове слово `red` на функцію `rgb`. Для функції `rgb` встановіть значення для червоного на `255`, для зеленого — `0` та для синього — `0`.

# --hints--

Ваше CSS-правило `.one` не повинне використовувати ключове слово `red` для встановлення властивості `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

Ваше CSS-правило `.one` повинне мати властивість `background-color` зі значенням `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
