---
id: 617b8e0d93a8d10d9a90e720
title: Крок 23
challengeType: 0
dashedName: step-23
---

# --description--

Зверніть увагу, що `background-color` для вашого маркера досі червоний. Це тому що ви встановили значення червоного в функції `rgb` на максимум `255`, або 100% червоного, а значення зеленого та синього встановили на `0`.

Тепер використовуйте функцію `rgb`, щоб встановити інші кольори.

В CSS-правилі `.two` використайте функцію `rgb`, щоб встановити `background-color` на максимальне значення для зеленого та `0` для інших значень. А в CSS-правилі `.three` використайте функцію `rgb`, щоб встановити `background-color` на максимальне значення для синього та `0` для інших значень.

# --hints--

Ваше CSS-правило `.two` не повинне використовувати ключове слово `green` для встановлення властивості `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor !== 'green');
```

Ваше CSS-правило `.two` повинне мати властивість `background-color` зі значенням `rgb(0, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 255, 0)');
```

Ваше CSS-правило `.three` не повинне використовувати ключове слово `blue` для встановлення властивості `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor !== 'blue');
```

Ваше CSS-правило `.three` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 255)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
