---
id: 618a00ed1ca871a2b3aca0cb
title: Крок 35
challengeType: 0
dashedName: step-35
---

# --description--

Ви ознайомилися з всіма основними, вторинними та третинними кольорами на колірному колі, отож вам буде легше зрозуміти інші концепції теорії кольору та їхній вплив на дизайн.

Спочатку налаштуйте значення в функції `rgb` в CSS-правилах `.one`, `.two` та `.three` таким чином, що `background-color` кожного елемента встановлено на чистий чорний. Пам’ятайте, що функція `rgb` використовує адитивну колірну модель, де кольори починаються з чорного та змінюються зі збільшенням значень червоного, зеленого та синього.

# --hints--

Ваше CSS-правило `.one` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(0, 0, 0)');
```

Ваше CSS-правило `.two` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'rgb(0, 0, 0)');
```

Ваше CSS-правило `.three` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(127, 255, 0);
}

.two {
  background-color: rgb(0, 127, 255);
}

.three {
  background-color: rgb(255, 0, 127);
}
--fcc-editable-region--

```
