---
id: 619b72a0db211f1c29afb906
title: Крок 42
challengeType: 0
dashedName: step-42
---

# --description--

Оновіть CSS-правило `.one`, щоб націлити новий клас `red`.

# --hints--

Вашому коду більше не потрібен селектор класу `.one`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.one'));
```

Ви повинні використати селектор класу, щоб націлити клас `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red'));
```

Ваше CSS-правило `.red` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
