---
id: 619b761916dac02643940022
title: Крок 46
challengeType: 0
dashedName: step-46
---

# --description--

Поширений спосіб застосування кольору до елемента в CSS містить <dfn>шістнадцяткові</dfn> значення. Хоча шістнадцяткові значення здаються важкими для розуміння, це просто інша форма значень RGB.

Шістнадцяткове значення кольору починається з символу `#` та містить шість символів 0-9 та A-F. Перша пара символів представляє червоний колір, друга пара – зелений, а третя пара – синій. Наприклад, `#4B5320`.

У селекторі класу `.green` встановіть для властивості `background-color` шістнадцятковий код кольору зі значеннями `00` для червоного, `FF` для зеленого та `00` для синього.

# --hints--

Ваше CSS-правило `.green` повинне мати властивість `background-color` зі значенням `#00FF00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
