---
id: 619b74fa777a2b2473c94f82
title: Крок 45
challengeType: 0
dashedName: step-45
---

# --description--

Оновіть селектор CSS-класу `.two`, щоб він був націлений на новий клас `green`. Та оновіть селектор класу `.three`, щоб він був націлений на новий клас `blue`.

# --hints--

Вашому коду більше не потрібен селектор класу `.two`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.two'));
```

Ви повинні використати селектор класу, щоб націлити клас `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green'));
```

Ваше CSS-правило `.green` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 0, 0)');
```

Вашому коду більше не потрібен селектор класу `.three`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.three'));
```

Ви повинні використати селектор класу, щоб націлити клас `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Ваше CSS-правило `.blue` повинне мати властивість `background-color` зі значенням `rgb(0, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
