---
id: 619b7c3c83de403126b69c1e
title: Крок 47
challengeType: 0
dashedName: step-47
---

# --description--

Мабуть, ви вже знайомі з десятковими значеннями, або значеннями з основою 10, які йдуть від 0 до 9. Шістнадцяткові значення, або значення з основою 16, йдуть від 0 до 9, потім від A до F:

```js
0, 1, 2, 3, 4, 5, 6, 7, 8, 9, A, B, C, D, E, F
```

Для шістнадцяткових кольорів `00` – це 0% цього кольору, а `FF` – це 100%. Отже, `#00FF00` означає 0% червоного, 100% зеленого та 0% синього, тобто те саме, що й `rgb(0, 255, 0)`.

Зменште інтенсивність зеленого, встановивши значення зеленого шістнадцяткового кольору на `7F`.

# --hints--

Ваше CSS-правило `.green` повинне мати властивість `background-color` зі значенням `#007F00`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor === 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.green {
  background-color: #00FF00;
}
--fcc-editable-region--

.blue {
  background-color: rgb(0, 0, 0);
}

```
