---
id: 619b7fd56aa2253778bcf5f7
title: Крок 48
challengeType: 0
dashedName: step-48
---

# --description--

Колірна модель <dfn>HSL</dfn>, або відтінок, насиченість та освітленість, – ще один спосіб представлення кольорів.

The CSS hsl function accepts 3 values: a number from 0 to 360 for hue, a percentage from 0 to 100 for saturation, and a percentage from 0 to 100 for lightness.

Якщо уявити колірне коло, то відтінок червоного – 0 градусів, зеленого – 120 градусів та синього – 240 градусів.

Saturation is the intensity of a color from 0%, or gray, to 100% for pure color. You must add the percent sign `%` to the saturation and lightness values.

Освітленість – наскільки яскравим здається колір від 0% (повністю чорний) до 100% (повністю білий), з 50% нейтральності.

В CSS-правилі `.blue` використайте функцію `hsl`, щоб змінити властивість `background-color` на чистий синій. Встановіть відтінок на `240`, насиченість на `100%` та освітленість на `50%`.

# --hints--

Ваше CSS-правило `.blue` повинне мати властивість `background-color` зі значенням `hsl(240, 100%, 50%)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor === 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

--fcc-editable-region--
.blue {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
