---
id: 61a493ead935c148d2b76312
title: Крок 50
challengeType: 0
dashedName: step-50
---

# --description--

Функція `linear-gradient` є дуже гнучкою – ось основний синтаксис, який ви використовуватимете в цьому туторіалі:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` – це напрямок лінії, використаної для переходу. `color1` та `color2` — це аргументи, тобто кольори, які використовуватимуться в самому переході. Це можуть бути будь-які типи кольорів, включаючи ключові слова, шістнадцяткові, `rgb` або `hsl`.

Тепер ви застосуєте градієнт від червоного до зеленого уздовж лінії під кутом 90 градусів до першого маркера.

Спочатку в CSS-правилі `.red` встановіть значення властивості `background` на `linear-gradient()` та передайте значення `90deg` як `gradientDirection`.

# --hints--

Ваше CSS-правило `.red` повинне мати властивість `background` зі значенням `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
