---
id: 61addaf7e83988b59a7d18ff
title: Крок 82
challengeType: 0
dashedName: step-82
---

# --description--

Скорочена властивість `border-left` дозволяє одночасно встановити ширину, стиль та колір лівого кордону.

Ось синтаксис:

```css
border-left: width style color;
```

В правилі CSS `.sleeve` замініть `border-left-width`, `border-left-style` та `border-left-color` на скорочену властивість `border-left`. Значення ширини, стилю та кольору лівого кордону повинні бути однаковими.

# --hints--

Ваше CSS-правило `.sleeve` не повинне мати властивості `border-left-width`.

```js
assert(!/border-left-width/g.test(code));
```

Ваше CSS-правило `.sleeve` не повинне мати властивості `border-left-style`.

```js
assert(!/border-left-style/g.test(code));
```

Ваше CSS-правило `.sleeve` не повинне мати властивості `border-left-color`.

```js
assert(!/border-left-color/g.test(code));
```

Ваше CSS-правило `.sleeve` повинне мати скорочену властивість `border-left` зі значенням `10px solid black`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sleeve')?.borderLeft, '10px solid black');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left-width: 10px;
  border-left-style: solid;
  border-left-color: black;
}
--fcc-editable-region--

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
