---
id: 61b093329e7fc020b43b1bb5
title: Крок 75
challengeType: 0
dashedName: step-75
---

# --description--

Ви вже знайомі з функцією `rgb` для встановлення кольорів. Щоб додати альфа-канал до кольору `rgb`, натомість використайте функцію `rgba`.

Функція `rgba` працює так само, як функція `rgb`, але приймає ще одне число від `0` до `1.0` для альфа-каналу:

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

Альфа-канал також можна використовувати з кольорами `hsl` та `hex`. Скоро ви дізнаєтесь, як це зробити.

В правилі `.sleeve` використайте функцію `rgba`, щоб встановити властивість `background-color` на чистий білий з 50% непрозорістю.

# --hints--

Ваше CSS-правило `.sleeve` повинне мати властивість `background-color` зі значенням `rgba(255, 255, 255, 0.5)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor === 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
