---
id: 61b0a44a6b865738ba49b9fb
title: Крок 86
challengeType: 0
dashedName: step-86
---

# --description--

Останнє, що ви зробите, це додасте легку тінь до кожного маркера, щоб вони виглядали ще реалістичніше.

Властивість `box-shadow` дозволяє застосовувати одну або більше тіней навколо елемента. Ось основний синтаксис:

```css
box-shadow: offsetX offsetY color;
```

Як працюють значення `offsetX` та `offsetY`:

* `offsetX` та `offsetY` приймають чисельні дані в `px` або інших одиницях CSS
* додатні значення `offsetX` пересувають тінь праворуч, а від’ємні – ліворуч
* додатні значення `offsetY` пересувають тінь вниз, а від’ємні – вверх
* якщо ви хочете додати нульове значення (`0`) до `offsetX` або `offsetY`, не потрібно вказувати одиницю. Кожен браузер розуміє, що нуль означає відсутність змін.

Висота і ширина тіні визначаються висотою і шириною елементу, до якого вони застосовуються. Ви також можете використати додаткове значення `spreadRadius`, щоб розтягнути тінь. Детальніше про це згодом.

Почніть з додавання простої тіні до червоного маркера.

В CSS-правилі `.red` додайте властивість `box-shadow` зі значеннями `5px` для `offsetX`, `5px` для `offsetY` та `red` для `color`.

# --hints--

Ваше CSS-правило `.red` повинне мати скорочену властивість `box-shadow` зі значенням `5px 5px red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow === 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

--fcc-editable-region--
.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}
--fcc-editable-region--

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
