---
id: 61b0936d9e7fc020b43b1bb8
title: Крок 78
challengeType: 0
dashedName: step-78
---

# --description--

Схоже, ваша насадка зникла, але не хвилюйтеся – вона все ще тут. Що сталось: ваш новий ковпачок `div` займає всю ширину маркера та штовхає насадку вниз до наступного рядка.

Це тому що за замовчуванням властивість `display` для елементів `div` є `block`. Тому якщо два елементи `block` знаходяться один біля одного, вони складаються як блоки. Наприклад, ваші маркери-елементи складені один на одному.

Щоб розмістити два елементи `div` в одному рядку, встановіть значення їхніх властивостей `display` на `inline-block`.

Створіть нове правило, щоб націлити обидва класи `cap` та `sleeve`, та встановіть значення `display` на `inline-block`.

# --hints--

Ви повинні використати селектор класу, щоб націлити класи `cap` та `sleeve`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

Ваше CSS-правило повинне мати властивість `display` зі значенням `inline-block`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
