---
id: 61968f8877c6720d6d61aaf5
title: 步骤 6
challengeType: 0
dashedName: step-6
---

# --description--

在 `body` 中，添加一个 `div`，并将 `class` 设置为 `ground`。

# --hints--

应该添加一个新的 `div`。

```js
assert.exists(document.querySelector('div'));
```

`div` 的 `class` 应该为 `ground`。

```js
assert.include(document.querySelector('div')?.className, 'ground');
```

`div` 应该在 `body` 内。

```js
assert.exists(document.querySelector('body > div.ground'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}
```
