---
id: 619692ff79f5770fc6d8c0b4
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

在 `.ground` 元素上面，添加一个 `div`，并设置其 `class` 为 `penguin`。 `div` 中将添加飞扬的企鹅。

# --hints--

应该在 `body` 内添加一个新的 `div`。

```js
assert.equal(document.querySelectorAll('body > div')?.length, 2);
```

`div` 的 `class` 应该为 `penguin`。

```js
assert.include(document.querySelector('body > div:not(.ground)')?.className, 'penguin');
```

应该将 `div.penguin` 放在 `div.ground` 前面。

```js
assert.strictEqual(document.querySelector('.ground')?.previousElementSibling, document.querySelector('.penguin'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
