---
id: 619665c9abd72906f3ad30f9
title: 步骤 1
challengeType: 0
dashedName: step-1
---

# --description--

你将创建一只快乐飞扬的企鹅，并在此过程中进一步学习 CSS transforms 和动画的知识。

先从基本的 HTML 模板开始。 需要包含 `DOCTYPE` 声明，语言设置为英文的 `html` 元素，相应的 `meta` 标签，`head`、`body` 以及 `title` 元素。 此外，还需要把样式表链接到页面里。

# --hints--

代码应该有一个 `<!DOCTYPE html>` 声明。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

代码应该有一个 `html` 元素。

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

你的 `html` 元素应该有一个带有值为 `en` 的 `lang` 属性的开始标签。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/gi));
```

你的代码应该在 `html` 元素中包含一个 `head` 元素。

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

你的代码应该在 `html` 元素中包含一个 `body` 元素。

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

你的 `head` 元素应该在你的 `body` 元素之前。

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

你应该有两个 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

一个 `meta` 元素应将 `name` 设置为 `viewport`，并将 `content` 设置为 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

另一个 `meta` 元素应将 `charset` 属性设置为 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

你的代码应该有一个 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的 `title` 应该有一些文本。

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

你的代码应该有一个 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

你的 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素应该有一个 `href` 属性，其值为 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
