---
id: 61993b72e874e709b8dfd666
title: 步骤 30
challengeType: 0
dashedName: step-30
---

# --description--

选择 `.penguin-head` 元素，使其 `width` 为父元素的一半，`height` 为 `45%`。 然后，设置其 `background` 为渐变轴为 `45deg` 从 `gray` 到 `rgb(239, 240, 228)` 的线性渐变。

# --hints--

应该使用 `.penguin-head` 选择器。

```js
assert.match(code, /\.penguin-head\s*\{/);
```

应该将 `.penguin-head` 的 `width` 设置为 `50%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin-head')?.width, '50%');
```

应该将 `.penguin-head` 的 `height` 设置为 `45%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.penguin-head')?.height, '45%');
```

`.penguin-head` 的 `background` 应该为 `linear-gradient(45deg, gray, rgb(239, 240, 228))`。

```js
assert.include(['linear-gradient(45deg,gray,rgb(239,240,228))', 'rgba(0,0,0,0)linear-gradient(45deg,gray,rgb(239,240,228))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.penguin-head')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
    <div class="penguin">
      <div class="penguin-head"></div>
      <div class="penguin-body"></div>
    </div>

    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
  z-index: 4;
  position: relative;
}

--fcc-editable-region--

--fcc-editable-region--

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
