---
id: 56533eb9ac21ba0edf2244a9
title: Initialisierung von Variablen mit dem Zuweisungsoperator
challengeType: 1
videoUrl: 'https://scrimba.com/c/cWJ4Bfb'
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

Es ist üblich, eine Variable in der gleichen Zeile, in der sie deklariert wird, mit einem Anfangswert zu <dfn>initialisieren</dfn>.

```js
var myVar = 0;
```

Erzeugt eine neue Variable mit dem Namen `myVar` und weist ihr einen Anfangswert von `0` zu.

# --instructions--

Definiere eine Variable `a` mit `var` und initialisiere sie mit einem Wert von `9`.

# --hints--

Du solltest `a` mit einem Wert von `9` initialisieren.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(code));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
