---
id: 60eebd07ea685b0e777b5583
title: Крок 1
challengeType: 0
dashedName: step-1
---

# --description--

Ласкаво просимо до проєкту «Реєстраційна форма»! Почніть з додавання декларації `!DOCTYPE html` в верхній частині документа, щоб браузер знав, який тип документа він читає.

# --hints--

Ваш код повинен містити посилання `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Ви повинні додати пробіл після посилання `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Ви повинні визначити тип документа як `html`.

```js
assert(code.match(/html/gi));
```

Ви повинні закрити декларацію `DOCTYPE` з допомогою `>` після типу.

```js
assert(code.match(/html\s*>/gi));
```

Ваша декларація `DOCTYPE` повинна бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
