---
id: 60f030d388cb74067cf291c3
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Під заголовком використайте наступний текст в елементі абзацу, щоб заохотити користувачів до реєстрації:

```md
Please fill out this form with the required information
```

# --hints--

Ви повинні додати елемент `p` в `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Ви повинні додати елемент `p` під `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Ви повинні надати елементу `p` текст `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>

  </body>
--fcc-editable-region--
</html>
```
