---
id: 60f02e7361b68405e27b62a5
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

В межах `body` надайте заголовок-контекст вмісту, додавши `h1` з текстом `Registration Form`.

# --hints--

Ви повинні додати елемент `h1` в `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Ви повинні надати `h1` текст `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
