---
id: 60f0286404aefb0562a4fdf9
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Додайте елементи `title` та `meta` до `head`. Назвіть свій проєкт `Registration Form` та надайте атрибут `charset` зі значенням `UTF-8` до свого елемента `meta`.

# --hints--

Ваш код повинен мати елемент `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Елемент `title` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Ваш проєкт повинен мати титул `Registration Form`.

```js
const title = document.querySelector('title');
assert.equal(title.text.toLowerCase(), 'registration form')
```

Пам’ятайте, що вживання великої літери та написання мають значення.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
