---
id: 60f027c87bc98f050395c139
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Потім додайте початковий та кінцевий теґи `head` та `body` в межах елемента `html`.

# --hints--

Ви повинні мати початковий теґ `head`.

```js
assert(code.match(/<head\s*>/i));
```

Ви повинні мати кінцевий теґ `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Ви повинні мати початковий теґ `body`.

```js
assert(code.match(/<body\s*>/i));
```

Ви повинні мати кінцевий теґ `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Елементи `head` та `body` повинні бути братськими.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

Елемент `head` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` повинен бути в межах елемента `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
