---
id: 60f027099a15b00485563dd2
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Під `DOCTYPE` додайте елемент `html` з атрибутом `lang` зі значенням `en`, щоб у вас було місце, щоб почати розміщувати код.

# --hints--

Ваша декларація `DOCTYPE` повинна бути на початку HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Ваш елемент `html` повинен мати початковий теґ та атрибут `lang` зі значенням `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ваш елемент `html` повинен мати кінцевий теґ.

```js
assert(code.match(/<\/html\s*>/));
```

Ваші теґи `html` повинні бути в правильному порядку.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Ви повинні мати тільки один елемент `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
