---
id: 60f1a9cbd23023082e149fee
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Так краще. Тепер зробіть фон приємним для очей, змінивши `body` `background-color` на `#1b1b32`. Потім, щоб побачити текст, змініть `color` на `#f5f6f7`.

# --hints--

Ви повинні додати `background-color` в селектор елемента `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Ви повинні надати `background-color` значення `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Ви повинні додати `color` в селектор елемента `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Ви повинні надати `color` значення `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
