---
id: 60f1922fcbd2410527b3bd89
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--
Одиниця `vh` означає висоту вюпорту (viewport height) та відноситься до 1% від `height` вюпорту.

Настав час причепурити проєкт за допомогою CSS. Спочатку надайте `body` властивості `width` зі значенням `100%` та `height` зі значенням `100vh`.

# --hints--

Ви повинні використати селектор елемента `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Ви повинні надати `body` властивість `width` зі значенням `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Ви повинні надати `body` властивість `height` зі значенням `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
