---
id: 60f5c3e399ff1a05629964e4
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Як випливає з назви, ви створюєте форму. Отже, після елемента `p` вставте `form` з атрибутом `action`, націленим на `https://register-demo.freecodecamp.org`.

# --hints--

Ви повинні додати елемент `form` поруч з елементом `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Ви повинні надати `form` атрибут `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Ви повинні надати `action` значення `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
