---
id: 60f81616cff80508badf9ad5
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

На цьому етапі ви зможете надіслати форму. Однак ви можете помітити, що майже нічого не відбувається.

Щоб зробити форму більш інтерактивною, додайте атрибут `required` до елементів `input` в першому `fieldset`.

Тепер, якщо ви спробуєте надіслати форму, не заповнивши обов’язкові поля, ви побачите повідомлення про помилку.

# --hints--

Ви повинні надати першому елементу `input` атрибут `required`.

```js
assert.equal(document.querySelector('input')?.required, true);
```

Ви повинні надати другому елементу `input` атрибут `required`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.required, true);
```

Ви повинні надати третьому елементу `input` атрибут `required`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.required, true);
```

Ви повинні надати четвертому елементу `input` атрибут `required`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.required, true);
```

Ви не повинні надавати `submit` `input` атрибут `required`.

```js
assert.equal(document.querySelector('input[type="submit"]')?.required, false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input type="text" /></label>
        <label>Enter Your Last Name: <input type="text" /></label>
        <label>Enter Your Email: <input type="email" /></label>
        <label>Create a New Password: <input type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
