---
id: 60f80e0081e0f2052ae5b505
title: Крок 19
challengeType: 0
dashedName: step-19
---

# --description--

Уточнення атрибута `type` елемента форми важливе для браузера, щоб знати, які дані йому варто очікувати. Якщо `type` не вказано, браузер за замовчуванням використовуватиме `text`.

Першим двом елементам `input` надайте атрибут `type` зі значенням `text`, третьому – атрибут `type` зі значенням `email`, четвертому – атрибут `type` зі значенням `password`.

Тип `email` дозволяє лише електронні адреси, що мають `@` та `.` в домені. Тип `password` приховує введення та попереджає, якщо сайт не використовує HTTPS.

# --hints--

Ви повинні надати першому елементу `input` атрибут `type` зі значенням `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Ви повинні надати другому елементу `input` атрибут `type` зі значенням `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Ви повинні надати третьому елементу `input` атрибут `type` зі значенням `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Ви повинні надати четвертому елементу `input` атрибут `type` зі значенням `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
