---
id: 60fabf0dd4959805dbae09e6
title: Крок 32
challengeType: 0
dashedName: step-32
---

# --description--

Додайте ще один `label` після першого, з текстом `Input your age (years):`. Потім вкладіть `input`, що має `type` зі значенням `number`.

Потім додайте атрибут `min` до `input` зі значенням `13`, оскільки користувачі молодші 13 років не можуть реєструватися. Крім того, ми можемо припустити, що користувачі старше 120 років не реєструватимуться; додайте атрибут `max` зі значенням `120`.

Тепер, якщо хтось спробує надіслати форму зі значеннями поза діапазоном, з’явиться попередження та форма не буде надіслана. Спробуйте.

# --hints--

Ви повинні додати `label` до третього `fieldset`, після наявного `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

Ви повинні надати `label` текст `Input your age (years):`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

Для `label` потрібно надати `input`, що має `type` зі значенням `number`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

Ви повинні надати `input` атрибут `min` зі значенням `13`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

Ви повинні надати `input` атрибут `max` зі значенням `120`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" /> Personal Account</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business Account</label>
        <label for="terms-and-conditions">
          <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>

      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
