---
id: 60fad1cafcde010995e15306
title: Крок 37
challengeType: 0
dashedName: step-37
---

# --description--

Під час відправки форм краще надавати кожному елементу, який можна надсилати, атрибут `name`. Цей атрибут використовують для ідентифікації елемента в поданні форми.

Надайте кожному елементу, який можна подати, унікальний атрибут `name` за своїм вибором, окрім двох перемикачів.

# --hints--

Ви повинні надати `input`, який очікує ім'я, атрибут `name`. _P.S. я б обрав `first-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

Ви повинні надати `input`, який очікує прізвище, атрибут `name`. _P.S. я б обрав `last-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

Ви повинні надати `email` атрибут `name`. _P.S. я б обрав `email`_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

Ви повинні надати `password` атрибут `name`. _P.S. я б обрав `password`_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

Ви повинні надати `checkbox` атрибут `name`. _P.S. я б обрав `terms`_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

Ви повинні надати `file` атрибут `name`. _P.S. я б обрав `file`_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

Ви повинні надати `number` атрибут `name`. _P.S. я б обрав `age`_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

Ви повинні надати елементу `select` атрибут `name`. _P.S. я б обрав `referrer`_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

Ви повинні надати елементу `textarea` атрибут `name`. _P.S. я б обрав `bio`_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

Ви не повинні надавати жодному з елементів `option` атрибут `name`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

Ви не повинні надавати жодному з елементів `label` атрибут `name`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

Ви не повинні надавати жодному з елементів `fieldset` атрибут `name`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input type="text" required /></label>
        <label>Enter Your Last Name: <input type="text" required /></label>
        <label>Enter Your Email: <input type="email" required /></label>
        <label>Create a New Password: <input type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <label><input type="radio" name="account-type" /> Personal Account</label>
        <label><input type="radio" name="account-type" /> Business Account</label>
        <label>
          <input type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
        </label>
      </fieldset>
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" />
        </label>
        <label>How did you hear about us?
          <select>
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label>Provide a bio:
          <textarea rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
