---
id: 62cc5b1779e4d313466f73c5
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Вкладіть самозакривний елемент `link` в елемент `head`. Надайте йому атрибут `rel` зі значенням `stylesheet` та атрибут `href` зі значенням `styles.css`.

# --hints--

Ваш код повинен мати елемент `link`.

```js
assert.exists(document.querySelector('link'));
```

Ви не повинні змінювати наявні теґи `head`. Переконайтеся, що не видалили кінцевий теґ.

```js
const heads = document.querySelectorAll('head');
assert.equal(heads?.length, 1);
```

Ви повинні мати один самозакривний елемент `link`.

```js
assert(document.querySelectorAll('link').length === 1);
```

Ваш елемент `link` повинен бути в межах елемента `head`.

```js
assert.exists(document.querySelector('head > link'));
```

Ваш елемент `link` повинен мати атрибут `rel` зі значенням `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ваш елемент `link` повинен мати атрибут `href` зі значенням `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
