---
id: 62ff8b9dab5ac88e4d3d43a3
title: Крок 17
challengeType: 0
dashedName: step-17
---

# --description--

Слідуючи за передовими практиками доступності, з'єднайте елементи `input` та `label` разом, використовуючи атрибут `for`.

Використайте такі значення для відповідних властивостей `id`: `first-name`, `last-name`, `email`, `new-password`

# --hints--

Перший елемент `input` повинен мати `id` зі значенням `first-name`.

```js
assert(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

Другий елемент `input` повинен мати `id` зі значенням `last-name`.

```js
assert(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

Третій елемент `input` повинен мати `id` зі значенням `email`.

```js
assert(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

Четвертий елемент `input` повинен мати `id` зі значенням `new-password`.

```js
assert(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

Перший елемент `label` повинен мати атрибут `for` зі значенням `first-name`.

```js
assert(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

Другий елемент `label` повинен мати атрибут `for` зі значенням `last-name`.

```js
assert(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

Третій елемент `label` повинен мати атрибут `for` зі значенням `email`.

```js
assert(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

Четвертий елемент `label` повинен мати атрибут `for` зі значенням `new-password`.

```js
assert(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
