---
id: 63541ef4f96cd82e8e6c788a
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Атрибут `method` визначає, як надсилати дані форми на URL-адресу, вказану в атрибуті `action`. Дані форми можна надіслати через запит `GET` як параметри URL-адреси (із `method="get"`) або через запит `POST` як дані у тілі запиту (із `method="post"`).

Встановіть атрибут `method` для надсилання даних форми через запит `POST`.

# --hints--

Ви не повинні додавати новий елемент `form`.

```js
assert.equal(document.querySelectorAll('form').length, 1 )
```

Елемент `form` повинен мати атрибут `method`.

```js
assert.exists(document.querySelector('form')?.getAttribute('method'));
```

Ваш атрибут `method` повинен мати значення `post`.

```js
assert.equal(document.querySelector('form')?.getAttribute('method'), 'post');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    --fcc-editable-region--
    <form action='https://register-demo.freecodecamp.org'></form>
--fcc-editable-region--
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
